---
id: 69f0a37c3bf194a7c9c62492
title: Costruire un registro prestiti dispositivi
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Sei uno sviluppatore incaricato di creare un'applicazione registro per gestire i dispositivi hardware IT forniti dal Service Desk della tua azienda. Il registro è un oggetto le cui chiavi sono i tag delle risorse e i cui valori sono oggetti:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti modellare il registro come un oggetto in cui ogni dispositivo è indicizzato dal suo tag risorsa e ha le proprietà `type`, `status`, `borrower` (con `name` e `email`) e `dueDate`.

1. Dovresti implementare una funzione `checkoutDevice(ledger, assetTag, borrower)` che prende un oggetto registro e ne crea una copia senza modificare l'originale. Sulla copia, deve impostare il `name` e l'`email` dell'oggetto `borrower`, aggiornare lo `status` a `"CheckedOut"` e restituire un oggetto della forma `{ ledger: updatedLedger, message: confirmationString }`, dove il `message` include il tag risorsa e il `name` dell'oggetto `borrower`.

   - Se il `assetTag` non esiste nel registro, la funzione deve restituire il registro invariato insieme a un `message` che include il tag risorsa e spiega che non è stato trovato.
   - Se il dispositivo è già `"CheckedOut"`, la funzione deve restituire il registro invariato lasciando il borrower attuale, insieme a un `message` che include il tag risorsa e spiega che il dispositivo è già in prestito.

1. Dovresti implementare una funzione `checkinDevice(ledger, assetTag)` che prende un oggetto registro e ne crea una copia senza modificare l'originale. Sulla copia, deve cancellare il `name` e l'`email` dell'oggetto `borrower`, resettare `dueDate` a `""`, aggiornare lo `status` a `"CheckedIn"` e restituire un oggetto della forma `{ ledger: updatedLedger, message: confirmationString }`, dove il `message` include il tag risorsa.

   - Se il `assetTag` non esiste nel registro, la funzione deve restituire il registro invariato insieme a un `message` che include il tag risorsa e spiega che non è stato trovato.

1. Dovresti implementare una funzione `listOverdueDevices(ledger, today)` che prende un oggetto registro e restituisce un array degli oggetti dispositivo che sono in ritardo, ordinati per `dueDate` in ordine crescente. Un dispositivo è in ritardo quando il suo `status` è `"CheckedOut"` e il suo `dueDate` è strettamente precedente a `today`. Sia `today` che `dueDate` sono stringhe nel formato mese/giorno/anno, e il mese e il giorno possono essere con o senza zero iniziale (per esempio, `"9/5/2025"` e `"09/05/2025"` sono entrambi validi).

1. Non dovresti usare l'oggetto JavaScript `Date` in nessuna parte del tuo codice.

1. Dovresti implementare una funzione `serializeLedger(ledger)` che converte l'oggetto registro in una stringa JSON, e una funzione `loadLedger(json)` che converte una stringa JSON di nuovo in un oggetto JavaScript.

# --hints--

Dovresti creare una funzione chiamata `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

Dovresti creare una funzione chiamata `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

Dovresti creare una funzione chiamata `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

Dovresti creare una funzione chiamata `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

Dovresti creare una funzione chiamata `loadLedger`.

```js
assert.isFunction(loadLedger);
```

La tua funzione `checkoutDevice` non dovrebbe modificare il registro originale.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

La tua funzione `checkoutDevice` dovrebbe impostare il `name` e l'`email` del borrower e impostare lo `status` a `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

La tua funzione `checkoutDevice` dovrebbe restituire un messaggio di conferma che include il tag risorsa e il `name` del borrower.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

La tua funzione `checkoutDevice` dovrebbe preservare i dispositivi non correlati e i campi non modificati nel registro restituito.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

La tua funzione `checkoutDevice` dovrebbe restituire il registro invariato e un messaggio che include il tag risorsa quando il tag risorsa non viene trovato.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

La tua funzione `checkoutDevice` dovrebbe restituire il registro invariato, lasciare il borrower attuale al suo posto e restituire un messaggio che include il tag risorsa quando il dispositivo è già in prestito.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

La tua funzione `checkinDevice` non dovrebbe modificare il registro originale.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

La tua funzione `checkinDevice` dovrebbe cancellare il `name` e l'`email` del borrower, resettare `dueDate` a `""` e impostare lo `status` a `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

La tua funzione `checkinDevice` dovrebbe restituire un messaggio di conferma che include il tag risorsa.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

La tua funzione `checkinDevice` dovrebbe preservare i dispositivi non correlati e i campi non modificati nel registro restituito.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

La tua funzione `checkinDevice` dovrebbe restituire il registro invariato e un messaggio che include il tag risorsa quando il tag risorsa non viene trovato.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

La tua funzione `listOverdueDevices` dovrebbe restituire un array.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

La tua funzione `listOverdueDevices` dovrebbe restituire solo i dispositivi in ritardo con uno `status` di `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

La tua funzione `listOverdueDevices` dovrebbe restituire i dispositivi in ritardo ordinati per `dueDate` in ordine crescente.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

La tua funzione `listOverdueDevices` dovrebbe ordinare correttamente anche quando il mese o il giorno non hanno lo zero iniziale.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

Non dovresti usare l'oggetto JavaScript `Date` in nessuna parte del tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

La tua funzione `serializeLedger` dovrebbe restituire una stringa JSON che rappresenta il registro passato.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

La tua funzione `loadLedger` dovrebbe restituire l'oggetto JavaScript che la stringa JSON rappresenta.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
