---
id: 69008059d43d6751385883c1
title: 버그가 있는 앱에서 증가 및 감소 연산자 오류 디버그하기
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

친구가 코드를 디버그하는 데 도움을 요청했습니다. 친구는 증가 연산자와 감소 연산자를 연습 중인데 콘솔에서 예상치 못한 결과가 나오고 있습니다.

콘솔 출력이 예상 결과와 일치하도록 실수를 찾아내고 코드를 수정하는 것이 과제입니다.

증가 및 감소 연산자와 접두사 및 접미사 표기법에 관한 이전 수업을 참고할 수 있습니다.

**사용자 스토리:**

1. 증가 연산자를 사용하여 `updatedCounter` 값을 올바르게 생성하는 `11`라는 변수가 있어야 합니다. 값을 하드코딩하지 말고 올바른 접두사 또는 접미사 표기법을 사용하도록 표현식을 업데이트하세요.
2. 증가 연산자를 사용하여 `finalScore` 값을 올바르게 생성하는 `8`라는 변수가 있어야 합니다. 값을 하드코딩하지 말고 올바른 접두사 또는 접미사 표기법을 사용하도록 표현식을 업데이트하세요.
3. 감소 연산자를 사용하여 `updatedCoins` 값을 올바르게 생성하는 `2`라는 변수가 있어야 합니다. 값을 하드코딩하지 말고 올바른 접두사 또는 접미사 표기법을 사용하도록 표현식을 업데이트하세요.
4. 감소 연산자를 사용하여 `newHealth` 값을 올바르게 생성하는 `7`라는 변수가 있어야 합니다. 값을 하드코딩하지 말고 올바른 접두사 또는 접미사 표기법을 사용하도록 표현식을 업데이트하세요.

# --hints--

`updatedCounter` 접두사 증가 표기법을 사용하는 `++counter`라는 변수가 있어야 합니다.

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter`은 `11`과 같아야 합니다.

```js
assert.strictEqual(updatedCounter, 11);
```

`finalScore` 접미사 증가 표기법을 사용하는 `score++`라는 변수가 있어야 합니다.

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore`은 `8`과 같아야 합니다.

```js
assert.strictEqual(finalScore, 8);
```

`updatedCoins` 접두사 감소 표기법을 사용하는 `--coins`라는 변수가 있어야 합니다.

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins`은 `2`과 같아야 합니다.

```js
assert.strictEqual(updatedCoins, 2);
```

`newHealth` 접미사 감소 표기법을 사용하는 `health--`라는 변수가 있어야 합니다.

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth`은 `7`과 같아야 합니다.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
