---
id: 6989ef885db3a3488a13b426
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Prima di implementare l'algoritmo principale, devi validare l'input. Il parametro `pairs` dovrebbe essere un intero, poiché rappresenta il numero di coppie di parentesi tonde da generare.

Aggiungi un'istruzione `if` all'inizio della tua funzione per verificare se `pairs` non è un intero. Usa il metodo `Number.isInteger()` per farlo.

Se la condizione è vera, restituisci la stringa `The number of pairs should be an integer`.

# --hints--

Dovresti avere un'istruzione if con la condizione `!Number.isInteger(pairs)` all'interno della funzione `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Dovresti restituire una stringa quando l'input non è un intero.

```js
let result = genParentheses("i");
assert.isString(result);
```

Dovresti restituire la stringa `The number of pairs should be an integer` quando l'input è una stringa.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
