---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Ora devi gestire l'aggiunta delle parentesi tonde di chiusura. La regola fondamentale è che puoi aggiungere una parentesi di chiusura solo se mantiene l'equilibrio, cioè deve esserci un numero maggiore di parentesi tonde di apertura usate rispetto a quelle di chiusura.

Sempre all'interno del blocco `else`, aggiungi una seconda istruzione `if` per verificare se `closesUsed` è minore di `opensUsed`.

Se questa condizione è vera, inserisci un altro nuovo array in `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Questo rappresenta lo stato dopo aver aggiunto una parentesi di chiusura.

# --hints--

Dovresti avere una seconda istruzione `if` con la condizione `closesUsed < opensUsed` all'interno del tuo blocco `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Dovresti inserire `[current + ')', opensUsed, closesUsed + 1]` nella coda quando la condizione è vera.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
