---
id: 6989f10a72ad424d4ad6d514
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Successivamente, devi verificare che il numero di coppie sia almeno uno, dato che non puoi generare combinazioni di parentesi con zero o coppie negative.

Aggiungi un'altra istruzione `if` per controllare se `pairs` è minore di `1`. Se questa condizione è vera, restituisci la stringa `The number of pairs should be at least 1`.

# --hints--

Dovresti avere una seconda istruzione `if` con la condizione `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Dovresti restituire la stringa `The number of pairs should be at least 1` quando `pairs` è minore di `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
