---
id: 6989eb65324d97775c4965fe
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, implementerai una funzione che genera tutte le combinazioni valide di parentesi usando un approccio di ricerca in ampiezza (BFS). Per esempio, le combinazioni valide di due coppie di parentesi sono `(())` e `()()`.

Inizia creando una funzione chiamata `genParentheses` con un singolo parametro `pairs`. Per ora, fai in modo che la funzione ritorni un array vuoto.

# --hints--

Dovresti definire una funzione chiamata `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
```

La tua funzione `genParentheses` dovrebbe avere un singolo parametro chiamato `pairs`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.equal(genParentheses?.parameters.length, 1);
assert.equal(genParentheses?.parameters[0].toString(), 'pairs');
```

La tua funzione dovrebbe restituire un array vuoto.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
