---
id: 6989f207306332c06b0add39
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Per l'approccio di ricerca in ampiezza, userai una coda per tenere traccia dei diversi stati mentre costruisci le combinazioni di parentesi tonde. Ogni stato sarà rappresentato come un array contenente tre elementi:

- La stringa corrente che si sta costruendo
- Il numero di parentesi tonde aperte usate finora
- Il numero di parentesi tonde chiuse usate finora

Crea una variabile chiamata `queue` e inizializzala con un array contenente un array: `['', 0, 0]`. Questo rappresenta lo stato iniziale con una stringa vuota e zero parentesi tonde usate.

# --hints--

Dovresti dichiarare una variabile chiamata `queue` nella tua funzione `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Dovresti inizializzare `queue` con un array contenente l'array `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  return result;
}
```
