---
id: 6989f48c3f0292c84f8d8938
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora devi identificare quando hai costruito una combinazione completa di parentesi tonde. Una combinazione completa ha una lunghezza uguale al doppio del numero di coppie (dato che ogni coppia contribuisce con una parentesi aperta e una chiusa).

All'interno del tuo loop `while`, aggiungi un'istruzione `if` per verificare se `current.length === 2 * pairs`. Quando questa condizione è vera, inserisci `current` nell'array `result`.

# --hints--

Dovresti avere un'istruzione `if` con la condizione `current.length === 2 * pairs` all'interno del tuo loop `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Dovresti inserire `current` nell'array `result` all'interno della tua nuova istruzione `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
