---
id: 6989f1ba358acba651d5ee6b
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora imposta la struttura dati per archiviare i tuoi risultati. Crea una variabile chiamata `result` e inizializzala con un array vuoto. Questo array conterrà tutte le combinazioni valide di parentesi che genererai.

Aggiorna l'istruzione return per restituire `result` invece di un array vuoto.

# --hints--

Dovresti dichiarare una variabile chiamata `result` nella tua funzione `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.result);
```

Dovresti inizializzare `result` con un array vuoto.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses?.variables ?? {};
assert.isFalse(result?.value.isEmpty());
const resultValue = eval(result?.value.toString());
assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

Dovresti restituire `result` alla fine della tua funzione.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.isTrue(genParentheses?.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
