---
id: 6989f3632f8b8bdf168aa321
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

All'interno del tuo loop `while`, usa `queue.shift()` per rimuovere e ottenere il primo elemento dalla coda. Questo implementa il comportamento first-in-first-out (FIFO) caratteristico del BFS.

Decompila questo array in tre variabili: `current`, `opensUsed` e `closesUsed`.

# --hints--

Dovresti usare `queue.shift()` per ottenere il primo elemento dalla coda.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Dovresti decomporre il risultato in tre variabili: `current`, `opensUsed` e `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
