---
id: 6989f555728e63ad8b349d84
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Se la stringa corrente non è ancora completa, devi esplorare i prossimi stati possibili. Aggiungi una clausola `else` per l'istruzione `if` corrente.

All'interno del blocco `else`, gestirai l'aggiunta delle parentesi tonde di apertura. Aggiungi un'istruzione `if` per verificare se `opensUsed` è minore di `pairs`. Questo garantisce che non usi più parentesi tonde di apertura di quelle consentite.

Se questa condizione è `true`, inserisci un nuovo array nella `queue`: `[current + '(', opensUsed + 1, closesUsed]`. Questo rappresenta lo stato dopo aver aggiunto una parentesi tonda di apertura.

# --hints--

Dovresti avere una clausola `else` per la tua istruzione `if` all'interno del loop `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Dovresti avere un'istruzione `if` all'interno del blocco `else` con la condizione `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Dovresti inserire `[current + '(', opensUsed + 1, closesUsed]` nella coda quando la condizione è `true`.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
