---
id: 6989f2e35e6941e13e445156
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora implementa il loop principale di BFS. Crea un loop `while` che continua finché `queue` non è vuoto, cioè finché `queue.length` è maggiore di `0`.

All'interno del loop, stampa `queue` sulla console.

# --hints--

Dovresti creare un loop `while` con la condizione `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Dovresti stampare `queue` sulla console all'interno del tuo loop `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
