---
id: 587d78ae367417b2b2512afe
title: flex省略形プロパティの使い方
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

複数のflexプロパティを一度に設定できる省略形があります。`flex-grow`、`flex-shrink`、`flex-basis`のプロパティはすべて`flex`プロパティを使ってまとめて設定できます。

例えば、`flex: 1 0 10px;`はアイテムの`flex-grow: 1;`、`flex-shrink: 0;`、`flex-basis: 10px;`を設定します。

デフォルトのプロパティ設定は`flex: 0 1 auto;`です。

# --instructions--

`flex`と`#box-1`の両方にCSSプロパティ`#box-2`を追加してください。`#box-1`には、`flex-grow`が`2`、`flex-shrink`が`2`、`flex-basis`が`150px`となるように値を設定してください。`#box-2`には、`flex-grow`が`1`、`flex-shrink`が`1`、`flex-basis`が`150px`となるように値を設定してください。

これらの値により、コンテナが300pxより大きい場合、`#box-1`は`#box-2`の2倍の速さで余白を埋めるように成長し、コンテナが300px未満の場合は`#box-2`の2倍の速さで縮みます。300pxは2つのボックスの`flex-basis`の合計サイズです。

# --hints--

`#box-1`要素には、`flex`プロパティを`2 2 150px`の値に設定してください。

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

`#box-2`要素には、`flex`プロパティを`1 1 150px`の値に設定してください。

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

コードでは、`flex`と`#box-1`に対して`#box-2`プロパティを使う必要があります。

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
