---
id: 587d78af367417b2b2512b00
title: align-self プロパティの使い方
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

フレックスアイテムに使う最後のプロパティは `align-self` です。このプロパティを使うと、すべてのアイテムを一括で設定するのではなく、各アイテムの配置を個別に調整できます。これは、CSSの `float`、`clear`、`vertical-align` といった他の一般的な調整手法がフレックスアイテムには効かないため、便利です。

`align-self` は `align-items` と同じ値を受け入れ、`align-items` プロパティで設定された値をオーバーライドします。

# --instructions--

`align-self` の CSS プロパティを `#box-1` と `#box-2` の両方に追加してください。`#box-1` には `center` の値を、`#box-2` には `flex-end` の値を設定してください。

# --hints--

`#box-1` 要素には `align-self` プロパティを `center` の値に設定してください。

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

`#box-2` 要素には `align-self` プロパティを `flex-end` の値に設定してください。

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
