---
id: 587d78ab367417b2b2512af2
title: "flex-direction プロパティを使って行を作る"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

要素に`display: flex`を追加すると、その要素がフレックスコンテナになります。これにより、その要素の子要素を行または列に揃えることができます。親要素に`flex-direction`プロパティを追加し、値をrowまたはcolumnに設定することで実現します。rowを作成すると子要素が水平方向に揃い、columnを作成すると子要素が垂直方向に揃います。

`flex-direction`の他のオプションは`row-reverse`と`column-reverse`です。

**注意:** `flex-direction`プロパティのデフォルト値は`row`です。

# --instructions--

`flex-direction`というCSSプロパティを`#box-container`要素に追加し、その値を`row-reverse`にしてください。

# --hints--

`#box-container`要素には`flex-direction`プロパティを`row-reverse`に設定してください。

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
