---
id: 587d78ae367417b2b2512aff
title: orderプロパティを使ってアイテムの順序を並べ替える
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

`order`プロパティは、フレックスアイテムがフレックスコンテナ内でどの順序で表示されるかをCSSに伝えるために使います。デフォルトでは、アイテムはソースHTMLに記述されている順番で表示されます。このプロパティは数値を値として取り、負の数も使うことができます。

# --instructions--

`order`のCSSプロパティを`#box-1`と`#box-2`の両方に追加してください。`#box-1`には`2`の値を、`#box-2`には`1`の値を設定してください。

# --hints--

`#box-1`要素には`order`プロパティを`2`の値で設定してください。

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

`#box-2`要素には`order`プロパティを`1`の値で設定してください。

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
