---
id: 587d78ad367417b2b2512afb
title: flex-shrink プロパティを使ってアイテムを縮小する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

これまでのチャレンジで使われてきたプロパティはすべてフレックスアイテムの親であるフレックスコンテナに適用されていました。しかし、フレックスアイテムに対しても便利なプロパティがいくつかあります。

最初のものは`flex-shrink`プロパティです。これを使うと、フレックスコンテナが小さすぎる場合にアイテムを縮小できます。親コンテナの幅が中のすべてのフレックスアイテムの幅の合計より小さいときにアイテムは縮小します。

`flex-shrink`プロパティは数値を値として取ります。数値が大きいほど、コンテナ内の他のアイテムと比べてより多く縮小します。たとえば、あるアイテムの`flex-shrink`の値が`1`で、もう一方の`flex-shrink`の値が`3`の場合、`3`の値を持つ方はもう一方の3倍縮小します。

# --instructions--

`flex-shrink`のCSSプロパティを`#box-1`と`#box-2`の両方に追加してください。`#box-1`には`1`の値を、`#box-2`には`2`の値を与えてください。

# --hints--

`#box-1`要素には`flex-shrink`プロパティを`1`の値で設定してください。

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

`#box-2`要素には`flex-shrink`プロパティを`2`の値で設定してください。

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
