---
id: 587d78ae367417b2b2512afc
title: flex-grow プロパティを使ってアイテムを拡大する
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

`flex-shrink`の反対は`flex-grow`プロパティです。`flex-shrink`はコンテナが縮むときのアイテムのサイズを制御することを思い出してください。`flex-grow`プロパティは親コンテナが拡大するときのアイテムのサイズを制御します。

前回のチャレンジの似た例を使うと、1つのアイテムの`flex-grow`の値が`1`で、もう1つのアイテムの`flex-grow`の値が`3`の場合、`3`の値を持つ方がもう一方の3倍の大きさに成長します。

# --instructions--

`flex-grow`のCSSプロパティを`#box-1`と`#box-2`の両方に追加してください。`#box-1`には`1`の値を、`#box-2`には`2`の値を設定してください。

# --hints--

`#box-1`要素には`flex-grow`プロパティを`1`の値に設定してください。

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

`#box-2`要素には`flex-grow`プロパティを`2`の値に設定してください。

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
