---
id: 587d78ad367417b2b2512afa
title: フレックスラッププロパティを使って行または列を折り返す
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSSのフレックスボックスには、フレックスコンテナを複数の行（または列）に分割する機能があります。デフォルトでは、フレックスコンテナはすべてのフレックスアイテムを一緒に収めます。例えば、行はすべて1行に並びます。

しかし、`flex-wrap`プロパティを使うと、CSSにアイテムを折り返すよう指示できます。つまり、余分なアイテムは新しい行または列に移動します。折り返しが起こるブレークポイントは、アイテムのサイズとコンテナのサイズによって決まります。

CSSには折り返しの方向を指定するオプションもあります：

<ul><li><code>nowrap</code>: これはデフォルトの設定で、アイテムを折り返しません。</li><li><code>wrap</code>: 行の場合は上から下へ、列の場合は左から右へ複数行に折り返します。</li><li><code>wrap-reverse</code>: 行の場合は下から上へ、列の場合は右から左へ複数行に折り返します。</li></ul>

# --instructions--

現在のレイアウトでは1行にボックスが多すぎます。`flex-wrap`要素に`#box-container`のCSSプロパティを追加し、その値を`wrap`にしてください。

# --hints--

`#box-container`要素に`flex-wrap`プロパティを`wrap`の値で設定してください。

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
