---
id: 587d78ab367417b2b2512af0
title: 'display: flex を使って2つのボックスを配置する'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

このセクションでは、CSSを使って要素を柔軟に配置する方法を示すために、交互にチャレンジスタイルを使います。最初に理論を説明するチャレンジがあり、その後、シンプルなツイートコンポーネントを使った実践的なチャレンジでflexboxの概念を適用します。

要素にCSSプロパティ`display: flex;`を設定すると、他のflexプロパティを使ってレスポンシブなページを作成できます。

# --instructions--

`display`のCSSプロパティを`#box-container`に追加し、その値を`flex`に設定してください。

# --hints--

`#box-container`には`display`プロパティを`flex`の値に設定する必要があります。

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
