---
id: 587d78ae367417b2b2512afd
title: アイテムの初期サイズを設定するためにflex-basisプロパティを使う
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

`flex-basis`プロパティは、CSSが`flex-shrink`や`flex-grow`で調整を行う前のアイテムの初期サイズを指定します。

`flex-basis`プロパティで使われる単位は、他のサイズ指定プロパティ（`px`、`em`、`%`など）と同じです。値`auto`はコンテンツに基づいてアイテムのサイズを決めます。

# --instructions--

`flex-basis`を使ってボックスの初期サイズを設定してください。`flex-basis`と`#box-1`の両方にCSSプロパティ`#box-2`を追加してください。`#box-1`には`10em`の値を、`#box-2`には`20em`の値を与えてください。

# --hints--

`#box-1`要素には`flex-basis`プロパティが必要です。

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-1`要素には`flex-basis`の値として`10em`を設定してください。

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

`#box-2`要素には`flex-basis`プロパティが必要です。

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-2`要素には`flex-basis`の値として`20em`を設定してください。

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
