---
id: 60a3e3396c7b40068ad6996f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

`div` に `body` 要素を追加してください。

`class` と等しい `canvas` 属性を設定してください。

これがペイント用のキャンバスとして機能します。

# --hints--

コードには `div` 要素が必要です。

```js
assert(document.querySelectorAll('div').length === 1)
```

`div` 要素には、値が `class` の `canvas` が必要です。

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
