---
id: 60a3e3396c7b40068ad69977
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

パディングを使って要素内の間隔を調整しましょう。

`.frame`では、`padding`省略形プロパティを使って、`.frame`と`.canvas`の要素間のスペースを`50px`だけ増やしてください。この省略形プロパティは、要素のボーダーとキャンバスの内側の上下左右のスペースを増やします。

# --hints--

`padding`プロパティを`50px`に設定してください。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame`要素は`padding`に`50px`の値を持つべきです。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
