---
id: 60a3e3396c7b40068ad6996e
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

マージンはボックスの外側のエリアで、他のボックスや要素との間のスペースを調整するために使えます。

ここでは、下の要素がより大きな上マージンを持っていて、ページの下の方に押し下げられています。

CSSボックスモデルをざっとレビューしたので、ロスコの絵画に取りかかりましょう。

`img` 要素を削除してください。

# --hints--

コード内に `img` 要素があってはいけません。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
