---
id: 60a3e3396c7b40068ad69979
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

`div` 要素の中に新しい `.canvas` 要素を追加してください。

新しい `div` に `class` 属性を与え、その値を `one` にしてください。これが最初の長方形になります。

# --hints--

新しい `div` 要素を作成してください。

```js
assert(document.querySelectorAll('div').length === 3);
```

新しい `div` 要素は `.canvas` 要素の中にネストしてください。

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

新しい `div` には `class` 属性を持ち、その値は `one` にしてください。

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
