---
id: 60a3e3396c7b40068ad6997f
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

キャンバスの上下左右に1ピクセルのパディングを追加したことで、サイズが502ピクセル×602ピクセルに変わりました。

`padding` プロパティを削除し、`overflow` を `hidden` に設定して、キャンバスを元のサイズに戻してください。

# --hints--

`padding` セレクターから `.canvas` プロパティを削除してください。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

`overflow` プロパティを `hidden` に設定してください。

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

`.canvas` 要素の `overflow` は `hidden` である必要があります。

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
