---
id: 60a3e3396c7b40068ad6996b
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

`img` 要素の中に、`src` の `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` を持つ `body` 要素を追加してプロジェクトを始めてください。

# --hints--

コードには `img` 要素が必要です。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 要素は `src` の `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` を持つべきです。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
