---
id: 60a3e3396c7b40068ad69972
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

`height`ルールに、値が`600px`の`.canvas`プロパティを追加してください。

# --hints--

`height`プロパティを`600px`に設定してください。

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

`.canvas`要素は`height`の`600px`を持つべきです。

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
