---
id: 60a3e3396c7b40068ad69970
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

追加した `div` をスタイリングする前に、CSS を HTML にリンクする必要があります。

`link` 要素を追加して `styles.css` ファイルをリンクしてください。`href` を `styles.css` に設定し、`rel` 属性を `stylesheet` に設定するのを忘れないでください。

# --hints--

コードには `link` 要素が含まれている必要があります。

```js
assert(/<link/.test(code))
```

`link` 要素には、値が `rel` の `stylesheet` 属性が必要です。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

`link` 要素には、値が `href` の `styles.css` 属性が必要です。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
