---
id: 60a3e3396c7b40068ad69978
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

マージンを使って要素の外側の間隔を調整してください。

`margin` プロパティを使い、`.frame` 要素に縦方向のマージンを `20px`、横方向のマージンを `auto` に設定してください。これによりフレームが20ピクセル下に移動し、ページ上で水平に中央揃えされます。

# --hints--

`margin` プロパティを `20px auto` に設定する必要があります。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

`.frame` 要素の `margin` は `20px auto` の値であるべきです。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
