---
id: 66e45c8140f9fda5c105ae26
title: Створіть запрошення на подію
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `header` у межах тіла.
1. Елемент `header` повинен містити зображення та `h1` з назвою події (в такому ж порядку). Можете використати це зображення: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. Ви повинні мати елемент `main` у межах тіла.
1. Елемент `main` повинен містити принаймні два елементи `section`, які надають інформацію про подію.
1. Елементи `section` повинні містити `h2`.
1. Тіло повинне мати верхній і нижній відступи зі значенням `50px`.
1. Тіло повинне мати верхнє та нижнє поля зі значенням `0`, а ліве і праве поля мають бути вирівняними за центром.
1. Ширину тіла потрібно встановити відносно ширини вікна перегляду.
1. Тіло має використовувати функцію `calc`, щоб встановити властивість `min-height` на 100% від ширини вікна перегляду, не враховуючи верхнього та нижнього відступів тіла.
1. Ви повинні мати принаймні один елемент `hr` у запрошенні.
1. Встановіть відсоткове значення для властивості `width` всіх елементів `hr` та `section` відносно батьківського елемента.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

В межах елемента `body` має бути елемент `header`.

```js
assert.exists(document.querySelector('body > header'));
```

Першим дочірнім елементом елемента `header` має бути `img`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

Другим дочірнім елементом елемента `header` має бути `h1`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Ви повинні мати тільки один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Елемент `h1` повинен містити назву події.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

В межах елемента `body` має бути елемент `main`.

```js
assert.exists(document.querySelector('body > main'));
```

В межах елемента `main` має бути принаймні два елементи `section`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Елементи `section` повинні містити `h2`.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Елементи `h2` не можуть бути порожніми.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Елемент `body` повинен мати верхній і нижній відступи зі значенням `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Елемент `body` повинен мати верхнє та нижнє поля зі значенням `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Елемент `body` повинен мати ліве та праве поля зі значенням `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Елемент `body` повинен використовувати `vw`, щоб встановити свою ширину.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Елемент `body` повинен використовувати `calc`, щоб встановити `min-height` як `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

Ви повинні мати принаймні один елемент `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

Ширина елементів `hr` має бути встановленою у відсотковому значенні.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

Ширина елементів `section` має бути встановленою у відсотковому значенні.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
