---
id: 66c07238b01053abaf812065
title: Побудувати калькулятор факторіала
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Факторіал — це добуток цілого числа та всіх цілих чисел, що менші за нього. Наприклад, факторіал `5` дорівнює `5 * 4 * 3 * 2 * 1 = 120`. У цій лабораторній роботі ви створите калькулятор факторіала, який прийматиме число від користувача і обчислюватиме факторіал цього числа.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні оголосити змінну `num` і призначити їй число на свій вибір. Призначене число має бути між `1` та `20` включно.
1. Створіть функцію з ім’ям `factorialCalculator`, яка прийматиме число як аргумент і повертатиме факторіал цього числа.
1. Усередині функції оголосіть змінну `result` і призначте їй значення `1`. Використовуючи цикл, пройдіть усі числа від `1` до вхідного числа (включно) і для кожного числа помножте змінну `result` на поточне число та призначте результат змінній `result`. Ви можете використати тут цикл `for`, `while` або `do...while`.
1. Ви повинні викликати функцію `factorialCalculator` з аргументом `num` і призначити результат змінній `factorial`.
1. Ви повинні зберегти фінальний вихід у форматі `Factorial of [num] is [factorial]` і призначити його змінній `resultMsg`.
1. Ви повинні вивести значення `resultMsg` у консоль.

# --hints--

Ви повинні мати змінну `num`.

```js
assert.isDefined(num);
```

Ви повинні призначити значення змінній `num`.

```js
assert.isNotNull(num);
```

Значення `num` має бути між `1` та `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

Значення `num` має бути числом.

```js
assert.isNumber(num);
```

Ви повинні мати функцію з ім’ям `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

Функція `factorialCalculator` має приймати число як аргумент.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

Факторіал `5` має повертати `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

Факторіал `7` має повертати `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

Ви повинні викликати вашу функцію `factorialCalculator` з аргументом `num`.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

Ви повинні оголосити змінну `factorial` і призначити їй результат функції `factorialCalculator`.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

Ваш `factorialCalculator` має давати правильний результат.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Ваш `resultMsg` має містити рядок у форматі `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

Ви повинні вивести значення `resultMsg` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
