---
id: 683ec95f2f8781355556ff81
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

你需要測試各種事物，這樣你才能在這個初始階段了解演算法的流程。

為了做到這一點，你需要先跳出迴圈。那是因為目前的實作只允許一次迭代，所以如果條件不符合，將會有無窮迴圈。

緊接在 `if` 述語之後，使用 `break` 關鍵字跳出 `while` 迴圈。然後，在 `while` 迴圈之後，傳回一個空的列表以表示未找到該值。

# --hints--

你應該在 `if` 述語之後使用 `break` 關鍵字來跳出 `while` 迴圈。

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

你應該在你的 `while` 迴圈之後傳回一個空的列表。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
