---
id: 683ec95f2f8781355556ff80
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

接下來要做的是創建一個條件，用來檢查標的 `value` 是否在中間。 

在 `while` 迴圈內，創建一個 `if` 述語來檢查標的 `value` 是否與 `value_at_middle` 相等。如果是，傳回 `path_to_target` 變數。

# --hints--

你應該創建一個 `if` 述語來檢查 `value` 是否與 `value_at_middle` 相等。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

你應該在 `if` 述語內傳回 `path_to_target` 變數。

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
