---
id: 683ec95f2f8781355556ff7c
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

既然你已經定義了搜尋的邊界，現在是時候創建將執行二分搜尋法的迴圈。

二分搜尋法透過反覆縮小搜尋範圍來運作。只要有有效的元素區間可供檢查，這個處理程序就會持續進行。

你可以用一個 `while` 迴圈來表示，只要你的 `low` 指標小於或相等於你的 `high` 指標，迴圈就會繼續。

如果 `low` 變得大於 `high`，表示搜尋空間已經變成空的，且 `value` 不在列表中。

在 `binary_search` 函式（程式）內，創建一個 `while` 迴圈，其條件是檢查 `low` 是否小於或相等於 `high`。在迴圈內，暫時將 `pass` 關鍵字作為佔位符添加。

# --hints--

你應該創建一個 `while` 迴圈，當 `low` 小於或相等於 `high` 時執行。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
