---
id: 683ec95f2f8781355556ff82
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

請記得該函式（程式）接受一個名為 `search_list` 的列表和你要搜尋的 `value`。

現在，使用 `binary_search([1, 2, 3, 4, 5], 3)` 和 `binary_search([1, 2, 3, 4, 5, 9], 4)` 呼叫函式，並立即列印這些呼叫。

# --hints--

你應該使用 `binary_search([1, 2, 3, 4, 5], 3)` 呼叫該函式，並立即列印它。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

你應該使用 `binary_search([1, 2, 3, 4, 5, 9], 4)` 呼叫該函式，並立即列印它。

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
