---
id: 683ec95f2f8781355556ff7d
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

要開始使用迴圈，你需要找到目前搜尋空間的中間索引。

你可以透過計算 `low` 和 `high` 變數的平均值來取得中間索引。你必須在將兩個值相加後使用向下取整除法來取得平均值，因此答案將永遠是向下取整的整數，因為索引必須是整數。

從你的 `while` 迴圈中移除 `pass` 關鍵字，並創建一個 `mid` 變數，設定為 `low` 和 `high` 的平均值。

# --hints--

你應該從 while 迴圈中移除 `pass` 關鍵字，並創建一個 `mid` 變數，設定為 `(low + high) // 2`。

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
