---
id: 683ec95f2f8781355556ff83
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

你可以看到第一次呼叫傳回 `[3]`。那是因為 `3` 是 `[1, 2, 3, 4, 5]` 中間的值。

另一方面，第二次呼叫傳回一個空列表。這是因為在檢查中間值之後，迴圈目前缺少透過更新 `low` 或 `high` 變數來縮小搜尋範圍的邏輯。它僅作為一次中點檢查，然後結束而不進行進一步搜尋。

為了讓二分搜尋法繼續縮小搜尋範圍，為你的程式碼添加一個 `elif` 區段，用來檢查 `value` 是否大於 `value_at_middle`。暫時在 `elif` 區段內添加 `pass` 關鍵字。

# --hints--

當 `value` 大於 `value_at_middle` 時，你應該為 `elif` 添加一個區塊。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
