---
id: 683ec95f2f8781355556ff7b
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

在二分搜尋法中，有最低可能的索引和最高可能的索引。

最低可能的索引表述目前在被搜尋的值部分中最左邊的邊界，在此情況下為 `search_list`。另一方面，最高可能的索引表述被搜尋值中最右邊的邊界。

為了考慮這兩個，你必須考慮整個被搜尋的列表。因此，在 `binary_search` 函式中，定義一個值為 `0` 的 `low` 變數，以及一個值為被搜尋列表中最後索引的 `high` 變數。

# --hints--

你應該在你的 `binary_search` 函式中有一個名為 `low` 的變數。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

你的 `low` 變數應該設定為 `0`。

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

你應該在你的 `binary_search` 函式中有一個名為 `high` 的變數。

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

你的 `high` 變數應該設定為 `len(search_list) - 1`。

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
