---
id: 683ec95f2f8781355556ff7e
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在你已經計算出 `mid` 索引，你需要實際從 `search_list` 中該索引取得值。這個值是你將用來與你的 `value` 參數，也就是你正在搜尋的標的，比較的。

創建一個 `value_at_middle` 變數，並將位於 `mid` 索引的 `search_list` 元素指定給它。 

# --hints--

你應該創建一個 `value_at_middle` 變數，並設定為 `search_list[mid]`。

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
