---
id: 683ec39565dbc7e3375263b1
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在之前的課程中，你學習了二分搜尋法以及它與線性搜尋的不同之處。

在這個工作坊中，你將透過建置二分搜尋法演算法來實作這個概念。

首先，創建一個帶有 `search_list` 和 `value` 參數的 `binary_search` 函式（程式）。

`search_list` 是函式將要在其中搜尋的元素列表，而 `value` 是函式將嘗試在 `search_list` 中尋找的標的。

暫時在函式內添加 `pass` 關鍵字，以便測試能通過。

# --hints--

你應該創建一個 `binary_search` 函式（程式）。

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

你的 `binary_search` 函式應該有 `search_list` 和 `value` 參數。

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
