---
id: 683ec95f2f8781355556ff7a
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

你正在建置的二分搜尋法演算法不僅會傳回 `True` 或 `False` 來顯示是否找到該值，還會傳回尋找該值所走過的路徑。為此，你會有一個 `path_to_target` 變數來追蹤路徑。

在你的 `binary_search` 函式中，將 `pass` 關鍵字取代為 `path_to_target` 變數，並初始化為空的列表，作為開始。

# --hints--

你應該在你的 `binary_search` 函式（程式）中創建一個設定為空列表的 `path_to_target` 變數。別忘了移除 `pass` 關鍵字。

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
