---
id: 683ec95f2f8781355556ff84
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

如果 `elif` 中的條件為真，則移除 `pass` 關鍵字，並將 `low` 變數的值更新為 `mid` 變數加上 `1`。

這將把搜尋範圍擴展到列表中目前搜尋區域的右半部，因為如果 `value` 大於 `value_at_middle`，表示 `value` 必須在目前搜尋區域的右半部。

# --hints--

你應該移除 `pass` 關鍵字並將 `low` 變數更新為 `mid + 1`。

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
