---
id: 684169b00a073c9b9f0a703c
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你有注意到第二次函式呼叫現在在列表中有 `3, 5, 4` 這些值嗎？那顯示二分搜尋法正在如預期運作。

事情是這樣發生的：演算法首先將 `3` 視為初始列表的中間。由於 `4` 大於 `3`，搜尋轉移到右半部。接著它將 `5` 視為新的中間。因為 `4` 小於 `5`，搜尋移向左邊，最終將 `4` 識別為最終範圍的中間。

要再次測試該函式，請使用 `[1, 3, 5, 9, 14, 22], 10` 來呼叫它，並立即列印該呼叫。這是值不會被找到的情況。

# --hints--

你應該使用 `binary_search([1, 3, 5, 9, 14, 22], 10)` 呼叫該函式，並立即列印它。

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
