---
id: 68416d66013635cb6325eb6d
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

你還可以添加的最後一個小改進是同時顯示找到標的 `value` 的索引。

所以，不是只在 `if` 述語中傳回 `path_to_target`，而是為你的 `if` 述語添加一個包含訊息 `Value found at index {mid}` 的 `f-string`。

這樣，你的二分搜尋法演算法工作坊就完成了！

# --hints--

你應該在 `if` 區塊內傳回 `path_to_target, f'Value found at index {mid}'`。

```js
({
    test: () => 
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target, f'Value found at index {mid}'")`
            )       
        )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            --fcc-editable-region--
            return path_to_target
            --fcc-editable-region--
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return [], "Value not found"

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```

# --solutions--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
          return path_to_target, f"Value found at index {mid}"
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return [], "Value not found"

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```
