---
id: 68c34d57ed874e4004f7b5d4
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

现在，文本内容为 `Break` 的 `td` 元素只跨越一列。但如果它能跨越所有三列，那就更好了。 

正如你从之前的研讨会和课程中回忆的那样，你可以使用 `colspan` 属性使表格单元格跨越多个列。

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

为你的 `td` 元素添加一个 `colspan` 属性并将其值设置为 `3`。

# --hints--

你的 `td` 元素中包含 `Break` 的文本内容，应将 `colspan` 属性设置为 `3`。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
