---
id: 68c34c608cf9e53f5d36b1d5
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

接下来，为表格添加第三行。首先添加另一个 `tr` 元素。在该 `tr` 元素内，添加一个带有 `scope` 属性且属性值为 `"row"` 的 `th` 元素，文本内容为 `11:00 AM`。然后在该 `th` 元素下方，添加一个文本内容为 `Break` 的 `td` 元素。 

# --hints--

你应该在 `tbody` 元素内有第三个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

你应该在第三个 `tr` 元素内有一个 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

你的 `th` 元素应具有 `11:00 AM` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

你的 `th` 元素应具有设置为 `"row"` 的 `scope` 属性。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

你应该在第三个 `tr` 元素内有一个 `td` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

你的 `td` 元素应具有 `Break` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
