---
id: 68c34a35c91b053dbae2a8c6
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

`scope` 属性的另一个值是 `row`，表示一个头部单元格是其整行的头部。

在你的 `th` 元素内，添加一个值为 `row` 的 `scope` 属性。

# --hints--

你的 `th` 元素应具有值为 `row` 的 `scope` 属性。

```js
assert.equal(document.querySelector("tbody tr th")?.getAttribute("scope"), "row");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        <th>9:00 AM</th>
        --fcc-editable-region--
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
