---
id: 68c353c4b114aa428c01f260
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

最后一步是为表格添加一行。

首先为你的表格添加另一个 `tr` 元素。在该 `tr` 元素内，添加一个带有 `scope` 属性且属性值为 `"row"` 的 `th` 元素，文本内容为 `2:00 PM`。 

然后，添加三个带有以下文本内容的 `td` 元素：

- `Voice UI Workshop`
- `Git & GitHub Essentials`
- `Color & Contrast in UI`

通过最后一组更改，你的表格现在已完成！

# --hints--

你应该在 `tbody` 元素内有第六个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 6);
```

你应该在第六个 `tr` 元素内有一个 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[5].querySelector("th"));
```

你的 `th` 元素应具有 `2:00 PM` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.textContent.trim(), "2:00 PM");
```

你的 `th` 元素应具有设置为 `"row"` 的 `scope` 属性。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.getAttribute("scope"), "row");
```

你的第六个 `tr` 元素应有三个 `td` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[5].querySelectorAll("td"), 3);
```

你的第一个 `td` 元素应包含 `Voice UI Workshop` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[0]?.textContent.trim(), "Voice UI Workshop");
```

你的第二个 `td` 元素应包含 `Git & GitHub Essentials` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[1]?.textContent.trim(), "Git & GitHub Essentials");
```

你的第三个 `td` 元素应具有 `Color & Contrast in UI` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[2]?.textContent.trim(), "Color & Contrast in UI");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

    --fcc-editable-region--
      
    --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

      <tr>
        <th scope="row">2:00 PM</th>
        <td>Voice UI Workshop</td>
        <td>Git & GitHub Essentials</td>
        <td>Color & Contrast in UI</td>
      </tr>
    </tbody>
  </table>
</body>
</html>
```

