---
id: 68c34ace730ca43e5fa7b8d6
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在是时候为表格添加另一行。 

首先为你的 `tr` 元素添加另一个 `tr` 元素。在该 `tr` 元素内，添加一个带有 `scope` 属性且属性值为 `"row"` 的 `th` 元素，文本内容为 `10:00 AM`。 

然后，添加三个带有以下文本内容的 `td` 元素：

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

你应该在 `tbody` 元素内有第二个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

你应该在第二个 `tr` 元素内有一个 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

你的 `th` 元素应具有 `10:00 AM` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

你的 `th` 元素应具有设置为 `"row"` 的 `scope` 属性。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

你的第二个 `tr` 元素应有三个 `td` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

你的第一个 `td` 元素应具有 `Accessibility Deep Dive` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

你的第二个 `td` 元素应包含 `CSS for Beginners` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

你的第三个 `td` 元素应具有 `Inclusive Design Principles` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
