---
id: 68c343692285523aef0fc6bf
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

`scope` 属性用于指定头部单元格是行、列或行组或列组的头部。以下是一个示例：

```html
<th scope="col">Example Header</th>
```

这有助于屏幕阅读器理解头部和数据单元格之间的关系。

对于所有 `th` 元素，为你的 `scope` 属性添加值为 `col`。

# --hints--

每个 `th` 元素都应具有值为 `"col"` 的 `scope` 属性。

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
