---
id: 68c341332d98fe38fae733af
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

现在是开始创建表格的时候了。 

首先在你的 `h1` 元素下方添加一个 `table` 元素。在 `table` 元素内部，添加一个带有文本 `Schedule by Track and Time` 的 `caption` 元素。

# --hints--

你应该有一个 `table` 元素。

```js
assert.exists(document.querySelector("table"));
```

你的 `table` 元素应包含一个 `caption` 元素。

```js
assert.exists(document.querySelector("table caption"));
```

你的 `caption` 元素应包含文本 `Schedule by Track and Time`。

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
