---
id: 68c34329d430583a6cd5f7f1
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

在你的 `tr` 元素内，添加四个 `th` 元素。第一个 `th` 元素应包含 `Time` 的文本，第二个应包含 `Track A` 的文本，第三个应包含 `Track B` 的文本，第四个应包含 `Track C` 的文本。

# --hints--

你的 `tr` 元素应包含四个 `th` 元素。

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

你的第一个 `th` 元素应包含 `Time` 的文本。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

你的第二个 `th` 元素应包含 `Track A` 的文本。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

你的第三个 `th` 元素的 `text` 应该是 `Track B`。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

你的第四个 `th` 元素应该有 `Track C` 的文本。

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
