---
id: 68c3482278a14e3cf225e897
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在你的 `tr` 元素内，添加一个带有文本 `9:00 AM` 的 `th` 元素。然后在该 `th` 元素下方，添加三个带有文本 `Keynote: Tech Future`、`Intro to Web Dev` 和 `UX for All` 的 `td` 元素。 

# --hints--

你的 `tr` 元素内应包含一个 `th` 元素。

```js
assert.exists(document.querySelector("tbody tr th"));
```

你的 `th` 元素应具有 `9:00 AM` 的文本内容。

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

你应该在你的 `th` 元素下方有总共三个 `td` 元素。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

你的第一个 `td` 元素应具有 `Keynote: Tech Future` 的文本内容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

你的第二个 `td` 元素应具有 `Intro to Web Dev` 的文本内容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

你的第三个 `td` 元素应具有 `UX for All` 的文本内容。

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
