---
id: 68c346d03709343bff132707
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在接下来的几个步骤中，你将构建表格的体。

首先在 `thead` 元素下方添加一个 `tbody` 元素。然后在 `tbody` 内部添加一个 `tr` 元素。

# --hints--

你应该在 `thead` 元素下方有一个 `tbody` 元素。

```js
assert.exists(document.querySelector("tbody"));
```

你应该在 `tbody` 元素内有一个 `tr` 元素。

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
