---
id: 68c33e502cd97f2fd52cb63b
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将学习如何通过创建一个技术会议的调度来制作可访问的表格元素。

首先为你的页面添加一个带有文本 `Tech Conference 2025 Schedule` 的 `h1` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

你的 `h1` 元素应包含文本 `Tech Conference 2025 Schedule`。

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
