---
id: 68c34e201b347c40acda8111
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

现在是时候为表格添加另一行。 

首先为你的 `tr` 元素添加另一个 `tr` 元素。在该 `tr` 元素内，添加一个带有 `scope` 属性且属性值为 `"row"` 的 `th` 元素，文本内容为 `11:30 AM`。 

然后，添加三个带有以下文本内容的 `td` 元素：

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

你应该在 `tbody` 元素内有第四个 `tr` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

你应该在第四个 `tr` 元素内有一个 `th` 元素。

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

你的 `th` 元素应具有 `11:30 AM` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

你的 `th` 元素应具有设置为 `"row"` 的 `scope` 属性。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

你的第四个 `tr` 元素应有三个 `td` 元素。

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

你的第一个 `td` 元素应具有 `AR/VR in Education` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

你的第二个 `td` 元素应具有 `JavaScript Fundamentals` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

你的第三个 `td` 元素应具有 `Design Systems at Scale` 的文本内容。

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
