---
id: a0b5010f579e69b815e7c5d6
title: بناء دالة ذكية لاستبدال الكلمات
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

في هذا المختبر، ستنشئ دالة تُجري عملية بحث واستبدال في الجملة باستخدام المعلمات المقدمة ثم تُرجع الجملة الجديدة.

**ملاحظة:** احفظ حالة الحرف الأول في الكلمة الأصلية عند استبدالها. على سبيل المثال، إذا كنت تريد استبدال الكلمة `Book` بالكلمة `dog`، يجب أن تُستبدَل بـ `Dog`

**الهدف**: حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `myReplace`.
1. يجب أن تأخذ دالة `myReplace` ثلاث معلمات: سلسلة نصية، كلمة ليتم استبدالها، والكلمة التي ستُستبدل بها.
1. يجب أن تُرجع دالة `myReplace` السلسلة النصية الجديدة مع الاستبدال.
1. يجب أن تحافظ على حالة الحرف الأول في الكلمة الأصلية عند الاستبدال.

# --hints--

يجب أن تنشئ دالة باسم `myReplace`.

```js
assert.isFunction(myReplace);
```

يجب أن تأخذ دالة `myReplace` ثلاث معلمات.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` يجب أن تُرجع السلسلة النصية `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` يجب أن تُرجع السلسلة النصية `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` يجب أن تُرجع السلسلة النصية `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` يجب أن تُرجع السلسلة النصية `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` يجب أن تُرجع السلسلة النصية `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` يجب أن تُرجع السلسلة النصية `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
