---
id: 6720b3191bbafd269bae0ae8
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Crie um elemento `div` que tenha um `id` de `greeting-card` e uma `class` de `card`.

Dentro do elemento `div`, adicione um `h1` com o texto `Happy Birthday!`. Então adicione um elemento de parágrafo com uma `class` chamada `message` e o texto `Wishing you all the happiness and joy on your special day!`.

# --hints--

Você deve ter um elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

O elemento `div` deve ter um `id` de `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

O elemento `div` deve ter uma `class` de `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Você deve ter um elemento `h1` dentro do elemento `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

O elemento `h1` deve conter o texto `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Você deve ter um elemento `p` dentro do `div`.

```js
assert.exists(document.querySelector("div > p"));
```

O elemento `p` deve ter uma `class` de `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

O elemento `p` deve conter o texto `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
