---
id: 6720b7f3155a9278d6b6a8b8
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Agora é hora de estilizar seu cartão de felicitações.

Adicione um seletor para o elemento `body`, então: 

- altere o `font-family` para ser `Arial` seguido do genérico `sans-serif`
- dê um padding em todos os lados de `40px`,
- defina a propriedade `text-align` para `center`.

# --hints--

Você deve ter um seletor `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Você deve definir `font-family` para `"Arial", sans-serif` no seletor `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Você deve definir `padding` para `40px` no seletor `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Você deve definir `text-align` como `center` no seletor `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
