---
id: 6720b1f2680ec706ed534acf
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará o uso de diferentes pseudo-classes e pseudo-elementos projetando um cartão de felicitações.O boilerplate HTML foi fornecido para você. 

Comece o workshop vinculando o arquivo `styles.css`.

# --hints--

O elemento `link` deve estar dentro do elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    
  </head>
  <body>
  </body>
</html>

--fcc-editable-region--

```

```css

```
