---
id: 6728fda085006772a247e779
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Dê a `.card` um `max-width` de `400px`, um `padding` de `40px` em todos os lados e uma `margin` de `0` para cima e para baixo e `auto` para esquerda e direita (use a propriedade abreviada).

# --hints--

`.card` deve ter um `max-width` de `400px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("max-width"), "400px");
```

`.card` deve ter um `padding` de `40px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("padding"), "40px");
```

`.card` deve ter um `margin` de `0 auto`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("margin"), ["0 auto", "0 auto 0 auto", "0 auto 0", "0px auto", "0px auto 0px auto", "0px auto 0px"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

--fcc-editable-region--

.card {
  background-color: white;

}

--fcc-editable-region--
```
