---
id: 672b4434dd6d6a1edc405419
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Adicione dois elementos `a` dentro do elemento `.card-links`.

O primeiro deve ter um texto `Send Card`, uma `class` `send-link` e um `href` `#send`.

O segundo deve ter um texto de `Share on Social Media`, um `class` de `share-link` e um `href` de `#share`.

# --hints--

O primeiro elemento `a` dentro de `.card-links` deve ter um texto `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

O primeiro elemento `a` deve ter uma `class` com o valor `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

O primeiro elemento `a` deve ter um `href` com o valor `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

O segundo elemento `a` dentro de `.card-links` deve ter um texto `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

O segundo elemento `a` deve ter um `class` com um valor de `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

O segundo elemento `a` deve ter um `href` com o valor `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
