---
id: 673332a6b63f9f1e1f81ba9b
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Agora você pode fazer algo semelhante para adicionar o emoji também após o título. 

Crie um seletor que aponte para o pseudo-elemento `::after` do elemento `h1`.Dê a ele uma propriedade `content` e defina seu valor como `" 🥳"` (observe que há um espaço antes do emoji).

# --hints--

Você deve ter um seletor `h1::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1::after"));
```

O seletor `h1::after` deve ter uma propriedade `content` definida como `" 🥳"`.

```js
assert.strictEqual(
  new __helpers.CSSHelp(document).getStyle("h1::after")?.getPropertyValue("content"),
  `" 🥳"`
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>

      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>

  <section id="send">
    <h2>Sending your card...</h2>
    <p>Card successfully sent to your recipient!</p>
  </section>

  <section id="share">
    <h2>Sharing your card...</h2>
    <p>Your card was shared on social media!</p>
  </section>

  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

--fcc-editable-region--
h1::before {
  content: "🥳 ";
}


--fcc-editable-region--

```
