---
id: 6626a060c4006f793e10cb33
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora el profesor necesita tu ayuda para convertir la puntuación del estudiante en una nota identificada por una letra.

Completa la función `getGrade` que toma un número `score` como parámetro. Tu función debe devolver una cadena que represente una calificación por letra basada en la puntuación numérica.

Aquí están los puntajes y sus correspondientes calificaciones por letra:

Rango de puntuaciónCalificación
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Consejos**

- Recuerda lo que aprendiste sobre sentencias condicionales (`if`, `else if` y `else`).
- Recuerda que aprendiste sobre los operadores de comparación (`>`, `<`, `>=`, `<=`, `===`).
- Recuerda que aprendiste sobre el operador lógico AND (`&&`).

# --hints--

Tu función `getGrade` debe devolver `"A++"` si la puntuación es `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Tu función `getGrade` debe devolver `"A"` si la puntuación es `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Tu función `getGrade` debe devolver `"B"` si la puntuación está entre `80` y `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Tu función `getGrade` debe devolver `"C"` si la puntuación es `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Tu función `getGrade` debe devolver `"D"` si la puntuación está entre `60` y `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Tu función `getGrade` debe devolver `"F"` si la puntuación es `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Tu función `getGrade` debe devolver `"F"` si la puntuación está entre `0` y `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
