---
id: a5229172f011153519423690
title: 构建一个奇数斐波那契和计算器
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

在本实验中，你将构建一个奇数 Fibonacci 和计算器，用于计算所有小于或等于给定正整数的奇数 Fibonacci 数的和。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事**

1. 你应该有一个接受数字作为参数的 `sumFibs` 函数。
1. `sumFibs` 函数应返回所有小于或等于给定数字的奇数斐波那契数的和。
1. 斐波那契序列以 `0` 和 `1` 开始，每个后续数字都是前两个数字的和。
1. 只有奇数的 Fibonacci 数应添加到总和中。

# --hints--

你应该有一个 `sumFibs` 函数。

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` 应返回一个数字。

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` 应该返回 `1785`。

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` 应该返回 `4613732`。

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` 应该返回 `5`。

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` 应该返回 `60696`。

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` 应该返回 `135721`。

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
