---
id: 6a57c331ae94fa4631b6086f
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

將編號為 `101` 的文物策展人更新為 `"Fran Sinclair"`。然後，使用 `console.log()` 呼叫並輸出 `collection[101].curator.name`，以記錄該文物的策展人名稱。

# --hints--

你應該使用 `101` 和 `"Fran Sinclair"` 呼叫 `updateCurator`。

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)/);
```

更新後，你應該再次記錄第一件文物策展人的 `name`。

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

第一件文物策展人的 `name` 應該是 `"Fran Sinclair"`。

```js
assert.strictEqual(collection[101].curator.name, "Fran Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
