---
id: 6a57c331ae94fa4631b60866
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

有時候工作人員會要求博物館沒有的文物。如果該 `id` 下沒有儲存任何文物，`collection[id]` 就會是 `undefined`，而從中讀取 `.title` 會丟出錯誤。

更新傳回述語以使用條件判斷。如果找到文物，傳回 `artifact.title`，如果沒有，則傳回字串 `"Artifact not found"`。

# --hints--

`getArtifactTitle(101)` 應該仍然傳回 `"Golden Mask"`。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

如果沒有任何文物符合該 `id`，`getArtifactTitle` 應該傳回 `"Artifact not found"`。

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` 應該傳回它找到的任何文物的標題，而不只是 `101`。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
