---
id: 6a57c331ae94fa4631b60872
title: 第 19 步驟
challengeType: 1
dashedName: step-19
---

# --description--

接著，傳回一個總結該文物的模板字串。第一行放標題，接著每行分別放類別、策展人姓名、目前展館和展示狀態。展館欄位請使用 `currentLocation.gallery`。

例如，如果你呼叫 `buildSummary(101)`，它應該傳回：

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)` 應該包含該文物的標題，`Golden Mask`。

```js
assert.include(buildSummary(101), "Golden Mask");
```

`buildSummary(101)` 應該包含一行 `Category: Ceremonial`。

```js
assert.include(buildSummary(101), "Category: Ceremonial");
```

`buildSummary(101)` 應該包含一行 `Curator: Fran Sinclair`。

```js
assert.include(buildSummary(101), "Curator: Fran Sinclair");
```

`buildSummary(101)` 應該包含一行 `Current Gallery: Hall C`，使用 `currentLocation.gallery`。

```js
assert.include(buildSummary(101), "Current Gallery: Hall C");
```

`buildSummary(101)` 應該包含一行 `On Display: true`。

```js
assert.include(buildSummary(101), "On Display: true");
```

`buildSummary(101)` 應該將每個欄位放在獨立的一行。

```js
assert.isAtLeast(buildSummary(101).split("\n").length, 5);
```

`buildSummary(102)` 應該使用最新的位置，顯示 `Current Gallery: Hall B`。

```js
assert.include(buildSummary(102), "Current Gallery: Hall B");
```

`buildSummary` 應該根據任何文物本身的值建置摘要，每行一個欄位且依序排列。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [
    { gallery: "Hall Y", year: 2001 },
    { gallery: "Hall Z", year: 2002 },
  ],
  tags: [],
  onDisplay: false,
};
const summary = buildSummary(457);
assert.deepEqual(
  summary.split("\n").map((line) => line.trim()),
  [
    "Test Artifact",
    "Category: Test Category",
    "Curator: Test Curator",
    "Current Gallery: Hall Z",
    "On Display: false",
  ]
);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
