---
id: 6a57c331ae94fa4631b6086c
title: 第 13 步驟
challengeType: 1
dashedName: step-13
---

# --description--

隨著時間推移，文物經常會在展示與非展示狀態之間切換。

創建一個名為 `toggleDisplayStatus` 的函式（程式），它接受一個 `id` 參數。使用 `collection[id]` 查找該文物。如果存在，將它的 `onDisplay` 屬性設為目前值的相反狀態。

# --hints--

你應該有一個名為 `toggleDisplayStatus` 的函式（程式），並帶有一個名為 `id` 的參數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

呼叫 `toggleDisplayStatus` 應該能在兩個方向上切換該文物的 `onDisplay` 屬性值。

```js
const before101 = collection[101].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, !before101);
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, before101);
const before102 = collection[102].onDisplay;
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, !before102);
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, before102);
```

當沒有文物符合該 `id` 時，`toggleDisplayStatus` 不應該丟出錯誤或更改該群集。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => toggleDisplayStatus(999));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

--fcc-editable-region--

--fcc-editable-region--
```
