---
id: 6a57c331ae94fa4631b6086a
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

隨著文物隨時間在不同展覽館之間移動，每次移動都會被記錄為 `locations` 陣列中的一個新項目。

創建一個名為 `moveArtifact` 的函式，帶有 `id`、`gallery` 和 `year` 這三個參數。使用 `collection[id]` 找到該文物。如果存在，將一個包含 `gallery` 和 `year` 的新物件推入它的 `locations` 陣列。

# --hints--

你應該有一個名為 `moveArtifact` 的函式，參數名稱為 `id`、`gallery` 和 `year`。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.moveArtifact);
const params = explorer.allFunctions.moveArtifact.parameters;
assert.equal(params.length, 3);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "gallery");
assert.equal(params[2].toString(), "year");
```

呼叫 `moveArtifact` 時，應該將一個新的 `{ gallery, year }` 物件推入該文物的 `locations` 陣列。

```js
const start = collection[101].locations.length;
moveArtifact(101, "Test Hall", 2030);
assert.lengthOf(collection[101].locations, start + 1);
assert.deepEqual(collection[101].locations[collection[101].locations.length - 1], { gallery: "Test Hall", year: 2030 });
```

`moveArtifact` 應該將新位置推入 `id` 所指向的任何文物。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
moveArtifact(457, "Test Vault", 2032);
assert.lengthOf(collection[457].locations, 2);
assert.deepEqual(collection[457].locations[1], { gallery: "Test Vault", year: 2032 });
```

當沒有文物符合該 `id` 時，`moveArtifact` 不應該丟出錯誤或更改群集。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => moveArtifact(999, "Test Hall", 2030));
assert.strictEqual(JSON.stringify(collection), before);
```

你應該將新位置推入該文物現有的 `locations` 陣列，而不是替換整個陣列。

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const locationsBefore = collection[458].locations;
const spy = __helpers.spyOn(locationsBefore, "push");
try {
  moveArtifact(458, "Push Hall", 2033);
  assert.deepEqual(spy.calls, [[{ gallery: "Push Hall", year: 2033 }]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].locations, locationsBefore);
assert.deepEqual(collection[458].locations[1], { gallery: "Push Hall", year: 2033 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

--fcc-editable-region--

--fcc-editable-region--
```
