---
id: 6a57c331ae94fa4631b60868
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

策展人會為文物加上標記，使其更容易被搜尋，而且標記絕不應該重複添加。在添加新標記之前，你可以使用 `includes()` 方法檢查該標記是否已存在，因為當陣列已包含該值時，它會傳回 `true`。

創建一個名為 `addTag` 的函式，接受 `id` 和 `tag` 兩個參數。使用 `collection[id]` 查找文物。如果該文物存在，且其 `tags` 陣列尚未包含 `tag`，就將 `tag` 推入 `artifact.tags`。

# --hints--

你應該有一個名為 `addTag` 的函式，且參數名稱為 `id` 和 `tag`。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

使用新的標記呼叫 `addTag` 時，應該將該標記添加到文物的 `tags` 陣列中。

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

使用文物已擁有的標記呼叫 `addTag` 時，不應該添加重複的標記。

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` 應該將標記添加到 `id` 所指向的任何文物上。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

當沒有文物符合 `id` 時，`addTag` 不應該丟出錯誤或更改群集。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

你應該將 `tag` 推入文物現有的 `tags` 陣列，而不是替換整個陣列。

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
