---
id: 6a57c265abbc95dfebc05098
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個文物來源稽核器。你會將博物館文物儲存在一個物件中，並撰寫函式來查詢它們、標記它們、在展覽廳之間移動它們，以及總結它們的歷史。

每件文物都以其 ID 儲存。創建一個名為 `collection` 的物件並新增一筆資料。將鍵設為 `101`，值設為一個物件，該物件有一個 `title` 為 `"Golden Mask"`，以及一個 `category` 為 `"Ceremonial"`。

# --hints--

你應該宣告一個名為 `collection` 的變數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` 應該是一個物件。

```js
assert.isObject(collection);
```

`collection` 應該只有一筆資料。

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

儲存在 `101` 底下的文物應該是一個物件。

```js
assert.isObject(collection[101]);
```

儲存在 `101` 底下的文物應該有一個 `title` 為 `"Golden Mask"`。

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

儲存在 `101` 底下的文物應該有一個 `category` 為 `"Ceremonial"`。

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
