---
id: 6a57c331ae94fa4631b60862
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

每件文物也會追蹤它曾經所在的位置、它的描述方式，以及是否正在展出。

為儲存在 `101` 底下的文物添加三個屬性：

- `locations`：一個陣列，包含兩個物件，分別是 `{ gallery: "Hall A", year: 2020 }` 和 `{ gallery: "Hall C", year: 2024 }`
- `tags`：一個陣列，包含字串 `"gold"` 和 `"egypt"`
- `onDisplay`：布林值 `true`

# --hints--

文物的 `locations` 應該是一個包含這兩個指定位置物件的陣列。

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

文物的 `tags` 應該是 `["gold", "egypt"]`。

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

文物的 `onDisplay` 應該是 `true`。

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
