---
id: 6a57c331ae94fa4631b60865
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

工作人員一天要查詢數十次文物，因此你的稽核員需要一個可重複使用的查詢功能。方括號也可以接受變數，所以 `collection[id]` 會傳回 `id` 所參考的那個文物。

創建一個名為 `getArtifactTitle` 的函式（程式），它接受一個 `id` 參數。使用 `collection[id]` 取得文物，將它存到名為 `artifact` 的變數，並傳回 `artifact.title`。

# --hints--

你應該有一個名為 `getArtifactTitle` 的函式（程式），並帶有一個名為 `id` 的參數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

你應該將 `collection[id]` 指定給一個名為 `artifact` 的變數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` 應該傳回 `artifact.title`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` 應該傳回 `"Golden Mask"`。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` 應該傳回 `"Bronze Tablet"`。

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
