---
id: 6a57c331ae94fa4631b60861
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

每個文物都有一個負責的人。

為儲存在鍵 `101` 下的物件添加一個 `curator` 屬性，並將其值設為一個物件，該物件具有 `id` 為 `201` 和 `name` 為 `"Earl Sinclair"`。

# --hints--

儲存在 `101` 下的文物應該有一個 `curator` 屬性。

```js
assert.property(collection[101], "curator");
```

`curator` 應該是一個物件。

```js
assert.isObject(collection[101].curator);
```

`curator` 應該有一個 `id` 為 `201`。

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator` 應該有一個 `name`，其值為 `"Earl Sinclair"`。

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
