---
id: 6a57c331ae94fa4631b60863
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

博物館有多於一件的文物記錄。為 `collection` 在鍵 `102` 下添加第二件文物，並遵循與第一件相同的結構。

使用這些數值：

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

`collection` 應該有兩筆記錄。

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

儲存在 `102` 下的文物應該符合所描述的物件。

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
