---
id: 6a57c331ae94fa4631b60864
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

稽核從閱讀已記錄的內容開始。由於每個文物都以其 ID 儲存，你可以透過將該 ID 放入中括號來存取它，例如 `collection[101]`。

使用 `console.log()` 呼叫 `collection[101].title` 來記錄第一個文物的標題。接著，使用 `console.log()` 呼叫 `collection[101].curator.name` 來記錄其策展人的姓名。

# --hints--

你應該將第一個文物的 `title` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

你應該將第一個文物策展人的 `name` 記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

你應該先記錄 `title`，再記錄策展人的 `name`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
