---
id: 6a57c331ae94fa4631b60870
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

最後，為任何 artifact 建置一個可讀的摘要。先處理 artifact 不存在的情況。

創建一個名為 `buildSummary` 的函式（程式），它接受一個 `id` 參數。使用 `collection[id]` 查找 artifact，並將其存入名為 `artifact` 的變數。如果找不到 artifact，傳回字串 `"Artifact not found"`。

# --hints--

你應該有一個名為 `buildSummary` 的函式（程式），並帶有一個名為 `id` 的參數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

你應該將 `collection[id]` 指定給一個名為 `artifact` 的變數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.artifact);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.artifact?.value.matches("collection[id]"));
```

當沒有 artifact 符合 `id` 時，`buildSummary` 應該傳回 `"Artifact not found"`。

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

當 artifact 存在時，`buildSummary` 不應該傳回 `"Artifact not found"`。

```js
assert.notStrictEqual(buildSummary(101), "Artifact not found");
assert.notStrictEqual(buildSummary(102), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
