---
id: 60a3e3396c7b40068ad69970
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Before you can start styling the `div` you added, you need to link your CSS to your HTML.

Додайте елемент `link`, щоб прив’язати свій файл `styles.css`. Встановіть `href` на `styles.css` та не забудьте встановити атрибут `rel` на `stylesheet`.

# --hints--

Ваш код повинен мати елемент `link`.

```js
assert(/<link/.test(code))
```

Елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
