---
id: 60a3e3396c7b40068ad6996e
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Margin is the area outside of the box, and can be used to control the space between other boxes or elements.

Тут нижній елемент має більше верхнє поле, штовхаючи його вниз по сторінці.

Тепер, коли ви розумієте блокову модель CSS, розпочнемо картину Ротко.

Видаліть елемент `<img>`.

# --hints--

Ви не повинні мати елемент `img` в своєму коді.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
