---
id: ab306dbdcc907c7ddfc30830
title: creare uno strumento per appiattire profondamente
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

In questo laboratorio implementerai un algoritmo per appiattire un array.

Appiattire un array significa trasformare un array annidato di qualsiasi profondità in un array monodimensionale unico. Il processo estrae tutti gli elementi in ordine, srotolando solo gli array. Altri tipi rimangono invariati.

Per esempio:

| Originale | Appiattito |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Obiettivo:** soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `steamrollArray`.
1. La funzione `steamrollArray` dovrebbe accettare un argomento: un array annidato.
1. La funzione dovrebbe appiattire l'array annidato, considerando livelli di annidamento variabili.
1. La tua soluzione non dovrebbe usare i metodi `Array.prototype.flat()` o `Array.prototype.flatMap()`.
1. Non dovresti usare variabili globali.

# --hints--

Dovresti avere una funzione `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` dovrebbe restituire `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` dovrebbe restituire `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` dovrebbe restituire `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` dovrebbe restituire `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

La tua soluzione non dovrebbe usare i metodi `Array.prototype.flat()` o `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Non dovresti usare variabili globali.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
