---
id: 66adc82170acb1464bd348eb
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

机器人发送的最后几条信息将集中在机器人最喜欢的主题上。

首先创建一个名为 `favoriteSubject` 的新变量，并将字符串 `"Computer Science"` 赋值给它。

# --hints--

你应该有一个名为 `favoriteSubject` 的变量。

```js
assert.isNotNull(favoriteSubject);
```

你的 `favoriteSubject` 变量应该是一个字符串。

```js
assert.isString(favoriteSubject);
```

你应该将 `"Computer Science"` 赋值给 `favoriteSubject` 变量。

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
