---
id: 66ad984e92ce213a66f7e51c
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

现在，是时候为你的 `bot` 变量赋值了。

在之前的课程中，你学习了如何像这样为变量赋值：

```js
variableName = "Here is the value";
```

请记住，赋值运算符 `=` 右侧的值就是你要赋给左侧变量的值。

将字串 `"teacherBot"` 赋值给 `bot` 变量。

# --hints--

你应该为 `bot` 变量赋值一个字符串。

```js
assert.isString(bot);
```

你的 `bot` 变量应被赋值为字符串 `"teacherBot"`。

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
