---
id: 66adc5919c1853448f119d11
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

现在看来，机器人想改变他们的昵称。

使用重新赋值，将字符串 `"awesomeTeacherBot"` 赋值给 `bot` 变量。

# --hints--

你应该将字符串 `"awesomeTeacherBot"` 赋值给 `bot` 变量。

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
