---
id: 66ada3f46945763dd97f43f8
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

之前，你创建了 `bot` 和 `botLocation` 变量。 现在，你将使用它们向控制台输出新信息。

在之前的课程中，你学习了如何使用 `+` 操作符进行字串串接，如下所示：

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

请记住，在将字符串与变量连接时需要注意空格。

创建一个名为 `botIntroduction` 的变量。

然后使用 `+` 运算符进行字符串连接，连接字符串 `"My name is "`，之后是 `bot` 变量，之后是一个句号（`.`）。 

将此值赋值给 `botIntroduction` 变量。

然后，将 `botIntroduction` 变量输出到控制台。

# --hints--

你应该有一个名为 `botIntroduction` 的变量。

```js
assert.isNotNull(botIntroduction);
```

你的 `botIntroduction` 变量应存储一个字符串的值。

```js
assert.isString(botIntroduction);
```

你应该使用 `+` 运算符进行字符串连接，连接字符串 `"My name is "` 和 `bot` 变量，之后是一个句号（`.`）。 注意空格。 将此连接字符串赋值给 `botIntroduction` 变量。

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

你的代码中应该有第四个 `console.log()` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

你的 `console` 语句应输出 `botIntroduction` 变量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
