---
id: 66adc6046acc18453decf577
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

要查看机器人的新昵称，需要在控制台中输出一条新信息。

创建一个名为 `newNicknameGreeting` 的新变量。

然后使用字符串连接，连接字符串 `"I love my nickname but I wish people would call me "` 和 `bot` 变量，之后跟一个句号。

将结果赋值给 `newNicknameGreeting` 变量。

然后，将 `newNicknameGreeting` 的值输出到控制台。

# --hints--

你应该有一个名为 `newNicknameGreeting` 的变量。

```js
assert.isNotNull(newNicknameGreeting);
```

你的 `newNicknameGreeting` 变量应该是一个字符串。

```js
assert.isString(newNicknameGreeting);
```

你应该使用字符串连接，连接字符串 `"I love my nickname but I wish people would call me "` 和 `bot` 变量，之后跟一个句号。

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

你的代码中应包含一个 `console` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

你应该将 `newNicknameGreeting` 的值输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
