---
id: 66ad8ab945d266383f318cbf
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

现在，你应该能在控制台中看到机器人发出的第一条信息。

现在是时候添加机器人的第二条信息了。

添加另一条 `console.log` 语句，向控制台输出消息 `"I am excited to talk to you."` 。

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你的代码中应该有第二个 `console.log()` 语句。

```js
assert.lengthOf(logSpy.calls, 2);
```

你的 `console` 语句应包含信息 `"I am excited to talk to you."`。 别忘了引号。

```js
assert.isTrue(
  logSpy.calls.some(
    (call) => /I\s+am\s+excited\s+to\s+talk\s+to\s+you\s*\./.test(call[0])
  )
);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
