---
id: 66ad8d150264db3926eccfeb
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的课程中，你学习了 `let` 关键字以及如何在 JavaScript 中声明变量。 

以下是提醒大家如何使用 `let` 关键字声明变量：

```js
let greeting;
```

使用 `let` 关键字声明一个名为 `bot` 的变量。

**注意**：在这里使用 `let` 是因为稍后在工作坊中，你将更改 `bot` 变量的值。

# --hints--

你应该在代码中使用 `let` 关键字。

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

你应该声明一个名为 `bot` 的变量。

```js
assert.isNotNull(bot);
```

你应该使用 `let` 关键字来声明变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
