---
id: 66ad984e92ce213a66f7e51c
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

当你需要用多个单词声明变量时，可以使用 <dfn>camelCase</dfn> 命名法。

使用 `camelCase` 时，第一个单词全部小写，随后每个单词的第一个字母大写。

这是一个例子：

```js
let thisIsCamelCase;
```

使用 `let` 声明一个名为 `botLocation` 的变量。

# --hints--

你应该使用 `let` 来声明变量。

```js
assert.lengthOf(code.match(/let/g), 2);
```

你的变量应命名为 `botLocation` 。

```js
assert.isNotNull(botLocation);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
