---
id: 66ad9c633fa26d3c1475eae3
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

现在，是时候再添加一条来自机器人的信息了。

在代码中添加另一条 `console` 语句，输出信息 `"Allow me to introduce myself."`。

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你的代码中应该有第三个 `console.log()` 语句。

```js
assert.lengthOf(logSpy.calls, 3);
```

你的 `console` 语句应包含信息 `"Allow me to introduce myself."`。 别忘了引号。

```js
assert.deepInclude(logSpy.calls, ["Allow me to introduce myself."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
