---
id: 66adc42868cab843ccee87d9
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

现在是时候查看新的 `bot` 值了。

首先创建一个名为 `nicknameIntroduction` 的新变量。

使用字符串连接，连接字符串 `"My nickname is "` 和 `bot` 变量，之后跟一个句号（`.`）。

将生成的字符串赋值给 `nicknameIntroduction` 变量。

然后，将 `nicknameIntroduction` 的值输出到控制台。 

# --hints--

你应该有一个名为 `nicknameIntroduction` 的变量。

```js
assert.isNotNull(nicknameIntroduction);
```

你的 `nicknameIntroduction` 变量应该是一个字符串。

```js
assert.isString(nicknameIntroduction);
```

你应该使用字符串连接，连接字符串 `"My nickname is "` 和 `bot` 变量，之后跟一个句号（`.`）。 注意空格。

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

你的代码中应包含一个 `console` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

你应该将 `nicknameIntroduction` 的值输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
