---
id: 66ad8ab945d266383f318cbf
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

现在，你应该能在控制台中看到机器人发出的第一条信息。 

现在是时候添加机器人的第二条信息了。

添加另一条 `console.log` 语句，向控制台输出消息 `"I am excited to talk to you."` 。

# --hints--

你的代码中应该有第二个 `console.log()` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

你的 `console` 语句应包含信息 `"I am excited to talk to you."`。 别忘了引号。

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
