---
id: 66adb844118ba74107ce771f
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

机器人的下一条信息将涉及机器人的位置。

创建一个名为 `botLocationSentence` 的变量。

然后使用 `+` 运算符进行字符串连接，连接字符串 `"I live in "` 和 `botLocation` 变量，之后是一个句号（`.`）。

将此值赋值给 `botLocationSentence` 变量。

然后，将 `botLocationSentence` 的值输出到控制台。

# --hints--

你应该有一个 `botLocationSentence` 变量。

```js
assert.isNotNull(botLocationSentence);
```

你的 `botLocationSentence` 变量应存储一个字符串的值。

```js
assert.isString(botLocationSentence);
```

你应该使用 `+` 运算符进行字符串连接，连接字符串 `"I live in "` 和 `botLocation` 变量，之后是一个句号（`.`）。 注意空格。

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

你的代码中应包含一个 `console` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

你的 `console` 语句应输出 `botLocationSentence` 变量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
