---
id: 66adbc8a9793d64250f7e609
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

在之前的课程中，你学习了如何像这样重新赋值给变量：

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

使用重新赋值，将字符串 `"professorBot"` 赋值给 `bot` 变量。

# --hints--

你应该将字符串 `"professorBot"` 赋值给 `bot` 变量。

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
