---
id: 66ad8294a0ad902f1b31b612
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本工作坊，你将通过构建一个问候机器人，学习如何使用 JavaScript 基础知识。

第一步，你需要由问候机器人向控制台输出一条信息。

请记住，你在之前的课程中学习了 `console.log()` 和字串。

下面是提示你如何使用带有字符串的 `console.log()` ：

```js
console.log("Hello, World!");
```

添加一条 `console.log()` 语句，将字符串 `"Hi there!"` 输出到控制台。 别忘了在信息周围加上引号！

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你的代码中应该有一个 `console.log()` 语句。

```js
assert.isNotEmpty(logSpy.calls);
```

你的 `console` 语句应包含信息 `"Hi there!"`。 别忘了引号。

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
