---
id: 66adcf383276814776aba3ca
title: 步骤14
challengeType: 1
dashedName: step-14
---

# --description--

接下来，创建一个名为 `favoriteSubjectSentence` 的变量。

使用字符串连接，连接字符串 `"My favorite subject is "` 和 `favoriteSubject` 变量，之后跟一个句号。

将结果赋值给 `favoriteSubjectSentence` 变量。

然后，将 `favoriteSubjectSentence` 的值记录到控制台日志中。

# --hints--

你应该有一个名为 `favoriteSubjectSentence` 的变量。

```js
assert.isNotNull(favoriteSubjectSentence);
```

你的 `favoriteSubjectSentence` 变量应该是一个字符串。

```js
assert.isString(favoriteSubjectSentence);
```

你应该使用字符串连接，连接字符串 `"My favorite subject is "` 和 `favoriteSubject` 变量，之后跟一个句号。 注意空格。

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

你的代码中应包含一个 `console` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

你应该将 `favoriteSubjectSentence` 的值记录到控制台日志中。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
