---
id: 66ad998285df023b23dacdd3
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在，是时候初始化 `botLocation` 变量。

当你需要用多个单词声明变量时，可以使用 <dfn>camelCase</dfn> 命名法。

使用 `camelCase` 时，第一个单词全部小写，随后每个单词的第一个字母大写。 

这是一个示例：

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

使用 `let` 关键字在同一行声明并赋值字串 `"the universe"` 给变量 `botLocation`。

# --hints--

使用 `let` 关键字声明一个变量。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

你的变量应命名为 `botLocation` 。

```js
assert.isNotNull(botLocation);
```

你应该为 `botLocation` 变量赋值一个字符串。

```js
assert.isString(botLocation);
```

你的 `botLocation` 变量应被赋值为字符串 `"the universe"`。

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
