---
id: 686daa7ed79ceacd0b264e80
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

你正在接近目标，现在看看这个例子：

```html
<svg viewBox="0 0 50 50">
</svg>
```

`viewBox` 属性控制 SVG 内部可见的图像部分。

- 前两个数字（`0 0`）设置了 `viewBox` 的起始位置——左上角（x 和 y）。
- 接下来的两个数字（`50 50`）定义了 `viewBox` 的宽度和高度。

将 `viewBox` 属性设置为 `0 0 24 24`。

# --hints--

你应该有一个 `viewBox` 属性。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

你应将 `viewBox` 的 x 坐标设置为 0。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

你应将 `viewBox` 的 y 坐标设置为 0。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

你应将 `viewBox` 的宽度设置为 24。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

你应将 `viewBox` 的高度设置为 24。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
