---
id: 686d749d6f32f708fdf2b7ad
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在之前的讲座中，你已经了解了 `svg` 元素及其在项目中常被用作图标的用法。 在实际代码库中，通常会依赖图标库，因此无需从头创建 `svg` 元素。 然而，在本工作坊中，你将通过构建一个心形图标来学习 `svg` 元素内部的核心属性。

首先在页面上创建一个 `svg` 元素。

# --hints--

页面上应包含一个 `svg` 元素。

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 

    --fcc-editable-region--
  </body>
</html>
```
