---
id: 686daa7ed79ceacd0b264e80
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

`path` 元素需要定义其形状。 这就该 `d` 属性发挥作用了。 它用于生成一系列命令字母和数字，从而绘制出形状。

这些字母代表诸如移动至、绘制直线和闭合等命令，而数字则代表坐标。

将心形图标的 `d` 属性设置为 `M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z`。

# --hints--

你的 `path` 元素应该有一个 `d` 属性。

```js
const path = document.querySelector('path');
const d = path.getAttribute('d');
assert.exists(d);
```

你应该将 `d` 属性设置为 `M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z`

```js
const path = document.querySelector('path');
const d = path.getAttribute('d');
assert.strictEqual(
  d,
  'M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z'
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24" viewBox="0 0 24 24">
      <path></path>
    </svg>
    --fcc-editable-region--
  </body>
</html>
```
