---
id: 686d749d6f32f708fdf2b7ad
title: 第 1 步
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在之前的课程中，你学习了 `svg` 元素以及它们如何经常被用作项目中的图标。在真实的代码库中，你通常会依赖图标库，因此不需要从头创建 `svg` 元素。不过，在本次工作坊中，你将构建一个心形图标，以了解 `svg` 元素内部使用的核心属性。

首先在页面上创建一个 `svg` 元素。

# --hints--

页面上应包含一个 `svg` 元素。

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
