---
id: 6814d8e1516e86b171929de4
title: "挑戰 1：母音平衡"
challengeType: 28
dashedName: challenge-1
---

# --description--

給定一個字串，判斷字串前半部的母音數量是否與後半部的母音數量相等。

- 該字串可以包含任何字元。
- 字母 `a`、`e`、`i`、`o` 和 `u`，無論是大寫還是小寫，都被視為母音。
- 如果字串中有奇數個字元，請忽略中間的字元。

# --hints--

`isBalanced("racecar")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` 應該傳回 `false`。

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` 應該傳回 `false`。

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` 應該傳回 `false`。

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

`isBalanced("a")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced("a"));
```

`isBalanced("E")` 應該傳回 `true`。

```js
assert.isTrue(isBalanced("E"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const chars = [...s];
  const half = Math.floor(chars.length / 2);

  const firstHalf = chars.slice(0, half);
  const secondHalf =
    chars.length % 2 === 0 ? chars.slice(half) : chars.slice(half + 1);

  const countVowels = arr =>
    arr.filter(c => vowels.includes(c.toLowerCase())).length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
