---
id: 681cb05adab50c87ddb2e513
title: "挑戰 2：基礎檢查"
challengeType: 28
dashedName: challenge-2
---

# --description--

給定一個表述數字的字串，以及一個介於 2 到 36 的整數進位，判斷該數字在該進位中是否有效。

- 這個字串可能包含數字和大寫或小寫字母。
- 檢查時不區分大小寫。
- 如果每個字元都是該進位制中有效的數字，則該數字有效。
- 不同進位制中有效數字的範例：
  - 2 進位制：0-1
  - 8 進位制：0-7
  - 10 進位制：0-9
  - 16 進位制：0-9 和 A-F
  - 36 進位制：0-9 和 A-Z

# --hints--

`isValidNumber("10101", 2)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("z", 36))
```

`isValidNumber("", 10)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("", 10))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
