---
id: 681cb05adab50c87ddb2e513
title: "挑戰 2：基礎檢查"
challengeType: 28
dashedName: challenge-2
---

# --description--

給定一個表述數字的字串，以及一個介於 2 到 36 的整數進位，判斷該數字在該進位中是否有效。

- 該字串可能包含整數，以及大寫或小寫字元。
- 該檢查應該不區分大小寫。
- 基數可以是 2 到 36 之間的任何數字。
- 如果每個字元都是給定進位制中有效的數字，則該數字是有效的。
- 有效進位制的數字範例：
  - 基數 2：0-1
  - 八進位：0-7
  - 十進位：0-9
  - 基數 16：0-9 和 A-F
  - 進位 36：0-9 和 A-Z

# --hints--

`isValidNumber("10101", 2)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` 應該傳回 `false`。

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` 應該傳回 `true`。

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
