---
id: 681cb1afdab50c87ddb2e517
title: "挑戰 6：字母重組檢查器"
challengeType: 28
dashedName: challenge-6
---

# --description--

給定兩個字串，判斷它們是否為彼此的變位字（包含相同字元但順序可不同）。

- 忽略大小寫和空白。

# --hints--

`areAnagrams("listen", "silent")` 應該傳回 `true`。

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` 應該傳回 `true`。

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` 應該傳回 `true`。

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` 應該傳回 `false`。

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` 應該傳回 `false`。

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` 應該傳回 `false`。

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
