---
id: 681cb1b0dab50c87ddb2e51b
title: "挑戰 10：3 次機會"
challengeType: 28
dashedName: challenge-10
---

# --description--

給定一個介於 1 和 10,000 之間的整數，傳回從 1 到該整數中，其平方數包含至少一個數字 3 的數字個數。

# --hints--

`squaresWithThree(1)` 應該傳回 `0`。

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` 應該傳回 `1`。

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` 應該傳回 `19`。

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` 應該傳回 `326`。

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` 應該傳回 `4531`。

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
