---
id: 681cb1afdab50c87ddb2e516
title: "挑戰 5：Jbelmud 文字"
challengeType: 28
dashedName: challenge-5
---

# --description--

給定一個字串，傳回該字串的亂序版本，其中每個字都依照以下條件進行轉換：

- 字的第一個和最後一個字母保持原位。
- 第一個和最後一個字母之間的所有字母皆按字母順序排序。
- 輸入字串將不包含標點符號，且全部為小寫。

# --hints--

`jbelmu("hello world")` 應該傳回 `"hello wlord"`。

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` 應該傳回 `"i love jbelmud text"`。

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` 應該傳回 `"faccdeeemorp is my faiortve pacle to laern to cdoe"`。

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` 應該傳回 `"the qciuk borwn fox jmpus oevr the lazy dog"`。

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
