---
id: 681cb1afdab50c87ddb2e515
title: "挑戰 4：S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

給定一個字串，移除字串中所有空格，並在每個字元之間插入兩個空格，將所有英文字母轉換為大寫，然後傳回結果。

- 非字母字元應保持不變（空格除外）。

# --hints--

`spaceJam("freeCodeCamp")` 應該傳回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` 應該傳回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` 應該傳回 `"H  E  L  L  O  W  O  R  L  D  ?  !"`。

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` 應該傳回 `"C  @  T  $  &  D  0  G  $"`。

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` 應該傳回 `"A  L  L  Y  O  U  R  B  A  S  E"`。

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
