---
id: 681cb1b0dab50c87ddb2e519
title: "挑戰 8：階乘器"
challengeType: 28
dashedName: challenge-8
---

# --description--

給定一個介於 0 到 20 的整數，傳回該數字的階乘。階乘是從 1 到該數字之間所有數字的乘積。

- 零的階乘是 1。

# --hints--

`factorial(0)` 應該傳回 `1`。

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` 應該傳回 `120`。

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` 應該傳回 `2432902008176640000`。

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
