---
id: 681cb1b0dab50c87ddb2e51a
title: "挑戰 9：平方和"
challengeType: 28
dashedName: challenge-9
---

# --description--

給定一個最大為 1,000 的正整數，傳回從 1 到該數字所有整數平方的總和。

# --hints--

`sumOfSquares(5)` 應該傳回 `55`。

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` 應該傳回 `385`。

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` 應該傳回 `5525`。

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` 應該傳回 `41791750`。

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` 應該傳回 `333833500`。

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
