---
id: 6821ebd4237de8297eaee791
title: "挑戰 15：camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

給定一個字串，使用以下規則傳回其駝峰式大小寫版本：

- 字串引數中的字、單字由以下集合中的一個或多個字元分隔：空格（` `）、破折號（`-`）或底線（`_`）。將這些字元的任意序列視為字、單字的分隔。
- 第一個字應該全部是小寫。
- 每個後續字、單字的首字母應為大寫，其餘字母則為小寫。
- 所有空格和分隔符應該被移除。

# --hints--

`toCamelCase("hello world")` 應該傳回 `"helloWorld"`。

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` 應該傳回 `"helloWorld"`。

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` 應該傳回 `"secretAgentX"`。

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` 應該傳回 `"freeCodeCamp"`。

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` 應該傳回 `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`。

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
