---
id: 68216ef80f957572e7c340c5
title: "挑戰 12：訊息解碼器"
challengeType: 28
dashedName: challenge-12
---

# --description--

給定一個秘密訊息字串，以及一個整數表示用來位移該訊息以編碼的字母數，傳回解碼後的字串。

- 正數表示訊息在字母表中向前移動。
- 負數表示訊息在字母表中向後移動。
- 大小寫有區別，解碼後的字元應保留其編碼對應字元的大小寫。
- 非字母字元不應被解碼。

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` 應該傳回 `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` 應該傳回 `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` 應該傳回 `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` 應該傳回 `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
