---
id: 6821ebea237de8297eaee795
title: "挑戰 19：燭光"
challengeType: 28
dashedName: challenge-19
---

# --description--

給定一個整數表示你開始時擁有的蠟燭數量，以及一個整數表示燃燒多少根蠟燭可以創建一根新的蠟燭，傳回在創建並燃燒盡可能多的蠟燭後你將使用的蠟燭數量。

例如，如果給定 7 根蠟燭且需要 2 根燃燒過的蠟燭來製作一根新的：

1. 燃燒 7 根蠟燭以獲得 7 個剩餘物，
2. 將 6 個剩餘物回收成 3 根新蠟燭（剩餘 1 個剩餘物），
3. 燃燒 3 根蠟燭以獲得 3 個剩餘物（共 4 個），
4. 將 4 個剩餘物回收成 2 根新蠟燭，
5. 燃燒 2 根蠟燭以獲得 2 個剩餘物，
6. 將 2 個剩餘物回收成 1 支新蠟燭，
7. 點燃 1 根蠟燭。

在範例中你將總共點燃 13 根蠟燭。

# --hints--

`burnCandles(7, 2)` 應該傳回 `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` 應該傳回 `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` 應該傳回 `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` 應該傳回 `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` 應該傳回 `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
