---
id: 6821ebee237de8297eaee796
title: "挑戰 20：陣列重複項目"
challengeType: 28
dashedName: challenge-20
---

# --description--

給定一個整數陣列，傳回一個整數陣列，該陣列包含在初始陣列中出現超過一次的整數，並以遞增順序排序。如果沒有值出現超過一次，則傳回一個空陣列。

- 只在傳回的陣列中包含每個值的一個實體。

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` 應該傳回 `[]`。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` 應該傳回 `[1, 2]`。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` 應該傳回 `[-6, 0, 2, 4, 5, 23]`。

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
