---
id: 6821ebdf237de8297eaee793
title: "挑戰 17：運算優先順序"
challengeType: 28
dashedName: challenge-17
---

# --description--

給定一個整數陣列和一個字串運算子陣列，依序從左到右將操作行為套用到數字上。重複操作行為直到所有數字都被使用。傳回最終結果。

例如，給定 `[1, 2, 3, 4, 5]` 和 `['+', '*']`，傳回從左到右計算 `1 + 2 * 3 + 4 * 5` 的結果，忽略標準的運算順序。

- 有效的運算子是 `+`、`-`、`*`、/` 和 `%`。

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` 應該傳回 `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` 應該傳回 `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` 應該傳回 `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` 應該傳回 `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` 應該傳回 `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
