---
id: 6821ebfd237de8297eaee799
title: "挑戰 23：RGB 轉 Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

給定一個 CSS `rgb(r, g, b)` 顏色字串，傳回其十六進位等價。

以下是給定輸入的一些範例輸出：

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- 將所有字母轉為小寫。
- 傳回一個 `#`，後面接著六個字元。不要使用任何縮寫值。

# --hints--

`rgbToHex("rgb(255, 255, 255)")` 應該傳回 `"#ffffff"`。

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` 應該傳回 `"#010b6f"`。

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` 應該傳回 `"#add8e6"`。

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` 應該傳回 `"#4f7bc9"`。

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
