---
id: 6821ebda237de8297eaee792
title: "挑戰 16：反轉括號"
challengeType: 28
dashedName: challenge-16
---

# --description--

給定一個包含正確巢狀小括弧的字串，使用以下規則傳回該字串的解碼版本：

- 每一對小括弧內的所有字元都應該反轉。
- 最終結果中應該移除小括弧。
- 如果小括弧巢狀，應該先反轉最內層的配對，然後將其結果包含在外層配對的反轉中。
- 假設所有小括弧均勻平衡且正確巢狀。

# --hints--

`decode("(f(b(dc)e)a)")` 應該傳回 `"abcdef"`。

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` 應該傳回 `"Can you read this?"`。

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` 應該傳回 `"freeCodeCamp"`。

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
