---
id: 6821ebc9237de8297eaee78f
title: "挑戰 13：不自然的質數"
challengeType: 28
dashedName: challenge-13
---

# --description--

給定一個整數，判斷該數字是否為質數或負質數。

- 質數是大於 1 且只能被 1 和自身整除的正整數。
- 負質數是正質數的負數版本。
- `1` 和 `0` 不被視為質數。

# --hints--

`isUnnaturalPrime(1)` 應該傳回 `false`。

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` 應該傳回 `false`。

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` 應該傳回 `true`。

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` 應該傳回 `true`。

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` 應該傳回 `false`。

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` 應該傳回 `true`。

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` 應該傳回 `true`。

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` 應該傳回 `false`。

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` 應該傳回 `false`。

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
