---
id: 6821ebe4237de8297eaee794
title: "挑戰 18：次佳"
challengeType: 28
dashedName: challenge-18
---

# --description--

給定一個表述不同筆記型電腦價格的整數陣列，以及一個表述你預算的整數，傳回：

1. 如果它在你的預算內，則為第二昂貴的筆記型電腦，或
2. 在你的預算範圍內最昂貴的筆記型電腦，或
3. `0` 如果沒有筆記型電腦在你的預算內。

- 重複的價格應該被忽略。

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` 應該傳回 `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` 應該傳回 `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` 應該傳回 `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` 應該傳回 `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` 應該傳回 `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
