---
id: 68adce01c0e1144d0a90295a
title: "挑戰 27：矩陣旋轉"
challengeType: 28
dashedName: challenge-27
---

# --description--

給定一個矩陣（一個陣列的陣列），將矩陣順時針旋轉 90 度並傳回它。例如，給定 `[[1, 2], [3, 4]]`，它看起來像這樣：

| 1 | 2 |
|---|---|
| 3 | 4 |

你應該傳回 `[[3, 1], [4, 2]]`，它看起來像這樣：

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` 應該傳回 `[[1]]`。

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` 應該傳回 `[[3, 1], [4, 2]]`。

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` 應該傳回 `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`。

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` 應該傳回 `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`。

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
