---
id: 68adce01c0e1144d0a90295e
title: "挑戰 29：縮寫詞建置器"
challengeType: 28
dashedName: challenge-29
---

# --description--

給定一個包含一個或多個字的字串，使用以下條件傳回這些字的縮寫詞：

- 縮寫應由每個字的第一個字母大寫組成，除非另有說明。
- 縮寫應忽略這些字的第一個字母，除非它們是給定字串的第一個字：`a`、`for`、`an`、`and`、`by` 和 `of`。
- 縮寫字母應以給定的順序傳回。
- 縮寫不應包含任何空格。

# --hints--

`buildAcronym("Search Engine Optimization")` 應該傳回 `"SEO"`。

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` 應該傳回 `"FAQ"`。

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` 應該傳回 `"NASA"`。

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` 應該傳回 `"FBI"`。

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` 應該傳回 `"FYI"`。

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` 應該傳回 `"BTW"`。

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` 應該傳回 `"AUHWPOTIMSH"`。

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
