---
id: 68adce01c0e1144d0a902958
title: "挑戰 26：IPv4 驗證器"
challengeType: 28
dashedName: challenge-26
---

# --description--

給定一個字串，判斷它是否為有效的 IPv4 位址。有效的 IPv4 位址由四個以點（`.`）分隔的整數所組成。每個數字必須符合以下條件：

- 數值介於 0 和 255（含）之間。
- 它沒有前導零（例如 0 是允許的，01 則不行）。
- 只允許數字字元。

# --hints--

`isValidIPv4("192.168.1.1")` 應該傳回 `true`。

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` 應該傳回 `true`。

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` 應該傳回 `false`。

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` 應該傳回 `false`。

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` 應該傳回 `false`。

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` 應該傳回 `false`。

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` 應該傳回 `false`。

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
