---
id: 68adce01c0e1144d0a90295c
title: "挑戰 28：羅馬數字解析器"
challengeType: 28
dashedName: challenge-28
---

# --description--

給定一個表述羅馬數字的字串，傳回它的整數值。

羅馬數字由以下符號和數值組成：

| Symbol | Value |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- 數字從左到右讀取。如果較小的數字出現在較大的數字之前，則該值會被減去。否則，數值會被相加。

# --hints--

`parseRomanNumeral("III")` 應該傳回 `3`。

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` 應該傳回 `4`。

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` 應該傳回 `26`。

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` 應該傳回 `99`。

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` 應該傳回 `460`。

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` 應該傳回 `504`。

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` 應該傳回 `2025`。

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
