---
id: 68adce01c0e1144d0a902956
title: "挑戰 25：母音重複器"
challengeType: 28
dashedName: challenge-25
---

# --description--

給定一個字串，傳回該字串的新版本，其中每個母音的重複次數比你遇到的前一個母音多一次。例如，句子中的第一個母音應保持不變。第二個母音應連續出現兩次。第三個母音應連續出現三次，依此類推。

- 字母 `a`、`e`、`i`、`o` 和 `u`，無論是大寫還是小寫，都被視為母音。
- 原始母音應保持其大小寫。
- 重複的母音應該是小寫。
- 所有非母音字元應保持其原始大小寫。

# --hints--

`repeatVowels("hello world")` 應該傳回 `"helloo wooorld"`。

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` 應該傳回 `"freeeCooodeeeeCaaaaamp"`。

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` 應該傳回 `"AEeIiiOoooUuuuu"`。

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` 應該傳回 `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`。

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
