---
id: 68af0687ef34c76c28ffa54b
title: "挑戰 32：反轉句子"
challengeType: 28
dashedName: challenge-32
---

# --description--

給定一個字串的字，傳回一個字串，其字的順序為反向。 例如，第一個字應該位於傳回字串的末尾，而最後一個字應該位於傳回字串的開頭。

- 在給定的字串中，字可以由一個或多個空格分隔。
- 傳回的字串中，字與字之間應該只有一個空格。

# --hints--

`reverseSentence("world hello")` 應該傳回 `"hello world"`。

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` 應該傳回 `"git commit push"`。

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` 應該傳回 `"sudo install npm"`。

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` 應該傳回 `"export function default import"`。

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
