---
id: 68af0687ef34c76c28ffa547
title: "挑戰 30：獨特字元"
challengeType: 28
dashedName: challenge-30
---

# --description--

給定一個字串，判斷字串中的所有字元是否唯一。

- 大寫和小寫字母應視為不同的字元。

# --hints--

`allUnique("abc")` 應該傳回 `true`。

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` 應該傳回 `true`。

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` 應該傳回 `true`。

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` 應該傳回 `true`。

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` 應該傳回 `false`。

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` 應該傳回 `false`。

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` 應該傳回 `false`。

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
