---
id: 68af0687ef34c76c28ffa54f
title: "挑戰 34：缺失的數字"
challengeType: 28
dashedName: challenge-34
---

# --description--

給定一個從 1 到 `n`（包含） 的整數陣列，傳回一個包含 1 到 `n` 之間所有缺失整數的陣列（其中 `n` 是給定陣列中的最大數字）。

- 給定的陣列可能未排序且可能包含重複項。
- 傳回的陣列應該是遞增順序。
- 如果沒有遺失的整數，傳回空的陣列。

# --hints--

`findMissingNumbers([1, 3, 5])` 應該傳回 `[2, 4]`。

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` 應該傳回 `[]`。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` 應該傳回 `[2, 3, 4, 5, 6, 7, 8, 9]`。

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` 應該傳回 `[2, 3, 4, 5, 6, 7, 8, 9]`。

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` 應該傳回 `[2, 6, 7, 8]`。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` 應該傳回 `[11]`。

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
