---
id: 68af0687ef34c76c28ffa54d
title: "挑戰 33：螢幕時間"
challengeType: 28
dashedName: challenge-33
---

# --description--

給定一個包含七個整數的輸入陣列，表述一週的時間，其中每個整數是該天使用手機的時數，根據這些條件判斷是否螢幕使用時間過長：

- 如果任何單一天有 10 小時或以上，這就太多了。
- 如果連續 3 天中任意一天的平均值大於或相等於 8 小時，則表示過多。
- 如果七天的平均值大於或相等於 6 小時，則太多了。

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` 應該傳回 `false`。

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` 應該傳回 `false`。

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` 應該傳回 `false`。

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` 應該傳回 `true`。

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` 應該傳回 `true`。

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` 應該傳回 `true`。

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` 應該傳回 `true`。

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
