---
id: 68b06e589bf2273243814777
title: "挑戰 39：填滿油箱"
challengeType: 28
dashedName: challenge-39
---

# --description--

給定油箱的大小、目前的燃料層次和每加侖的價格，傳回將油箱加滿的費用。

- `tankSize` 是油箱的總容量，單位為加侖。
- `fuelLevel` 是油箱中目前的燃料量，單位為加侖。
- `pricePerGallon` 是一加侖燃料的成本。
- 傳回的值應該以 `"$d.dd"` 格式四捨五入到小數點後兩位。

# --hints--

`costToFill(20, 0, 4.00)` 應該傳回 `"$80.00"`。

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` 應該傳回 `"$17.50"`。

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` 應該傳回 `"$29.25"`。

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` 應該傳回 `"$0.00"`。

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` 應該傳回 `"$21.89"`。

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
