---
id: 68b06e589bf2273243814773
title: "挑戰 37：句子首字母大寫器"
challengeType: 28
dashedName: challenge-37
---

# --description--

給定一段文字，傳回一段新的文字，其中每個句子的第一個字母皆為大寫。

- 所有其他字元應保持不變。
- 句子可以以句點（`.`）、一個或多個問號（`?`）或一個或多個驚嘆號（`!`）結尾。

# --hints--

`capitalize("this is a simple sentence.")` 應該傳回 `"This is a simple sentence."`。

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` 應該傳回 `"Hello world. How are you?"`。

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` 應該傳回 `"I did today's coding challenge... It was fun!!"`。

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` 應該傳回 `"Crazy!!!Strange???Unconventional...Sentences."`。

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` 應該傳回 `"There's a space before this period . Why is there a space before that period ?"`。

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
