---
id: 68b06e589bf227324381476f
title: "挑戰 35：字頻率"
challengeType: 28
dashedName: challenge-35
---

# --description--

給定一段文字，傳回出現頻率最高的三個字陣列。

- 段落中的字會以空格分隔。
- 忽略給定段落中的大小寫。例如，將 `Hello` 和 `hello` 視為相同的字、單字。
- 忽略給定段落中的標點符號。標點符號包括逗號（`,`）、句號（`.`）和驚嘆號（`!`）。
- 傳回的陣列應該包含所有小寫字。
- 傳回的陣列應該以降冪排序，最常出現的字放在最前面。

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` 應該傳回 `["coding", "python", "in"]`。

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` 應該傳回 `["i", "like", "coding"]`。

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` 應該傳回 `["debug", "test", "deploy"]`。

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
