---
id: 68b06e589bf2273243814775
title: "挑戰 38：Slug 產生器"
challengeType: 28
dashedName: challenge-38
---

# --description--

給定一個字串，使用以下條件傳回一個適合 URL 的字串版本：

- 所有字母應該是小寫。
- 所有非字母、數字或空格的字元都應該被移除。
- 所有空格都應該被取代為 URL 編碼的空格代碼 `%20`。
- 連續空格應該取代為單一 `%20`。
- 傳回的字串不應該有前置或後置的 `%20`。

# --hints--

`generateSlug("helloWorld")` 應該傳回 `"helloworld"`。

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` 應該傳回 `"hello%20world"`。

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` 應該傳回 `"helloworld"`。

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` 應該傳回 `"hello%20world"`。

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` 應該傳回 `"h3110%20w0r1d"`。

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
