---
id: 68b1f72371a5ac895ac70a0a
title: "挑戰 44：字串鏡像"
challengeType: 28
dashedName: challenge-44
---

# --description--

給定兩個字串，判斷第二個字串是否為第一個字串的鏡像。

- 如果一個字串包含相同的字母且順序相反，則該字串被視為鏡像字串。
- 將大寫字母和小寫字母視為不同。
- 忽略所有非字母字元。

# --hints--

`isMirror("helloworld", "helloworld")` 應該傳回 `false`。

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` 應該傳回 `true`。

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` 應該傳回 `true`。

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` 應該傳回 `false`。

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` 應該傳回 `false`。

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` 應該傳回 `true`。

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` 應該傳回 `true`。

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
