---
id: 68b1f72371a5ac895ac70a06
title: "挑戰 42：影片儲存"
challengeType: 28
dashedName: challenge-42
---

# --description--

給定一個影片大小、影片大小的單位、一個硬碟容量，以及硬碟的單位，傳回硬碟可以儲存的影片數量，並使用以下條件：

- 影片大小的單位可以是位元組（`"B"`）、千位元組（`"KB"`）、百萬位元組（`"MB"`）或千兆位元組（`"GB"`）。
- 如果沒有給定上述其中一個影片單位，傳回 `"Invalid video unit"`。
- 硬碟容量的單位可以是 gigabytes (`"GB"`) 或 terabytes (`"TB"`)。
- 如果沒有給出上述硬碟單位之一，傳回 `"Invalid drive unit"`。
- 傳回磁碟機可容納的完整影片數量。
- 使用以下轉換：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

例如，給定 `500`、`"MB"`、`100` 和 `"GB"` 作為引數，判斷在 100 GB 硬碟中可以容納多少個 500 MB 影片。

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` 應該傳回 `200`。

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` 應該傳回 `"Invalid video unit"`。

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` 應該傳回 `"Invalid drive unit"`。

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` 應該傳回 `4000`。

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` 應該傳回 `1466`。

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
