---
id: 68b1f72371a5ac895ac70a08
title: "挑戰 43：數字對字母"
challengeType: 28
dashedName: challenge-43
---

# --description--

給定一個字串，如果字串中數字多於字母，傳回 `"digits"`，如果字母多於數字，傳回 `"letters"`，如果數字和字母數量相同，傳回 `"tie"`。

- 數字由 `0-9` 組成。
- 字母由大寫或小寫的 `a-z` 組成。
- 忽略任何其他字元。

# --hints--

`digitsOrLetters("abc123")` 應該傳回 `"tie"`。

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` 應該傳回 `"letters"`。

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` 應該傳回 `"digits"`。

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` 應該傳回 `"letters"`。

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` 應該傳回 `"digits"`。

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` 應該傳回 `"tie"`。

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
