---
id: 68b7687dded630607aceccb1
title: "挑戰 48：垃圾郵件偵測器"
challengeType: 28
dashedName: challenge-48
---

# --description--

給定一個格式為 `"+A (BBB) CCC-DDDD"` 的電話號碼，其中每個字母代表一個數字，如下所示：

- `A` 表示國家代碼，可以是任意位數的數字。
- `BBB` 表示區域代碼，且永遠是三位數字。
- `CCC` 和 `DDDD` 表述區域的號碼，分別永遠是三位數和四位數長。

根據以下條件判斷是否為垃圾電話號碼：

- 國家代碼長度超過 2 位數或未以零 (`0`) 開頭。
- 區域代碼大於 900 或小於 200。
- 區域的號碼前 3 個數字的總和出現在區域的號碼後 4 個數字中。
- 該數字中有相同的字元連續出現四次或以上（忽略格式字元）。

# --hints--

`isSpam("+0 (200) 234-0182")` 應該傳回 `false`。

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` 應該傳回 `true`。

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` 應該傳回 `false`。

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
