---
id: 68b7cadffed0e75a517da677
title: "挑戰 54：P@ssw0rd Str3ngth！"
challengeType: 28
dashedName: challenge-54
---

# --description--

給定一個密碼字串，根據密碼強度傳回 `"weak"`、`"medium"` 或 `"strong"`。

密碼會根據以下規則進行評估：

- 長度至少為 8 個字元。
- 它包含大寫和小寫字母。
- 它至少包含 1 個數字。
- 它至少包含來自此集合的 1 個特殊字元：`!`、`@`、`#`、`$`、`%`、`^`、`&` 或 `*`。

如果密碼符合少於兩項規則，傳回 `"weak"`。  
如果密碼符合 2 或 3 項規則，傳回 `"medium"`。  
如果密碼符合全部 4 項規則，傳回 `"strong"`。

# --hints--

`checkStrength("123456")` 應該傳回 `"weak"`。

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` 應該傳回 `"weak"`。

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` 應該傳回 `"weak"`。

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` 應該傳回 `"weak"`。

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` 應該傳回 `"medium"`。

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` 應該傳回 `"medium"`。

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` 應該傳回 `"medium"`。

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` 應該傳回 `"strong"`。

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` 應該傳回 `"strong"`。

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
