---
id: 68b7687dded630607aceccab
title: "挑戰 45：完全平方數"
challengeType: 28
dashedName: challenge-45
---

# --description--

給定一個整數，判斷它是否為完全平方數。

- 如果你可以將一個整數乘以它自己來得到該數字，則該數字是完全平方數。 例如，9 是完全平方數，因為你可以將 3 乘以它自己來得到它。

# --hints--

`isPerfectSquare(9)` 應該傳回 `true`。

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` 應該傳回 `true`。

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` 應該傳回 `true`。

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` 應該傳回 `false`。

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` 應該傳回 `false`。

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` 應該傳回 `false`。

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` 應該傳回 `true`。

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` 應該傳回 `true`。

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
