---
id: 68b7687dded630607aceccaf
title: "挑戰 47：超速被抓"
challengeType: 28
dashedName: challenge-47
---

# --description--

給定一個表述車輛行駛速度的數字陣列，以及一個表述速限的數字，傳回一個包含兩個條款的陣列，第一個條款是超速車輛的數量，第二個條款是這些車輛超出速限的平均速度。

- 如果沒有車輛超速，傳回 `[0, 0]`。

# --hints--

`speeding([50, 60, 55], 60)` 應該傳回 `[0, 0]`。

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` 應該傳回 `[2, 4]`。

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` 應該傳回 `[4, 8.5]`。

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` 應該傳回 `[2, 3.5]`。

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` 應該傳回 `[1, 57]`。

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
