---
id: 68b7cadffed0e75a517da675
title: "挑戰 53：十進位轉二進位"
challengeType: 28
dashedName: challenge-53
---

# --description--

給定一個非負整數，傳回其二進位表述的字串。

二進位數字只使用數字 `0` 和 `1` 來表述任何數字。要將十進位數字轉換為二進位，反覆將該數字除以 `2` 並記錄餘數。重複此動作直到數字為零。從最後記錄的餘數讀取到第一個餘數。 例如，要將 `12` 轉換為二進位：

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` 的二進位是 `1100`。

# --hints--

`toBinary(5)` 應該傳回 `"101"`。

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` 應該傳回 `"1100"`。

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` 應該傳回 `"110010"`。

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` 應該傳回 `"1100011"`。

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
