---
id: 68b7cadffed0e75a517da673
title: "挑戰 52：二進位轉十進位"
challengeType: 28
dashedName: challenge-52
---

# --description--

給定一個表述二進位數字的字串，傳回其等價的十進位數字。

二進位數字只使用數字 `0` 和 `1` 來表述任何數字。要將二進位轉換為十進位，將每個數字乘以 `2` 的冪次方，然後將它們相加。從最右邊的數字乘以 `2^0` 開始，左邊的下一個數字乘以 `2^1`，依此類推。當所有數字都乘以 `2` 的冪次方後，將結果相加。

例如，二進位數字 `101` 在十進位中等於 `5`，因為：

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` 應該傳回 `5`。

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` 應該傳回 `10`。

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` 應該傳回 `18`。

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` 應該傳回 `85`。

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
