---
id: 68c1a929005bf54d342aa8d6
title: "挑戰 59：太空週第 5 天：適居帶"
challengeType: 28
dashedName: challenge-59
---

# --description--

在太空週的第 5 天，你將計算恆星的「適居帶」——恆星周圍條件「恰到好處」以維持液態水存在的區域。

給定恆星的質量，傳回一個包含其黃金圈區域起始和結束距離（以天文單位為單位）的陣列。

計算 Goldilocks Zone：

1. 透過將恆星的質量提升到 3.5 次方來求出其光度。
2. 該區域的起點是其亮度平方根的 0.95 倍。
3. 該區域的末端是其亮度平方根的 1.37 倍。

- 傳回四捨五入到小數點後兩位的距離。

例如，給定 `1` 作為質量，傳回 `[0.95, 1.37]`。

# --hints--

`goldilocksZone(1)` 應該傳回 `[0.95, 1.37]`。

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` 應該傳回 `[0.28, 0.41]`。

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` 應該傳回 `[21.85, 31.51]`。

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` 應該傳回 `[9.38, 13.52]`。

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` 應該傳回 `[179.69, 259.13]`。

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
