---
id: 68c1a929005bf54d342aa8d4
title: "挑戰 57：太空週第 3 天：打電話回家"
challengeType: 28
dashedName: challenge-57
---

# --description--

在太空週的第三天，你會得到一個陣列，該陣列表述你自己、衛星和你的母星之間通訊路徑中的距離（以公里為單位）。請根據以下條件判斷透過該路徑傳送的訊息到達目的星球所需的時間：

- 陣列中的第一個值是你的位置到第一顆衛星的距離。
- 每個後續值（最後一個除外）都是到下一顆衛星的距離。
- 陣列中的最後一個值是從前一顆衛星到你的母星的距離。
- 該訊息以 300,000 公里／秒 的速度傳遞。
- 每個衛星訊息**通過**時都會增加 0.5 秒的傳輸延遲。
- 傳回一個四捨五入到 4 位小數且移除尾隨零的數字。

# --hints--

`sendMessage([300000, 300000])` 應該傳回 `2.5`。

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` 應該傳回 `3.0627`。

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` 應該傳回 `364.5`。

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` 應該傳回 `1674.3333`。

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` 應該傳回 `2.4086`。

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` 應該傳回 `21.1597`。

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
