---
id: 68c1a929005bf54d342aa8d3
title: "挑戰 56：太空週第 2 天：系外行星搜尋"
challengeType: 28
dashedName: challenge-56
---

# --description--

在太空週的第二天，你會得到一個字串，其中每個字元表述一顆恆星的亮度讀數。判斷這些讀數是否使用過境方法偵測到系外行星。過境方法是指當行星經過恆星前方時，會降低觀測到的亮度。

- 亮度讀數僅包含字元 `0-9` 和 `A-Z`，每個讀數對應以下數值：
- 字元 `0-9` 對應到亮度層次 `0-9`。
- 字元 `A-Z` 對應到亮度層次 `10-35`。

如果任何單一讀數小於或相等於所有讀數平均值的 80％，則該恆星被視為擁有一顆系外行星。 例如，如果恆星的平均亮度為 10，則如果任何單一讀數為 8 或更低，該恆星將被視為擁有一顆系外行星。

# --hints--

`hasExoplanet("665544554")` 應該傳回 `false`。

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` 應該傳回 `true`。

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` 應該傳回 `false`。

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` 應該傳回 `true`。

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` 應該傳回 `false`。

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` 應該傳回 `true`。

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
