---
id: 68c1a929005bf54d342aa8d2
title: "挑戰 55：太空週第 1 天：恆星分類"
challengeType: 28
dashedName: challenge-55
---

# --description--

10 月 4 日標誌著世界太空週的開始。接下來的 7 天將為你帶來天文主題的程式設計挑戰。

在今天的挑戰中，你會得到一顆恆星的表面溫度，單位為開爾文（K），並需要根據以下範圍判定其恆星分類：

- "O"：30,000 K 或更高
- "B"：10,000 K - 29,999 K
- "A"：7,500 K - 9,999 K
- "F"：6,000 K - 7,499 K
- "G"：5,200 K ～ 5,999 K
- "K"：3,700 K - 5,199 K
- "M"：0 K - 3,699 K

- 傳回給定恆星的分類。

# --hints--

`classification(5778)` 應該傳回 `"G"`。

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` 應該傳回 `"M"`。

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` 應該傳回 `"A"`。

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` 應該傳回 `"K"`。

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` 應該傳回 `"M"`。

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` 應該傳回 `"O"`。

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` 應該傳回 `"F"`。

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` 應該傳回 `"B"`。

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
