---
id: 68c497f3aaefc9fd9f1b0e26
title: "挑戰 62：十六進位轉十進位"
challengeType: 28
dashedName: challenge-62
---

# --description--

給定一個表述十六進位數字（基數 16）的字串，傳回其十進位（基數 10）值作為整數。

十六進位是一種使用 16 個數字的數字系統：

- `0-9` 表述值 `0` 到 `9`。
- `A-F` 表述值 `10` 到 `15`。

這是一個部分轉換表：

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- 該字串只會包含字元 `0–9` 和 `A–F`。

# --hints--

`hexToDecimal("A")` 應該傳回 `10`。

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` 應該傳回 `21`。

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` 應該傳回 `46`。

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` 應該傳回 `255`。

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` 應該傳回 `2623`。

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
