---
id: 68c497f3aaefc9fd9f1b0e24
title: "挑戰 60：太空週第 6 天：月相"
challengeType: 28
dashedName: challenge-60
---

# --description--

在太空週的第 6 天，你將會收到一個格式為 `"YYYY-MM-DD"` 的日期，並且需要使用以下規則判斷該日期的月相：

使用 28 天的簡化農曆週期，分為四個相等的階段：

- 「New」：第 1 天至第 7 天
- 「蠟狀期」：第 8 天至第 14 天
- 「Full」：第 15 天至第 21 天
- 「Waning」：第 22 天至第 28 天

在第 28 天之後，週期會從第 1 天（新月）重新開始。

- 使用 `"2000-01-06"` 作為參考新月（週期的第 1 天）來判斷指定日期的月相。
- 你不會被給予任何早於引用日期的日期。
- 傳回正確階段的字串。

**注意：** 第 1 天表述新月當天，表示自上一次新月以來尚未過任何天數。

# --hints--

`moonPhase("2000-01-12")` 應該傳回 `"New"`。

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` 應該傳回 `"Waxing"`。

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` 應該傳回 `"Full"`。

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` 應該傳回 `"Waning"`。

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` 應該傳回 `"New"`。

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
