---
id: 68cae5b538ff798bbd4da009
title: "挑戰 71：小費計算機"
challengeType: 28
dashedName: challenge-71
---

# --description--

給定你的餐點價格和訂製的小費百分比，傳回一個包含三個小費數值的陣列；15％、20％ 和訂製金額。

- 價格將以 `"$N.NN"` 格式顯示。
- 訂製的小費百分比將以此格式提供：`"25%"`。
- 以相同的 `"$N.NN"` 格式傳回金額，四捨五入至小數點後兩位。

例如，給定一個 `"$10.00"` 的餐點價格，以及一個 `"25%"` 的訂製小費值，傳回 `["$1.50", "$2.00", "$2.50"]`。

# --hints--

`calculateTips("$10.00", "25%")` 應該傳回 `["$1.50", "$2.00", "$2.50"]`。

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` 應該傳回 `["$13.45", "$17.93", "$23.31"]`。

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` 應該傳回 `["$2.98", "$3.97", "$1.79"]`。

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
