---
id: 68cae5b538ff798bbd4da007
title: "挑戰 69：缺少的襪子"
challengeType: 28
dashedName: challenge-69
---

# --description--

給定一個整數，表述你一開始擁有的襪子配對數量，以及另一個整數，表述你已經經歷的洗滌週期數，請使用以下條件傳回你目前擁有的完整襪子配對數量：

- 每 2 次洗衣週期，你會少一隻襪子。
- 每 3 次洗衣週期，你會發現一隻襪子遺失。
- 每 5 次洗滌週期後，一隻襪子會磨損並必須丟棄。
- 每 10 次洗滌週期，你會買一雙襪子。
- 你的襪子總數絕不會少於零。
- 規則可以重疊。例如，在洗衣週期 10，你會遺失一隻襪子、丟棄一隻襪子，並購買一雙新的襪子。
- 傳回完整襪子配對的數量。

# --hints--

`sockPairs(2, 5)` 應該傳回 `1`。

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` 應該傳回 `0`。

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` 應該傳回 `4`。

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` 應該傳回 `3`。

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` 應該傳回 `0`。

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
