---
id: 68cae5b538ff798bbd4da006
title: "挑戰 68：信用卡遮罩器"
challengeType: 28
dashedName: challenge-68
---

# --description--

給定一個信用卡號字串，使用以下條件傳回其遮罩版本：

- 該字串將包含四組四位數字（`0-9`），所有組別皆以單一空格或單一連字號（`-`）分隔。
- 將所有數字（除了最後 4 個）取代為星號（`*`）。
- 保留剩餘的字元不變。

例如，給定 `"4012-8888-8888-1881"` 傳回 `"****-****-****-1881"`。

# --hints--

`mask("4012-8888-8888-1881")` 應該傳回 `"****-****-****-1881"`。

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` 應該傳回 `"**** **** **** 5100"`。

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` 應該傳回 `"**** **** **** 1117"`。

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` 應該傳回 `"****-****-****-0010"`。

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
