---
id: 68cae5b538ff798bbd4da008
title: "挑戰 70：HTML 屬性擷取器"
challengeType: 28
dashedName: challenge-70
---

# --description--

給定一個有效 HTML 元素的字串，根據以下條件傳回該元素的屬性：

- 你只會被給予一個元素。
- 屬性將採用 `attribute="value"` 格式。
- 傳回一個字串陣列，每個屬性屬性和值以逗號分隔，格式為：`["attribute1, value1", "attribute2, value2"]`。
- 以給定的順序傳回屬性。
- 如果找不到任何屬性，傳回空陣列。

# --hints--

`extractAttributes('<span class="red"></span>')` 應該傳回 `["class, red"]`。

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` 應該傳回 `["charset, UTF-8"]`。

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` 應該傳回 `[]`。

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` 應該傳回 `["name, email", "type, email", "required, true"]`。

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` 應該傳回 `["id, submit", "class, btn btn-primary"]`。

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
