---
id: 68cae5b538ff798bbd4da004
title: "挑戰 66：HTML 頁籤剝除器"
challengeType: 28
dashedName: challenge-66
---

# --description--

給定一個 HTML `code` 的 `string`，移除 `tag` 並傳回純文字內容。

- 輸入字串將只包含有效的 HTML。
- HTML 頁籤可能會巢狀排列。
- 移除頁籤和任何屬性。

例如，`'<a href="#">Click here</a>'` 應該傳回 `"Click here"`。

# --hints--

`stripTags('<a href="#">Click here</a>')` 應該傳回 `"Click here"`。

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` 應該傳回 `"Hello World!"`。

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` 應該傳回空字串（`""`）。

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` 應該傳回 `sectionsection`。

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
