---
id: 68cae5b538ff798bbd4da005
title: "挑戰 67：電子郵件驗證器"
challengeType: 28
dashedName: challenge-67
---

# --description--

給定一個字串，判斷它是否為有效的電子郵件位址，使用以下條件：

- 它必須包含恰好一個 `@` 符號。
- 區域的零件（`@` 之前）：
  - 只能包含字母（`a-z`、`A-Z`）、數字（`0-9`）、點（`.`）、底線（`_`）或連字號（`-`）。
  - 不能以點號開始或結束。
- `@` 之後的網域零件
  - 必須包含至少一個點。
  - 必須以句點後接至少兩個字母結尾。
- 區域的或網域零件都不能有連續兩個點。

# --hints--

`validate("a@b.cd")` 應該傳回 `true`。

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` 應該傳回 `true`。

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` 應該傳回 `false`。

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` 應該傳回 `false`。

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` 應該傳回 `false`。

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` 應該傳回 `true`。

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` 應該傳回 `false`。

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` 應該傳回 `false`。

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` 應該傳回 `false`。

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` 應該傳回 `false`。

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
