---
id: 68cae5b538ff798bbd4da003
title: "挑戰 65：字串計數"
challengeType: 28
dashedName: challenge-65
---

# --description--

給定兩個字串，判斷第二個字串在第一個字串中出現了多少次。

- 樣式字串可以在第一個字串中重疊。例如，`"aaa"` 包含兩次 `"aa"`。第一次是前兩個 `a`，第二次是後兩個。

# --hints--

`count('abcdefg', 'def')` 應該傳回 `1`。

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` 應該傳回 `0`。

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` 應該傳回 `2`。

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` 應該傳回 `3`。

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` 應該傳回 `10`。

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
