---
id: 68cae5b538ff798bbd4da001
title: "挑戰 63：字之戰"
challengeType: 28
dashedName: challenge-63
---

# --description--

給定兩個句子，分別表述你的團隊和對手團隊，其中你團隊的每個字與對手團隊對應的字進行對戰，根據以下規則判斷哪個團隊獲勝：

- 給定的句子將始終包含相同數量的字。
- 字、單字 由單一空格分隔，且只包含字母。
- 每個字的值是其字母的總和。
- 字母 `a` 到 `z` 對應值為 `1` 到 `26`。例如，`a` 是 `1`，而 `z` 是 `26`。
- 大寫字母會將該字母的值加倍。 例如，`A` 是 `2`，而 `Z` 是 `52`。
- 字、單字依序對戰：你隊伍的第一個字、單字與對方隊伍的第一個字、單字對戰，依此類推。
- 如果一個字的值大於對方字的值，該字即獲勝。
- 獲勝字數較多的隊伍即為贏家。

傳回 `"We win"` 如果你的隊伍是贏家，傳回 `"We lose"` 如果你的隊伍輸了，並且如果雙方隊伍的勝場數相同，傳回 `"Draw"`。

# --hints--

`battle("hello world", "hello word")` 應該傳回 `"We win"`。

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` 應該傳回 `"We win"`。

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` 應該傳回 `"We lose"`。

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` 應該傳回 `"Draw"`。

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` 應該傳回 `"We win"`。

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` 應該傳回 `"We lose"`。

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` 應該傳回 `"Draw"`。

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
