---
id: 68d2ba1468508398389487cf
title: "挑戰 74：最愛歌曲"
challengeType: 28
dashedName: challenge-74
---

# --description--

還記得 iPod 嗎？第一款機型在 2001 年 10 月 23 日，也就是今天，推出，距今已經 24 年。

給定一個表述你 iPod 播放列表的歌曲物件陣列，傳回一個包含播放次數最多的兩首歌曲標題的陣列，並以播放次數最多的歌曲為首。

- 每個物件都會有一個 `"title"` 屬性（字串），以及一個 `"plays"` 屬性（整數）。

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` 應該傳回 `["Sync or Swim", "Earbud Blues"]`。

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` 應該傳回 `["Clickwheel Love", "99 Downloads"]`。

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` 應該傳回 `["Song B", "Song C"]`。

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
