---
id: 68d2ba1468508398389487ce
title: "挑戰 73：你必須明智地說話"
challengeType: 28
dashedName: challenge-73
---

# --description--

給定一句話，傳回一個聽起來像智慧老師建議的版本，並遵循以下規則：

- 字、單字 以單一空格分隔。
- 尋找句子中以下字、單字之一的第一次出現位置：`"have"`、`"must"`、`"are"`、`"will"`、`"can"`。
- 將該字及其之前的所有字移動到句尾，然後：
  - 關鍵規則 - 嚴格遵守：  
1. 絕對不要移除文字中的反引號（`）  
2. 絕對不要翻譯反引號內的文字 - 保持原文英文不變  
3. 反引號表示程式碼／技術術語，必須保持不變  
4. 不要在回應中加入引號  
5. 完整保留所有格式（**、*、<
  - 將它們全部轉成小寫。
  - 並在它們之前添加逗號和空格。
- 將句子中新第一個字的首字母大寫。
- 所有給定的句子都會以單一標點符號結尾。請保留原始標點符號並將其移至新句子的末尾。
- 傳回新的句子，確保每個字之間只有一個空格，且句子開頭和結尾沒有空格。

例如，給定 `"You must speak wisely."` 傳回 `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` 應該傳回 `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` 應該傳回 `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` 應該傳回 `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` 應該傳回 `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` 應該傳回 `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
