---
id: 68d30845cc08266018fc46bd
title: "挑戰 77：持續時間格式化器"
challengeType: 28
dashedName: challenge-77
---

# --description--

給定一個整數秒數，傳回一個字串，表述相同的持續時間，格式為 `"H:MM:SS"`，其中 `"H"` 是小時數，`"MM"` 是分鐘數，`"SS"` 是秒數。請依照以下規則傳回時間：

- 秒數：應始終為兩位數。
- 分鐘：當不需要時應省略前導零。如果持續時間少於 1 分鐘，則使用 `"0"`。
- 小時：只有當其大於零時才應包含。

# --hints--

`format(500)` 應該傳回 `"8:20"`。

```js
assert.equal(format(500), "8:20");
```

`format(4000)` 應該傳回 `"1:06:40"`。

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` 應該傳回 `"0:01"`。

```js
assert.equal(format(1), "0:01");
```

`format(5555)` 應該傳回 `"1:32:35"`。

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` 應該傳回 `"27:46:39"`。

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
