---
id: 68d30845cc08266018fc46bc
title: "挑戰 76：互補 DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

給定一個表述 DNA 序列的字串，使用以下規則傳回其互補鏈：

- DNA 由字母 `"A"`、`"C"`、`"G"` 和 `"T"` 組成。
- 字母 `"A"` 和 `"T"` 互補。
- 字母 `"C"` 和 `"G"` 互補。

例如，給定 `"ACGT"`，傳回 `"TGCA"`。

# --hints--

`complementaryDNA("ACGT")` 應該傳回 `"TGCA"`。

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` 應該傳回 `"TACGCATGCAATCG"`。

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` 應該傳回 `"CCGAATGCTAGCTTC"`。

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` 應該傳回 `"CTAGATCGATCCGATCGATC"`。

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
