---
id: 68d30845cc08266018fc46be
title: "挑戰 78：整數序列"
challengeType: 28
dashedName: challenge-78
---

# --description--

給定一個正整數，傳回一個字串，該字串包含從 `1` 到該數字（包含該數字）所有整數，且依數字順序排列。

例如，給定 `5`，傳回 `"12345"`。

# --hints--

`sequence(5)` 應該傳回 `"12345"`。

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` 應該傳回 `"12345678910"`。

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` 應該傳回 `"1"`。

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` 應該傳回 `"123456789101112131415161718192021222324252627"`。

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
