---
id: 68d30fc57588d97fd3027b30
title: "挑戰 79：Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

1994 年 10 月 28 日，Netscape Navigator 發布，幫助數百萬人探索早期的網路。

給定你在 Netscape Navigator 上執行過的一組瀏覽器指令陣列，請根據以下規則傳回執行所有指令後你所在的當前頁面：

- 你總是從 `"Home"` 頁面開始，該頁面不會包含在命令陣列中。
- 有效指令有：
  - 「Visit Page」：其中「Page」是你正在造訪的頁面名稱。 例如，「Visit About」會帶你到「About」頁面。 當你造訪新頁面時，請確保捨棄任何前進歷史。
  - 「Back」：帶你回到歷史記錄中的上一頁，若沒有上一頁則停留在目前頁面。
  - 「Forward」：將你帶到歷史紀錄中你來自的頁面，若沒有則停留在目前頁面。

例如，給定 `["Visit About Us", "Back", "Forward"]`，傳回 `"About Us"`。

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` 應該傳回 `"About Us"`。

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` 應該傳回 `"Home"`。

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` 應該傳回 `"Home"`。

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` 應該傳回 `"Gallery"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` 應該傳回 `"Home"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` 應該傳回 `"Contact"`。

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` 應該傳回 `"Visit Us"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
