---
id: 68e39ed6106dac2f0a98fd65
title: "挑戰 83：標記驗證"
challengeType: 28
dashedName: challenge-83
---

# --description--

給定一個訊息字串、一個秘密金鑰字串和一個標記數字，使用此編碼方法判斷標記是否有效：

- 訊息和秘密金鑰中的字母具有這些值：
  - `a` 到 `z` 分別有值 `1` 到 `26`。
  - `A` 到 `Z` 分別有值 `27` 到 `52`。
- 所有其他字元沒有價值。
- 透過將訊息的總和加上密鑰的總和來計算標記。

例如，給定訊息 `"foo"` 和祕密金鑰 `"bar"`，標記將會是 `57`：

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

檢查計算出的 `signature` 是否與提供的 `signature` 相符。

# --hints--

`verify("foo", "bar", 57)` 應該傳回 `true`。

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` 應該傳回 `false`。

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` 應該傳回 `true`。

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` 應該傳回 `false`。

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` 應該傳回 `true`。

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` 應該傳回 `true`。

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
