---
id: 68e39ed6106dac2f0a98fd63
title: "挑戰 81：第 N 個質數"
challengeType: 28
dashedName: challenge-81
---

# --description--

質數是大於 1 的正整數，且只能被 1 和它本身整除。前五個質數是 `2`、`3`、`5`、`7` 和 `11`。

給定一個正整數 `n`，傳回第 `n` 個質數。例如，給定 `5` 傳回第 5 個質數：`11`。

# --hints--

`nthPrime(5)` 應該傳回 `11`。

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` 應該傳回 `29`。

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` 應該傳回 `53`。

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` 應該傳回 `523`。

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` 應該傳回 `7919`。

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
