---
id: 68e39ed6106dac2f0a98fd64
title: "挑戰 82：SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

給定一個表述變數名稱的字串，使用以下條件將其轉換為「spooky case」：

- 將所有底線（`_`）和連字號（`-`）取代為波浪號（`~`）。
- 將字串的第一個字元及之後每隔一個字元的大寫。計數時忽略波浪號字元。其他字元皆轉為小寫。

例如，給定 `hello_world`，傳回 `HeLlO~wOrLd`。

# --hints--

`spookify("hello_world")` 應該傳回 `"HeLlO~wOrLd"`。

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` 應該傳回 `"SpOoKy~CaSe"`。

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` 應該傳回 `"TrIcK~oR~tReAt"`。

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` 應該傳回 `"C~a~N~d~Y~~b~O~w~L"`。

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` 應該傳回 `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`。

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
