---
id: 68ee9e3066cfd4eb2328e8a7
title: "挑戰 88：星期幾查找器"
challengeType: 28
dashedName: challenge-88
---

# --description--

給定格式為 `YYYY-MM-DD` 的字串日期，傳回星期幾。

有效的傳回天數為：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

請務必忽略時區。

# --hints--

`getWeekday("2025-11-06")` 應該傳回 `Thursday`。

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` 應該傳回 `Friday`。

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` 應該傳回 `Saturday`。

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` 應該傳回 `Wednesday`。

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` 應該傳回 `Monday`。

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
