---
id: 68ee9e3066cfd4eb2328e8a4
title: "挑戰 85：字數計數器"
challengeType: 28
dashedName: challenge-85
---

# --description--

給定一個句子字串，傳回句子中字的數量。

- 字、單字是任何非空白字元的序列，並以單一空格分隔。

# --hints--

`countWords("Hello world")` 應該傳回 `2`。

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` 應該傳回 `9`。

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` 應該傳回 `4`。

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` 應該傳回 `7`。

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` 應該傳回 `7`。

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
