---
id: 68ee9e3066cfd4eb2328e8a5
title: "挑戰 86：影像搜尋"
challengeType: 28
dashedName: challenge-86
---

# --description--

在 2001 年 11 月 4 日，Google 推出了它的影像搜尋，讓人們可以使用搜尋詞來尋找影像。在這個挑戰中，你將模仿影像搜尋。

給定一個影像名稱的陣列和一個搜尋詞，傳回包含該搜尋詞的影像名稱陣列。

- 比對搜尋詞時忽略大小寫。
- 傳回影像在輸入陣列中出現的相同順序。

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` 應該傳回 `["dog.png"]`。

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` 應該傳回 `["Sunset.jpg", "sunflower.jpeg"]`。

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` 應該傳回 `["Moon.png", "stars.png"]`。

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` 應該傳回 `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`。

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
