---
id: 68ee9e3066cfd4eb2328e8a8
title: "挑戰 89：計算牌數"
challengeType: 28
dashedName: challenge-89
---

# --description--

一副標準撲克牌的每個花色有 13 張獨特的牌。給定一個整數表示從牌堆中抽取的牌數，傳回你可以抽取的獨特牌組合數量。

- 順序不重要。先選擇卡片 A 再選擇卡片 B 與先選擇卡片 B 再選擇卡片 A 是相同的。

例如，給定 `52`，傳回 `1`。從 52 張牌的牌組中選擇 52 張牌只有一種組合。並且給定 `2`，傳回 `1326`，從牌組中選擇 2 張牌時，你可以得到 1326 種牌組合。

# --hints--

`combinations(52)` 應該傳回 `1`。

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` 應該傳回 `52`。

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` 應該傳回 `1326`。

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` 應該傳回 `2598960`。

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` 應該傳回 `15820024220`。

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` 應該傳回 `1326`。

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
