---
id: 68ee9e3066cfd4eb2328e8a6
title: "挑戰 87：矩陣建置器"
challengeType: 28
dashedName: challenge-87
---

# --description--

給定兩個整數（列數和行數），傳回一個指定大小、以零（`0`）填滿的矩陣（陣列的陣列）。

例如，給定 `2` 和 `3`，傳回：

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` 應該傳回 `[[0, 0, 0], [0, 0, 0]]`。

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` 應該傳回 `[[0, 0], [0, 0], [0, 0]]`。

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` 應該傳回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` 應該傳回 `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`。

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
