---
id: 68f6587287ad1f4ad39b0c84
title: "挑戰 98：矩形計數"
challengeType: 28
dashedName: challenge-98
---

# --description--

給定兩個正整數，表述矩形的寬度和高度，判斷在該矩形中可以容納多少個矩形。

- 只計算寬度和高度為整數的矩形。

例如，給定 `1` 和 `3`，傳回 `6`。三個 1x1 矩形、兩個 1x2 矩形和一個 1x3 矩形。

# --hints--

`countRectangles(1, 3)` 應該傳回 `6`。

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` 應該傳回 `18`。

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` 應該傳回 `3`。

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` 應該傳回 `150`。

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` 應該傳回 `12540`。

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
