---
id: 68f6587287ad1f4ad39b0c7c
title: "挑戰 90：字元限制"
challengeType: 28
dashedName: challenge-90
---

# --description--

在這個挑戰中，你會得到一個字串，並且需要判斷它是否適合用於社群媒體貼文。根據給定的規則傳回以下字串：

- 「短貼文」如果符合 40 字元限制。
- 「long post」如果它超過 40 個字元且符合 80 個字元限制內。
- 「invalid post」如果它太長而無法符合任一限制。

# --hints--

`canPost("Hello world")` 應該傳回 `"short post"`。

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` 應該傳回 `"long post"`。

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` 應該傳回 `"invalid post"`。

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
