---
id: 68f6587287ad1f4ad39b0c81
title: "挑戰 95：陣列位移"
challengeType: 28
dashedName: challenge-95
---

# --description--

給定一個陣列和一個整數，表述要將陣列位移多少個位置，傳回位移後的陣列。

- 正整數將陣列向左移動。
- 負整數會將陣列向右移動。
- 位移會在陣列中環繞。

例如，給定 `[1, 2, 3]` 和 `1`，將陣列向左移動 1 個位置，傳回 `[2, 3, 1]`。

# --hints--

`shiftArray([1, 2, 3], 1)` 應該傳回 `[2, 3, 1]`。

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` 應該傳回 `[3, 1, 2]`。

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` 應該傳回 `["charlie", "alpha", "bravo"]`。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` 應該傳回 `["bravo", "charlie", "alpha"]`。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` 應該傳回 `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`。

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
