---
id: 68f6587287ad1f4ad39b0c80
title: "挑戰 94：電子郵件標記產生器"
challengeType: 28
dashedName: challenge-94
---

# --description--

給定一個人的名字、職稱和公司字串，依照以下規則傳回一個電子郵件標記作為單一字串：

- 名稱應該先出現，前面加上取決於名稱第一個字母的前置式。對於以（不區分大小寫）開頭的名稱：
  - `A-I`：使用 `>>` 作為前置式。
  - `J-R`：使用 `--` 作為前置式。
  - `S-Z`：使用 `::` 作為前置式。
- 名稱後面應該跟著逗號和空格（`, `）。
- 標題和公司名稱應該跟在逗號和空格之後，以 `" at "`（兩側有空格）分隔。

例如，給定 `"Quinn Waverly"`、`"Founder and CEO"` 和 `"TechCo"` 傳回 `"--Quinn Waverly, Founder and CEO at TechCo"`。

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` 應該傳回 `"--Quinn Waverly, Founder and CEO at TechCo"`。

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` 應該傳回 `">>Alice Reed, Engineer at TechCo"`。

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` 應該傳回 `"::Tina Vaughn, Developer at example.com"`。

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` 應該傳回 `">>B. B., Product Tester at AcmeCorp"`。

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` 應該傳回 `"::windstorm, Cloud Architect at Atmospheronics"`。

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
