---
id: 68f6587287ad1f4ad39b0c7f
title: "挑戰 93：母音和子音"
challengeType: 28
dashedName: challenge-93
---

# --description--

給定一個字串，傳回一個陣列，該陣列包含字串中母音和子音的數量。

- 母音由任何大小寫的 `a`、`e`、`i`、`o`、`u` 組成。
- 子音由所有其他字母（不論大小寫）組成。
- 忽略任何非字母字元。

例如，給定 `"Hello World"`，傳回 `[3, 7]`。

# --hints--

`count("Hello World")` 應該傳回 `[3, 7]`。

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` 應該傳回 `[3, 7]`。

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` 應該傳回 `[1, 5]`。

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` 應該傳回 `[5, 7]`。

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` 應該傳回 `[3, 7]`。

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` 應該傳回 `[11, 24]`。

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
