---
id: 68f6587287ad1f4ad39b0c82
title: "挑戰 96：今天是週末嗎？"
challengeType: 28
dashedName: challenge-96
---

# --description--

給定格式為 `"YYYY-MM-DD"` 的日期，傳回距離週末剩餘的天數。

- 週末從星期六開始。
- 如果給定的日期是星期六或星期日，傳回 `"It's the weekend!"`。
- 否則，傳回 `"X days until the weekend."`，其中 `X` 是距離星期六的天數。
- 如果 `X` 是 `1`，請使用 `"day"`（單數）而非 `"days"`（複數）。
- 請確保計算忽略你的區域的時區。

# --hints--

`daysUntilWeekend("2025-11-14")` 應該傳回 `"1 day until the weekend."`。

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` 應該傳回 `"3 天直到週末。"`。

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` 應該傳回 `"It's the weekend!"`。

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` 應該傳回 `"4 days until the weekend."`。

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` 應該傳回 `"5 days until the weekend."`。

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` 應該傳回 `"It's the weekend!"`。

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
