---
id: 68f6587287ad1f4ad39b0c7e
title: "挑戰 92：擴充套件擷取器"
challengeType: 28
dashedName: challenge-92
---

# --description--

給定一個表述檔名的字串，傳回該檔案的擴充套件。

- 擴充套件是檔名中最後一個句點（`.`）之後的零件。
- 如果檔案名稱不包含句點或以句點結尾，傳回 `"none"`。
- 擴充套件應保持原樣傳回，並保留大小寫。

# --hints--

`getExtension("document.txt")` 應該傳回 `"txt"`。

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` 應該傳回 `"none"`。

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` 應該傳回 `"PNG"`。

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` 應該傳回 `"gitignore"`。

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` 應該傳回 `"gz"`。

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` 應該傳回 `"none"`。

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
