---
id: 68ffb91507a5b645769328cc
title: "挑戰 109：我到底幾歲了？"
challengeType: 28
dashedName: challenge-109
---

# --description--

給定某人生日的日期格式為 `YYYY-MM-DD`，傳回該人在 2025 年 11 月 27 日的年齡。

- 假設所有生日皆為 2025 年 11 月 27 日之前的有效日期。
- 傳回年齡為整數。
- 請務必考慮該人在 2025 年是否已經過生日。

# --hints--

`calculateAge("2000-11-20")` 應該傳回 `25`。

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` 應該傳回 `24`。

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` 應該傳回 `11`。

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` 應該傳回 `31`。

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` 應該傳回 `30`。

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
