---
id: 68ffb91507a5b645769328c5
title: "挑戰 102：最長字、單字"
challengeType: 28
dashedName: challenge-102
---

# --description--

給定一個句子字串，傳回句子中最長的單字。

- 字、單字 以單一空格分隔。
- 只有字母（`a-z`，不分大小寫）計入字的長度。
- 如果有多個相同長度的字，傳回第一個出現的字。
- 傳回該字串中出現的字，並移除標點符號。

# --hints--

`longestWord("The quick red fox")` 應該傳回 `"quick"`。

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` 應該傳回 `"challenge"`。

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` 應該傳回 `"This"`。

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` 應該傳回 `"exclamation"`。

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` 應該傳回 `"tie"`。

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` 應該傳回 `"Wouldnt"`。

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
