---
id: 68ffb91507a5b645769328c8
title: "挑戰 105：字元計數"
challengeType: 28
dashedName: challenge-105
---

# --description--

給定一個句子字串，傳回一個依字母順序排列且包含每個字元計數的陣列。

- 計數時將大寫和小寫字母視為相同字母。
- 忽略數字、空格、標點符號等。
- 以 `"letter count"` 格式傳回計數和字母。例如，`"a 3"`。
- 所有傳回的字母都應為小寫。
- 不要傳回不在給定 `string` 中的字母數量。

# --hints--

`countCharacters("hello world")` 應該傳回 `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`。

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` 應該傳回 `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`。

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` 應該傳回 `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`。

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
