---
id: 68ffb91507a5b645769328c9
title: "挑戰 106：訊息驗證器"
challengeType: 28
dashedName: challenge-106
---

# --description--

給定一個訊息字串和一個驗證字串，判斷該訊息是否有效。

- 如果訊息中的每個字都依序以驗證字串中對應的字母開頭，則該訊息為有效。
- 字母不區分大小寫。
- 訊息中的字、單字以單一空格分隔。

# --hints--

`isValidMessage("hello world", "hw")` 應該傳回 `true`。

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` 應該傳回 `true`。

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` 應該傳回 `false`。

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` 應該傳回 `true`。

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` 應該傳回 `false`。

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
