---
id: 69162d64f96574d9bb629f04
title: "挑戰 119：字串壓縮"
challengeType: 28
dashedName: challenge-119
---

# --description--

給定一個字串句子，傳回句子的壓縮版本，其中連續重複的字被該字及其重複次數（小括弧內）所取代。

- 只有連續重複項會被壓縮。
- 字、單字由單一空格分隔。

例如，給定 `"yes yes yes please"`，傳回 `"yes(3) please"`。

# --hints--

`compressString("yes yes yes please")` 應該傳回 `"yes(3) please"`。

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` 應該傳回 `"I have(3) apples"`。

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` 應該傳回 `"one(2) three and to the(4)"`。

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` 應該傳回 `"route(6) tee(6)"`。

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
