---
id: 69162d64f96574d9bb629efd
title: "挑戰 112：AI 偵測器"
challengeType: 28
dashedName: challenge-112
---

# --description--

今天的挑戰靈感來自於 2022 年 11 月 30 日 ChatGPT 的發布。

給定一個包含一個或多個句子的字串，判斷它是否很可能是使用以下規則由 AI 產生。

- 它包含兩個或更多的破折號（`-`）。
- 它包含兩組或以上的括號（`()`）。文字可以在括號內。
- 它包含三個或更多字、單字且長度為 7 個或以上字母。

- 字以單一空格分隔，且僅由字母（`A-Z`）組成。不要將標點符號或其他非字母視為字的一部分。

如果給定的句子符合上述任何規則，請傳回 `"AI"`，否則，請傳回 `"Human"`。

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` 應該傳回 `"Human"`。

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` 應該傳回 `"Human"`。

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` 應該傳回 `"AI"`。

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` 應該傳回 `"AI"`。

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` 應該傳回 `"AI"`。

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
