---
id: 69162d64f96574d9bb629efe
title: "挑戰 113：英里轉公里"
challengeType: 28
dashedName: challenge-113
---

# --description--

給定以英里為單位的距離數字，傳回等價的公里距離。

- 輸入將永遠是非負數字。
- 1 英里相等於 1.60934 公里。
- 將結果四捨五入到兩位小數。
- 從四捨五入的結果中移除不必要的尾隨零。

# --hints--

`convertToKm(1)` 應該傳回 `1.61`。

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` 應該傳回 `33.8`。

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` 應該傳回 `5.63`。

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` 應該傳回 `0`。

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` 應該傳回 `1`。

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
