---
id: 69162d64f96574d9bb629f01
title: "挑戰 116：排列數量"
challengeType: 28
dashedName: challenge-116
---

# --description--

給定一個字串，傳回可以由其字元組成的不同排列數量。

- 排列是字串中字元的任何重新排序。
- 不要計算重複的排列組合。
- 如果字串包含重複的字元，重複的排列組合應該只計算一次。
- 該字串將只包含字母（`A-Z`、`a-z`）。

例如，給定 `"abb"`，傳回 `3`，因為有三種不同的排列方式：`"abb"`、`"bab"` 和 `"bba"`。

# --hints--

`countPermutations("abb")` 應該傳回 `3`。

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` 應該傳回 `6`。

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` 應該傳回 `630`。

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` 應該傳回 `39916800`。

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
