---
id: 69162d64f96574d9bb629eff
title: "挑戰 114：駝峰式轉蛇形"
challengeType: 28
dashedName: challenge-114
---

# --description--

給定一個駝峰式字串，請使用以下規則傳回該字串的蛇底線式版本：

- 輸入字串將只包含字母（`A-Z` 和 `a-z`），且總是以小寫字母開頭。
- 駝峰式字串中的每個大寫字母都代表一個新字。
- 將所有字母轉換為小寫。
- 以底線（`_`）分隔字、單字。

# --hints--

`toSnake("helloWorld")` 應該傳回 `"hello_world"`。

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` 應該傳回 `"my_variable_name"`。

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` 應該傳回 `"freecodecamp_daily_challenges"`。

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
