---
id: 69162d64f96574d9bb629f00
title: "挑戰 115：Markdown 有序列表項目轉換器"
challengeType: 28
dashedName: challenge-115
---

# --description--

給定一個表述 Markdown 中有序列表項目的字串，傳回等價的 HTML 字串。

Markdown 中有效的有序列表項目必須：

- 以零個或多個空格開始，接著
- 一個數字（1 或以上）和一個句點（`.`），接著
- 至少一個空格，然後
- 列表項目文字。

如果字串沒有上述的正確格式，傳回 `"Invalid format"`。否則，將列表項目的文字包裹在 `li` 頁籤中並傳回該字串。

例如，給定 `"1. My item"`，傳回 `"<li>My item</li>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`convertListItem("1. My item")` 應該傳回 `"<li>My item</li>"`。

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` 應該傳回 `"<li>Another item</li>"`。

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` 應該傳回 `"Invalid format"`。

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` 應該傳回 `"<li>list item text</li>"`。

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` 應該傳回 `"Invalid format"`。

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` 應該傳回 `"Invalid format"`。

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
