---
id: 691b559495c5cb5a37b9b488
title: "挑戰 128：子音計數"
challengeType: 28
dashedName: challenge-128
---

# --description--

給定一個字串和一個標的數字，判斷該字串是否恰好包含標的數字的子音。

- 子音是除任何大小寫的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"` 以外的所有字母字元。
- 計算字元數時，請忽略數字、標點符號、空格和其他非字母字元。

# --hints--

`hasConsonantCount("helloworld", 7)` 應該傳回 `true`。

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` 應該傳回 `false`。

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` 應該傳回 `true`。

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` 應該傳回 `false`。

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` 應該傳回 `true`。

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
