---
id: 691b559495c5cb5a37b9b486
title: "挑戰 126：將首字母大寫"
challengeType: 28
dashedName: challenge-126
---

# --description--

給定一個字串 title，傳回一個使用以下規則格式化為標題大小寫的新字串：

- 將每個字的第一個字母大寫。
- 將每個字中所有其他字母轉為小寫。
- 字、單字總是以單一空格分隔。

# --hints--

`titleCase("hello world")` 應該傳回 `"Hello World"`。

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` 應該傳回 `"The Quick Brown Fox"`。

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` 應該傳回 `"Javascript And Python"`。

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` 應該傳回 `"Avocado Toast For Breakfast"`。

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
