---
id: 691b559495c5cb5a37b9b484
title: "挑戰 124：庫存更新"
challengeType: 28
dashedName: challenge-124
---

# --description--

給定一個表述你商店庫存的二維陣列，以及另一個表述你收到的貨運的二維陣列，傳回你更新後的庫存。

- 陣列中的每個元素都將具有以下格式：`[quantity, "item"]`，其中 `quantity` 是整數，`"item"` 是字串。
- 透過將貨運中任何匹配項目的數量添加到庫存中，更新項目。
- 如果收到的項目在目前的庫存中不存在，則將它作為新條款添加到庫存的末尾。
- 以給定的順序傳回庫存，並將新項目依出貨順序添加到最後。

例如，給定一個庫存為 `[[2, "apples"], [5, "bananas"]]` 和一個貨運為 `[[1, "apples"], [3, "bananas"]]`，傳回 `[[3, "apples"], [8, "bananas"]]`。

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` 應該傳回 `[[3, "apples"], [8, "bananas"]]`。

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` 應該傳回 `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`。

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` 應該傳回 `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`。

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` 應該傳回 `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`。

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
