---
id: 691b559495c5cb5a37b9b480
title: "挑戰 120：磅轉公斤"
challengeType: 28
dashedName: challenge-120
---

# --description--

給定一個以磅為單位的重量數字，傳回字串 `"(lbs) pounds equals (kgs) kilograms."`。

- 將 `"(lbs)"` 取代為輸入的數字。
- 將 `"(kgs)"` 取代為轉換成公斤的輸入值，四捨五入至小數點後兩位，且數值始終包含兩位小數。
- 1 磅相等於 0.453592 公斤。
- 如果輸入是 `1`，請使用 `"pound"` 而非 `"pounds"`。
- 如果轉換後的值是 `1`，請使用 `"kilogram"` 而非 `"kilograms"`。

# --hints--

`convertToKgs(1)` 應該傳回 `"1 pound equals 0.45 kilograms."`。

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` 應該傳回 `"0 pounds equals 0.00 kilograms."`。

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` 應該傳回 `"100 pounds equals 45.36 kilograms."`。

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` 應該傳回 `"2.5 pounds equals 1.13 kilograms."`。

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` 應該傳回 `"2.20462 pounds 相等 1.00 kilogram."`。

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
