---
id: 691b559495c5cb5a37b9b482
title: "挑戰 122：Markdown 粗體字解析器"
challengeType: 28
dashedName: challenge-122
---

# --description--

給定一個可能包含部分 Markdown 粗體文字的字串，傳回等價的 HTML 字串。

- Markdown 中的粗體文字是任何以兩個星號（`**`）或兩個底線（`__`）開始和結束的文字。
- 文字和星號或底線之間不能有任何空格，但文字本身可以有空格。
- 將所有 <b> 粗體出現的部分轉換為 HTML `b` 標籤並傳回該字串。

例如，給定 `"**This is bold**"`，傳回 `"<b>This is bold</b>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseBold("**This is bold**")` 應該傳回 `"<b>This is bold</b>"`。

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` 應該傳回 `"<b>This is also bold</b>"`。

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` 應該傳回 `"**This is not bold **"`。

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` 應該傳回 `"__ This is also not bold__"`。

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` 應該傳回 `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`。

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
