---
id: 6925e2068081f40f549ced1c
title: "挑戰 138：除數總和"
challengeType: 28
dashedName: challenge-138
---

# --description--

給定一個正整數，傳回它所有因數的總和。

- 除數是任何能夠整除該數字（餘數為 `0`）的整數。
- 每個除數只計算一次。

例如，給定 `6`，傳回 `12`，因為 `6` 的因數是 `1`、`2`、`3` 和 `6`，這些數字的總和是 `12`。

# --hints--

`sumDivisors(6)` 應該傳回 `12`。

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` 應該傳回 `14`。

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` 應該傳回 `56`。

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` 應該傳回 `224`。

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` 應該傳回 `806`。

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` 應該傳回 `23520`。

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
