---
id: 6925e2068081f40f549ced1a
title: "挑戰 136：Markdown 影像解析器"
challengeType: 28
dashedName: challenge-136
---

# --description--

給定 Markdown 中一個影像的字串，傳回等價的 HTML 字串。

Markdown 影像具有以下格式：`"![alt text](image_url)"`。其中：

- `alt text` 是影像的描述（`alt` 屬性值）。
- `image_url` 是影像的原始碼 URL（`src` 屬性值）。

傳回一個 HTML `img` 頁籤的字串，`src` 設定為影像 URL，`alt` 設定為替代文字。

例如，給定 `"![Cute cat](cat.png)"` 傳回 `'<img src="cat.png" alt="Cute cat">'`；

- 確保頁籤、屬性的順序、間距和引號使用與上述相同。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤 — 請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseImage("![Cute cat](cat.png)")` 應該傳回 `'<img src="cat.png" alt="Cute cat">'`。

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` 應該傳回 `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`。

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` 應該傳回 `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`。

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
