---
id: 6925e2068081f40f549ced1d
title: "挑戰 139：剪刀、石頭、布"
challengeType: 28
dashedName: challenge-139
---

# --description--

給定兩個字串，第一個表述玩家 1，第二個表述玩家 2，判斷石頭、剪刀、布比賽的勝者。

- 輸入字串將永遠是 `"Rock"`、`"Paper"` 或 `"Scissors"`。
- 「Rock」贏過「Scissors」。
- 「Paper」贏過「Rock」。
- 「Scissors」贏過「Paper」。

傳回：

- 如果 Player 1 獲勝，則顯示「Player 1 wins」。
- 如果 Player 2 獲勝，則顯示「Player 2 wins」。
- 「Tie」如果雙方玩家選擇相同的選項。

# --hints--

`rockPaperScissors("Rock", "Rock")` 應該傳回 `"Tie"`。

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` 應該傳回 `"Player 2 wins"`。

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` 應該傳回 `"Player 1 wins"`。

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` 應該傳回 `"Player 1 wins"`。

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` 應該傳回 `"Tie"`。

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` 應該傳回 `"Player 2 wins"`。

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
