---
id: 6925e2068081f40f549ced1b
title: "挑戰 137：Snowflake 產生器"
challengeType: 28
dashedName: challenge-137
---

# --description--

給定一個使用換行字元（`\n`）表述換行的多行字串，傳回一個新字串，其中每一行都水平鏡像並附加到原始行的末尾。

- 透過反轉所有字元（包括空格）來鏡像一行。

例如，給定 `"* \n *\n* "`，它會記錄到主控台為：

```sh
* 
 *
* 
```

傳回 `"*  *\n ** \n*  *"`，該字串會在主控台記錄為：

```sh
*  *
 ** 
*  *
```

請仔細注意給定字串和傳回字串中的空白字元。務必不要刪除任何空白字元。

# --hints--

`generateSnowflake("* \n *\n* ")` 應該傳回 `"*  *\n ** \n*  *"`。

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` 應該傳回 `"X=~~=X"`。

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` 應該傳回 `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`。

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` 應該傳回 `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`。

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` 應該傳回 `"*  --  *\n * -- * \n*  --  *"`。

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
