---
id: 69272dcf1c24b44fd79137c3
title: "挑戰 140：SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

給定一個表述變數名稱的字串，傳回轉換為 SCREAMING_SNAKE_CASE 的變數名稱。

給定的變數名稱將以以下其中一種格式書寫：

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

在上述格式中，字、單字由底線（`_`）、連字號（`-`）分隔，或新字以大寫字母開頭。

要轉換為 SCREAMING_SNAKE_CASE：

- 將所有字母轉為大寫
- 用底線（`_`）分隔字、單字

# --hints--

`toScreamingSnakeCase("userEmail")` 應該傳回 `"USER_EMAIL"`。

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` 應該傳回 `"USER_PASSWORD"`。

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` 應該傳回 `"USER_ID"`。

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` 應該傳回 `"USER_ADDRESS"`。

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` 應該傳回 `"USERNAME"`。

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();;
}
```
