---
id: 69272dcf1c24b44fd79137c5
title: "挑戰 142：字串求和"
challengeType: 28
dashedName: challenge-142
---

# --description--

給定一個包含數字和其他字元的字串，傳回字串中所有數字的總和。

- 將連續的數字視為一個整數。例如，`"13"` 計算為 13，而不是 1 + 3。
- 忽略任何非數字字元。

# --hints--

`stringSum("3apples2bananas")` 應該傳回 `5`。

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` 應該傳回 `17`。

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` 應該傳回 `125344`。

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` 應該傳回 `321`。

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` 應該傳回 `1653`。

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
