---
id: 69272dcf1c24b44fd79137c6
title: "挑戰 143：Markdown 斜體字解析器"
challengeType: 28
dashedName: challenge-143
---

# --description--

給定一個可能包含部分 Markdown 斜體文字的字串，傳回等價的 HTML 字串。

- Markdown 中的斜體文字是任何以單一星號（`*`）或單一底線（`_`）開始和結束的文字。
- 文字和星號或底線之間不能有任何空格，但文字本身可以有空格。
- 將所有斜體字改為 HTML 的 `i` 標籤並傳回該字串。

例如，給定 `"*This is italic*"`，傳回 `"<i>This is italic</i>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseItalics("*This is italic*")` 應該傳回 `"<i>This is italic</i>"`。

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` 應該傳回 `"<i>This is also italic</i>"`。

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` 應該傳回 `"*This is not italic *"`。

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` 應該傳回 `"_ This is also not italic_"`。

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` 應該傳回 `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`。

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
