---
id: 69272dcf1c24b44fd79137c4
title: "挑戰 141：起飛燃油"
challengeType: 28
dashedName: challenge-141
---

# --description--

根據你飛機中目前的加侖數量以及到達目的地所需的公升數量，判斷你應該為你的飛機添加多少額外的加侖燃料。

- 1 加侖相等於 3.78541 公升。
- 如果飛機已經有足夠的燃料，傳回 `0`。
- 你只能添加整加侖。
- 不要在傳回數字中包含小數。

# --hints--

`fuelToAdd(0, 1)` 應該傳回 `1`。

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` 應該傳回 `6`。

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` 應該傳回 `0`。

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` 應該傳回 `4520`。

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` 應該傳回 `12209`。

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
