---
id: 69306364df283fcaff2e1ad6
title: "挑戰 145：第 N 個費氏數列數字"
challengeType: 28
dashedName: challenge-145
---

# --description--

給定一個整數 `n`，傳回費氏數列中的第 `n` 個數字。

費波那契數列是一串數字，其中每個數字都是前兩個數字的總和。數列中的前 10 個數字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

# --hints--

`nthFibonacci(4)` 應該傳回 `2`。

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` 應該傳回 `34`。

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` 應該傳回 `377`。

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` 應該傳回 `63245986`。

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` 應該傳回 `1304969544928657`。

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
