---
id: 69306364df283fcaff2e1ada
title: "挑戰 149：vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

給定一個字串，傳回一個新字串，其中所有母音都轉換為大寫，所有其他字母字元都轉換為小寫。

- 母音是任何大小寫的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"`。
- 非字母字元應保持不變。

# --hints--

`vowelCase("vowelcase")` 應該傳回 `"vOwElcAsE"`。

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` 應該傳回 `"cOdIng Is fUn"`。

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` 應該傳回 `"hEllO, wOrld!"`。

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` 應該傳回 `"gIt chErry-pIck"`。

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` 應該傳回 `"hEAd~1"`。

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
