---
id: 69306364df283fcaff2e1ad9
title: "挑戰 148：輪胎壓力"
challengeType: 28
dashedName: challenge-148
---

# --description--

給定一個包含四個數字的陣列，表述你車輛四個輪胎的胎壓（單位為 psi），以及另一個包含兩個數字的陣列，表述你輪胎的最低和最高壓力（單位為 bar），傳回一個包含四個字串的陣列，描述每個輪胎的狀態。

- 1 bar 相等 14.5038 psi。

傳回一個陣列，其中包含每個輪胎的以下值：

- 當輪胎壓力低於允許的最低值時顯示「Low」。
- 如果它介於允許的最小值和最大值之間，則為「Good」。
- 「High」如果它高於允許的最大值。

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` 應該傳回 `["Good", "Low", "Good", "Low"]`。

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` 應該傳回 `["Good", "Low", "High", "Good"]`。

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` 應該傳回 `["Low", "Low", "Good", "Low"]`。

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` 應該傳回 `["High", "High", "High", "Good"]`。

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` 應該傳回 `["Good", "Good", "Good", "Good"]`。

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
