---
id: 69306364df283fcaff2e1ad8
title: "挑戰 147：閏年計算器"
challengeType: 28
dashedName: challenge-147
---

# --description--

給定一個整數年份，判斷它是否為閏年。

如果符合以下規則，該年即為閏年：

- 該年份能被 4 整除，且
- 該年份不能被 100 整除，除非
- 該年份可被 400 整除。

# --hints--

`isLeapYear(2024)` 應該傳回 `true`。

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` 應該傳回 `false`。

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` 應該傳回 `false`。

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` 應該傳回 `true`。

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` 應該傳回 `false`。

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` 應該傳回 `true`。

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` 應該傳回 `false`。

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
