---
id: 69306364df283fcaff2e1ad5
title: "挑戰 144：決議過程連勝"
challengeType: 28
dashedName: challenge-144
---

# --description--

給定一個陣列的陣列，其中每個子陣列表述你在某一天的決議活動，並包含三個整數：該天走的步數、該天的螢幕使用分鐘數，以及該天閱讀的頁數；判斷你是否有遵守你的決議。

- 第一個子陣列是第 1 天，第二個是第 2 天，依此類推。

如果達成以下三個目標，則該天被視為成功：

- 你走了至少 10,000 步。
- 你最多只有 120 分鐘的螢幕時間。
- 你閱讀至少 5 頁。

如果所有給定的天數都成功，傳回 `"Resolution on track: N day streak."`，其中 `N` 是成功天數的數目。

如果一或多天未成功，傳回 `"Resolution failed on day X: N day streak."`。其中 `X` 是第一個不成功天數的天數編號，`N` 是第一個不成功天數之前的成功天數。

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` 應該傳回 `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` 應該傳回 `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` 應該傳回 `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` 應該傳回 `"決議過程在第 4 天失敗：連續失敗 3 天。"`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` 應該傳回「決議過程順利：連續 10 天。」

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
