---
id: 69373793f5a867f769cde135
title: "挑戰 150：Markdown 無序列表解析器"
challengeType: 28
dashedName: challenge-150
---

# --description--

給定 Markdown 中有效無序列表的字串，傳回等價的 HTML 字串。

無序列表由一個或多個列表項目組成。有效的列表項目必須獨立成行，並且：

- -
- 至少一個空格，然後
- 列表項目文字。

該列表以單一字串形式給出，換行由換行字元（`"\n"`）分隔。請勿在條款文字中包含換行字元。

將每個列表項目包裹在 HTML 的 `li` 標記中，並將整個列表包裹在 `ul` 標記中。

例如，給定 `"- Item A\n- Item B"`，傳回 `"<ul><li>Item A</li><li>Item B</li></ul>"`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseUnorderedList("- Item A\n- Item B")` 應該傳回 `"<ul><li>Item A</li><li>Item B</li></ul>"`。

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` 應該傳回 `"<ul><li>JavaScript</li><li>Python</li></ul>"`。

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` 應該傳回 `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`。

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` 應該傳回 `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`。

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
