---
id: 69373793f5a867f769cde13a
title: "挑戰 155：種植作物"
challengeType: 28
dashedName: challenge-155
---

# --description--

給定一個整數代表你的農場欄的大小，以及 `"acres"` 或 `"hectares"` 代表你的農場欄大小的單位，還有一種作物的型別，判斷你可以在你的欄中種植多少該型別的植物。

- 1 英畝相等於 4046.86 平方公尺。
- 1 公頃相等於 10,000 平方公尺。

以下是將作為輸入的作物列表以及單株植物佔用的空間：

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

傳回可放入欄中的植物數量，並向下取整至最接近的整株植物。

# --hints--

`getNumberOfPlants(1, "acres", "corn")` 應該傳回 `4046`。

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` 應該傳回 `100000`。

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` 應該傳回 `161874`。

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` 應該傳回 `150000`。

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` 應該傳回 `271139`。

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
