---
id: 69373793f5a867f769cde137
title: "挑戰 152：循環質數"
challengeType: 28
dashedName: challenge-152
---

# --description--

給定一個整數，判斷它是否為循環質數。

循環質數是指其所有數字旋轉後的結果皆為質數的整數。

例如，`197` 是一個循環質數，因為它的所有數字旋轉：`197`、`971` 和 `719`，都是質數。

# --hints--

`isCircularPrime(197)` 應該傳回 `true`。

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` 應該傳回 `false`。

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` 應該傳回 `true`。

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` 應該傳回 `false`。

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` 應該傳回 `true`。

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
