---
id: 6939b873185d8e00d453563b
title: "挑戰 156：奇數或偶數？"
challengeType: 28
dashedName: challenge-156
---

# --description--

給定一個正整數，如果它是奇數，傳回 `"Odd"`，如果是偶數，傳回 `"Even"`。

# --hints--

`oddOrEven(1)` 應該傳回 `"Odd"`。

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` 應該傳回 `"Even"`。

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` 應該傳回 `"Odd"`。

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` 應該傳回 `"Even"`。

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` 應該傳回 `"Odd"`。

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
