---
id: 6939b873185d8e00d453563c
title: "挑戰 157：Markdown 聯結解析器"
challengeType: 28
dashedName: challenge-157
---

# --description--

給定 Markdown 中的聯結字串，傳回等價的 HTML 字串。

Markdown 影像、圖片具有以下格式：`"[link_text](link_url)"`。傳回 HTML `a` 頁籤的字串，其 `href` 設為 `link_url`，並以 `link_text` 作為頁籤內容。

例如，給定 `"[freeCodeCamp](https://freecodecamp.org/)"` 傳回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`；

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤 — 請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` 應該傳回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`。

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` 應該傳回 `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`。

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` 應該傳回 `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`。

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
