---
id: 6939b873185d8e00d453563f
title: "挑戰 160：騎士移動"
challengeType: 28
dashedName: challenge-160
---

# --description--

給定騎士在棋盤上的位置，傳回騎士可以移動到的有效格子數量。

標準的棋盤是 8x8，行從左到右標記為 `A` 到 `H`，列從下到上標記為 `1` 到 `8`。它看起來像這樣：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

騎士以「L」形移動：在一個方向（水平方向或垂直方向）移動兩格，並在垂直方向移動一格。

這表示騎士最多可以移動到八個可能的位置，但當靠近棋盤邊緣時會較少。例如，如果騎士在 `A1`，它只能移動到 `B3` 或 `C2`。

# --hints--

`knightMoves("A1")` 應該傳回 `2`。

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` 應該傳回 `8`。

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` 應該傳回 `6`。

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` 應該傳回 `3`。

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` 應該傳回 `4`。

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
