---
id: 6939b873185d8e00d453563d
title: "挑戰 158：陣列交換"
challengeType: 28
dashedName: challenge-158
---

# --description--

給定一個包含兩個值的陣列，傳回一個值已交換的陣列。

例如，給定 `["A", "B"]` 傳回 `["B", "A"]`。

# --hints--

`arraySwap(["A", "B"])` 應該傳回 `["B", "A"]`。

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` 應該傳回 `[20, 25]`。

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` 應該傳回 `[false, true]`。

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` 應該傳回 `[1, "1"]`。

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
