---
id: 6939b873185d8e00d453563e
title: "挑戰 159：整數斜邊"
challengeType: 28
dashedName: challenge-159
---

# --description--

給定兩個正整數，代表直角三角形兩條直角邊（兩條較短的邊）的長度，判斷斜邊是否為整數。

斜邊的長度是透過將兩個直角邊長的平方相加，然後取該總和的平方根來計算（a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>）。

# --hints--

`isIntegerHypotenuse(3, 4)` 應該傳回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` 應該傳回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` 應該傳回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` 應該傳回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` 應該傳回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` 應該傳回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
