---
id: 694596b0585c11170ac7c7fb
title: "挑戰 163：子音大小寫"
challengeType: 28
dashedName: challenge-163
---

# --description--

給定一個表述變數名稱的字串，使用以下規則將其轉換為子音大小寫：

- 所有子音字母應該轉換為大寫。
- 所有母音（任何大小寫的 `a`、`e`、`i`、`o`、`u`）都應該轉換成小寫。
- 所有連字號（`-`）都應該轉換為底線（`_`）。

# --hints--

`toConsonantCase("helloworld")` 應該傳回 `"HeLLoWoRLD"`。

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` 應該傳回 `"HeLLoWoRLD"`。

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` 應該傳回 `"_HeLLo_WoRLD_"`。

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` 應該傳回 `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`。

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
