---
id: 694596b0585c11170ac7c7fc
title: "挑戰 164：Markdown 行內程式碼解析器"
challengeType: 28
dashedName: challenge-164
---

# --description--

給定一個包含一個或多個行內程式碼區塊的 Markdown 字串，傳回等價的 HTML 字串。

Markdown 中的行內程式碼區塊使用單個反引號（```）作為程式碼區塊文字的開始和結束。

傳回給定的字串，並將所有程式碼區塊轉換為 HTML 的 `code` 頁籤。

例如，給定字串 `"Use `let` to declare the variable."`，傳回 `"Use <code>let</code> to declare the variable."`。

注意：主控台在記錄訊息時可能不會顯示字串中的 HTML 頁籤。請檢查瀏覽器主控台以查看包含頁籤的記錄。

# --hints--

`parseInlineCode("Use `let` to declare the variable.")` 應該傳回 `"Use <code>let</code> to declare the variable."`。

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

`parseInlineCode("Use `let` or `const` to declare a variable.")` 應該傳回 `"Use <code>let</code> or <code>const</code> to declare a variable."`。

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

`parseInlineCode("Run `npm install` then `npm start`.")` 應該傳回 `"Run <code>npm install</code> then <code>npm start</code>."`。

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
