---
id: 694596b0585c11170ac7c7fa
title: "挑戰 162：能量消耗"
challengeType: 28
dashedName: challenge-162
---

# --description--

根據運動期間燃燒的卡路里數量，以及你的電子裝置在該運動期間使用的瓦特時數量，判斷哪一個使用的能量較多。

要比較它們，請使用以下轉換將兩個值都轉換為焦耳：

- 1 卡路里相等於 4184 焦耳。
- 1 瓦特時相等於 3600 焦耳。

傳回：

- 「Workout」如果該 workout 使用了更多能量。
- 「裝置」如果該裝置使用了更多能量。
- 「相等」如果兩者使用了相同量的能量。

# --hints--

`compareEnergy(250, 50)` 應該傳回 `"Workout"`。

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` 應該傳回 `"Devices"`。

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` 應該傳回 `"Equal"`。

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` 應該傳回 `"Workout"`。

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` 應該傳回 `"Devices"`。

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` 應該傳回 `"Equal"`。

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
