---
id: 696655d24b614176d4c9b789
title: "挑戰 166：十六進位驗證器"
challengeType: 28
dashedName: challenge-166
---

# --description--

給定一個字串，判斷它是否為有效的 CSS 十六進位顏色。有效的 CSS 十六進位顏色必須：

- 以 `#` 開頭，並
- 後面必須接 3 個或 6 個十六進位字元。

十六進位字元是數字 `0` 到 `9` 和字母 `a` 到 `f`（不區分大小寫）。

# --hints--

`isValidHex("#123")` 應該傳回 `true`。

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` 應該傳回 `true`。

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` 應該傳回 `true`。

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` 應該傳回 `true`。

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` 應該傳回 `false`。

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` 應該傳回 `false`。

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` 應該傳回 `false`。

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` 應該傳回 `false`。

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
