---
id: 696655d24b614176d4c9b78c
title: "挑戰 169：FizzBuzz 迷你"
challengeType: 28
dashedName: challenge-169
---

# --description--

給定一個整數，根據以下規則傳回一個字串：

- 如果數字能被 3 整除，傳回 `"Fizz"`。
- 如果數字能被 5 整除，傳回 `"Buzz"`。
- 如果數字同時可被 3 和 5 整除，傳回 `"FizzBuzz"`。
- 否則，將給定的數字傳回為字串。

# --hints--

`fizzBuzzMini(3)` 應該傳回 `"Fizz"`。

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` 應該傳回 `"4"`。

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` 應該傳回 `"Buzz"`。

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` 應該傳回 `"FizzBuzz"`。

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` 應該傳回 `"98"`。

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
