---
id: 69738771fb5a7b8b24cca2a3
title: "挑戰 177：字串鏡像"
challengeType: 28
dashedName: challenge-177
---

# --description--

給定一個字串，傳回一個新字串，該字串由給定字串與其反轉副本附加在末尾所組成。

# --hints--

`mirror("freeCodeCamp")` 應該傳回 `"freeCodeCamppmaCedoCeerf"`。

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` 應該傳回 `"RaceCarraCecaR"`。

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` 應該傳回 `"helloworlddlrowolleh"`。

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` 應該傳回 `"The quick brown fox......xof nworb kciuq ehT"`。

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
