---
id: 69738771fb5a7b8b24cca29e
title: "挑戰 172：字母-數字"
challengeType: 28
dashedName: challenge-172
---

# --description--

給定一個只包含字母和數字的字串，傳回一個新字串，當字串從字母切換到數字，或從數字切換到字母時，每次插入一個連字號（`-`）。

# --hints--

`separateLettersAndNumbers("ABC123")` 應該傳回 `"ABC-123"`。

```js
assert.equal(separateLettersAndNumbers("ABC123"), "ABC-123");
```

`separateLettersAndNumbers("Route66")` 應該傳回 `"Route-66"`。

```js
assert.equal(separateLettersAndNumbers("Route66"), "Route-66");
```

`separateLettersAndNumbers("H3LL0W0RLD")` 應該傳回 `"H-3-LL-0-W-0-RLD"`。

```js
assert.equal(separateLettersAndNumbers("H3LL0W0RLD"), "H-3-LL-0-W-0-RLD");
```

`separateLettersAndNumbers("a1b2c3d4")` 應該傳回 `"a-1-b-2-c-3-d-4"`。

```js
assert.equal(separateLettersAndNumbers("a1b2c3d4"), "a-1-b-2-c-3-d-4");
```

# --seed--

## --seed-contents--

```js
function separateLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function separateLettersAndNumbers(str) {
  if (str.length === 0) return str;

  let result = str[0];

  for (let i = 1; i < str.length; i++) {
    const prev = str[i - 1];
    const curr = str[i];

    const prevIsLetter = /[a-zA-Z]/.test(prev);
    const currIsLetter = /[a-zA-Z]/.test(curr);

    if (prevIsLetter !== currIsLetter) {
      result += "-";
    }

    result += curr;
  }

  return result;
}
```
