---
id: 69738771fb5a7b8b24cca2a4
title: "挑戰 178：截斷 `Text`"
challengeType: 28
dashedName: challenge-178
---

# --description--

給定一個字串，如果它的長度在 20 個字元以內，則原樣傳回。如果超過 20 個字元，則將它截斷為前 17 個字元，並在結尾附加 `"..."`（總共 20 個字元），然後傳回結果。

# --hints--

`truncateText("Hello, world!")` 應該傳回 `"Hello, world!"`。

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` 應該傳回 `"This string shoul..."`。

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` 應該傳回 `"Exactly twenty chars"`。

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` 應該傳回 `"...................."`。

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
