---
id: 69738771fb5a7b8b24cca2a2
title: "挑戰 176：土撥鼠日"
challengeType: 28
dashedName: challenge-176
---

# --description--

今天是土撥鼠日，土撥鼠會根據是否看到自己的影子來預測天氣。

給定一個表述土撥鼠外觀的值，傳回正確的預測：

- 如果給定的值是布林值 `true`（土撥鼠看到了牠的影子），傳回 `"Looks like we'll have six more weeks of winter."`。
- 如果值是布林值 `false`（土撥鼠沒有看到牠的影子），傳回 `"It's going to be an early spring."`。
- 如果值是其他任何東西（土撥鼠沒有出現），傳回 `"No prediction this year."`。

# --hints--

`groundhogDayPrediction(true)` 應該傳回 `"Looks like we'll have six more weeks of winter."`。

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` 應該傳回 `"It's going to be an early spring."`。

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` 應該傳回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` 應該傳回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` 應該傳回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
