---
id: 69738771fb5a7b8b24cca29f
title: "挑戰 173：有效的兵的移動"
challengeType: 28
dashedName: challenge-173
---

# --description--

給定你其中一個棋子在棋盤上的位置，傳回一個包含它可以移動到的所有有效格子且按升冪排序的陣列。

標準的棋盤是 8x8，行從左到右標記為 `A` 到 `H`，列從下到上標記為 `1` 到 `8`。它看起來像這樣：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

對於這個挑戰：

- 你是棋盤底部的玩家。
- 兵只能「向上」移動一格。
- 除非兵在起始列（第 2 列），否則它只能向上移動一格或兩格。

例如，給定 `"D4"`，傳回 `["D5"]`，這是你的兵唯一能移動到的格子。  
給定 `"B2"`，傳回 `["B3", "B4"]`，因為它在起始列且需要以遞增順序排列。

# --hints--

`findPawnMoves("D4")` 應該傳回 `["D5"]`。

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` 應該傳回 `["B3", "B4"]`。

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` 應該傳回 `["A8"]`。

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` 應該傳回 `["G3", "G4"]`。

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` 應該傳回 `["E4"]`。

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
