---
id: 69738771fb5a7b8b24cca2a5
title: "挑戰 179：零錢"
challengeType: 28
dashedName: challenge-179
---

# --description--

給定一個整數陣列，表述你口袋中硬幣的數量，每個整數為硬幣的分值，傳回總金額，格式為 `"$D.CC"`。

- 100 分相等於 1 元。
- 在返回值中，對於小於 1 元的金額包含前導零，並且分 Always exactly two digits for the cents。

# --hints--

`countChange([25, 10, 5, 1])` 應該傳回 `"$0.41"`。

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` 應該傳回 `"$1.43"`。

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` 應該傳回 `"$37.55"`。

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` 應該傳回 `"$0.70"`。

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` 應該傳回 `"$0.01"`。

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` 應該傳回 `"$1.00"`。

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
