---
id: 697a49e6ff50d756c9b6935d
title: "æŒ‘æˆ° 180ï¼š2026 å†¬å­£éŠæˆ²ç¬¬ 1 å¤©ï¼šé–‹å¹•æ—¥"
challengeType: 28
dashedName: challenge-180
---

# --description--

ä»Šå¤©æ¨™èªŒè‘— 2026 å†¬å­£éŠæˆ²çš„é–‹å§‹ã€‚æŽ¥ä¸‹ä¾†çš„ 17 å¤©å°‡ç‚ºä½ å¸¶ä¾†å—å…¶å•Ÿç™¼çš„ç¨‹å¼æŒ‘æˆ°ã€‚

å°æ–¼ç¬¬ä¸€å€‹ï¼Œä½ æœƒå¾—åˆ°ä¸€å€‹å…©å€‹å­—æ¯çš„åœ‹å®¶ä»£ç¢¼ï¼Œä¸¦ä¸”éœ€è¦å‚³å›žè©²åœ‹å®¶çš„æ——æ¨™è¡¨æƒ…ç¬¦è™Ÿã€‚

ä½¿ç”¨æ­¤åˆ—è¡¨ï¼š

| Country | Code | Flag |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"ðŸ‡¦ðŸ‡±"` |
| Andorra                    |            `"AD"` | `"ðŸ‡¦ðŸ‡©"` |
| Argentina                  |            `"AR"` | `"ðŸ‡¦ðŸ‡·"` |
| Armenia                    |            `"AM"` | `"ðŸ‡¦ðŸ‡²"` |
| Australia                  |            `"AU"` | `"ðŸ‡¦ðŸ‡º"` |
| Austria                    |            `"AT"` | `"ðŸ‡¦ðŸ‡¹"` |
| Azerbaijan                 |            `"AZ"` | `"ðŸ‡¦ðŸ‡¿"` |
| Belgium                    |            `"BE"` | `"ðŸ‡§ðŸ‡ª"` |
| Benin                      |            `"BJ"` | `"ðŸ‡§ðŸ‡¯"` |
| Bolivia                    |            `"BO"` | `"ðŸ‡§ðŸ‡´"` |
| Bosnia and Herzegovina     |            `"BA"` | `"ðŸ‡§ðŸ‡¦"` |
| Brazil                     |            `"BR"` | `"ðŸ‡§ðŸ‡·"` |
| Bulgaria                   |            `"BG"` | `"ðŸ‡§ðŸ‡¬"` |
| Canada                     |            `"CA"` | `"ðŸ‡¨ðŸ‡¦"` |
| Chile                      |            `"CL"` | `"ðŸ‡¨ðŸ‡±"` |
| China                      |            `"CN"` | `"ðŸ‡¨ðŸ‡³"` |
| Colombia                   |            `"CO"` | `"ðŸ‡¨ðŸ‡´"` |
| Croatia                    |            `"HR"` | `"ðŸ‡­ðŸ‡·"` |
| Cyprus                     |            `"CY"` | `"ðŸ‡¨ðŸ‡¾"` |
| Czech Republic             |            `"CZ"` | `"ðŸ‡¨ðŸ‡¿"` |
| Denmark                    |            `"DK"` | `"ðŸ‡©ðŸ‡°"` |
| Ecuador                    |            `"EC"` | `"ðŸ‡ªðŸ‡¨"` |
| Eritrea                    |            `"ER"` | `"ðŸ‡ªðŸ‡·"` |
| Estonia                    |            `"EE"` | `"ðŸ‡ªðŸ‡ª"` |
| Finland                    |            `"FI"` | `"ðŸ‡«ðŸ‡®"` |
| France                     |            `"FR"` | `"ðŸ‡«ðŸ‡·"` |
| Georgia                    |            `"GE"` | `"ðŸ‡¬ðŸ‡ª"` |
| Germany                    |            `"DE"` | `"ðŸ‡©ðŸ‡ª"` |
| Great Britain              |            `"GB"` | `"ðŸ‡¬ðŸ‡§"` |
| Greece                     |            `"GR"` | `"ðŸ‡¬ðŸ‡·"` |
| Guinea-Bissau              |            `"GW"` | `"ðŸ‡¬ðŸ‡¼"` |
| Haiti                      |            `"HT"` | `"ðŸ‡­ðŸ‡¹"` |
| Hong Kong                  |            `"HK"` | `"ðŸ‡­ðŸ‡°"` |
| Hungary                    |            `"HU"` | `"ðŸ‡­ðŸ‡º"` |
| Iceland                    |            `"IS"` | `"ðŸ‡®ðŸ‡¸"` |
| India                      |            `"IN"` | `"ðŸ‡®ðŸ‡³"` |
| Iran                       |            `"IR"` | `"ðŸ‡®ðŸ‡·"` |
| Ireland                    |            `"IE"` | `"ðŸ‡®ðŸ‡ª"` |
| Israel                     |            `"IL"` | `"ðŸ‡®ðŸ‡±"` |
| Italy                      |            `"IT"` | `"ðŸ‡®ðŸ‡¹"` |
| Jamaica                    |            `"JM"` | `"ðŸ‡¯ðŸ‡²"` |
| Japan                      |            `"JP"` | `"ðŸ‡¯ðŸ‡µ"` |
| Kazakhstan                 |            `"KZ"` | `"ðŸ‡°ðŸ‡¿"` |
| Kenya                      |            `"KE"` | `"ðŸ‡°ðŸ‡ª"` |
| Kosovo                     |            `"XK"` | `"ðŸ‡½ðŸ‡°"` |
| Kyrgyzstan                 |            `"KG"` | `"ðŸ‡°ðŸ‡¬"` |
| Latvia                     |            `"LV"` | `"ðŸ‡±ðŸ‡»"` |
| Lebanon                    |            `"LB"` | `"ðŸ‡±ðŸ‡§"` |
| Liechtenstein              |            `"LI"` | `"ðŸ‡±ðŸ‡®"` |
| Lithuania                  |            `"LT"` | `"ðŸ‡±ðŸ‡¹"` |
| Luxembourg                 |            `"LU"` | `"ðŸ‡±ðŸ‡º"` |
| Madagascar                 |            `"MG"` | `"ðŸ‡²ðŸ‡¬"` |
| Malaysia                   |            `"MY"` | `"ðŸ‡²ðŸ‡¾"` |
| Malta                      |            `"MT"` | `"ðŸ‡²ðŸ‡¹"` |
| Mexico                     |            `"MX"` | `"ðŸ‡²ðŸ‡½"` |
| Moldova                    |            `"MD"` | `"ðŸ‡²ðŸ‡©"` |
| Monaco                     |            `"MC"` | `"ðŸ‡²ðŸ‡¨"` |
| Mongolia                   |            `"MN"` | `"ðŸ‡²ðŸ‡³"` |
| Montenegro                 |            `"ME"` | `"ðŸ‡²ðŸ‡ª"` |
| Morocco                    |            `"MA"` | `"ðŸ‡²ðŸ‡¦"` |
| Netherlands                |            `"NL"` | `"ðŸ‡³ðŸ‡±"` |
| New Zealand                |            `"NZ"` | `"ðŸ‡³ðŸ‡¿"` |
| Nigeria                    |            `"NG"` | `"ðŸ‡³ðŸ‡¬"` |
| North Macedonia            |            `"MK"` | `"ðŸ‡²ðŸ‡°"` |
| Norway                     |            `"NO"` | `"ðŸ‡³ðŸ‡´"` |
| Pakistan                   |            `"PK"` | `"ðŸ‡µðŸ‡°"` |
| Philippines                |            `"PH"` | `"ðŸ‡µðŸ‡­"` |
| Poland                     |            `"PL"` | `"ðŸ‡µðŸ‡±"` |
| Portugal                   |            `"PT"` | `"ðŸ‡µðŸ‡¹"` |
| Puerto Rico                |            `"PR"` | `"ðŸ‡µðŸ‡·"` |
| Romania                    |            `"RO"` | `"ðŸ‡·ðŸ‡´"` |
| San Marino                 |            `"SM"` | `"ðŸ‡¸ðŸ‡²"` |
| Saudi Arabia               |            `"SA"` | `"ðŸ‡¸ðŸ‡¦"` |
| Serbia                     |            `"RS"` | `"ðŸ‡·ðŸ‡¸"` |
| Singapore                  |            `"SG"` | `"ðŸ‡¸ðŸ‡¬"` |
| Slovakia                   |            `"SK"` | `"ðŸ‡¸ðŸ‡°"` |
| Slovenia                   |            `"SI"` | `"ðŸ‡¸ðŸ‡®"` |
| South Africa               |            `"ZA"` | `"ðŸ‡¿ðŸ‡¦"` |
| South Korea                |            `"KR"` | `"ðŸ‡°ðŸ‡·"` |
| Spain                      |            `"ES"` | `"ðŸ‡ªðŸ‡¸"` |
| Sweden                     |            `"SE"` | `"ðŸ‡¸ðŸ‡ª"` |
| Switzerland                |            `"CH"` | `"ðŸ‡¨ðŸ‡­"` |
| Thailand                   |            `"TH"` | `"ðŸ‡¹ðŸ‡­"` |
| Trinidad & Tobago          |            `"TT"` | `"ðŸ‡¹ðŸ‡¹"` |
| Turkey                     |            `"TR"` | `"ðŸ‡¹ðŸ‡·"` |
| Ukraine                    |            `"UA"` | `"ðŸ‡ºðŸ‡¦"` |
| United Arab Emirates       |            `"AE"` | `"ðŸ‡¦ðŸ‡ª"` |
| United States              |            `"US"` | `"ðŸ‡ºðŸ‡¸"` |
| Uruguay                    |            `"UY"` | `"ðŸ‡ºðŸ‡¾"` |
| Uzbekistan                 |            `"UZ"` | `"ðŸ‡ºðŸ‡¿"` |
| Venezuela                  |            `"VE"` | `"ðŸ‡»ðŸ‡ª"` |

# --hints--

`getFlag("AL")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("AL"), "ðŸ‡¦ðŸ‡±");
```

`getFlag("AD")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡©"`ã€‚

```js
assert.equal(getFlag("AD"), "ðŸ‡¦ðŸ‡©");
```

`getFlag("AR")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("AR"), "ðŸ‡¦ðŸ‡·")
```

`getFlag("AM")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("AM"), "ðŸ‡¦ðŸ‡²")
```

`getFlag("AU")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("AU"), "ðŸ‡¦ðŸ‡º")
```

`getFlag("AT")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("AT"), "ðŸ‡¦ðŸ‡¹")
```

`getFlag("AZ")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("AZ"), "ðŸ‡¦ðŸ‡¿")
```

`getFlag("BE")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("BE"), "ðŸ‡§ðŸ‡ª")
```

`getFlag("BJ")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡¯"`ã€‚

```js
assert.equal(getFlag("BJ"), "ðŸ‡§ðŸ‡¯")
```

`getFlag("BO")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("BO"), "ðŸ‡§ðŸ‡´")
```

`getFlag("BA")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("BA"), "ðŸ‡§ðŸ‡¦")
```

`getFlag("BR")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("BR"), "ðŸ‡§ðŸ‡·")
```

`getFlag("BG")` æ‡‰è©²å‚³å›ž `"ðŸ‡§ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("BG"), "ðŸ‡§ðŸ‡¬")
```

`getFlag("CA")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("CA"), "ðŸ‡¨ðŸ‡¦")
```

`getFlag("CL")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("CL"), "ðŸ‡¨ðŸ‡±")
```

`getFlag("CN")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("CN"), "ðŸ‡¨ðŸ‡³")
```

`getFlag("CO")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("CO"), "ðŸ‡¨ðŸ‡´")
```

`getFlag("HR")` æ‡‰è©²å‚³å›ž `"ðŸ‡­ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("HR"), "ðŸ‡­ðŸ‡·")
```

`getFlag("CY")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("CY"), "ðŸ‡¨ðŸ‡¾")
```

`getFlag("CZ")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("CZ"), "ðŸ‡¨ðŸ‡¿")
```

`getFlag("DK")` æ‡‰è©²å‚³å›ž `"ðŸ‡©ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("DK"), "ðŸ‡©ðŸ‡°")
```

`getFlag("EC")` æ‡‰è©²å‚³å›ž `"ðŸ‡ªðŸ‡¨"`ã€‚

```js
assert.equal(getFlag("EC"), "ðŸ‡ªðŸ‡¨")
```

`getFlag("ER")` æ‡‰è©²å‚³å›ž `"ðŸ‡ªðŸ‡·"`ã€‚

```js
assert.equal(getFlag("ER"), "ðŸ‡ªðŸ‡·")
```

`getFlag("EE")` æ‡‰è©²å‚³å›ž `"ðŸ‡ªðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("EE"), "ðŸ‡ªðŸ‡ª")
```

`getFlag("FI")` æ‡‰è©²å‚³å›ž `"ðŸ‡«ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("FI"), "ðŸ‡«ðŸ‡®")
```

`getFlag("FR")` æ‡‰è©²å‚³å›ž `"ðŸ‡«ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("FR"), "ðŸ‡«ðŸ‡·")
```

`getFlag("GE")` æ‡‰è©²å‚³å›ž `"ðŸ‡¬ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("GE"), "ðŸ‡¬ðŸ‡ª")
```

`getFlag("DE")` æ‡‰è©²å‚³å›ž `"ðŸ‡©ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("DE"), "ðŸ‡©ðŸ‡ª")
```

`getFlag("GB")` æ‡‰è©²å‚³å›ž `"ðŸ‡¬ðŸ‡§"`ã€‚

```js
assert.equal(getFlag("GB"), "ðŸ‡¬ðŸ‡§")
```

`getFlag("GR")` æ‡‰è©²å‚³å›ž `"ðŸ‡¬ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("GR"), "ðŸ‡¬ðŸ‡·")
```

`getFlag("GW")` æ‡‰è©²å‚³å›ž `"ðŸ‡¬ðŸ‡¼"`ã€‚

```js
assert.equal(getFlag("GW"), "ðŸ‡¬ðŸ‡¼")
```

`getFlag("HT")` æ‡‰è©²å‚³å›ž `"ðŸ‡­ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("HT"), "ðŸ‡­ðŸ‡¹")
```

`getFlag("HK")` æ‡‰è©²å‚³å›ž `"ðŸ‡­ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("HK"), "ðŸ‡­ðŸ‡°")
```

`getFlag("HU")` æ‡‰è©²å‚³å›ž `"ðŸ‡­ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("HU"), "ðŸ‡­ðŸ‡º")
```

`getFlag("IS")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("IS"), "ðŸ‡®ðŸ‡¸")
```

`getFlag("IN")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("IN"), "ðŸ‡®ðŸ‡³")
```

`getFlag("IR")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("IR"), "ðŸ‡®ðŸ‡·")
```

`getFlag("IE")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("IE"), "ðŸ‡®ðŸ‡ª")
```

`getFlag("IL")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("IL"), "ðŸ‡®ðŸ‡±")
```

`getFlag("IT")` æ‡‰è©²å‚³å›ž `"ðŸ‡®ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("IT"), "ðŸ‡®ðŸ‡¹")
```

`getFlag("JM")` æ‡‰è©²å‚³å›ž `"ðŸ‡¯ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("JM"), "ðŸ‡¯ðŸ‡²")
```

`getFlag("JP")` æ‡‰è©²å‚³å›ž `"ðŸ‡¯ðŸ‡µ"`ã€‚

```js
assert.equal(getFlag("JP"), "ðŸ‡¯ðŸ‡µ")
```

`getFlag("KZ")` æ‡‰è©²å‚³å›ž `"ðŸ‡°ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("KZ"), "ðŸ‡°ðŸ‡¿")
```

`getFlag("KE")` æ‡‰è©²å‚³å›ž `"ðŸ‡°ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("KE"), "ðŸ‡°ðŸ‡ª")
```

`getFlag("XK")` æ‡‰è©²å‚³å›ž `"ðŸ‡½ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("XK"), "ðŸ‡½ðŸ‡°")
```

`getFlag("KG")` æ‡‰è©²å‚³å›ž `"ðŸ‡°ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("KG"), "ðŸ‡°ðŸ‡¬")
```

`getFlag("LV")` æ‡‰è©²å‚³å›ž `"ðŸ‡±ðŸ‡»"`ã€‚

```js
assert.equal(getFlag("LV"), "ðŸ‡±ðŸ‡»")
```

`getFlag("LB")` æ‡‰è©²å‚³å›ž `"ðŸ‡±ðŸ‡§"`ã€‚

```js
assert.equal(getFlag("LB"), "ðŸ‡±ðŸ‡§")
```

`getFlag("LI")` æ‡‰è©²å‚³å›ž `"ðŸ‡±ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("LI"), "ðŸ‡±ðŸ‡®")
```

`getFlag("LT")` æ‡‰è©²å‚³å›ž `"ðŸ‡±ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("LT"), "ðŸ‡±ðŸ‡¹")
```

`getFlag("LU")` æ‡‰è©²å‚³å›ž `"ðŸ‡±ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("LU"), "ðŸ‡±ðŸ‡º")
```

`getFlag("MG")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("MG"), "ðŸ‡²ðŸ‡¬")
```

`getFlag("MY")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("MY"), "ðŸ‡²ðŸ‡¾")
```

`getFlag("MT")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("MT"), "ðŸ‡²ðŸ‡¹")
```

`getFlag("MX")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡½"`ã€‚

```js
assert.equal(getFlag("MX"), "ðŸ‡²ðŸ‡½")
```

`getFlag("MD")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡©"`ã€‚

```js
assert.equal(getFlag("MD"), "ðŸ‡²ðŸ‡©")
```

`getFlag("MC")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡¨"`ã€‚

```js
assert.equal(getFlag("MC"), "ðŸ‡²ðŸ‡¨")
```

`getFlag("MN")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("MN"), "ðŸ‡²ðŸ‡³")
```

`getFlag("ME")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("ME"), "ðŸ‡²ðŸ‡ª")
```

`getFlag("MA")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("MA"), "ðŸ‡²ðŸ‡¦")
```

`getFlag("NL")` æ‡‰è©²å‚³å›ž `"ðŸ‡³ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("NL"), "ðŸ‡³ðŸ‡±")
```

`getFlag("NZ")` æ‡‰è©²å‚³å›ž `"ðŸ‡³ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("NZ"), "ðŸ‡³ðŸ‡¿")
```

`getFlag("NG")` æ‡‰è©²å‚³å›ž `"ðŸ‡³ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("NG"), "ðŸ‡³ðŸ‡¬")
```

`getFlag("MK")` æ‡‰è©²å‚³å›ž `"ðŸ‡²ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("MK"), "ðŸ‡²ðŸ‡°")
```

`getFlag("NO")` æ‡‰è©²å‚³å›ž `"ðŸ‡³ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("NO"), "ðŸ‡³ðŸ‡´")
```

`getFlag("PK")` æ‡‰è©²å‚³å›ž `"ðŸ‡µðŸ‡°"`ã€‚

```js
assert.equal(getFlag("PK"), "ðŸ‡µðŸ‡°")
```

`getFlag("PH")` æ‡‰è©²å‚³å›ž `"ðŸ‡µðŸ‡­"`ã€‚

```js
assert.equal(getFlag("PH"), "ðŸ‡µðŸ‡­")
```

`getFlag("PL")` æ‡‰è©²å‚³å›ž `"ðŸ‡µðŸ‡±"`ã€‚

```js
assert.equal(getFlag("PL"), "ðŸ‡µðŸ‡±")
```

`getFlag("PT")` æ‡‰è©²å‚³å›ž `"ðŸ‡µðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("PT"), "ðŸ‡µðŸ‡¹")
```

`getFlag("PR")` æ‡‰è©²å‚³å›ž `"ðŸ‡µðŸ‡·"`ã€‚

```js
assert.equal(getFlag("PR"), "ðŸ‡µðŸ‡·")
```

`getFlag("RO")` æ‡‰è©²å‚³å›ž `"ðŸ‡·ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("RO"), "ðŸ‡·ðŸ‡´")
```

`getFlag("SM")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("SM"), "ðŸ‡¸ðŸ‡²")
```

`getFlag("SA")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("SA"), "ðŸ‡¸ðŸ‡¦")
```

`getFlag("RS")` æ‡‰è©²å‚³å›ž `"ðŸ‡·ðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("RS"), "ðŸ‡·ðŸ‡¸")
```

`getFlag("SG")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("SG"), "ðŸ‡¸ðŸ‡¬")
```

`getFlag("SK")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("SK"), "ðŸ‡¸ðŸ‡°")
```

`getFlag("SI")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("SI"), "ðŸ‡¸ðŸ‡®")
```

`getFlag("ZA")` æ‡‰è©²å‚³å›ž `"ðŸ‡¿ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("ZA"), "ðŸ‡¿ðŸ‡¦")
```

`getFlag("KR")` æ‡‰è©²å‚³å›ž `"ðŸ‡°ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("KR"), "ðŸ‡°ðŸ‡·")
```

`getFlag("ES")` æ‡‰è©²å‚³å›ž `"ðŸ‡ªðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("ES"), "ðŸ‡ªðŸ‡¸")
```

`getFlag("SE")` æ‡‰è©²å‚³å›ž `"ðŸ‡¸ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("SE"), "ðŸ‡¸ðŸ‡ª")
```

`getFlag("CH")` æ‡‰è©²å‚³å›ž `"ðŸ‡¨ðŸ‡­"`ã€‚

```js
assert.equal(getFlag("CH"), "ðŸ‡¨ðŸ‡­")
```

`getFlag("TH")` æ‡‰è©²å‚³å›ž `"ðŸ‡¹ðŸ‡­"`ã€‚

```js
assert.equal(getFlag("TH"), "ðŸ‡¹ðŸ‡­")
```

`getFlag("TT")` æ‡‰è©²å‚³å›ž `"ðŸ‡¹ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("TT"), "ðŸ‡¹ðŸ‡¹")
```

`getFlag("TR")` æ‡‰è©²å‚³å›ž `"ðŸ‡¹ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("TR"), "ðŸ‡¹ðŸ‡·")
```

`getFlag("UA")` æ‡‰è©²å‚³å›ž `"ðŸ‡ºðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("UA"), "ðŸ‡ºðŸ‡¦")
```

`getFlag("AE")` æ‡‰è©²å‚³å›ž `"ðŸ‡¦ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("AE"), "ðŸ‡¦ðŸ‡ª")
```

`getFlag("US")` æ‡‰è©²å‚³å›ž `"ðŸ‡ºðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("US"), "ðŸ‡ºðŸ‡¸")
```

`getFlag("UY")` æ‡‰è©²å‚³å›ž `"ðŸ‡ºðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("UY"), "ðŸ‡ºðŸ‡¾")
```

`getFlag("UZ")` æ‡‰è©²å‚³å›ž `"ðŸ‡ºðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("UZ"), "ðŸ‡ºðŸ‡¿")
```

`getFlag("VE")` æ‡‰è©²å‚³å›ž `"ðŸ‡»ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("VE"), "ðŸ‡»ðŸ‡ª")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
