---
id: 697a49e6ff50d756c9b69362
title: "挑戰 185：2026 冬季遊戲第 6 天：花式滑冰"
challengeType: 28
dashedName: challenge-185
---

# --description--

給定一個評審分數的陣列和選用的罰分，計算花式滑冰常式的最終分數。

第一個引數是包含 10 個評審分數的陣列，每個分數都是 0 到 10 的數字。移除最高和最低的評審分數，將剩下的 8 個分數相加以取得基礎分數。

傳遞給函式的任何額外引數都是懲罰。從基本分數中扣除所有懲罰以取得最終分數。

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` 應該傳回 `64`。

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` 應該傳回 `80`。

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` 應該傳回 `67`。

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` 應該傳回 `67.5`。

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` 應該傳回 `59`。

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```
