---
id: 697a49e9860d24853adef67b
title: "挑戰 190：2026 冬季遊戲第 11 天：冰球"
challengeType: 28
dashedName: challenge-190
---

# --description--

給定 6 支冰球隊的陣列及其循環賽後的記錄，判定準決賽的對戰組合。

- 每個陣列項目都會有一個隊伍及其記錄，格式為 `"TEAM: W-OTW-OTL-L"`。其中：
  - 「W」是正規賽中的勝場數，每場勝利價值 3 分
  - 「OTW」是加時勝利的次數，每次價值 2 分。
  - 「OTL」是加時賽失利的次數，每次價值 1 分。
  - 「L」是損失次數，每次損失值 0 分。
  
例如，`"FIN: 2-2-1-0"` 在加總他們的記錄後會有 11 分。
  
找出每隊的總得分並傳回 `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`。例如，`"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` 應該傳回 `"The semi-final games will be FIN vs SWE and CAN vs USA."`。

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` 應該傳回 `"半決賽將是 USA 對 CZE 以及 CAN 對 FIN。"`。

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` 應該傳回 `"半決賽將是 CAN 對 ITA 和 USA 對 CZE。"`。

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` 應該傳回 `"The semi-final games will be JPN vs ITA and AUT vs KOR."`。

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
