---
id: 697a49e9860d24853adef67c
title: "挑戰 191：2026 冬季遊戲第 12 天：雪橇競速"
challengeType: 28
dashedName: challenge-191
---

# --description--

給定一個表述雪橇隊運動員體重的陣列和一個表述雪橇重量的數字，判斷該隊是否有資格參賽。

- 陣列的長度決定隊伍人數：1 人、2 人或 4 人隊伍。
- 所有給定的重量值均以公斤（kg）為單位。

雪橇（單獨的雪橇）必須有最低重量：

- 1 人團隊的 162 kg
- 2 人團隊的 170 kg
- 4 人團隊的 210 kg

雪橇總重量（運動員加上雪橇）不得超過：

- 1 人團隊的 247 公斤
- 2 人團隊的 390 kg
- 4 人團隊的 630 kg

傳回 `"Eligible"` 如果團隊符合所有要求，或傳回 `"Not Eligible"` 如果團隊未能符合一項或多項要求。

# --hints--

`checkEligibility([78], 165)` 應該傳回 `"Eligible"`。

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` 應該傳回 `"Eligible"`。

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` 應該傳回 `"Eligible"`。

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` 應該傳回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
