---
id: 697a49e6ff50d756c9b6935d
title: "挑戰 180：冬季運動會第 1 天：開幕日"
challengeType: 28
dashedName: challenge-180
---

# --description--

接下來的 17 天將帶給你受冬季運動會啟發的程式碼挑戰。

對於第一個，你會得到一個兩個字母的國家代碼，並且需要傳回該國家的旗標表情符號。

使用此列表：

| 國家 | 代碼 | 旗標 |
| - | -: | :--: |
| 阿爾巴尼亞                    |            `"AL"` | `"🇦🇱"` |
| 安道爾                      |            `"AD"` | `"🇦🇩"` |
| 阿根廷                      |            `"AR"` | `"🇦🇷"` |
| 亞美尼亞                    |            `"AM"` | `"🇦🇲"` |
| 澳洲                        |            `"AU"` | `"🇦🇺"` |
| 奧地利                      |            `"AT"` | `"🇦🇹"` |
| 亞塞拜然                    |            `"AZ"` | `"🇦🇿"` |
| 比利時                      |            `"BE"` | `"🇧🇪"` |
| 貝寧                        |            `"BJ"` | `"🇧🇯"` |
| 玻利維亞                    |            `"BO"` | `"🇧🇴"` |
| 波士尼亞與赫塞哥維納        |            `"BA"` | `"🇧🇦"` |
| 巴西                        |            `"BR"` | `"🇧🇷"` |
| 保加利亞                    |            `"BG"` | `"🇧🇬"` |
| 加拿大                      |            `"CA"` | `"🇨🇦"` |
| 智利                        |            `"CL"` | `"🇨🇱"` |
| 中國                        |            `"CN"` | `"🇨🇳"` |
| 哥倫比亞                    |            `"CO"` | `"🇨🇴"` |
| 克羅埃西亞                  |            `"HR"` | `"🇭🇷"` |
| 賽普勒斯                    |            `"CY"` | `"🇨🇾"` |
| 捷克共和國                  |            `"CZ"` | `"🇨🇿"` |
| 丹麥                        |            `"DK"` | `"🇩🇰"` |
| 厄瓜多爾                    |            `"EC"` | `"🇪🇨"` |
| 厄立特里亞                  |            `"ER"` | `"🇪🇷"` |
| 愛沙尼亞                    |            `"EE"` | `"🇪🇪"` |
| 芬蘭                        |            `"FI"` | `"🇫🇮"` |
| 法國                        |            `"FR"` | `"🇫🇷"` |
| 喬治亞                      |            `"GE"` | `"🇬🇪"` |
| 德國                        |            `"DE"` | `"🇩🇪"` |
| 大不列顛                    |            `"GB"` | `"🇬🇧"` |
| 希臘                        |            `"GR"` | `"🇬🇷"` |
| 幾內亞比索                  |            `"GW"` | `"🇬🇼"` |
| 海地                        |            `"HT"` | `"🇭🇹"` |
| 香港                        |            `"HK"` | `"🇭🇰"` |
| 匈牙利                      |            `"HU"` | `"🇭🇺"` |
| 冰島                        |            `"IS"` | `"🇮🇸"` |
| 印度                        |            `"IN"` | `"🇮🇳"` |
| 伊朗                        |            `"IR"` | `"🇮🇷"` |
| 愛爾蘭                      |            `"IE"` | `"🇮🇪"` |
| 以色列                      |            `"IL"` | `"🇮🇱"` |
| 義大利                      |            `"IT"` | `"🇮🇹"` |
| 牙買加                      |            `"JM"` | `"🇯🇲"` |
| 日本                        |            `"JP"` | `"🇯🇵"` |
| 哈薩克                      |            `"KZ"` | `"🇰🇿"` |
| 肯亞                        |            `"KE"` | `"🇰🇪"` |
| 科索沃                      |            `"XK"` | `"🇽🇰"` |
| 吉爾吉斯                    |            `"KG"` | `"🇰🇬"` |
| 拉脫維亞                    |            `"LV"` | `"🇱🇻"` |
| 黎巴嫩                      |            `"LB"` | `"🇱🇧"` |
| 列支敦斯登                  |            `"LI"` | `"🇱🇮"` |
| 立陶宛                      |            `"LT"` | `"🇱🇹"` |
| 盧森堡                      |            `"LU"` | `"🇱🇺"` |
| 馬達加斯加                  |            `"MG"` | `"🇲🇬"` |
| 馬來西亞                    |            `"MY"` | `"🇲🇾"` |
| 馬爾他                      |            `"MT"` | `"🇲🇹"` |
| 墨西哥                      |            `"MX"` | `"🇲🇽"` |
| 摩爾多瓦                    |            `"MD"` | `"🇲🇩"` |
| 摩納哥                      |            `"MC"` | `"🇲🇨"` |
| 蒙古                        |            `"MN"` | `"🇲🇳"` |
| 黑山                        |            `"ME"` | `"🇲🇪"` |
| 摩洛哥                      |            `"MA"` | `"🇲🇦"` |
| 荷蘭                        |            `"NL"` | `"🇳🇱"` |
| 紐西蘭                      |            `"NZ"` | `"🇳🇿"` |
| 尼日利亞                    |            `"NG"` | `"🇳🇬"` |
| 北馬其頓                    |            `"MK"` | `"🇲🇰"` |
| 挪威                        |            `"NO"` | `"🇳🇴"` |
| 巴基斯坦                    |            `"PK"` | `"🇵🇰"` |
| 菲律賓                      |            `"PH"` | `"🇵🇭"` |
| 波蘭                        |            `"PL"` | `"🇵🇱"` |
| 葡萄牙                      |            `"PT"` | `"🇵🇹"` |
| 波多黎各                    |            `"PR"` | `"🇵🇷"` |
| 羅馬尼亞                    |            `"RO"` | `"🇷🇴"` |
| 聖馬利諾                  |            `"SM"` | `"🇸🇲"` |
| 沙烏地阿拉伯                |            `"SA"` | `"🇸🇦"` |
| 塞爾維亞                    |            `"RS"` | `"🇷🇸"` |
| 新加坡                      |            `"SG"` | `"🇸🇬"` |
| 斯洛伐克                    |            `"SK"` | `"🇸🇰"` |
| 斯洛維尼亞                  |            `"SI"` | `"🇸🇮"` |
| 南非                        |            `"ZA"` | `"🇿🇦"` |
| 南韓                        |            `"KR"` | `"🇰🇷"` |
| 西班牙                      |            `"ES"` | `"🇪🇸"` |
| 瑞典                        |            `"SE"` | `"🇸🇪"` |
| 瑞士                        |            `"CH"` | `"🇨🇭"` |
| 泰國                        |            `"TH"` | `"🇹🇭"` |
| 千里達及托巴哥              |            `"TT"` | `"🇹🇹"` |
| 土耳其                      |            `"TR"` | `"🇹🇷"` |
| 烏克蘭                      |            `"UA"` | `"🇺🇦"` |
| 阿拉伯聯合大公國            |            `"AE"` | `"🇦🇪"` |
| 美國                        |            `"US"` | `"🇺🇸"` |
| 烏拉圭                      |            `"UY"` | `"🇺🇾"` |
| 烏茲別克                    |            `"UZ"` | `"🇺🇿"` |
| 委內瑞拉                    |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` 應該傳回 `"🇦🇱"`。

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` 應該傳回 `"🇦🇩"`。

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` 應該傳回 `"🇦🇷"`。

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` 應該傳回 `"🇦🇲"`。

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` 應該傳回 `"🇦🇺"`。

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` 應該傳回 `"🇦🇹"`。

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` 應該傳回 `"🇦🇿"`。

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` 應該傳回 `"🇧🇪"`。

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` 應該傳回 `"🇧🇯"`。

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` 應該傳回 `"🇧🇴"`。

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` 應該傳回 `"🇧🇦"`。

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` 應該傳回 `"🇧🇷"`。

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` 應該傳回 `"🇧🇬"`。

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` 應該傳回 `"🇨🇦"`。

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` 應該傳回 `"🇨🇱"`。

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` 應該傳回 `"🇨🇳"`。

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` 應該傳回 `"🇨🇴"`。

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` 應該傳回 `"🇭🇷"`。

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` 應該傳回 `"🇨🇾"`。

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` 應該傳回 `"🇨🇿"`。

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` 應該傳回 `"🇩🇰"`。

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` 應該傳回 `"🇪🇨"`。

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` 應該傳回 `"🇪🇷"`。

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` 應該傳回 `"🇪🇪"`。

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` 應該傳回 `"🇫🇮"`。

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` 應該傳回 `"🇫🇷"`。

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` 應該傳回 `"🇬🇪"`。

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` 應該傳回 `"🇩🇪"`。

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` 應該傳回 `"🇬🇧"`。

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` 應該傳回 `"🇬🇷"`。

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` 應該傳回 `"🇬🇼"`。

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` 應該傳回 `"🇭🇹"`。

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` 應該傳回 `"🇭🇰"`。

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` 應該傳回 `"🇭🇺"`。

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` 應該傳回 `"🇮🇸"`。

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` 應該傳回 `"🇮🇳"`。

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` 應該傳回 `"🇮🇷"`。

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` 應該傳回 `"🇮🇪"`。

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` 應該傳回 `"🇮🇱"`。

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` 應該傳回 `"🇮🇹"`。

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` 應該傳回 `"🇯🇲"`。

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` 應該傳回 `"🇯🇵"`。

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` 應該傳回 `"🇰🇿"`。

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` 應該傳回 `"🇰🇪"`。

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` 應該傳回 `"🇽🇰"`。

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` 應該傳回 `"🇰🇬"`。

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` 應該傳回 `"🇱🇻"`。

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` 應該傳回 `"🇱🇧"`。

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` 應該傳回 `"🇱🇮"`。

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` 應該傳回 `"🇱🇹"`。

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` 應該傳回 `"🇱🇺"`。

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` 應該傳回 `"🇲🇬"`。

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` 應該傳回 `"🇲🇾"`。

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` 應該傳回 `"🇲🇹"`。

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` 應該傳回 `"🇲🇽"`。

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` 應該傳回 `"🇲🇩"`。

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` 應該傳回 `"🇲🇨"`。

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` 應該傳回 `"🇲🇳"`。

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` 應該傳回 `"🇲🇪"`。

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` 應該傳回 `"🇲🇦"`。

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` 應該傳回 `"🇳🇱"`。

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` 應該傳回 `"🇳🇿"`。

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` 應該傳回 `"🇳🇬"`。

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` 應該傳回 `"🇲🇰"`。

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` 應該傳回 `"🇳🇴"`。

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` 應該傳回 `"🇵🇰"`。

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` 應該傳回 `"🇵🇭"`。

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` 應該傳回 `"🇵🇱"`。

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` 應該傳回 `"🇵🇹"`。

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` 應該傳回 `"🇵🇷"`。

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` 應該傳回 `"🇷🇴"`。

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` 應該傳回 `"🇸🇲"`。

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` 應該傳回 `"🇸🇦"`。

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` 應該傳回 `"🇷🇸"`。

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` 應該傳回 `"🇸🇬"`。

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` 應該傳回 `"🇸🇰"`。

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` 應該傳回 `"🇸🇮"`。

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` 應該傳回 `"🇿🇦"`。

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` 應該傳回 `"🇰🇷"`。

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` 應該傳回 `"🇪🇸"`。

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` 應該傳回 `"🇸🇪"`。

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` 應該傳回 `"🇨🇭"`。

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` 應該傳回 `"🇹🇭"`。

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` 應該傳回 `"🇹🇹"`。

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` 應該傳回 `"🇹🇷"`。

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` 應該傳回 `"🇺🇦"`。

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` 應該傳回 `"🇦🇪"`。

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` 應該傳回 `"🇺🇸"`。

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` 應該傳回 `"🇺🇾"`。

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` 應該傳回 `"🇺🇿"`。

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` 應該傳回 `"🇻🇪"`。

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
