---
id: 697a49e6ff50d756c9b69365
title: "挑戰 188：2026 冬季遊戲第 9 天：Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

給定一個表述骨架賽道曲線的字串，判斷該賽道的難度。

- 給定的字串將只包含以下字母：
  - 「L」代表左轉
  - 「R」表示右轉
  - 「S」代表直線段

- 每次方向改變會增加 15 分（`"L"` 後接 `"R"` 或反之亦然）。
- 所有其他曲線各加 5 點（所有其他 `"L"` 或 `"R"` 字元）。
- 直線段加 0 分。

賽道的難度是根據總分計算。傳回：

- 「Easy」如果總數是 0 - 100
- 「Medium」如果總數在 101-200 之間
- 如果總數超過 200，則為「Hard」

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` 應該傳回 `"Easy"`。

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` 應該傳回 `"Hard"`。

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` 應該傳回 `"Medium"`。

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` 應該傳回 `"Hard"`。

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` 應該傳回 `"Medium"`。

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` 應該傳回 `"Easy"`。

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
