---
id: 697a49e6ff50d756c9b69360
title: "挑戰 183：2026 冬季遊戲第 4 天：跳台滑雪"
challengeType: 28
dashedName: challenge-183
---

# --description--

給定距離分數、風力補償值和 K 點獎勵值，計算你的跳台滑雪分數並判斷你是否獲得獎牌。

- 你的分數是將上述四個數值相加計算而得。

其他跳躍者目前的總分為：

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- 如果你的分數是最佳，傳回 `"Gold"`
- 如果是第二名，傳回 `"Silver"`
- 如果是第三名，傳回 `"Bronze"`
- 否則，傳回 `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` 應該傳回 `"Gold"`。

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` 應該傳回 `"Bronze"`。

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` 應該傳回 `"Silver"`。

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` 應該傳回 `"No Medal"`。

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` 應該傳回 `"Gold"`。

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` 應該傳回 `"No Medal"`。

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
