---
id: 698a1a73ade5ac0e19180fa7
title: "挑戰 204：字母總和"
challengeType: 28
dashedName: challenge-204
---

# --description--

給定一個字串，傳回其字母的總和。

- 字母為大寫或小寫的 A-Z
- 字母值為：`"A"` = 1、`"B"` = 2、…、`"Z"` = 26
- 大寫和小寫字母具有相同的值。
- 忽略所有非字母字元。

# --hints--

`sumLetters("Hello")` 應該傳回 `52`。

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` 應該傳回 `94`。

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` 應該傳回 `473`。

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` 應該傳回 `1681`。

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` 應該傳回 `0`。

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
