---
id: 698a1a863194f1f4e63f645e
title: "挑戰 207：最小間距"
challengeType: 28
dashedName: challenge-207
---

# --description--

給定一個字串，傳回兩個相同字元之間字元數最少（最小間距）的子字串。

- 總會至少有一組匹配的字元配對。
- 傳回的子字串應該排除匹配的字元。
- 如果兩個或以上的空格長度相同，則傳回第一個空格中的字元。

例如，給定 `"ABCDAC"`，傳回 `"DA"`。

- 只有 `"A"` 和 `"C"` 在字串中重複。
- 兩個 `"A"` 字元之間的字元數是 3，兩個 `"C"` 字元之間的字元數是 2。
- 所以傳回兩個 `"C"` 字元之間的字串。

# --hints--

`smallestGap("ABCDAC")` 應該傳回 `"DA"`。

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` 應該傳回 `"e"`。

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` 應該傳回 `"#q6e&rkf(p"`。

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` 應該傳回 `""`。

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` 應該傳回 `"fox"`。

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
