---
id: 698a1a73ade5ac0e19180fa0
title: "挑戰 197：血型相容性"
challengeType: 28
dashedName: challenge-197
---

# --description--

給定一個捐贈者的血型和一個接受者的血型，判斷捐贈者是否可以捐血給接受者。

每種血型由以下組成：

- 一個字母：`"A"`、`"B"`、`"AB"` 或 `"O"`
- 以及 Rh 因子：`"+"` 或 `"-"`

血型將會是有效字母之一，後面接著 Rh 因子。例如，`"AB+"` 和 `"O-"` 是有效的血型。

字母規則：

- `"O"` 可以捐贈給其他字母型別。
- `"A"` 可以捐血給 `"A"` 和 `"AB"`。
- `"B"` 可以捐贈給 `"B"` 和 `"AB"`。
- `"AB"` 只能捐給 `"AB"`。

Rh 規則：

- 負號 (`"-"`) 可以同時捐贈給 `"-"` 和 `"+"`。
- 正電荷 (`"+"`) 只能捐贈給 `"+"`。

捐贈者必須同時通過字母規則和 Rh 規則，才能向受贈者捐贈。

# --hints--

`canDonate("B+", "B+")` 應該傳回 `true`。

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` 應該傳回 `true`。

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` 應該傳回 `false`。

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` 應該傳回 `true`。

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` 應該傳回 `false`。

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` 應該傳回 `true`。

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` 應該傳回 `false`。

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` 應該傳回 `true`。

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` 應該傳回 `false`。

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` 應該傳回 `false`。

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
