---
id: 698a1a863194f1f4e63f645f
title: "挑戰 208：路徑遍歷"
challengeType: 28
dashedName: challenge-208
---

# --description--

給定一個表述你路線圖的字串陣列，傳回一個到達目標所需移動的字串。

給定的字串將包含以下值：

- `"C"`：你目前的位置
- `"G"`：你的目標
- `"T"`：路徑中可遍歷的零件
- `"-"`：地圖中不可通行的零件

傳回一個字串，內容是從你的位置到目標所需的移動步驟，條件如下：

- `"R"` 是向右移動
- `"D"` 是向下移動
- `"L"` 是向左移動
- `"U"` 是向上移動

- 從你目前的位置到你的目標，將始終有一條單一且連續的路徑，且不會有任何分支。
- 每個路徑位置最多會有兩個可通行的位置相鄰。

例如，給定：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

傳回 `"RDDRDD"`。

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` 應該傳回 `"RDDRDD"`。

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` 應該傳回 `"RRUUURR"`。

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` 應該傳回 `"DLDDRRRRD"`。

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` 應該傳回 `"RRRDDL"`。

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` 應該傳回 `"ULLLUUURRRRRRDDDR"`。

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
