---
id: 698a1a863194f1f4e63f6460
title: "挑戰 209：元素大小"
challengeType: 28
dashedName: challenge-209
---

# --description--

給定一個視窗大小、一個元素以視口寬度單位 `"vw"` 表示的寬度，以及一個元素以視口高度單位 `"vh"` 表示的高度，判斷該元素的大小（以像素為單位）。

- 給定的視窗大小和傳回的元素大小是以 `"width x height"` 格式的字串，例如 `"1200 x 800"`。

- `"vw"` 單位是視窗寬度的百分比。`"50vw"` 例如，是視窗寬度的 50%。
- `"vh"` 單位是視窗高度的百分比。`"50vh"` 例如，是視窗高度的 50%。

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` 應該傳回 `"600 x 400"`。

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` 應該傳回 `"80 x 240"`。

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` 應該傳回 `"70 x 15"`。

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` 應該傳回 `"1824 x 1080"`。

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` 應該傳回 `"0 x 0"`。

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` 應該傳回 `"1440 x 1026"`。

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
