---
id: 698a1a73ade5ac0e19180fa9
title: "挑戰 206：撲克牌點數"
challengeType: 28
dashedName: challenge-206
---

# --description--

給定一個撲克牌的陣列，傳回一個包含每張牌數值的新陣列。

卡牌數值：

- 一張 Ace (`"A"`) 的值為 1。
- 數字牌（`"2"` - `"10"`）的點數即為其面值：分別為 2 - 10。
- 人頭牌：Jack (`"J"`)、Queen (`"Q"`) 和 King (`"K"`) 各值 10。

適用於：

- 每張牌都有一個花色：黑桃（`"S"`）、梅花（`"C"`）、方塊（`"D"`）或紅心（`"H"`）。

卡片格式：

- 每張牌以字串表示：`"valueSuit"`。例如：`"AS"` 是黑桃 A，`"10H"` 是紅心 10，`"QC"` 是梅花 Q。

# --hints--

`cardValues(["3H", "4D", "5S"])` 應該傳回 `[3, 4, 5]`。

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` 應該傳回 `[1, 10, 10, 6, 7]`。

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` 應該傳回 `[8, 10, 2, 10, 9]`。

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` 應該傳回 `[1, 10]`。

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` 應該傳回 `[10, 10, 10, 10, 1]`。

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
