---
id: 698a1a73ade5ac0e19180fa8
title: "挑戰 205：完美立方計數"
challengeType: 28
dashedName: challenge-205
---

# --description--

給定兩個整數，判斷在這兩個數字之間（包含兩端）有多少個完美立方數存在。

- 較低的數字不保證是第一個 `argument`。
- 如果存在一個整數（`n`）使得 `n * n * n = number`，則該數字是完美立方體。 例如，27 是完美立方體，因為 `3 * 3 * 3 = 27`。

# --hints--

`countPerfectCubes(3, 30)` 應該傳回 `2`。

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` 應該傳回 `3`。

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` 應該傳回 `4`。

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` 應該傳回 `9`。

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` 應該傳回 `41`。

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
