---
id: 698a1a73ade5ac0e19180fa3
title: "挑戰 200：字母與數字計數"
challengeType: 28
dashedName: challenge-200
---

# --description--

給定一個字串，傳回一則包含其中有多少字母和數字的訊息。

- 字母是 `A-Z` 和 `a-z`。
- 數字是 `0-9`。
- 忽略所有其他字元。

傳回 `"The string has X letters and Y numbers."`，其中 `"X"` 是字母的數量，`"Y"` 是數字的數量。如果任一數量為 1，則該項目使用單數形式。例如：使用 `"1 letter"` 而非 `"1 letters"`，以及使用 `"1 number"` 而非 `"1 numbers"`。

# --hints--

`countLettersAndNumbers("helloworld123")` 應該傳回 `"The string has 10 letters and 3 numbers."`。

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` 應該傳回 `"The string has 5 letters and 2 numbers."`。

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` 應該傳回 `"The string has 1 letter and 1 number."`。

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` 應該傳回 `"The string has 0 letters and 5 numbers."`。

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` 應該傳回 `"The string has 8 letters and 0 numbers."`。

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
