---
id: 698a1a73ade5ac0e19180fa6
title: "挑戰 203：扁平化"
challengeType: 28
dashedName: challenge-203
---

# --description--

給定一個陣列，判斷它是否為平坦。

- 如果陣列中沒有任何元素是陣列，則該陣列是平坦的。

# --hints--

`isFlat([1, 2, 3, 4])` 應該傳回 `true`。

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` 應該傳回 `false`。

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` 應該傳回 `true`。

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` 應該傳回 `false`。

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` 應該傳回 `false`。

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
