---
id: 698a1a73ade5ac0e19180fa2
title: "挑戰 199：連續差異"
challengeType: 28
dashedName: challenge-199
---

# --description--

給定一個數字陣列，傳回一個新陣列，該陣列包含從每個數字到下一個數字所需的值。

- 對於最後一個數字，使用 `0`，因為沒有下一個數字。

例如，給定 `[1, 2, 4, 7]`，傳回 `[1, 2, 3, 0]`。

# --hints--

`findDifferences([1, 2, 4, 7])` 應該傳回 `[1, 2, 3, 0]`。

```js
assert.deepEqual(findDifferences([1, 2, 4, 7]), [1, 2, 3, 0]);
```

`findDifferences([10, 15, 19, 22, 24, 25])` 應該傳回 `[5, 4, 3, 2, 1, 0]`。

```js
assert.deepEqual(findDifferences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0]);
```

`findDifferences([25, 20, 16, 13, 11, 10])` 應該傳回 `[-5, -4, -3, -2, -1, 0]`。

```js
assert.deepEqual(findDifferences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0]);
```

`findDifferences([0, 1, 2, 2, 3, 3, 4, 5])` 應該傳回 `[ 1, 1, 0, 1, 0, 1, 1, 0 ]`。

```js
assert.deepEqual(findDifferences([0, 1, 2, 2, 3, 3, 4, 5]), [ 1, 1, 0, 1, 0, 1, 1, 0 ]);
```

`findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])` 應該傳回 `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`。

```js
assert.deepEqual(findDifferences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]);
```

# --seed--

## --seed-contents--

```js
function findDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function findDifferences(arr) {
  const result = [];

  for (let i = 0; i < arr.length; i++) {
    result.push(arr[i + 1] - arr[i] || 0);
  }

  return result;
}
```
