---
id: 698a1a73ade5ac0e19180fa1
title: "挑戰 198：商務日計數"
challengeType: 28
dashedName: challenge-198
---

# --description--

給定一個開始日期和結束日期，傳回兩者之間的工作日天數。

- 給定的日期格式為 `"YYYY-MM-DD"`。
- 平日是商務日（週一至週五）。
- 週末不是商務日（星期六和星期日）。
- 計算時包含起始和結束日期。

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` 應該傳回 `3`。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` 應該傳回 `4`。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` 應該傳回 `5`。

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` 應該傳回 `7`。

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` 應該傳回 `262`。

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
