---
id: 6994cff2290543b3aec9f511
title: "挑戰 212：陣列插入"
challengeType: 28
dashedName: challenge-212
---

# --description--

給定一個陣列、一個要插入陣列的值，以及要插入該值的索引，傳回一個在指定索引插入該值的新陣列。

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)` 應該傳回 `[2, 4, 6, 8, 10]`。

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)` 應該傳回 `["the", "quick", "brown", "fox"]`。

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)` 應該傳回 `[0]`。

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)` 應該傳回 `[0, 1, 1, 2, 3, 5, 8, 13]`。

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
