---
id: 6994cff2290543b3aec9f50f
title: "挑戰 210：HSL 驗證器"
challengeType: 28
dashedName: challenge-210
---

# --description--

給定一個字串，判斷它是否為有效的 CSS `hsl()` 顏色值。

- 有效的 HSL 值必須符合格式 `"hsl(h, s%, l%)"`，其中：
  - `h`（色相）必須是介於 0 和 360（含）之間的數字。
  - `s`（飽和度）必須是介於 0% 和 100% 之間的百分比。
  - `l`（亮度）必須是介於 0% 和 100% 之間的百分比。

- 只允許空格：
  - 在開啟括號之後
  - 逗號之前和／或之後
  - 在閉合括號之前和／或之後

- 值可以選用以分號 (`";"`) 結尾。

例如，`"hsl(240, 50%, 50%)"` 是一個有效的 HSL 值。

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` 應該傳回 `true`。

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` 應該傳回 `true`。

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` 應該傳回 `true`。

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` 應該傳回 `true`。

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` 應該傳回 `true`。

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` 應該傳回 `false`。

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` 應該傳回 `false`。

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` 應該傳回 `false`。

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` 應該傳回 `false`。

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
