---
id: 699c8e045ee7cb94ed2322d9
title: "挑戰 218：均勻可整除"
challengeType: 28
dashedName: challenge-218
---

# --description--

給定兩個整數，判斷你是否能將第一個整數整除以第二個整數。

# --hints--

`isEvenlyDivisible(4, 2)` 應該傳回 `true`。

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` 應該傳回 `false`。

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` 應該傳回 `false`。

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` 應該傳回 `true`。

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` 應該傳回 `true`。

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` 應該傳回 `false`。

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
