---
id: 699c8e045ee7cb94ed2322d6
title: "挑戰 215：停車費計算器"
challengeType: 28
dashedName: challenge-215
---

# --description--

給定兩個表示你停車時間和取車時間的字串，計算停車費用。

- 給定的字串將會是使用 24 小時制的 `"HH:MM"` 格式。例如，`"14:00"` 是下午 2 點。
- 第一個字串將是你停車的時間，第二個將是你取車的時間。
- 如果取車時間早於進場時間，表示停車時間跨越午夜進入隔天。

費用規則：

- 每停車 1 小時收費 3 美元。
- 部分小時會進位到下一個整點小時。
- 如果停車時間跨越過夜（過午夜），將額外收取 10 美元過夜費。
- 最低費用為 5 美元（僅在總額低於 5 美元時使用）。

以 `"$cost"` 格式傳回總費用，`"$5"` 例如。

# --hints--

`calculateParkingFee("09:00", "11:00")` 應該傳回 `"$6"`。

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` 應該傳回 `"$5"`。

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` 應該傳回 `"$9"`。

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` 應該傳回 `"$27"`。

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` 應該傳回 `"$34"`。

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` 應該傳回 `"$82"`。

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
