---
id: 699c8e045ee7cb94ed2322da
title: "挑戰 219：週年里程碑"
challengeType: 28
dashedName: challenge-219
---

# --description--

給定一個整數，表示一對夫妻結婚的年數，傳回根據此表格他們最近的周年紀念里程碑：

| Years Married | Milestone    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- 如果他們還沒達到第一個里程碑，傳回 `"Newlyweds"`。

# --hints--

`getMilestone(0)` 應該傳回 `"Newlyweds"`。

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` 應該傳回 `"Paper"`。

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` 應該傳回 `"Wood"`。

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` 應該傳回 `"Tin"`。

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` 應該傳回 `"Silver"`。

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` 應該傳回 `"Ruby"`。

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` 應該傳回 `"Gold"`。

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` 應該傳回 `"Diamond"`。

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` 應該傳回 `"Platinum"`。

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
