---
id: 699c8e045ee7cb94ed2322d4
title: "挑戰 213：字長轉換器"
challengeType: 28
dashedName: challenge-213
---

# --description--

給定一個字串，傳回一個新的字串，其中每個字都被其長度取代。

- 給定字串中的字、單字將以單一空格分隔
- 保持傳回字串中的空格。

例如，給定 `"hello world"`，傳回 `"5 5"`。

# --hints--

`convertWords("hello world")` 應該傳回 `"5 5"`。

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` 應該傳回 `"6 3 5 6"`。

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` 應該傳回 `"3 5 5 3 5 4 3 4 3"`。

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` 應該傳回 `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`。

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```
