---
id: 699c8e045ee7cb94ed2322db
title: "挑戰 220：最大數字"
challengeType: 28
dashedName: challenge-220
---

# --description--

給定一個由各種標點符號分隔的數字字串，傳回最大的數字。

- 給定的字串將只包含數字和分隔符。
- 分隔符可以是逗號（`","`）、驚嘆號（`"!"`）、問號（`"?"`）、冒號（`":"`）或分號（`";"`）。

# --hints--

`largestNumber("1,2")` 應該傳回 `2`。

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` 應該傳回 `60`。

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` 應該傳回 `-402`。

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` 應該傳回 `99.9`。

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
