---
id: 69a890af247de743333bd4cc
title: "挑戰 223：QR 解碼器"
challengeType: 28
dashedName: challenge-223
---

# --description--

給定一個 6x6 矩陣（陣列的陣列），表述一個 QR code，傳回程式碼中的二進位資料字串。

- QR 碼可以以任意 90 度增量旋轉的角度給出。
- 正確定向的程式碼在左下、左上和右上角有一個 2x2 的 `1`（定向標記）群組。
- 這三個 2x2 方向標記不是二進位資料的零件。
- 當 `QR code` 正確定位時，二進位資料會從左至右、從上至下（像閱讀書本一樣）讀取。
- 一個程式碼將永遠只有一個有效的方向。

例如，給定：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

或者給定相同的 `code` 但方向不同：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

傳回 `"000000000000000000000001"`，所有二進位資料不包含三個 2x2 定向標記。

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` 應該傳回 `"000000000000000000000001"`。

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` 應該傳回 `"000000000000000000000001"`。

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` 應該傳回 `"001101000011000000110100"`。

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` 應該傳回 `"010001000100010101010110"`。

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` 應該傳回 `"010000100100100101001110"`。

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
