---
id: 69a890af247de743333bd4ce
title: "挑戰 225：禁止連續重複"
challengeType: 28
dashedName: challenge-225
---

# --description--

給定一個字串，判斷它是否沒有重複的字元。

- 如果一個字串沒有相同的字元連續出現兩次或以上，則該字串沒有重複。

# --hints--

`hasNoRepeats("hi world")` 應該傳回 `true`。

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` 應該傳回 `false`。

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` 應該傳回 `true`。

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` 應該傳回 `false`。

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` 應該傳回 `true`。

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` 應該傳回 `false`。

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
