---
id: 69a890af247de743333bd4d1
title: "挑戰 228：電影之夜"
challengeType: 28
dashedName: challenge-228
---

# --description--

給定一個表示星期幾的字串、一個表示場次時間的字串，以及一個整數票數，傳回該場次電影票的總費用。

給定的日期將會是以下之一：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

演出時間將以 `"H:MMam"` 或 `"H:MMpm"` 格式給出。例如 `"10:00am"` 或 `"10:00pm"`。

以 `"$D.CC"` 格式傳回總成本，並使用以下規則：

- 週末（星期五－星期日）：每張票 $12.00。
- 平日（星期一至星期四）：每張票 10.00 美元。
- Matinee（下午 5:00 前）：每張票減 2.00 美元（星期二除外）。
- 星期二：所有票券每張 $5.00。

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` 應該傳回 `"$12.00"`。

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` 應該傳回 `"$10.00"`。

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` 應該傳回 `"$10.00"`。

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` 應該傳回 `"$30.00"`。

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` 應該傳回 `"$32.00"`。

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` 應該傳回 `"$50.00"`。

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` 應該傳回 `"$5.00"`。

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
