---
id: 69a890af247de743333bd4d0
title: "挑戰 227：冷卻時間"
challengeType: 28
dashedName: challenge-227
---

# --description--

給定兩個時間戳記，第一個表述使用者完成考試的時間，第二個表述當前時間，判斷使用者是否可以再次參加考試。

- 兩個時間戳記都將使用格式：`"YYYY-MM-DDTHH:MM:SS"`，例如 `"2026-03-25T14:00:00"`。請注意時間為 24 小時制。
- 使用者必須等待至少 48 小時才能重新參加考試。

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` 應該傳回 `true`。

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` 應該傳回 `false`。

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` 應該傳回 `true`。

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` 應該傳回 `false`。

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
