---
id: 69a890af247de743333bd4d2
title: "挑戰 229：截斷 `text` 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

給定一個字串，傳回一個新的字串，使字元的總寬度不超過 50 單位。

每個字元都有特定的寬度：

| Letters | Width |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

上表包含所有大寫和小寫字母。此外：

- 空格 (`" "`) 的寬度為 2
- 句號（`"."`）的寬度為 1

- 如果給定的字串長度在 50 個單位以內，則原樣傳回該字串，否則
- 將字串截斷並在末尾添加三個句點（`"..."`），使其總寬度（包括三個句點）盡可能接近 60 單位且不超過。

# --hints--

`truncateText("The quick brown fox")` 應該傳回 `"The quick brown f..."`。

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` 應該傳回 `"The silky smooth s..."`。

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` 應該傳回 `"THE LOUD BRIG..."`。

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` 應該傳回 `"The fast striped z..."`。

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` 應該傳回 `"The big black bear"`。

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
