---
id: 69b1028d6e265413d0198a2a
title: "挑戰 231：ISBN-10 驗證器"
challengeType: 28
dashedName: challenge-231
---

# --description--

給定一個字串，判斷它是否為有效的 ISBN-10。

ISBN-10 由連字號（`"-"`）和其他 10 個字元組成。移除連字號後（`"-"`）：

- 前 9 個字元必須是數字，且
- 最後一個字元可以是數字或字母 `"X"`，代表數字 10。

驗證方法：

- 將每個數字（或數值）乘以它的位置（第一個數字乘以 1，第二個乘以 2，以此類推）。
- 將所有結果相加。
- 如果總和能被 11 整除，則為有效。

# --hints--

`isValidIsbn10("0-306-40615-2")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` 應該傳回 `false`。

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` 應該傳回 `false`。

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
