---
id: 69b1028d6e265413d0198a29
title: "挑戰 230：帕斯卡三角形的列"
challengeType: 28
dashedName: challenge-230
---

# --description--

給定一個整數 `n`，傳回帕斯卡三角形的第 `n` 列，並以陣列形式呈現。

在帕斯卡三角形中，每一列的開頭和結尾都是 1，而中間的每個數值都是正上方兩個數值的和。

以下是三角形的前 5 列：

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` 應該傳回 `[1, 4, 6, 4, 1]`。

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` 應該傳回 `[1, 2, 1]`。

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` 應該傳回 `[1]`。

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` 應該傳回 `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`。

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` 應該傳回 `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`。

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
