---
id: 69b1028d6e265413d0198a2c
title: "挑戰 233：起床鬧鐘"
challengeType: 28
dashedName: challenge-233
---

# --description--

給定一個字串，表示你設定鬧鐘的時間，以及另一個字串，表示你實際起床的時間，判斷你是早起、準時起床，還是遲到。

- 兩個時間都會以 `"HH:MM"` 24 小時制格式給出。

傳回：

- `"early"` 如果你在鬧鐘時間之前起床。
- `"on time"` 如果你在鬧鐘時間起床，或是在鬧鐘時間後 10 分鐘的延遲時間內起床。
- `"late"` 如果你在鬧鐘時間後超過 10 分鐘才起床。

兩個時間都在同一天。

# --hints--

`alarmCheck("07:00", "06:45")` 應該傳回 `"early"`。

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` 應該傳回 `"on time"`。

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` 應該傳回 `"on time"`。

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` 應該傳回 `"late"`。

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` 應該傳回 `"on time"`。

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` 應該傳回 `"late"`。

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` 應該傳回 `"early"`。

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
