---
id: 69b1028d6e265413d0198a2b
title: "挑戰 232：到期日"
challengeType: 28
dashedName: challenge-232
---

# --description--

給定一個日期字串，傳回未來 9 個月的日期。

- 給定的字串和傳回的字串格式為 `"YYYY-MM-DD"`。
- 如果未來 9 個月的月份中沒有原本的日期數字，請傳回該月份的最後一天。

# --hints--

`getDueDate("2025-03-30")` 應該傳回 `"2025-12-30"`。

```js
assert.equal(getDueDate("2025-03-30"), "2025-12-30");
```

`getDueDate("2025-04-27")` 應該傳回 `"2026-01-27"`。

```js
assert.equal(getDueDate("2025-04-27"), "2026-01-27");
```

`getDueDate("2025-05-29")` 應該傳回 `"2026-02-28"`。

```js
assert.equal(getDueDate("2025-05-29"), "2026-02-28");
```

`getDueDate("2026-06-30")` 應該傳回 `"2027-03-30"`。

```js
assert.equal(getDueDate("2026-06-30"), "2027-03-30");
```

`getDueDate("2026-10-11")` 應該傳回 `"2027-07-11"`。

```js
assert.equal(getDueDate("2026-10-11"), "2027-07-11");
```

# --seed--

## --seed-contents--

```js
function getDueDate(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getDueDate(dateStr) {
  let [year, month, day] = dateStr.split("-").map(Number);

  month += 9;

  year += Math.floor((month - 1) / 12);
  month = ((month - 1) % 12) + 1;

  const daysInMonth = new Date(year, month, 0).getDate();

  if (day > daysInMonth) {
    day = daysInMonth;
  }

  const mm = String(month).padStart(2, "0");
  const dd = String(day).padStart(2, "0");

  return `${year}-${mm}-${dd}`;
}
```
