---
id: 69b1028d6e265413d0198a2e
title: "挑戰 235：大寫費波那契"
challengeType: 28
dashedName: challenge-235
---

# --description--

給定一個字串，傳回一個新字串，其中每個字元如果其索引是費波那契數字，則轉成大寫，否則轉成小寫。

費波那契數列是一系列數字，其中每個數字都是前兩個數字的和。數列的前 10 個數字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

- 第一個字元的索引是 `0`。
- 如果非字母字元的索引是費波那契數字，則保持不變。

# --hints--

`capitalizeFibonacci("hello world")` 應該傳回 `"HELLo woRld"`。

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` 應該傳回 `"HELLo woRld"`。

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` 應該傳回 `"HELLo, wOrld!"`。

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` 應該傳回 `"THE qUicK broWn fox jUmped over thE lazy dog."`。

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` 應該傳回 `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."`。

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```
