---
id: 69b559d2903b9e4afe9075f8
title: "挑戰 238：數字旋轉逃脫"
challengeType: 28
dashedName: challenge-238
---

# --description--

給定一個正整數，判斷它本身或其任一旋轉後的數字，是否能被其位數整除。

旋轉指的是將第一個數字移到最後面。例如，旋轉一次後，123 會變成 231。

- 先檢查旋轉 `0`（給定的數字）。
- 給定的數字不會包含任何零。
- 如果找到符合條件的旋轉數字，傳回該數字；否則傳回 `"none"`。

# --hints--

`getRotation(123)` 應該傳回 `0`。

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` 應該傳回 `3`。

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` 應該傳回 `"none"`。

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` 應該傳回 `6`。

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
