---
id: 69b5b2be76ec8135a7fbe973
title: "挑戰 244：車與象的攻擊"
challengeType: 28
dashedName: challenge-244
---

# --description--

給定一個字串表示車在棋盤上的位置，另一個字串表示象的位置，判斷其中一個棋子是否能攻擊另一個。

標準棋盤是 8x8，行標籤從 `A` 到 `H`（從左到右），列標籤從 `1` 到 `8`（從下到上）。棋盤長這樣：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- 車可以在水平方向或垂直方向移動任意格數。
- 象可以在任意斜角方向移動任意格數。
- 如果一個棋子能移動到另一個棋子的位置，則該棋子可以攻擊另一個。

傳回：

- `"rook"` 如果車能攻擊象。
- `"bishop"` 如果象能攻擊車。
- `"neither"` 如果兩個棋子都無法攻擊對方。

# --hints--

`rookBishopAttack("A1", "A5")` 應該傳回 `"rook"`。

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` 應該傳回 `"bishop"`。

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` 應該傳回 `"rook"`。

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` 應該傳回 `"bishop"`。

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` 應該傳回 `"neither"`。

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` 應該傳回 `"neither"`。

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
