---
id: 69b58ce40693f140c84c855b
title: "挑戰 242：下一個賓果號碼"
challengeType: 28
dashedName: challenge-242
---

# --description--

給定一個賓果號碼，傳回下一個依序排列的賓果號碼。

賓果號碼由一個字母接著一個數字組成，數字的範圍依照下表：

| 字母 | 數字範圍 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

例如，給定 `"B10"`，傳回下一個賓果號碼 `"B11"`。如果給定的是最後一個賓果號碼，則傳回 `"B1"`。

# --hints--

`getNextBingoNumber("B10")` 應該傳回 `"B11"`。

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` 應該傳回 `"N34"`。

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` 應該傳回 `"N31"`。

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` 應該傳回 `"O61"`。

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` 應該傳回 `"B1"`。

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
