---
id: 69b58ce40693f140c84c855c
title: "挑戰 243：車的攻擊"
challengeType: 28
dashedName: challenge-243
---

# --description--

給定兩個字串，代表棋盤上兩個車的位置，判斷它們是否能互相攻擊。

標準棋盤是 8x8，行標籤從 `A` 到 `H`（由左至右），列標籤從 `1` 到 `8`（由下至上）。棋盤看起來像這樣：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

車可以沿著水平或垂直方向移動任意多格。因此，如果它們位於同一列或同一行，就能互相攻擊。

# --hints--

`rookAttack("A1", "A8")` 應該傳回 `true`。

```js
assert.isTrue(rookAttack("A1", "A8"));
```

`rookAttack("B4", "F4")` 應該傳回 `true`。

```js
assert.isTrue(rookAttack("B4", "F4"));
```

`rookAttack("E3", "D4")` 應該傳回 `false`。

```js
assert.isFalse(rookAttack("E3", "D4"));
```

`rookAttack("H7", "F6")` 應該傳回 `false`。

```js
assert.isFalse(rookAttack("H7", "F6"));
```

# --seed--

## --seed-contents--

```js
function rookAttack(rook1, rook2) {

  return rook1;
}
```

# --solutions--

```js
function rookAttack(rook1, rook2) {
  return rook1[0] === rook2[0] || rook1[1] === rook2[1];
}
```
