---
id: 69b559d2903b9e4afe9075f9
title: "挑戰 239：今天星期幾？"
challengeType: 28
dashedName: challenge-239
---

# --description--

給定一個以毫秒為單位的 Unix 時間戳，傳回該日期是星期幾。

有效的傳回星期為：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

請務必忽略時區。

# --hints--

`getDayOfWeek(1775492249000)` 應該傳回 `"Monday"`。

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` 應該傳回 `"Saturday"`。

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` 應該傳回 `"Tuesday"`。

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` 應該傳回 `"Sunday"`。

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` 應該傳回 `"Thursday"`。

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
