---
id: 69b58ce40693f140c84c8559
title: "挑戰 240：回文字元"
challengeType: 28
dashedName: challenge-240
---

# --description--

給定一個字串，判斷它是否為回文，並傳回中間的字元（如果長度為奇數）或中間兩個字元（如果長度為偶數）。

- 回文是指正著讀和反著讀都相同的字串。
- 如果不是回文，傳回 `"none"`。

# --hints--

`palindromeLocator("racecar")` 應該傳回 `"e"`。

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` 應該傳回 `"v"`。

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` 應該傳回 `"none"`。

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` 應該傳回 `"oo"`。

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` 應該傳回 `"00"`。

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
