---
id: 69b58ce40693f140c84c855a
title: "挑戰 241：FizzBuzz 驗證器"
challengeType: 28
dashedName: challenge-241
---

# --description--

給定一個連續整數的陣列，其中 3 和 5 的倍數已被替換，判斷它是否為有效的 FizzBuzz 序列。

在有效的 FizzBuzz 序列中：

- 3 的倍數會被替換成「Fizz」。
- 5 的倍數會被替換成「Buzz」。
- 同時是 3 和 5 的倍數會被替換成「FizzBuzz」。
- 其他所有數字保持為整數。

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` 應該傳回 `true`。

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` 應該傳回 `true`。

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` 應該傳回 `false`。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` 應該傳回 `true`。

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` 應該傳回 `false`。

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` 應該傳回 `false`。

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` 應該傳回 `true`。

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
