---
id: 69bc6cb30c1d112a2e110a08
title: "挑戰 251：陣列加總尋找器"
challengeType: 28
dashedName: challenge-251
---

# --description--

給定一個數字陣列和一個標的數字，傳回第一組兩個或以上數字的子集合，其加總等於該標的。

- 「第一」組子集合是指其元素擁有最低可能索引值的集合，優先考慮最早的索引。
- 陣列中的每個數字只能使用一次。
- 若不存在有效的子集合，傳回 `"Sum not found"`。

傳回符合條件的數字陣列，順序依照它們在原始陣列中出現的順序。

# --hints--

`findSum([1, 3, 5, 7], 6)` 應該傳回 `[1, 5]`。

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` 應該傳回 `[1, 4]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` 應該傳回 `[1, 2, 3]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` 應該傳回 `[-1, -2, 4]`。

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` 應該傳回 `[3, 1, 4, 2]`。

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` 應該傳回 `[1, 2, 3, 5, 9]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` 應該傳回 `"Sum not found"`。

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
