---
id: 69bc6cb30c1d112a2e110a05
title: "挑戰 248：排序陣列交換"
challengeType: 28
dashedName: challenge-248
---

# --description--

給定一個整數陣列，請依照以下規則傳回一個新的陣列：

1. 將整數依照遞增順序排序
2. 接著將所有索引為 3 的倍數的值與它前一個值交換

# --hints--

`sortAndSwap([3, 1, 2, 4, 6, 5])` 應該傳回 `[1, 2, 4, 3, 5, 6]`。

```js
assert.deepEqual(sortAndSwap([3, 1, 2, 4, 6, 5]), [1, 2, 4, 3, 5, 6]);
```

`sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8])` 應該傳回 `[1, 2, 4, 3, 5, 7, 6, 8, 9]`。

```js
assert.deepEqual(sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9])` 應該傳回 `[1, 2, 4, 3, 5, 7, 6, 8, 9]`。

```js
assert.deepEqual(sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11])` 應該傳回 `[1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]`。

```js
assert.deepEqual(sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11]), [1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]);
```

`sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25])` 應該傳回 `[-75, -50, 0, -25, 25, 75, 50, 100]`。

```js
assert.deepEqual(sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25]), [-75, -50, 0, -25, 25, 75, 50, 100]);
```

`sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2])` 應該傳回 `[-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]`。

```js
assert.deepEqual(sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2]), [-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]);
```

# --seed--

## --seed-contents--

```js
function sortAndSwap(arr) {

  return arr;
}
```

# --solutions--

```js
function sortAndSwap(arr) {
  const result = [...arr].sort((a, b) => a - b);

  for (let i = 3; i < result.length; i += 3) {
    [result[i], result[i - 1]] = [result[i - 1], result[i]];
  }

  return result;
}
```
