---
id: 69bc6cb30c1d112a2e110a07
title: "挑戰 250：隱藏的金鑰"
challengeType: 28
dashedName: challenge-250
---

# --description--

歡迎來到第 250 個每日挑戰！

給定一個加密字串，請使用加密金鑰將其解碼並傳回。

尋找金鑰的方法：

- 查看截至今天所有挑戰編號為 25 的倍數的每日挑戰（包含本次）。
- 從這些挑戰標題中取出第一個字母並組合成一個字串。如果標題以非字母字元開頭，則找到它的第一個字母。

要解碼訊息，請逐字檢視加密訊息中的每個字母，並：

  1. 查看金鑰中對應位置的字母（如果訊息比金鑰長，則重複使用金鑰）。
  2. 將金鑰字母轉換成對應的數字：`"A"` = 1，`"B"` = 2，…，`"Z"` = 26。
  3. 將加密字母在字母表中向後移動該數字的位數。
  4. 如果移動超過 `"A"`，則從 `"Z"` 繞回。

例如，如果加密訊息以 `"Y"` 開頭，且第一個金鑰字母是 `"V"`（22），則將 `"Y"` 向後移動 22 位得到 `"C"`。對每個字母重複此過程以解碼完整訊息。

- 只有字母會被移動，空格保持不變。
- 所有給定和傳回的字母皆為大寫。

# --hints--

`decode("YAVJYNXE")` 應該傳回 `"CONGRATS"`。

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` 應該傳回 `"CODING LEGEND"`。

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` 應該傳回 `"YOU ARE AWESOME"`。

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` 應該傳回 `"KEEP GOING"`。

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` 應該傳回 `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`。

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
