---
id: 69bc6cb30c1d112a2e110a09
title: "挑戰 252：獨特的爬樓梯方法"
challengeType: 28
dashedName: challenge-252
---

# --description--

給定樓梯的階數，傳回有人每次可以跨 1 或 2 階時，有多少種不同的爬法。

# --hints--

`getUniqueClimbs(4)` 應該傳回 `5`。

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` 應該傳回 `8`。

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` 應該傳回 `89`。

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` 應該傳回 `4181`。

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` 應該傳回 `832040`。

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` 應該傳回 `20365011074`。

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
