---
id: 69bc6cb30c1d112a2e110a06
title: "挑戰 249：字串數學"
challengeType: 28
dashedName: challenge-249
---

# --description--

給定一個包含數字和其他字元的字串，根據數字之間非數字字元的數量對數字進行數學運算。

- 如果分隔兩個數字的字元數是偶數，則使用加法。
- 如果是奇數，則使用減法。
- 連續的數字組成一個整數。
- 運算從左到右依序進行。
- 忽略開頭和結尾非數字的字元。

例如，給定 `"3ab10c8"`，傳回 `5`。將 3 和 10 相加得到 13，因為它們之間有偶數個字元。接著從 13 減去 8，因為結果和 8 之間有奇數個字元。

# --hints--

`doMath("3ab10c8")` 應該傳回 `5`。

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` 應該傳回 `2`。

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` 應該傳回 `12`。

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` 應該傳回 `15`。

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` 應該傳回 `67`。

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
