---
id: 69bc6cb30c1d112a2e110a04
title: "挑戰 247：最後一個字母"
challengeType: 28
dashedName: challenge-247
---

# --description--

給定一個字串，傳回該字串中字母表中最後出現的字母。

- 如果有兩個或以上的字母並列為字母表中最後，傳回第一個出現的字母。
- 忽略所有非字母字元。

# --hints--

`getLastLetter("world")` 應該傳回 `"w"`。

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` 應該傳回 `"W"`。

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` 應該傳回 `"z"`。

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` 應該傳回 `"L"`。

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` 應該傳回 `"T"`。

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
