---
id: 69bc6cb30c1d112a2e110a03
title: "挑戰 246：姓名縮寫"
challengeType: 28
dashedName: challenge-246
---

# --description--

給定一個完整姓名的字串，傳回其縮寫。

- 要縮寫的名字以空格分隔。
- 縮寫應該使用大寫字母。
- 縮寫之間應以點號分隔。

例如，`"Tommy Millwood"` 傳回 `"T.M."`。

# --hints--

`getInitials("Tommy Millwood")` 應該傳回 `"T.M."`。

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` 應該傳回 `"S.P."`。

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` 應該傳回 `"F.C.C."`。

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` 應該傳回 `"D."`。

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` 應該傳回 `"D.V.C.H.N."`。

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
