---
id: 69c5f3d787b1725d5f00c8ba
title: "挑戰 256：最接近的時間方向"
challengeType: 28
dashedName: challenge-256
---

# --description--

給定兩個時間，判斷從第一個時間到第二個時間，向前走還是向後走比較快。

- 時間以 24 小時制表示（`"HH:MM"`）
- 時鐘會環繞（向前走時，23:59 會到 00:00，向後走時，00:00 會到 23:59）

傳回：

- 如果向前走較短，傳回 `"forward"`
- 如果向後走較短，傳回 `"backward"`
- 如果兩個方向所需時間相同，傳回 `"equal"`

# --hints--

`getDirection("10:00", "12:00")` 應該傳回 `"forward"`。

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` 應該傳回 `"backward"`。

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` 應該傳回 `"equal"`。

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` 應該傳回 `"forward"`。

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` 應該傳回 `"backward"`。

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` 應該傳回 `"equal"`。

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
