---
id: 69c5f3d787b1725d5f00c8bd
title: "挑戰 259：FizzBuzz 爆炸"
challengeType: 28
dashedName: challenge-259
---

# --description--

給定一個整數，傳回將字 `"fizzbuzz"` 轉換成包含至少該數量 `"z"` 的字串所需的步驟數，規則如下：

- 從字串 `"fizzbuzz"` 開始。
- 每一步，根據字串中字母的位置套用標準 FizzBuzz 規則（第一個 `"f"` 的位置為 1）。
  - 如果字母位置能被 3 整除，將該字母取代為 `"fizz"`
  - 如果能被 5 整除，將該字母取代為 `"buzz"`
  - 如果能同時被 3 和 5 整除，將該字母取代為 `"fizzbuzz"`

因此，經過 1 步後，`"fizzbuzz"` 會變成 `"fifizzzbuzzfizzzz"`，其中包含 9 個 `"z"`。

# --hints--

`explodeFizzbuzz(9)` 應該傳回 `1`。

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` 應該傳回 `2`。

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` 應該傳回 `3`。

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` 應該傳回 `4`。

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` 應該傳回 `5`。

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` 應該傳回 `6`。

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` 應該傳回 `11`。

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` 應該傳回 `14`。

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
