---
id: 69c5f3d787b1725d5f00c8b9
title: "挑戰 255：地球日清理隊"
challengeType: 28
dashedName: challenge-255
---

# --description--

今天是地球日。給定一個你清理過的項目陣列，根據以下規則傳回你的總清理分數。

給定的項目將是以下之一：

| 項目 | 基本分數 |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

稀有項目以 `["rare", value]` 表示。舉例來說，`["rare", 80]`。稀有項目不會獲得連續加分。

- 連續加分：如果相同項目連續出現，會獲得遞增的加分。
  - 第一次連續出現：基本分數
  - 第二次：基本分數 + 1
  - 第三次：基本分數 + 2
  - 依此類推。
- 第五項目倍數：每收集到第 5 項目會有倍數加成。
  - 第五項目：乘以 2
  - 第十項目：乘以 3
  - 依此類推。

- 先計算任何加分後，再套用倍數。

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` 應該傳回 `44`。

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` 應該傳回 `58`。

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` 應該傳回 `79`。

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` 應該傳回 `358`。

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` 應該傳回 `383`。

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
