---
id: 69c5f3d787b1725d5f00c8b8
title: "挑戰 254：奇數字的單字"
challengeType: 28
dashedName: challenge-254
---

# --description--

給定一個由多個單字組成的字串，傳回只包含字母數為奇數的單字。

- 給定字串中的單字會以單一空格分隔。
- 傳回的單字也以單一空格分隔。

# --hints--

`getOddWords("This is a super good test")` 應該傳回 `"a super"`。

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` 應該傳回 `"one two three"`。

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` 應該傳回 `"split rainbow sprinkles top"`。

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` 應該傳回 `"The quick brown fox the river"`。

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
