---
id: 69c6ff713a52713463aa7927
title: "挑戰 261：數字單字"
challengeType: 28
dashedName: challenge-261
---

# --description--

給定一個介於 0 到 99 的整數，傳回它的英文單字表述。

- `0` 傳回 `"zero"`。
- 數字 1 到 19 有獨特的名稱（`"one"`、`"two"`、…、`"ten"`、`"eleven"`、…、`"eighteen"`、`"nineteen"`）。
- 20 到 90 的 10 的倍數有自己的名稱（`"twenty"`、`"thirty"`、…、`"eighty"`、`"ninety"`）。
- 21 到 99 中非 10 的倍數的數字，寫成兩個單字並用連字號連接。例如 `"forty-two"` 和 `"fifty-three"`。

# --hints--

`getNumberWords(0)` 應該傳回 `"zero"`。

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` 應該傳回 `"ten"`。

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` 應該傳回 `"nineteen"`。

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` 應該傳回 `"thirty"`。

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` 應該傳回 `"fifty-three"`。

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` 應該傳回 `"seven"`。

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` 應該傳回 `"twelve"`。

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` 應該傳回 `"sixty"`。

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` 應該傳回 `"sixty-seven"`。

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` 應該傳回 `"ninety-eight"`。

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
