---
id: 69c6ff713a52713463aa7926
title: "挑戰 260：單字分數"
challengeType: 28
dashedName: challenge-260
---

# --description--

給定一個單字，使用標準字母分值表傳回它的分數：

| Letter | Value |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- 大寫和小寫字母的分值相同。

# --hints--

`getWordScore("hi")` 應該傳回 `17`。

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` 應該傳回 `52`。

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` 應該傳回 `169`。

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` 應該傳回 `94`。

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
