---
id: 69c6ff713a52713463aa7929
title: "挑戰 263：二進位填字遊戲"
challengeType: 28
dashedName: challenge-263
---

# --description--

給定一個字元，判斷它的 8 位元二進位表述是否能在下方的格子中找到，無論是水平或垂直方向，且方向可為任一方向：

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

例如，`"A"` 的二進位表述是 `01000001`，它出現在第一列從左到右的方向。

# --hints--

`isInCrossword("I")` 應該傳回 `true`。

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` 應該傳回 `true`。

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` 應該傳回 `true`。

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` 應該傳回 `true`。

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` 應該傳回 `false`。

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` 應該傳回 `false`。

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` 應該傳回 `false`。

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` 應該傳回 `false`。

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
