---
id: 69cfca90e8a0a6d4d6871c51
title: "挑戰 267：Parsec 轉換器"
challengeType: 28
dashedName: challenge-267
---

# --description--

在遙遠的銀河系中，parsec 用來測量時間和距離。給定一個整數 parsec，傳回其等價的時間或距離。

- 如果給定的整數是奇數，代表時間。如果是偶數，代表距離。

使用以下轉換率：

| Parsecs | 時間／距離 |
| - | - |
| 1 | 2 小時 |
| 2 | 6 光年 |

傳回轉換後的整數值。

# --hints--

`convertParsecs(1)` 應該傳回 `2`。

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` 應該傳回 `6`。

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` 應該傳回 `62`。

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` 應該傳回 `264`。

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` 應該傳回 `34`。

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` 應該傳回 `42`。

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
