---
id: 69cfca90e8a0a6d4d6871c52
title: "挑戰 268：自戀數字"
challengeType: 28
dashedName: challenge-268
---

# --description--

給定一個正整數，判斷它是否為自戀數字。

- 如果一個數字的每個位數的數字，分別以位數總數為冪次方後相加，結果等於該數字本身，則該數字為自戀數字。

例如，153 有 3 位數，且 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153，因此它是自戀數字。

# --hints--

`isNarcissistic(153)` 應該傳回 `true`。

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` 應該傳回 `false`。

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` 應該傳回 `true`。

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` 應該傳回 `false`。

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` 應該傳回 `true`。

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` 應該傳回 `false`。

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` 應該傳回 `true`。

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` 應該傳回 `false`。

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
