---
id: 69cfca90e8a0a6d4d6871c4e
title: "挑戰 264：字謎群組"
challengeType: 28
dashedName: challenge-264
---

# --description--

給定一個字陣列，傳回一個二維陣列，將字依字謎分組。

- 如果字包含相同字母且順序可不同，則它們是字謎。
- 每個字只屬於一個群組。
- 傳回的順序不重要。

例如，給定 `["listen", "silent", "hello", "enlist", "world"]`，傳回 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`。

# --hints--

`groupAnagrams(["listen", "silent", "hello", "enlist", "world"])` 應該傳回 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`。

```js
const groups = groupAnagrams(["listen", "silent", "hello", "enlist", "world"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["enlist", "listen", "silent"], ["hello"], ["world"]]);
```

`groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` 應該傳回 `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`。

```js
const groups = groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]);
```

`groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` 應該傳回 `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`。

```js
const groups = groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]);

```

`groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` 應該傳回 `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`。

```js
const groups = groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]);
```

# --seed--

## --seed-contents--

```js
function groupAnagrams(words) {

  return words;
}
```

# --solutions--

```js
function groupAnagrams(words) {
  const map = {};
  for (const word of words) {
    const key = word.split('').sort().join('');
    if (!map[key]) map[key] = [];
    map[key].push(word);
  }
  return Object.values(map);
}
```
