---
id: 69cfca90e8a0a6d4d6871c4f
title: "挑戰 265：最深層的括號"
challengeType: 28
dashedName: challenge-265
---

# --description--

給定一個包含平衡括號的字串，傳回最深層巢狀括號的內容。

- 括號可以是三種型別中的任一種：`()`、`[]` 和 `{}`。
- 輸入字串中一定只有一個最深層的群組。

例如，給定 `"(hello (world))"`，傳回 `"world"`。

# --hints--

`getDeepestBrackets("(hello (world))")` 應該傳回 `"world"`。

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` 應該傳回 `"inner"`。

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` 應該傳回 `"e"`。

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` 應該傳回 `"fox"`。

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` 應該傳回 `"C"`。

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
