---
id: 69cfca90e8a0a6d4d6871c50
title: "挑戰 266：美好的一天"
challengeType: 28
dashedName: challenge-266
---

# --description--

給定一個 `"HH:MM"` 格式（24 小時制）的時間字串，傳回：

- `"Good morning"`，時間介於 `05:00` 到 `11:59` 之間
- `"Good afternoon"`，時間介於 `12:00` 到 `17:59` 之間
- `"Good evening"`，時間介於 `18:00` 到 `21:59` 之間
- `"Good night"`，時間介於 `22:00` 到 `04:59` 之間

# --hints--

`getGreeting("06:30")` 應該傳回 `"Good morning"`。

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` 應該傳回 `"Good afternoon"`。

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` 應該傳回 `"Good evening"`。

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` 應該傳回 `"Good night"`。

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` 應該傳回 `"Good morning"`。

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
