---
id: 69e2383af7832c8032603b92
title: "挑戰 276：錯置元素"
challengeType: 28
dashedName: challenge-276
---

# --description--

給定一個整數陣列，該陣列除了有一個元素位置錯置外，其餘元素皆依升冪排序，請傳回該錯置元素的索引。

- 如果有多個元素可能被視為錯置，請傳回第一個錯置元素的索引。

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` 應該傳回 `1`。

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` 應該傳回 `3`。

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` 應該傳回 `0`。

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` 應該傳回 `2`。

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` 應該傳回 `6`。

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
