---
id: 69e2383af7832c8032603b93
title: "挑戰 277：鏡像"
challengeType: 28
dashedName: challenge-277
---

# --description--

給定兩個字串，判斷第二個字串是否為第一個字串的鏡像。

鏡像是透過反轉字串並將每個字元替換成其鏡像等價字元所形成。

- 對稱字元在鏡中看起來和自己一樣：

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, 和空白 (` `)。

- 鏡像配對字元會互相交換：

| 字元 | 互換字元 |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

如果任一字串包含不在上述列表中的字元，則無法從這些字元產生鏡像。

例如，`"[HOW]"` 的鏡像是 `"[WOH]"`。

# --hints--

`isMirrorImage("[HOW]", "[WOH]")` 應該傳回 `true`。

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")` 應該傳回 `true`。

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")` 應該傳回 `true`。

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")` 應該傳回 `false`。

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")` 應該傳回 `false`。

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")` 應該傳回 `false`。

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")` 應該傳回 `true`。

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` 應該傳回 `true`。

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
