---
id: 69e2383af7832c8032603b94
title: "挑戰 278：咖啡訂單解析器"
challengeType: 28
dashedName: challenge-278
---

# --description--

給定一個咖啡訂單的字串，識別任何選單項目並傳回格式化的訂單。

使用以下選單項目和價格：

| 項目、條款 | 價格 | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

傳回一個字串，將符合的項目用 `" + "` 連接，後面接冒號和空格（`": "`），以及總價格。

例如，給定 `"I'd like an oat latte with vanilla syrup and an extra shot please."`，傳回 `"oat latte + vanilla syrup + extra shot: $6.25"`

項目應該依照它們在選單中出現的順序排列，且總價格應始終保留兩位小數。

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` 應該傳回 `"oat latte + vanilla syrup + extra shot: $6.25"`。

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` 應該傳回 `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`。

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` 應該傳回 `"cold brew + extra shot + cream: $5.75"`。

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` 應該傳回 `"espresso: $3.00"`。

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` 應該傳回 `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`。

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
