---
id: 69e2383af7832c8032603b8f
title: "挑戰 273：ISBN-13 驗證器"
challengeType: 28
dashedName: challenge-273
---

# --description--

給定一個字串，判斷它是否為有效的 ISBN-13 編號。

有效的 ISBN-13：

- 僅包含數字和連字號
- 移除連字號後恰好有 13 個數字
- 通過以下檢查：
  1. 對每個數字交替乘以 1 或 3（第一個數字乘以 1，第二個乘以 3，第三個乘以 1，依此類推）。
  2. 結果的總和必須能被 10 整除。

# --hints--

`isValidIsbn13("9780306406157")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` 應該傳回 `false`。

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` 應該傳回 `false`。

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` 應該傳回 `true`。

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
