---
id: 69e2383af7832c8032603b90
title: "挑戰 274：最年長的人"
challengeType: 28
dashedName: challenge-274
---

# --description--

給定一個物件陣列，每個物件都有 `"name"` 和 `"age"` 屬性，傳回一個包含最年長者姓名的陣列。

如果有多個人擁有相同的最高年齡，請依照他們在輸入中出現的順序，傳回所有人的姓名。

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` 應該傳回 `["Brenda"]`。

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` 應該傳回 `["Alice"]`。

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` 應該傳回 `["Bill", "Carol"]`。

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` 應該傳回 `["George", "Holly", "Zach"]`。

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
