---
id: 69f35a5bb823ed620fcb7cbc
title: "挑戰 283：字串交錯"
challengeType: 28
dashedName: challenge-283
---

# --description--

給定兩個字串，傳回一個新的字串，將它們的字元依序交錯排列。如果其中一個字串較長，則將剩餘的字元附加在最後。

從第一個字串的第一個字元開始。

# --hints--

`zipStrings("abc", "123")` 應該傳回 `"a1b2c3"`。

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` 應該傳回 `"abcdefghijklmnopqrstuvwxyz"`。

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` 應該傳回 `"dnaiyght"`。

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` 應該傳回 `"pjyatvhaosncript"`。

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` 應該傳回 `"freeCodeCamp"`。

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
