---
id: 69f35a5bb823ed620fcb7cbd
title: "挑戰 284：I 在 E 之前"
challengeType: 28
dashedName: challenge-284
---

# --description--

給定一個字或句子，傳回一個修正版本，其中每個字都遵守「I 在 E 之前，除非在 C 之後」的規則。

- 如果一個字包含 `"ei"`，且前面沒有 `"c"`，就將它取代為 `"ie"`。
- 如果一個字包含 `"ie"`，且前面有 `"c"`，就將它取代為 `"ei"`。
- 其他所有字都保持不變。

# --hints--

`iBeforeE("beleive")` 應該傳回 `"believe"`。

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` 應該傳回 `"receive"`。

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` 應該傳回 `"we received a brief"`。

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` 應該傳回 `"she believed the friendly niece could perceive the grief"`。

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` 應該傳回 `"we received relief after the thief gave us a brief piece of fierce deceit"`。

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
