---
id: 69f35a5bb823ed620fcb7cb8
title: "挑戰 279：最長的多米諾串鏈"
challengeType: 28
dashedName: challenge-279
---

# --description--

給定一個二維陣列，代表一組多米諾骨牌，傳回最長的有效串鏈。

- 每個多米諾骨牌是由 0 到 6 的兩個數字配對組成，例如 `[3, 2]`。
- 當每個多米諾骨牌的第二個數字與下一個骨牌的第一個數字相符時，該串鏈即為有效。
- 第一個多米諾骨牌的第一個數字和最後一個骨牌的第二個數字不需要與任何數字相符。
- 任何多米諾骨牌都可以翻轉，因此 `[3, 2]` 可以當作 `[2, 3]` 使用。
- 總是有且只有一個最長的有效串鏈。

例如，給定 `[[1, 2], [4, 5], [2, 3]]`，傳回 `[[1, 2], [2, 3]]`。

# --hints--

`getLongestChain([[1, 2], [4, 5], [2, 3]])` 應該傳回 `[[1, 2], [2, 3]]`。

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [4, 5], [2, 3]]));
const result1 = JSON.stringify([[1, 2], [2, 3]]);
const result2 = JSON.stringify([[3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[2, 1], [4, 3], [5, 3]])` 應該傳回 `[[4, 3], [3, 5]]`。

```js
const chain = JSON.stringify(getLongestChain([[2, 1], [4, 3], [5, 3]]));
const result1 = JSON.stringify([[4, 3], [3, 5]]);
const result2 = JSON.stringify([[5, 3], [3, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]])` 應該傳回 `[[1, 2], [2, 3], [3, 4], [4, 0]]`。

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]]));
const result1 = JSON.stringify([[1, 2], [2, 3], [3, 4], [4, 0]]);
const result2 = JSON.stringify([[0, 4], [4, 3], [3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])` 應該傳回 `[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]`。

```js
const chain = JSON.stringify(getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]]));
const result1 = JSON.stringify([[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]);
const result2 = JSON.stringify([[5, 6], [6, 6], [6, 1], [1, 1], [1, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])` 應該傳回 `[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]`。

```js
const chain = JSON.stringify(getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]]));
const result1 = JSON.stringify([[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]);
const result2 = JSON.stringify([[6, 5], [5, 5], [5, 4], [4, 4], [4, 0], [0, 3], [3, 3]]);
assert.isTrue(chain === result1 || chain === result2);
```

# --seed--

## --seed-contents--

```js
function getLongestChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function getLongestChain(dominoes) {
  function search(chain, remaining) {
    let best = chain;
    const last = chain[chain.length - 1][1];

    for (let i = 0; i < remaining.length; i++) {
      const [a, b] = remaining[i];
      const rest = remaining.filter((_, j) => j !== i);

      if (a === last) {
        const result = search([...chain, [a, b]], rest);
        if (result.length > best.length) best = result;
      }
      if (b === last && a !== b) {
        const result = search([...chain, [b, a]], rest);
        if (result.length > best.length) best = result;
      }
    }
    return best;
  }

  let best = [];
  for (let i = 0; i < dominoes.length; i++) {
    const [a, b] = dominoes[i];
    const rest = dominoes.filter((_, j) => j !== i);

    const r1 = search([[a, b]], rest);
    if (r1.length > best.length) best = r1;

    if (a !== b) {
      const r2 = search([[b, a]], rest);
      if (r2.length > best.length) best = r2;
    }
  }
  return best;
}
```
