---
id: 69f35a5bb823ed620fcb7cbb
title: "挑戰 282：睡眠負債"
challengeType: 28
dashedName: challenge-282
---

# --description--

給定一個陣列，內容是從過去到今天每晚睡眠的時數，以及每晚的目標睡眠時數，傳回你今晚需要睡多少小時才能消除你的睡眠負債。

- 包含今晚的睡眠時數在需要補足的總時數中。
- 如果你已經睡足今晚的目標時數或更多，傳回 `0`。

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` 應該傳回 `20`。

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` 應該傳回 `10`。

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` 應該傳回 `4`。

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` 應該傳回 `0`。

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` 應該傳回 `0`。

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
