---
id: 69f8c998d78ad3171a0713bb
title: "挑戰 288：秘密數字"
challengeType: 28
dashedName: challenge-288
---

# --description--

給定一個秘密數字和一個猜測，判斷猜測是否正確。

傳回：

- `"higher"` 如果秘密數字比猜測大。
- `"lower"` 如果秘密數字比猜測小。
- `"you got it!"` 如果猜測正確。

# --hints--

`guessNumber(50, 30)` 應該傳回 `"higher"`。

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` 應該傳回 `"lower"`。

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` 應該傳回 `"you got it!"`。

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` 應該傳回 `"higher"`。

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` 應該傳回 `"lower"`。

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` 應該傳回 `"you got it!"`。

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
