---
id: 69f8c998d78ad3171a0713bd
title: "挑戰 290：披薩派對"
challengeType: 28
dashedName: challenge-290
---

# --description--

給定一個陣列，內容是每個人今天工作的時數，傳回訂購披薩派對所需的披薩數量。

- 將每個人的工作時數除以 3，以取得他們需要的披薩片數。
- 你不能吃半片披薩，所以要將每個人的披薩片數向上取整到最接近的整數。
- 每個人至少要有兩片披薩。
- 每個披薩有 8 片。將披薩總數向上取整到最接近的整個披薩數。

# --hints--

`getPizzasToOrder([8, 8, 8])` 應該傳回 `2`。

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` 應該傳回 `3`。

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` 應該傳回 `2`。

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` 應該傳回 `3`。

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` 應該傳回 `1`。

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` 應該傳回 `5`。

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
