---
id: 69f8c998d78ad3171a0713bc
title: "挑戰 289：差值總和"
challengeType: 28
dashedName: challenge-289
---

# --description--

給定一個數字陣列，傳回每個數字與其後一個數字之間差值的總和。

例如，給定 `[1, 3, 4]`，傳回 `3`（2 + 1）。

# --hints--

`sumOfDifferences([1, 3, 4])` 應該傳回 `3`。

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` 應該傳回 `5`。

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` 應該傳回 `-16`。

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` 應該傳回 `25`。

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
