---
id: 69f8c998d78ad3171a0713ba
title: "挑戰 287：羅馬數字修正器"
challengeType: 28
dashedName: challenge-287
---

# --description--

給定一個格式錯誤的羅馬數字字串，傳回標準羅馬數字表示法的數值。

輸入只會使用加法表示法，因此每個符號都將其數值加到總和中。提醒你，以下是符號和數值：

| 符號 | 數值 |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

重新編碼時，每一步都使用最大的可能符號，並在需要時使用減法配對（`"IV"`、`"IX"`、`"XL"`、`"XC"`、`"CD"`、`"CM"`）。

# --hints--

`fixNumerals("XIIIII")` 應該傳回 `"XV"`。

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` 應該傳回 `"LXIV"`。

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` 應該傳回 `"XL"`。

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` 應該傳回 `"MDCCXCIX"`。

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` 應該傳回 `"MCDLXIV"`。

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` 應該傳回 `"MMCMLXXXIV"`。

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
