---
id: 69f8c998d78ad3171a0713be
title: "挑戰 291：FizzBuzz 計數"
challengeType: 28
dashedName: challenge-291
---

# --description--

給定一個起始數字和結束數字，計算在該範圍（包含端點）內 fizz 和 buzz 出現的次數。

- 可被 3 整除的數字算作 fizz。
- 可被 5 整除的數字算作 buzz。
- 同時可被 3 和 5 整除的數字同時算作 fizz 和 buzz。

傳回一個物件或字典，格式為 `{ fizz, buzz }`，其中包含計數結果。

# --hints--

`fizzBuzzCount(1, 11)` 應該傳回 `{fizz: 3, buzz: 2}`。

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` 應該傳回 `{fizz: 9, buzz: 6}`。

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` 應該傳回 `{fizz: 26, buzz: 16}`。

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` 應該傳回 `{fizz: 207, buzz: 125}`。

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` 應該傳回 `{fizz: 4074, buzz: 2444}`。

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
