---
id: 69f8c998d78ad3171a0713b9
title: "挑戰 286：尚未解決的議題"
challengeType: 28
dashedName: challenge-286
---

# --description--

給定一個議題編號的陣列和另一個拉取請求（PR）編號的陣列，傳回所有 PR 合併後仍然開啟的議題陣列。

- 如果 PR 編號的數字是議題編號的旋轉組合，該 PR 就會關閉該議題。例如，議題 123 會被 PR 231 或 312 關閉。
- PR 不會關閉與它完全相同編號的議題。例如，PR 123 不會關閉議題 123。因此，所有數字都相同的議題無法被關閉。
- 兩者的編號都可能會去除前導零。例如，PR 201 會關閉議題 12（012，是 201 的旋轉組合）。同理，議題 201 會被 PR 12 關閉。

請依照原本給定的順序，傳回剩下仍開啟的議題。

# --hints--

`getOpenIssues([123, 234], [231])` 應該傳回 `[234]`。

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` 應該傳回 `[345, 16]`。

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` 應該傳回 `[456, 332, 15]`。

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` 應該傳回 `[115, 296, 24]`。

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` 應該傳回 `[95, 395, 754, 296, 709, 237, 1802]`。

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
