---
id: 69f8c998d78ad3171a0713bf
title: "挑戰 292：更寬的寬高比"
challengeType: 28
dashedName: challenge-292
---

# --description--

給定兩個代表不同影像尺寸的字串，傳回寬高比較大的影像的寬高比。

- 給定的字串格式為 `"WxH"`，例如 `"1920x1080"`。
- 寬高比是寬度與高度的比例，需約分為最簡整數比。例如 `"1920x1080"` 約分後為 `"16:9"`。
- 傳回格式為 `"W:H"` 的字串，例如 `"16:9"`。

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` 應該傳回 `"16:9"`。

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` 應該傳回 `"4:3"`。

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` 應該傳回 `"2:1"`。

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` 應該傳回 `"9:16"`。

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` 應該傳回 `"43:18"`。

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` 應該傳回 `"2089:8539"`。

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
