---
id: 69f90c1329a94b37e2a2086d
title: "挑戰 294：小括號組合"
challengeType: 28
dashedName: challenge-294
---

# --description--

給定一個整數 `n`，傳回 `n` 對小括號的有效組合數量。

- 有效組合是指每個開啟小括號都有對應的關閉小括號，且沒有任何關閉小括號出現在其對應的開啟小括號之前。

例如，給定 `2`，有 2 種有效組合：

```json
(())
()()
```

# --hints--

`getCombinations(2)` 應該傳回 `2`。

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` 應該傳回 `5`。

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` 應該傳回 `42`。

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` 應該傳回 `1430`。

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` 應該傳回 `742900`。

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
