---
id: 69f90c1329a94b37e2a2086c
title: "挑戰 293：最佳牌型"
challengeType: 28
dashedName: challenge-293
---

# --description--

給定一個包含五張撲克牌的字串陣列，傳回最佳牌型的名稱。

- 每張牌以兩個字元的字串表示：先是等級，接著是花色，`"2h"` 例如。
  - 等級由低到高依序為：`"2"`、`"3"`、`"4"`、`"5"`、`"6"`、`"7"`、`"8"`、`"9"`、`"T"`、`"J"`、`"Q"`、`"K"` 和 `"A"`。
  - 花色有：`"h"`、`"d"`、`"c"` 和 `"s"`。
- A（`"A"`）在順子中可以當作最大或最小使用。

牌型由差到好依序為：

| 名稱 | 說明 |
| - | - |
| `"High Card"` | 沒有任何配對或更好的牌型 |
| `"Pair"` | 一對 |
| `"Two Pair"` | 兩對 |
| `"Three of a Kind"` | 三條 |
| `"Straight"` | 順子（五張連續等級） |
| `"Flush"` | 同花（五張同花色） |
| `"Full House"` | 葫蘆（三條加一對） |
| `"Four of a Kind"` | 鐵支（四條） |
| `"Straight Flush"` | 同花順（五張同花色且連續等級） |
| `"Royal Flush"` | 皇家同花順（`"A"`、`"K"`、`"Q"`、`"J"`、`"T"` 同花色） |

傳回最佳牌型的名稱。

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` 應該傳回 `"Three of a Kind"`。

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` 應該傳回 `"Full House"`。

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` 應該傳回 `"Flush"`。

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` 應該傳回 `"Four of a Kind"`。

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` 應該傳回 `"Two Pair"`。

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` 應該傳回 `"Straight Flush"`。

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` 應該傳回 `"High Card"`。

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` 應該傳回 `"Straight"`。

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` 應該傳回 `"Pair"`。

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` 應該傳回 `"Royal Flush"`。

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
