---
id: 6a0dcd03ee4e68698080ef6c
title: "挑戰 308：信用卡驗證器"
challengeType: 28
dashedName: challenge-308
---

# --description--

給定一組信用卡號碼的數字字串，使用以下方法判斷它是否為有效的卡號：

- 從倒數第二個數字開始，向左每隔一個數字將其加倍。
- 如果加倍後的數字大於 9，則減去 9。
- 將所有數字（加倍和未加倍的）相加。
- 如果總和能被 10 整除，則該號碼有效。

# --hints--

`isValidCard("4532015112830366")` 應該傳回 `true`。

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` 應該傳回 `true`。

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` 應該傳回 `true`。

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` 應該傳回 `true`。

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` 應該傳回 `false`。

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` 應該傳回 `false`。

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` 應該傳回 `false`。

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
