---
id: 6a0dcc730cb92a616f86f0bf
title: "挑戰 295：結構驗證器第一部分"
challengeType: 28
dashedName: challenge-295
---

# --description--

給定一個物件（JavaScript）或字典（Python），判斷它是否符合以下結構：

```json
{
  username: string
}
```

- 允許有額外的鍵

# --hints--

`isValidSchema({ username: "bob" })` 應該傳回 `true`。

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` 應該傳回 `true`。

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` 應該傳回 `true`。

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
