---
id: 6a0dcd03ee4e68698080ef6b
title: "挑戰 307：分區規定"
challengeType: 28
dashedName: challenge-307
---

# --description--

給定一個二維網格（陣列的陣列），代表一個城市的建築布局，傳回所有違反分區規定的建築座標。

網格中的每個格子包含下表中的一個標籤。如果建築物的任一（最多）4 個鄰居（水平或垂直方向）是它不能相鄰的型別，該建築即為違規。

| 標籤 | 型別 | 不能相鄰的型別 |
| - | - | - |
| `"i"` | 工業區 | `"R"`, `"I"` |
| `"A"` | 農業區 | `"C"` |
| `"R"` | 住宅區 | `"i"`, `"C"` |
| `"I"` | 公共機構區 | `"i"` |
| `"C"` | 商業區 | `"R"`, `"A"` |
| `""`（空字串） | 未開發區 | 無限制 |

傳回所有違規格子的座標，格式為 `[row, col]` 對組的陣列，順序不限。如果沒有違規，傳回空陣列。

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` 應該傳回 `[[0, 0], [0, 1]]`。

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` 應該傳回 `[[0, 1], [1, 1]]`。

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` 應該傳回 `[]`。

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` 應該傳回 `[[0, 1], [0, 2], [0, 3]]`。

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` 應該傳回 `[[2, 3], [2, 4], [3, 1], [3, 2]]`。

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
