---
id: 6a0dcd03ee4e68698080ef68
title: "挑戰 304：行程安排"
challengeType: 28
dashedName: challenge-304
---

# --description--

給定一個至少包含兩個選用停靠點的當日行程陣列，傳回有效行程安排的數量。

行程總是包含 `"breakfast"`、`"lunch"` 和 `"dinner"`，這些不會作為引數傳入。選用停靠點可以放在行程中的任何位置，但必須遵守以下規則：

- `"breakfast"` 總是第一個，且在 `"lunch"` 之前至少有一個停靠點。
- `"lunch"` 必須出現在 `"dinner"` 之前，且兩者之間至少有一個停靠點。
- 在 `"dinner"` 之後最多只能有一個選用停靠點。

傳回有效安排的數量。

# --hints--

`getItineraryCount(["library", "park"])` 應該傳回 `2`。

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` 應該傳回 `18`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` 應該傳回 `120`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` 應該傳回 `840`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` 應該傳回 `55440`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
