---
id: 6a0dcd03ee4e68698080ef6d
title: "挑戰 309：數字排序"
challengeType: 28
dashedName: challenge-309
---

# --description--

給定一個以逗號分隔的數字字串，傳回一個從小到大排序的數字陣列。

# --hints--

`sortNumbers("3,1,2")` 應該傳回 `[1, 2, 3]`。

```js
assert.deepEqual(sortNumbers("3,1,2"), [1, 2, 3]);
```

`sortNumbers("5,3,8,1,9,2")` 應該傳回 `[1, 2, 3, 5, 8, 9]`。

```js
assert.deepEqual(sortNumbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9]);
```

`sortNumbers("12,61,49,80,19,50,77,38")` 應該傳回 `[12, 19, 38, 49, 50, 61, 77, 80]`。

```js
assert.deepEqual(sortNumbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80]);
```

`sortNumbers("0,6,-19,44,-2,7,0")` 應該傳回 `[-19, -2, 0, 0, 6, 7, 44]`。

```js
assert.deepEqual(sortNumbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44]);
```

# --seed--

## --seed-contents--

```js
function sortNumbers(str) {

  return str;
}
```

# --solutions--

```js
function sortNumbers(str) {
  return str.split(",").map(Number).sort((a, b) => a - b);
}
```
