---
id: 6a0dcc730cb92a616f86f0c0
title: "挑戰 296：結構驗證器第二部分"
challengeType: 28
dashedName: challenge-296
---

# --description--

給定一個物件（JavaScript）或字典（Python），判斷它是否符合以下結構：

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- 允許有額外的鍵

# --hints--

`isValidSchema({ username: "alice", posts: 10, verified: false })` 應該傳回 `true`。

```js
assert.isTrue(isValidSchema({ username: "alice", posts: 10, verified: false }));
```

`isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 })` 應該傳回 `true`。

```js
assert.isTrue(isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 }));
```

`isValidSchema({ username: "frank", posts: "21", verified: true })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: "frank", posts: "21", verified: true }));
```

`isValidSchema({ username: "sam", posts: 17, verified: "false" })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: "sam", posts: 17, verified: "false" }));
```

`isValidSchema({ username: "bill", verified: true })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: "bill", verified: true }));
```

`isValidSchema({ username: "fred", verified: true })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: "fred", verified: true }));
```

`isValidSchema({ username: 5, posts: 10, verified: true })` 應該傳回 `false`。

```js
assert.isFalse(isValidSchema({ username: 5, posts: 10, verified: true }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean'
  );
}
```
