---
id: 6a0dcd03ee4e68698080ef67
title: "挑戰 303：室友分配"
challengeType: 28
dashedName: challenge-303
---

# --description--

給定一個包含人員及其室友群組的陣列，請根據以下規則傳回飯店住宿的房間分配：

- 每個人都有一個 `name` 和 `group` 屬性：

```json
[
  { "name": "Alice", "group": "A" },
  { "name": "Bob", "group": "B" },
  { "name": "Carol", "group": "A" }
]
```

- 人員只能與同一群組的人共用房間，且配對順序依照給定的順序。
- 傳回一個字串陣列，若為共用房間，名字間以 `" and "` 分隔；若為單人房，則只傳回名字。名字必須依照配對的順序出現。以上述範例為例，傳回 `["Alice and Carol", "Bob"]`。

# --hints--

`getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }])` 應該傳回 `["Alice and Carol", "Bob"]`。

```js
assert.deepEqual(getRoommates([{ "name": "Alice", "group": "A" }, { "name": "Bob", "group": "B" }, { "name": "Carol", "group": "A" }]).sort(), ["Alice and Carol", "Bob"].sort());
```

`getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }])` 應該傳回 `["John and Julia", "Jim"]`。

```js
assert.deepEqual(getRoommates([{ "name": "John", "group": "C" }, { "name": "Julia", "group": "C" }, { "name": "Jim", "group": "C" }]).sort(), ["John and Julia", "Jim"].sort());
```

`getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }])` 應該傳回 `["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"]`。

```js
assert.deepEqual(getRoommates([{ "name": "Adam", "group": "D" }, { "name": "Abraham", "group": "E" }, { "name": "Austin", "group": "E" }, { "name": "Augustus", "group": "D" }, { "name": "Angelica", "group": "D" }, { "name": "Aaron", "group": "E" }]).sort(), ["Adam and Augustus", "Angelica", "Abraham and Austin", "Aaron"].sort());
```

`getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }])` 應該傳回 `["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"]`。

```js
assert.deepEqual(getRoommates([{ "name": "Frank", "group": "A" }, { "name": "Emitt", "group": "B" }, { "name": "Daria", "group": "F" }, { "name": "Charles", "group": "D" }, { "name": "Bailey", "group": "A" }, { "name": "Albert", "group": "F" }]).sort(), ["Frank and Bailey", "Emitt", "Daria and Albert", "Charles"].sort());
```

`getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }])` 應該傳回 `["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"]`。

```js
assert.deepEqual(getRoommates([{ "name": "Kevin", "group": "A" }, { "name": "Yuri", "group": "A" }, { "name": "Hugo", "group": "B" }, { "name": "Violet", "group": "A" }, { "name": "Brett", "group": "A" }, { "name": "Wayne", "group": "B" }]).sort(), ["Kevin and Yuri", "Violet and Brett", "Hugo and Wayne"].sort());
```

# --seed--

## --seed-contents--

```js
function getRoommates(people) {

  return people;
}
```

# --solutions--

```js
function getRoommates(people) {
  const groups = {};

  for (const { name, group } of people) {
    (groups[group] = groups[group] || []).push(name);
  }

  return Object.values(groups).flatMap(group => {
    const rooms = [];
    for (let i = 0; i < group.length; i += 2) {
      rooms.push(group.slice(i, i + 2).join(" and "));
    }
    return rooms;
  });
}
```
