---
id: 6a0dcd03ee4e68698080ef66
title: "挑戰 302：時差疲勞"
challengeType: 28
dashedName: challenge-302
---

# --description--

給定出發城市、抵達城市、飛行時長（小時）和旅行方向，傳回旅客所經歷的時差疲勞小時數。

所給城市將來自以下列表，包含它們的 UTC 偏移量：

| 城市 | 偏移量 |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

計算時差疲勞小時數的方法：

1. 找出兩個城市之間的時區差（小時）。
2. 確定方向乘數。如果是 `"east"` 方向，乘數為 1.5，否則為 1.0。
3. 使用公式計算時差疲勞小時數：時區差 +（飛行時長 * 0.1）* 方向乘數

傳回四捨五入到小數點第一位的時差疲勞小時數。

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` 應該傳回 `6.5`。

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` 應該傳回 `5.8`。

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` 應該傳回 `1.6`。

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` 應該傳回 `4.7`。

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` 應該傳回 `17.5`。

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` 應該傳回 `5.9`。

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` 應該傳回 `9.5`。

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
