---
id: 6a0dcc730cb92a616f86f0c5
title: "挑戰 301：最後的洗衣精用量"
challengeType: 28
dashedName: challenge-301
---

# --description--

給定你剩餘的洗衣精勺數和一個陣列，該陣列記錄你過去每天使用的勺數，傳回你還能使用多少整天的洗衣精。

從使用紀錄計算你的每日平均用量，並假設未來每天都使用這個量。

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` 應該傳回 `5`。

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` 應該傳回 `7`。

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` 應該傳回 `11`。

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` 應該傳回 `12`。

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` 應該傳回 `1`。

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
