---
id: 6a15cadf5f240d05a264955b
title: "挑戰 317：BMI 計算器"
challengeType: 28
dashedName: challenge-317
---

# --description--

給定一個以磅為單位的體重和一個以英吋為單位的身高，傳回四捨五入到小數點後一位的 BMI（身體質量指數）。

計算 BMI：將體重除以身高的平方，然後將結果乘以 703。

# --hints--

`calculateBmi(180, 70)` 應該傳回 `25.8`。

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` 應該傳回 `24.0`。

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` 應該傳回 `19.5`。

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` 應該傳回 `39.1`。

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` 應該傳回 `22.8`。

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
