---
id: 6a15cadf5f240d05a264955a
title: "挑戰 316：1337 語言"
challengeType: 28
dashedName: challenge-316
---

# --description--

給定一個小寫字串，將它轉換成 1337 語言，方法是用下表中的 1337 替代字元取代字母：

| 字母 | 1337 |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 沒有替代字元的字元保持不變。

# --hints--

`makeLeet("cool")` 應該傳回 `"c001"`。

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` 應該傳回 `"1337"`。

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` 應該傳回 `"h4ck3r"`。

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` 應該傳回 `"547311173"`。

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` 應該傳回 `"4bcd3f9h1jk1mn0pqr57uvwxyz"`。

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
