---
id: 6a15cadf5f240d05a264955f
title: "挑戰 321：週期性拼字"
challengeType: 28
dashedName: challenge-321
---

# --description--

給定一個字，判斷是否能用週期表中的元素符號拼出來。

- 拼字時忽略大小寫。`"neon"` 可以用符號 `"Ne"`、`"O"` 和 `"N"` 拼出來。

這裡是完整的元素符號列表：

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

傳回一個用來拼出該字的元素陣列，保持原始大小寫並依拼字順序排列。如果無法拼出，則傳回空陣列。

# --hints--

`getPeriodicSpelling("neon")` 應該傳回 `["Ne", "O", "N"]`。

```js
assert.deepEqual(getPeriodicSpelling("neon"), ["Ne", "O", "N"]);
```

`getPeriodicSpelling("rational")` 應該傳回 `["Ra", "Ti", "O", "N", "Al"]`。

```js
assert.deepEqual(getPeriodicSpelling("rational"), ["Ra", "Ti", "O", "N", "Al"]);
```

`getPeriodicSpelling("yarn")` 應該傳回 `["Y", "Ar", "N"]`。

```js
assert.deepEqual(getPeriodicSpelling("yarn"), ["Y", "Ar", "N"]);
```

`getPeriodicSpelling("carbon")` 應該傳回 `["C", "Ar", "B", "O", "N"]` 或 `["Ca", "Rb", "O", "N"]`。

```js
const result = JSON.stringify(getPeriodicSpelling("carbon"));
const path1 = JSON.stringify(["C", "Ar", "B", "O", "N"]);
const path2 = JSON.stringify(["Ca", "Rb", "O", "N"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("noisy")` 應該傳回 `["N", "O", "I", "S", "Y"]` 或 `["No", "I", "S", "Y"]`。

```js
const result = JSON.stringify(getPeriodicSpelling("noisy"));
const path1 = JSON.stringify(["N", "O", "I", "S", "Y"]);
const path2 = JSON.stringify(["No", "I", "S", "Y"]);
assert.isTrue(result === path1 || result === path2);
```


`getPeriodicSpelling("bicycles")` 應該傳回 `["B", "I", "C", "Y", "Cl", "Es"]` 或 `["Bi", "C", "Y", "Cl", "Es"]`。

```js
const result = JSON.stringify(getPeriodicSpelling("bicycles"));
const path1 = JSON.stringify(["B", "I", "C", "Y", "Cl", "Es"]);
const path2 = JSON.stringify(["Bi", "C", "Y", "Cl", "Es"]);
assert.isTrue(result === path1 || result === path2);
```

`getPeriodicSpelling("optics")` 應該傳回 `["O", "P", "Ti", "C", "S"]`、`["O", "P", "Ti", "Cs"]`、`["O", "Pt", "I", "C", "S"]`，或 `["O", "Pt", "I", "Cs"]`。

```js
const result = JSON.stringify(getPeriodicSpelling("optics"));
const path1 = JSON.stringify(["O", "P", "Ti", "C", "S"]);
const path2 = JSON.stringify(["O", "P", "Ti", "Cs"]);
const path3 = JSON.stringify(["O", "Pt", "I", "C", "S"]);
const path4 = JSON.stringify(["O", "Pt", "I", "Cs"]);
assert.isTrue(result === path1 || result === path2 || result === path3 || result === path4);
```

`getPeriodicSpelling("value")` 應該傳回 `[]`。

```js
assert.deepEqual(getPeriodicSpelling("value"), []);
```

# --seed--

## --seed-contents--

```js
function getPeriodicSpelling(word) {

  return word;
}
```

# --solutions--

```js
function getPeriodicSpelling(word) {
  const elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];

  const lower = new Map(elements.map(e => [e.toLowerCase(), e]));

  function spell(word, i) {
    if (i === word.length) return [];
    const one = word.slice(i, i + 1);
    const two = word.slice(i, i + 2);
    if (lower.has(two)) {
      const rest = spell(word, i + 2);
      if (rest !== null) return [lower.get(two), ...rest];
    }
    if (lower.has(one)) {
      const rest = spell(word, i + 1);
      if (rest !== null) return [lower.get(one), ...rest];
    }
    return null;
  }

  return spell(word.toLowerCase(), 0) ?? [];
}
```
