---
id: 6a15cadf5f240d05a2649556
title: "挑戰 312：串流費用"
challengeType: 28
dashedName: challenge-312
---

# --description--

給定一個陣列，表述你串流服務購物車中的電影，以及一個字串表示你的訂閱等級，傳回電影的總費用。

購物車中的每個條款都是一個物件，包含 `"format"`（`"HD"` 或 `"4K"`）和 `"type"`（`"rent"` 或 `"buy"`）。它們的費用如下：

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

套用以下訂閱等級折扣：

- `"none"`：全價
- `"basic"`：九折
- `"premium"`：七五折

傳回以兩位小數四捨五入的總費用，格式為 `"$D.CC"`。

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` 應該傳回 `"$3.99"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "buy" }], "premium")` 應該傳回 `"$9.74"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "buy" }], "premium"), "$9.74");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` 應該傳回 `"$18.87"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` 應該傳回 `"$44.98"`。

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$44.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` 應該傳回 `"$42.96"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic")` 應該傳回 `"$38.66"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic"), "$38.66");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
