---
id: 6a15cadf5f240d05a2649559
title: "挑戰 315：夏至"
challengeType: 28
dashedName: challenge-315
---

# --description--

今天是夏至，北半球一年中白天最長的一天，南半球則是最短的一天。給定一個緯度，傳回一個表述白天和夜晚時數的字串。

- 緯度介於 90（北極）和 -90（南極）之間，包含端點
- 白天時數 = 12 + (緯度 / 90) * 12
- 將白天時數四捨五入到最接近的偶數

傳回一個 24 字元長的字串，使用 `"☀️"` 表示白天時數，`"🌑"` 表示夜晚時數，其中：

- 每個字元代表一小時，從午夜（第 0 小時）開始
- 日出和日落對稱分布在中午兩側

例如，緯度為 `0`（赤道）時有 12 小時的日照，因此日出時間是上午 6 點，日落時間是下午 6 點。傳回：`"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`。

# --hints--

`getDaytimeHours(0)` 應該傳回 `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`。

```js
assert.equal(getDaytimeHours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(90)` 應該傳回 `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`。

```js
assert.equal(getDaytimeHours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️");
```

`getDaytimeHours(-90)` 應該傳回 `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
assert.equal(getDaytimeHours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(-33)` 應該傳回 `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
assert.equal(getDaytimeHours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(66.5)` 應該傳回 `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`。

```js
assert.equal(getDaytimeHours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑");
```

`getDaytimeHours(40)` 應該傳回 `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`。

```js
assert.equal(getDaytimeHours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑");
```

`getDaytimeHours(-50)` 應該傳回 `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`。

```js
assert.equal(getDaytimeHours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

# --seed--

## --seed-contents--

```js
function getDaytimeHours(latitude) {

  return latitude;
}
```

# --solutions--

```js
function getDaytimeHours(latitude) {
  const daylightHours = Math.round((12 + (latitude / 90) * 12) / 2) * 2;
  const sunrise = 12 - daylightHours / 2;
  const sunset = 12 + daylightHours / 2;

  let result = "";
  for (let hour = 0; hour < 24; hour++) {
    result += hour >= sunrise && hour < sunset ? "☀️" : "🌑";
  }

  return result;
}
```
