---
id: 6a151d32d772271dd2448c2d
title: "挑戰 311：魔法施放者"
challengeType: 28
dashedName: challenge-311
---

# --description--

給定一串你正在施放的魔法程式碼字串，計算總分數。

字串中的每個字元代表一個魔法：

| 程式碼 | 魔法 | 類別 | 基本分數 |
| - | - | - | - |
| `"f"` | 火焰 | 破壞 | 3 |
| `"l"` | 閃電 | 破壞 | 3 |
| `"i"` | 冰霜 | 控制元件 | 2 |
| `"w"` | 風 | 控制元件 | 2 |
| `"h"` | 治療 | 恢復 | 1 |
| `"s"` | 盾牌 | 恢復 | 1 |

根據連續施放的魔法類別不同，會套用連擊倍數：

- 第一個魔法總是以基本分數計算。
- 每當施放的魔法類別與前一個不同，連擊倍數會增加 1。
- 若施放的魔法類別與前一個相同，連擊倍數會重設為 1。
- 每個魔法的分數為其基本分數乘以當前連擊倍數。

傳回整串魔法序列的總分數。

# --hints--

`cast("fihwl")` 應該傳回 `33`。

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` 應該傳回 `45`。

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` 應該傳回 `37`。

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` 應該傳回 `50`。

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` 應該傳回 `101`。

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
