---
id: 6a15cadf5f240d05a264955d
title: "挑戰 319：Frontmatter 解析器"
challengeType: 28
dashedName: challenge-319
---

# --description--

給定一個代表 frontmatter 區段的字串，解析它並傳回一個包含鍵和值的物件（JavaScript）或字典（Python）。

Frontmatter 包裹在 `---` 分隔符中，並包含每行一組的 `key: value` 配對。舉例來說：

```md
---
title: My Post
draft: false
views: 100
---
```

應該傳回：

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 數字、布林值和字串都應該以它們各自的型別傳回。
- 給定的字串會以換行字元（`"\n"`）分隔新行。上述範例會以 `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"` 的形式給出。

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` 應該傳回 `{ title: "My Post", draft: false, views: 100 }`。

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` 應該傳回 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`。

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` 應該傳回 `{ version: "1.0.0", url: "https://example.com", private: true }`。

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` 應該傳回 `{ rating: 4.5, price: 9.99 }`。

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
