---
id: 6a15cadf5f240d05a2649557
title: "挑戰 313：租賃費用"
challengeType: 28
dashedName: challenge-313
---

# --description--

給定租賃時間戳、歸還時間戳和租賃等級，傳回包含任何逾期費用的租賃總費用。

- 給定的時間戳是 UTC ISO 字串，例如：`"2026-06-18T18:30:00Z"`。
- 租賃等級是指租賃到期前的天數：`1`、`3` 或 `7`。
- 租賃必須在租期最後一天的 UTC 中午 12:00 或之前歸還。例如，3 月 15 日任何時間租出的 1 天租賃，必須在 3 月 16 日 UTC 中午 12:00 前歸還。
- 每超過到期日期和時間一天，將收取逾期費用。

價格如下：

| 等級 | 基本費用 | 每日逾期費用 |
| - | - | - |
| 1 天 | $4.99 | $3.99 |
| 3 天 | $3.99 | $2.99 |
| 7 天 | $2.99 | $0.99 |

傳回以兩位小數四捨五入的總費用，格式為 `"$D.CC"`。

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` 應該傳回 `"$4.99"`。

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` 應該傳回 `"$12.97"`。

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` 應該傳回 `"$3.99"`。

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` 應該傳回 `"$9.97"`。

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` 應該傳回 `"$2.99"`。

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` 應該傳回 `"$358.40"`。

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
