---
id: 6a15cadf5f240d05a2649558
title: "挑戰 314：質因數分解"
challengeType: 28
dashedName: challenge-314
---

# --description--

給定一個大於 1 的整數，傳回其質因數分解，結果為一個由數字組成且依升冪排序的陣列。

質因數分解是將一組質數相乘以產生該整數的集合。每個數字都有且只有一組質因數分解。舉例來說，20 的質因數分解是 `[2, 2, 5]`，因為 2 * 2 * 5 = 20。

如果給定的整數本身就是質數，則以單元素陣列傳回它。

# --hints--

`primeFactorization(20)` 應該傳回 `[2, 2, 5]`。

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` 應該傳回 `[17]`。

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` 應該傳回 `[3, 5]`。

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` 應該傳回 `[5, 7]`。

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` 應該傳回 `[3, 3, 3, 37]`。

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` 應該傳回 `[2, 2, 2, 3, 3, 5]`。

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` 應該傳回 `[2, 3, 5, 7, 11, 13, 17]`。

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
