---
id: 6a15cadf5f240d05a264955e
title: "挑戰 320：血庫"
challengeType: 28
dashedName: challenge-320
---

# --description--

給定一個血庫庫存的陣列和一個病人血型需求的陣列，傳回格式為 `"X of Y patients served"` 的字串。其中 `X` 是能從血庫庫存中供應血液的最大病人數，`Y` 是病人總數。

兩個陣列中的每個項目都是以下血型之一：`"AB"`、`"A"`、`"B"` 或 `"O"`。

相容性規則：

- `"AB"` 可以接受任何血型。
- `"A"` 可以接受 `"A"` 和 `"O"`。
- `"B"` 可以接受 `"B"` 和 `"O"`。
- `"O"` 只能接受 `"O"`。

給定陣列中的重複項目代表數量。

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` 應該傳回 `"4 of 4 patients served"`。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` 應該傳回 `"3 of 5 patients served"`。

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` 應該傳回 `"4 of 5 patients served"`。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` 應該傳回 `"4 of 4 patients served"`。

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` 應該傳回 `"8 of 11 patients served"`。

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` 應該傳回 `"13 of 14 patients served"`。

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
