---
id: 6a19b1062d1b153d8ac76d70
title: "挑戰 322：連線 3"
challengeType: 28
dashedName: challenge-322
---

# --description--

給定一個字串矩陣，代表遊戲格子上的棋子，判斷是否有玩家連成三子。

- 每個格子包含 `"R"`、`"Y"` 或 `""`（空字串）。
- 連成三子指的是同一型別的三個連續非空格子，橫向、縱向或斜向皆可。

傳回：

- 傳回一個扁平陣列，包含贏家和他們三個連線格子的座標，格式為 `["R", [0,2], [1,3], [2,4]]`。座標依照由上到下、由左到右的順序傳回。
- 若無贏家，傳回空陣列。

# --hints--

`connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])` 應傳回 `["R", [3, 1], [3, 2], [3, 3]]`。

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]]);
```

`connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])` 應傳回 `["Y", [1, 1], [2, 1], [3, 1]]`。

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]]);
```

`connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])` 應傳回 `["R", [0, 3], [1, 2], [2, 1]]`。

```js
assert.deepEqual(connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]]);
```

`connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])` 應傳回 `["Y", [0, 1], [1, 2], [2, 3]]`。

```js
assert.deepEqual(connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]]);
```

`connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])` 應傳回 `[]`。

```js
assert.deepEqual(connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), []);
```

# --seed--

## --seed-contents--

```js
function connectThree(matrix) {

  return matrix;
}
```

# --solutions--

```js
function connectThree(matrix) {
  const rows = matrix.length;
  const cols = matrix[0].length;

  const directions = [[0, 1], [1, 0], [1, 1], [1, -1]];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const piece = matrix[r][c];
      if (!piece) continue;

      for (const [dr, dc] of directions) {
        const cells = [[r, c]];

        for (let i = 1; i < 3; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols) break;
          if (matrix[nr][nc] !== piece) break;
          cells.push([nr, nc]);
        }

        if (cells.length === 3) {
          cells.sort(([r1, c1], [r2, c2]) => r1 !== r2 ? r1 - r2 : c1 - c2);
          return [piece, ...cells];
        }
      }
    }
  }

  return [];
}
```
