---
id: 6a19b1062d1b153d8ac76d72
title: "挑戰 324：重複字元計數"
challengeType: 28
dashedName: challenge-324
---

# --description--

給定兩個字串，傳回第二個字串中可在第一個字串找到的字元數量。

- 第二個字串中重複的字元會分別計算。

# --hints--

`duplicateCharacterCount("aloha", "hei")` 應該傳回 `1`。

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` 應該傳回 `4`。

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` 應該傳回 `3`。

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` 應該傳回 `0`。

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` 應該傳回 `5`。

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` 應該傳回 `2`。

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` 應該傳回 `26`。

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
