---
id: 6a19b1062d1b153d8ac76d71
title: "挑戰 323：歌曲心情判斷器"
challengeType: 28
dashedName: challenge-323
---

# --description--

給定歌曲的類別字串和 BPM 數字，使用下表判斷歌曲的心情：

| 心情 | 類別 | BPM 範圍 |
|------|-------|-----------|
| `"focus"` | `"classical"` | 60–109 |
| `"focus"` | `"electronic"` | 60–89 |
| `"happy"` | `"pop"` | 60–180 |
| `"happy"` | `"classical"` | 110–180 |
| `"happy"` | `"rock"` | 60–129 |
| `"happy"` | `"electronic"` | 90–134 |
| `"hype"` | `"rock"` | 130–180 |
| `"hype"` | `"electronic"` | 135–180 |

# --hints--

`getMood("rock", 111)` 應該傳回 `"happy"`。

```js
assert.equal(getMood("rock", 111), "happy");
```

`getMood("electronic", 74)` 應該傳回 `"focus"`。

```js
assert.equal(getMood("electronic", 74), "focus");
```

`getMood("classical", 180)` 應該傳回 `"happy"`。

```js
assert.equal(getMood("classical", 180), "happy");
```

`getMood("rock", 155)` 應該傳回 `"hype"`。

```js
assert.equal(getMood("rock", 155), "hype");
```

`getMood("electronic", 90)` 應該傳回 `"happy"`。

```js
assert.equal(getMood("electronic", 90), "happy");
```

`getMood("classical", 67)` 應該傳回 `"focus"`。

```js
assert.equal(getMood("classical", 67), "focus");
```

`getMood("pop", 100)` 應該傳回 `"happy"`。

```js
assert.equal(getMood("pop", 100), "happy");
```

`getMood("electronic", 135)` 應該傳回 `"hype"`。

```js
assert.equal(getMood("electronic", 135), "hype");
```

# --seed--

## --seed-contents--

```js
function getMood(genre, bpm) {

  return genre;
}
```

# --solutions--

```js
function getMood(genre, bpm) {
  const table = [
    { mood: "focus",   genre: "classical",  min: 60,  max: 109 },
    { mood: "focus",   genre: "electronic", min: 60,  max: 89  },
    { mood: "happy",   genre: "pop",        min: 60,  max: 180 },
    { mood: "happy",   genre: "classical",  min: 110, max: 180 },
    { mood: "happy",   genre: "rock",       min: 60,  max: 129 },
    { mood: "happy",   genre: "electronic", min: 90,  max: 134 },
    { mood: "hype",    genre: "rock",       min: 130, max: 180 },
    { mood: "hype",    genre: "electronic", min: 135, max: 180 },
  ];

  return table.find(row => row.genre === genre && bpm >= row.min && bpm <= row.max).mood;
}
```
