---
id: 6a19dd69388e5a6e59e4f2ef
title: "挑戰 326：最大利潤"
challengeType: 28
dashedName: challenge-326
---

# --description--

給定一個每日股票價格的陣列和一個預算（以美元計），計算在指定期間內買賣股票所能獲得的最大利潤。

- 你只能在買入後賣出。
- 你最多只能執行一次買入和一次賣出交易。賣出後就不能再買入。
- 你只能購買整股。
- 傳回最大可能利潤，格式為字串，向下取整到最接近的分，並顯示兩位小數。

# --hints--

`getMaxProfit([5, 6], 50)` 應該傳回 `"10.00"`。

```js
assert.equal(getMaxProfit([5, 6], 50), "10.00");
```

`getMaxProfit([8, 2, 5, 10], 20)` 應該傳回 `"80.00"`。

```js
assert.equal(getMaxProfit([8, 2, 5, 10], 20), "80.00");
```

`getMaxProfit([4, 5, 3, 6], 20)` 應該傳回 `"18.00"`。

```js
assert.equal(getMaxProfit([4, 5, 3, 6], 20), "18.00");
```

`getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` 應該傳回 `"8.31"`。

```js
assert.equal(getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31");
```

`getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` 應該傳回 `"0.00"`。

```js
assert.equal(getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00");
```

`getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` 應該傳回 `"73.80"`。

```js
assert.equal(getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80");
```

# --seed--

## --seed-contents--

```js
function getMaxProfit(prices, budget) {

  return prices;
}
```

# --solutions--

```js
function getMaxProfit(prices, budget) {
  let maxProfit = 0;

  for (let buy = 0; buy < prices.length - 1; buy++) {
    for (let sell = buy + 1; sell < prices.length; sell++) {
      const shares = Math.floor(budget / prices[buy]);
      const profit = shares * (prices[sell] - prices[buy]);
      if (profit > maxProfit) maxProfit = profit;
    }
  }

  return (Math.floor(maxProfit * 100) / 100).toFixed(2);
}
```
