---
id: 6a1d9f98e819ed70a0e994e0
title: "挑戰 334：精確找零"
challengeType: 28
dashedName: challenge-334
---

# --description--

給定一個以分為單位的整數金額，傳回使用一分硬幣（1 分）、五分硬幣（5 分）、十分硬幣（10 分）和二十五分硬幣（25 分）組成精確找零的不同方法數量。

# --hints--

`exactChange(3)` 應該傳回 `1`。

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` 應該傳回 `2`。

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` 應該傳回 `6`。

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` 應該傳回 `24`。

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` 應該傳回 `73`。

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` 應該傳回 `213`。

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
