---
id: 6a1d9f98e819ed70a0e994da
title: "挑戰 328：Kaprekar 的常式"
challengeType: 28
dashedName: challenge-328
---

# --description--

給定一個 4 位數字，傳回你需要套用 Kaprekar 常式直到達到 6174 的次數。

Kaprekar 常式的運作方式如下：

- 將數字的各位數字依降冪排列，組成最大數字
- 將數字的各位數字依升冪排列，組成最小數字（必要時補零）
- 用較大的數字減去較小的數字
- 用新數字重複此過程

# --hints--

`kaprekar(1234)` 應該傳回 `3`。

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` 應該傳回 `6`。

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` 應該傳回 `4`。

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` 應該傳回 `7`。

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` 應該傳回 `2`。

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
