---
id: 6a1d9f98e819ed70a0e994dc
title: "挑戰 330：小寫字"
challengeType: 28
dashedName: challenge-330
---

# --description--

給定一個字串，傳回其中完全是小寫的字，保持原本順序並以空格分隔每個字。

# --hints--

`getLowercaseWords("hello GOOD world")` 應該傳回 `"hello world"`。

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` 應該傳回 `"these are all lowercase"`。

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` 應該傳回 `"less is more"`。

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` 應該傳回 `"eat pizza every day"`。

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` 應該傳回 `"the quick brown fox jumped over the lazy dog"`。

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
