---
id: 6a1d9f98e819ed70a0e994db
title: "挑戰 329：桶狀填充"
challengeType: 28
dashedName: challenge-329
---

# --description--

給定一個二維網格、一個起始位置（`[row, col]`）和一個新值，將起始位置及所有相同值且相連的儲存格的值取代為新值。

- 儲存格若在水平方向或垂直方向相鄰（不包含斜角相鄰）即視為相連。

傳回更新後的網格。

# --hints--

`bucketFill([["R", "G"], ["R", "G"]], [0, 1], "B")` 應傳回 `[["R", "B"], ["R", "B"]]`。

```js
assert.deepEqual(bucketFill([["R", "G"], ["R", "G"]], [0, 1], "B"), [["R", "B"], ["R", "B"]]);
```

`bucketFill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B")` 應傳回 `[["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]`。

```js
assert.deepEqual(bucketFill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B"), [["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]);
```

`bucketFill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R")` 應傳回 `[["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]`。

```js
assert.deepEqual(bucketFill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R"), [["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]);
```

`bucketFill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y")` 應傳回 `[["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]`。

```js
assert.deepEqual(bucketFill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y"), [["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]);
```

`bucketFill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G")` 應傳回 `[["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]`。

```js
assert.deepEqual(bucketFill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G"), [["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, [row, col], newValue) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, [row, col], newValue) {
  const target = grid[row][col];
  if (target === newValue) return grid;

  function fill(r, c) {
    if (r < 0 || r >= grid.length) return;
    if (c < 0 || c >= grid[0].length) return;
    if (grid[r][c] !== target) return;

    grid[r][c] = newValue;
    fill(r - 1, c);
    fill(r + 1, c);
    fill(r, c - 1);
    fill(r, c + 1);
  }

  fill(row, col);
  return grid;
}
```
