---
id: 6a1d9f98e819ed70a0e994de
title: "挑戰 332：議題分流"
challengeType: 28
dashedName: challenge-332
---

# --description--

給定自上次在議題上發文以來的毫秒數，以及該議題上最後發佈的訊息，根據以下規則判斷你應該如何處理該議題：

- 如果最後一則訊息距今不到 7 天，傳回 `"leave it"`
- 如果最後一則訊息距今 7 天或以上，且內容包含 `"bump"`（不分大小寫），傳回 `"close it"`
- 否則，傳回 `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` 應該傳回 `"leave it"`。

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` 應該傳回 `"bump it"`。

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` 應該傳回 `"close it"`。

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` 應該傳回 `"bump it"`。

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` 應該傳回 `"close it"`。

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` 應該傳回 `"leave it"`。

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
