---
id: 6a1d9f98e819ed70a0e994dd
title: "挑戰 331：最接近的倍數"
challengeType: 28
dashedName: challenge-331
---

# --description--

給定兩個整數，將第一個數四捨五入到最接近的第二個數的倍數。

# --hints--

`roundToNearestMultiple(5, 3)` 應該傳回 `6`。

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` 應該傳回 `16`。

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` 應該傳回 `45`。

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` 應該傳回 `33`。

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` 應該傳回 `96`。

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
