---
id: 6a1d9f98e819ed70a0e994e2
title: "挑戰 336：星座配對"
challengeType: 28
dashedName: challenge-336
---

# --description--

給定兩個星座字串，傳回它們的相容性百分比。

星座依照這個順序排列在一個 12 個位置的輪盤上：`"Aries"`、`"Taurus"`、`"Gemini"`、`"Cancer"`、`"Leo"`、`"Virgo"`、`"Libra"`、`"Scorpio"`、`"Sagittarius"`、`"Capricorn"`、`"Aquarius"`、`"Pisces"`，在`"Pisces"`之後會繞回`"Aries"`。找出兩個星座之間的最短距離並傳回相容性：

| 距離 | 相容性 |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscopeMatch("Libra", "Sagittarius")` 應該傳回 `"80%"`。

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

`horoscopeMatch("Gemini", "Scorpio")` 應該傳回 `"20%"`。

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

`horoscopeMatch("Pisces", "Aries")` 應該傳回 `"40%"`。

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

`horoscopeMatch("Capricorn", "Cancer")` 應該傳回 `"50%"`。

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

`horoscopeMatch("Aquarius", "Aquarius")` 應該傳回 `"100%"`。

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

`horoscopeMatch("Virgo", "Taurus")` 應該傳回 `"90%"`。

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

`horoscopeMatch("Leo", "Scorpio")` 應該傳回 `"30%"`。

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
