---
id: 6a1d9f98e819ed70a0e994e1
title: "挑戰 335：五顆骰子"
challengeType: 28
dashedName: challenge-335
---

# --description--

給定一個包含五顆骰子（數值介於 1 到 6）的陣列，傳回最佳可能的牌型。

以下是牌型由低到高的等級排序：

| 牌型 | 說明 |
|------|-------------|
| `"no pair"` | 沒有配對或更好的牌型 |
| `"pair"` | 兩顆骰子數值相同 |
| `"two pair"` | 兩組不同的配對 |
| `"three of a kind"` | 三顆骰子數值相同 |
| `"small straight"` | 四個連續數字 |
| `"large straight"` | 五個連續數字 |
| `"full house"` | 三顆相同加一對 |
| `"four of a kind"` | 四顆骰子數值相同 |
| `"five of a kind"` | 五顆骰子數值全部相同 |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` 應該傳回 `"five of a kind"`。

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` 應該傳回 `"four of a kind"`。

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` 應該傳回 `"large straight"`。

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` 應該傳回 `"three of a kind"`。

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` 應該傳回 `"pair"`。

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` 應該傳回 `"no pair"`。

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` 應該傳回 `"small straight"`。

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` 應該傳回 `"full house"`。

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` 應該傳回 `"two pair"`。

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
