---
id: 6a2037a68a0bc2aef0006002
title: "挑戰 340：Pig Latin 轉換器"
challengeType: 28
dashedName: challenge-340
---

# --description--

給定一個字串，使用以下規則將它轉換成 Pig Latin：

- 如果一個字以母音（`"a"`、`"e"`、`"i"`、`"o"` 或 `"u"`）開頭，則在字尾加上 `"way"`。例如，`"universe"` 會轉換成 `"universeway"`。
- 如果一個字以一個或多個子音開頭，將這些子音移到字尾並加上 `"ay"`。例如，`"hello"` 會轉換成 `"ellohay"`。
- 保持第一個字母的大小寫。例如，`"Hello"` 會轉換成 `"Ellohay"`。

# --hints--

`pigLatin("universe")` 應該傳回 `"universeway"`。

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` 應該傳回 `"ellohay"`。

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` 應該傳回 `"ellohay universeway"`。

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` 應該傳回 `"Ellohay universeway"`。

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` 應該傳回 `"Igpay Atinlay isway unfay"`。

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` 應該傳回 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`。

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
