---
id: 6a2037a68a0bc2aef0006000
title: "挑戰 338：寵物年齡計算器"
challengeType: 28
dashedName: challenge-338
---

# --description--

給定寵物的種類和以人類年齡計算的年齡，使用以下換算表傳回等價的寵物年齡：

| 寵物 | 乘數 |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` 應該傳回 `35`。

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` 應該傳回 `54`。

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` 應該傳回 `24`。

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` 應該傳回 `120`。

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` 應該傳回 `60`。

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` 應該傳回 `12`。

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` 應該傳回 `5`。

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
