---
id: 6a2037a68a0bc2aef0006006
title: "挑戰 344：黃金比例"
challengeType: 28
dashedName: challenge-344
---

# --description--

給定兩個數字，判斷它們的比例是否接近黃金比例。

- 使用黃金比例 `1.618`
- 容許誤差為 `0.01`

# --hints--

`isGoldenRatio(21, 34)` 應該傳回 `true`。

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` 應該傳回 `false`。

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` 應該傳回 `true`。

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` 應該傳回 `false`。

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` 應該傳回 `true`。

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` 應該傳回 `false`。

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
