---
id: 6a2037a68a0bc2aef0006005
title: "挑戰 343：電梯停靠順序"
challengeType: 28
dashedName: challenge-343
---

# --description--

給定電梯目前所在樓層的數字和一個請求樓層的陣列，傳回電梯應該拜訪這些樓層的順序陣列，以最小化行駛樓層數。

- 若有平手，先往上
- 電梯第一次經過有請求的樓層時必須停靠

# --hints--

`elevatorStops(5, [2, 8, 3, 9])` 應該傳回 `[3, 2, 8, 9]`。

```js
assert.deepEqual(elevatorStops(5, [2, 8, 3, 9]), [3, 2, 8, 9]);
```

`elevatorStops(6, [2, 10, 8, 3, 1, 9])` 應該傳回 `[8, 9, 10, 3, 2, 1]`。

```js
assert.deepEqual(elevatorStops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1]);
```

`elevatorStops(1, [4, 8, 3, 6, 9])` 應該傳回 `[3, 4, 6, 8, 9]`。

```js
assert.deepEqual(elevatorStops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9]);
```

`elevatorStops(12, [6, 10, 7, 3, 1, 4])` 應該傳回 `[10, 7, 6, 4, 3, 1]`。

```js
assert.deepEqual(elevatorStops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1]);
```

`elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` 應該傳回 `[10, 9, 8, 6, 5, 2, 12, 19, 23]`。

```js
assert.deepEqual(elevatorStops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23]);
```

# --seed--

## --seed-contents--

```js
function elevatorStops(currentFloor, stops) {

  return currentFloor;
}
```

# --solutions--

```js
function elevatorStops(currentFloor, stops) {
  stops = stops.sort((a, b) => a - b);
  const above = stops.filter(s => s > currentFloor);
  const below = stops.filter(s => s < currentFloor).reverse();
  const upDistance = above.length ? above[above.length - 1] - currentFloor : Infinity;
  const downDistance = below.length ? currentFloor - below[below.length - 1] : Infinity;
  return downDistance < upDistance ? [...below, ...above] : [...above, ...below];
}
```
