---
id: 6a2037a68a0bc2aef0006001
title: "挑戰 339：陣列分塊"
challengeType: 28
dashedName: challenge-339
---

# --description--

給定一個陣列和分塊大小，傳回將陣列拆分成該大小子陣列的結果。

- 如果陣列無法平均分割，最後一個分塊可能會比較小。

# --hints--

`chunkArray([1, 2, 3, 4, 5, 6], 3)` 應該傳回 `[[1, 2, 3], [4, 5, 6]]`。

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]]);
```

`chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` 應該傳回 `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`。

```js
assert.deepEqual(chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]);
```

`chunkArray([1, 2, 3, 4, 5], 3)` 應該傳回 `[[1, 2, 3], [4, 5]]`。

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]]);
```

`chunkArray(["a", "b", "c", "d", "e"], 1)` 應該傳回 `[["a"], ["b"], ["c"], ["d"], ["e"]]`。

```js
assert.deepEqual(chunkArray(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]]);
```

`chunkArray([1, 2, 3], 5)` 應該傳回 `[[1, 2, 3]]`。

```js
assert.deepEqual(chunkArray([1, 2, 3], 5), [[1, 2, 3]]);
```

# --seed--

## --seed-contents--

```js
function chunkArray(arr, size) {

  return arr;
}
```

# --solutions--

```js
function chunkArray(arr, size) {
  const result = [];
  for (let i = 0; i < arr.length; i += size) {
    result.push(arr.slice(i, i + size));
  }
  return result;
}
```
