---
id: 6a2037a68a0bc2aef0006007
title: "挑戰 345：字詞混合器"
challengeType: 28
dashedName: challenge-345
---

# --description--

給定兩個字，傳回一個新字，方法是將第一個字的前半部與第二個字的後半部結合起來。

- 對於長度為奇數的字，前半部為較短的那一半。

# --hints--

`blendWords("turtle", "toucan")` 應該傳回 `"turcan"`。

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` 應該傳回 `"chipingo"`。

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` 應該傳回 `"falican"`。

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` 應該傳回 `"hyana"`。

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` 應該傳回 `"scorilla"`。

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` 應該傳回 `"platerine"`。

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
