---
id: 6a2037a68a0bc2aef0006004
title: "挑戰 342：骰子機率"
challengeType: 28
dashedName: challenge-342
---

# --description--

給定要擲的六面骰數量和一個目標和，傳回以字串格式 `"1 in X"` 表示擲出該和的機率。

- 骰子數量介於 1 到 6 之間。
- 目標和在給定的骰子數量下總是可達成的。
- 將 `"X"` 四捨五入到最接近的整數。

# --hints--

`getOdds(1, 5)` 應該傳回 `"1 in 6"`。

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` 應該傳回 `"1 in 12"`。

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` 應該傳回 `"1 in 8"`。

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` 應該傳回 `"1 in 65"`。

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` 應該傳回 `"1 in 111"`。

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` 應該傳回 `"1 in 7776"`。

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
