---
id: 6a2037a68a0bc2aef0005fff
title: "挑戰 337：計數器"
challengeType: 28
dashedName: challenge-337
---

# --description--

給定一串計數記號字串，傳回所表述的總數。

- 每個豎線 `"|"` 代表一個計數。
- 每第五個記號用斜線 `"/"` 表示，完成一個五個一組的群組（`"||||/"`）。
- 群組之間以空格分隔。

# --hints--

`getTallyCount("||||")` 應該傳回 `4`。

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` 應該傳回 `5`。

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` 應該傳回 `8`。

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` 應該傳回 `17`。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` 應該傳回 `41`。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
