---
id: 6a22d77ddf034bc4e35b1d57
title: "挑戰 350：字母距離"
challengeType: 28
dashedName: challenge-350
---

# --description--

給定兩個長度相等的字串，傳回每對字元之間最短距離的總和。

- 輸入只會包含小寫字母
- 字母表被視為一個圓圈，因此 `a` 和 `z` 之間的距離是 1。

# --hints--

`letterDistance("abc", "bcd")` 應該傳回 `3`。

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` 應該傳回 `9`。

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` 應該傳回 `10`。

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` 應該傳回 `43`。

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` 應該傳回 `47`。

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` 應該傳回 `55`。

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
