---
id: 6a22d77ddf034bc4e35b1d5d
title: "挑戰 356：魔術方塊解題器"
challengeType: 28
dashedName: challenge-356
---

# --description--

給定一個 3x3 的格子，其中有一個數字遺失（以 `0` 表示），傳回能完成魔術方塊的遺失數字，若不存在有效數字則傳回 `"impossible"`。

魔術方塊是一種格子，其中每一列、行和對角線的數字加總都相同。

# --hints--

`solveMagicSquare([[2, 7, 6], [9, 0, 1], [4, 3, 8]])` 應該傳回 `5`。

```js
assert.equal(solveMagicSquare([[2, 7, 6], [9, 0, 1], [4, 3, 8]]), 5);
```

`solveMagicSquare([[0, 14, 12], [18, 10, 2], [8, 6, 16]])` 應該傳回 `4`。

```js
assert.equal(solveMagicSquare([[0, 14, 12], [18, 10, 2], [8, 6, 16]]), 4);
```

`solveMagicSquare([[12, 17, 16], [19, 0, 10], [14, 13, 18]])` 應該傳回 `"impossible"`。

```js
assert.equal(solveMagicSquare([[12, 17, 16], [19, 0, 10], [14, 13, 18]]), "impossible");
```

`solveMagicSquare([[15, 35, 31], [43, 27, 11], [23, 19, 0]])` 應該傳回 `39`。

```js
assert.equal(solveMagicSquare([[15, 35, 31], [43, 27, 11], [23, 19, 0]]), 39);
```

`solveMagicSquare([[26, 41, 14], [47, 35, 0], [32, 29, 44]])` 應該傳回 `"impossible"`。

```js
assert.equal(solveMagicSquare([[26, 41, 14], [47, 35, 0], [32, 29, 44]]), "impossible");
```

# --seed--

## --seed-contents--

```js
function solveMagicSquare(grid) {

  return grid;
}
```

# --solutions--

```js
function solveMagicSquare(grid) {
  const flat = grid.flat();
  const zeroIndex = flat.indexOf(0);
  const row = Math.floor(zeroIndex / 3);
  const col = zeroIndex % 3;

  const target = grid.find(r => !r.includes(0))
    .reduce((sum, n) => sum + n, 0);

  const missing = target - grid[row].reduce((sum, n) => sum + n, 0);
  const filled = grid.map((r, ri) =>
    r.map((n, ci) => (ri === row && ci === col) ? missing : n)
  );

  const check = (arr) => arr.reduce((sum, n) => sum + n, 0) === target;

  for (let i = 0; i < 3; i++) {
    if (!check(filled[i])) return "impossible";
    if (!check(filled.map(r => r[i]))) return "impossible";
  }
  if (!check([filled[0][0], filled[1][1], filled[2][2]])) return "impossible";
  if (!check([filled[0][2], filled[1][1], filled[2][0]])) return "impossible";

  return missing;
}
```
