---
id: 6a22d77ddf034bc4e35b1d5a
title: "挑戰 353：對比評級 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

給定兩個相對亮度值和一個布林值表示文字是否為大型，使用以下方法傳回 WCAG 對比評級：

透過在每個亮度值上加 0.05，然後將較亮的值除以較暗的值來計算對比比率。較亮的值永遠是第一個引數。

根據對比比率使用以下表格傳回評級：

| 評級 | 一般文字 | 大型文字 |
|--------|-------------|------------|
| `"AAA"` | 7.0 以上 | 4.5 以上 |
| `"AA"` | 4.5 以上 | 3.0 以上 |
| `"Fail"` | 低於 4.5 | 低於 3.0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
