---
id: 6a22d77ddf034bc4e35b1d54
title: "挑戰 347：博弈論"
challengeType: 28
dashedName: challenge-347
---

# --description--

給定兩個相同長度的字串，分別表述兩位玩家在一個遊戲中的策略，傳回分數陣列 `[player1, player2]`。

- 給定的字串只會包含兩個字母之一：`"C"`（合作）或 `"D"`（背叛）。
- 每個字元代表一回合，分數計算如下：
  - 若雙方都合作，雙方各得 3 分。
  - 若雙方都背叛，雙方各得 1 分。
  - 若一方背叛另一方合作，背叛者得 5 分，合作方得 0 分。

# --hints--

`playGame("CCCC", "CCCC")` 應傳回 `[12, 12]`。

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` 應傳回 `[4, 4]`。

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` 應傳回 `[5, 10]`。

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` 應傳回 `[24, 34]`。

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` 應傳回 `[66, 21]`。

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
