---
id: 6a22d77ddf034bc4e35b1d56
title: "挑戰 349：手機信號"
challengeType: 28
dashedName: challenge-349
---

# --description--

給定一個包含三個基地台讀數的網格，判斷手機的位置。

- 網格中的每個格子要麼是 `0`（無基地台），要麼是正整數，表示從該基地台到手機的格子數，計算方式為直線距離：水平、垂直或對角線。
- 傳回距離三個基地台格子數都正確的那個格子的 `[row, col]`。
- 總是只有一個解法。

# --hints--

`findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` 應該傳回 `[1, 2]`。

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

`findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` 應該傳回 `[2, 1]`。

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

`findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` 應該傳回 `[2, 2]`。

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

`findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` 應該傳回 `[3, 4]`。

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

`findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` 應該傳回 `[3, 3]`。

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
