---
id: 6a22d77ddf034bc4e35b1d59
title: "挑戰 352：對比評級 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

給定一個對比度比率和一個布林值表示文字是否為大字，請根據下表傳回 WCAG 評級：

| 評級 | 一般文字 | 大字 |
|--------|-------------|------------|
| `"AAA"` | 7.0 以上 | 4.5 以上 |
| `"AA"` | 4.5 以上 | 3.0 以上 |
| `"Fail"` | 低於 4.5 | 低於 3.0 |

# --hints--

`getContrastRating("7.5", false)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
