---
id: 6a22d77ddf034bc4e35b1d5b
title: "挑戰 354：對比評級 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

給定兩個代表 RGB 值的陣列和一個布林值表示文字是否為大字，使用以下方法傳回 WCAG 對比評級：

首先，將每個 RGB 值轉換為相對亮度：

- 將每個通道 `[R, G, B]` 除以 255，取得介於 `0` 和 `1` 之間的值
- 對每個通道套用伽瑪校正公式：
  - 如果通道值小於或等於 `0.04045`：`channel / 12.92`
  - 否則：`((channel + 0.055) / 1.055) ^ 2.4`
- 計算亮度：`0.2126 * R + 0.7152 * G + 0.0722 * B`

接著，將每個亮度值加上 0.05，再用較亮的亮度除以較暗的亮度。較亮的亮度值永遠是第一個引數。

根據對比比率，使用下表傳回評級：

| 評級 | 一般文字 | 大字文字 |
|--------|-------------|------------|
| `"AAA"` | 7.0 以上 | 4.5 以上 |
| `"AA"` | 4.5 以上 | 3.0 以上 |
| `"Fail"` | 低於 4.5 | 低於 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` 應該傳回 `"AAA"`。

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` 應該傳回 `"AA"`。

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` 應該傳回 `"Fail"`。

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
