---
id: 6a22d77ddf034bc4e35b1d58
title: "挑戰 351：連乘數"
challengeType: 28
dashedName: challenge-351
---

# --description--

給定一個數字，判斷它是否為連乘數。

連乘數是兩個連續整數的乘積。例如，6 是連乘數，因為 2 * 3 = 6。

# --hints--

`isPronic(6)` 應該傳回 `true`。

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` 應該傳回 `false`。

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` 應該傳回 `true`。

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` 應該傳回 `true`。

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` 應該傳回 `false`。

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` 應該傳回 `true`。

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
