---
id: 6a26df3e2988bcdded204895
title: "挑戰 359：高爾夫差點計算器"
challengeType: 28
dashedName: challenge-359
---

# --description--

給定一個高爾夫得分的陣列和一個對應的球場標準桿陣列，使用以下方法傳回球員的差點指數：

- 計算每一輪的差異值，方法是用得分減去標準桿，然後傳回所有差異值的平均數，並四捨五入到小數點第一位。

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` 應該傳回 `0`。

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` 應該傳回 `6`。

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` 應該傳回 `8.3`。

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` 應該傳回 `8.8`。

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` 應該傳回 `11.7`。

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
