---
id: 6a26df3e2988bcdded204894
title: "挑戰 358：表情符號翻譯器"
challengeType: 28
dashedName: challenge-358
---

# --description--

給定一串表情符號，請使用下表傳回對應的片語：

| Emoji | 單字 |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

將單字以空格分隔後傳回。

# --hints--

`getEmojiPhrase("🪨⭐")` 應該傳回 `"rock star"`。

```js
assert.equal(getEmojiPhrase("🪨⭐"), "rock star");
```

`getEmojiPhrase("🥵🐕")` 應該傳回 `"hot dog"`。

```js
assert.equal(getEmojiPhrase("🥵🐕"), "hot dog");
```

`getEmojiPhrase("👶🦈")` 應該傳回 `"baby shark"`。

```js
assert.equal(getEmojiPhrase("👶🦈"), "baby shark");
```

`getEmojiPhrase("⭐🐟")` 應該傳回 `"star fish"`。

```js
assert.equal(getEmojiPhrase("⭐🐟"), "star fish");
```

`getEmojiPhrase("🧊🧊👶")` 應該傳回 `"ice ice baby"`。

```js
assert.equal(getEmojiPhrase("🧊🧊👶"), "ice ice baby");
```

`getEmojiPhrase("🐱🐟🍲")` 應該傳回 `"cat fish soup"`。

```js
assert.equal(getEmojiPhrase("🐱🐟🍲"), "cat fish soup");
```

# --seed--

## --seed-contents--

```js
function getEmojiPhrase(str) {

  return str;
}
```

# --solutions--

```js
function getEmojiPhrase(str) {
  const table = {
    '👶': 'baby',
    '🐱': 'cat',
    '🐕': 'dog',
    '🐟': 'fish',
    '🥵': 'hot',
    '🧊': 'ice',
    '🪨': 'rock',
    '🦈': 'shark',
    '🍲': 'soup',
    '⭐': 'star',
  };

  return [...str].map(emoji => table[emoji]).join(' ');
}
```
