---
id: 6a26df3e2988bcdded204897
title: "挑戰 361：口語時間"
challengeType: 28
dashedName: challenge-361
---

# --description--

給定一個指針式時鐘的時針和分針角度（以 12 點方向順時針為單位的度數），傳回口語化的時間。

將分針角度轉換成分鐘（360° = 60 分鐘），然後使用以下規則：

| 分鐘 | 口語說法 |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29（不含 15） | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59（不含 45） | `"X minutes to Z"`（其中 `X` 是 60 減去分鐘數）|

`Y` 是目前的時數，`Z` 是下一個時數，兩者皆由時針角度推算（360° = 12 小時）。

注意：指針角度可能不會剛好落在整數上，請考慮以某種方式進行四捨五入。

# --hints--

`getSpokenTime(90, 0)` 應該傳回 `"3 o'clock"`。

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` 應該傳回 `"20 minutes past 5"`。

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` 應該傳回 `"half past 8"`。

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` 應該傳回 `"quarter past 2"`。

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` 應該傳回 `"20 minutes to 7"`。

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` 應該傳回 `"quarter to 11"`。

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` 應該傳回 `"5 minutes to 4"`。

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
