---
id: 6a26df95efa55a2524399744
title: "挑戰 363：桶子填色 2"
challengeType: 28
dashedName: challenge-363
---

# --description--

給定一個由單一字母顏色字串組成的二維網格和一個標的顏色，傳回使整個網格變成標的顏色所需的最少洪水填色「點擊」次數。

- 每次點擊會將被點擊的格子顏色和整個與該格子顏色相連的區域改成標的顏色。
- 格子只會水平和垂直相連（不包含斜角相連）。

# --hints--

`bucketFill([["R", "R"], ["R", "R"]], "G")` 應該傳回 `1`。

```js
assert.equal(bucketFill([["R", "R"], ["R", "R"]], "G"), 1);
```

`bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` 應該傳回 `0`。

```js
assert.equal(bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0);
```

`bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` 應該傳回 `3`。

```js
assert.equal(bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3);
```

`bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` 應該傳回 `5`。

```js
assert.equal(bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5);
```

`bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` 應該傳回 `12`。

```js
assert.equal(bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, targetColor) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, targetColor) {
  const rows = grid.length;
  const cols = grid[0].length;
  const visited = new Set();

  function dfs(r, c, color) {
    const key = `${r},${c}`;
    if (r < 0 || r >= rows || c < 0 || c >= cols) return;
    if (visited.has(key)) return;
    if (grid[r][c] !== color) return;
    visited.add(key);
    dfs(r + 1, c, color);
    dfs(r - 1, c, color);
    dfs(r, c + 1, color);
    dfs(r, c - 1, color);
  }

  let clicks = 0;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const key = `${r},${c}`;
      if (!visited.has(key) && grid[r][c] !== targetColor) {
        dfs(r, c, grid[r][c]);
        clicks++;
      }
    }
  }

  return clicks;
}
```
