---
id: 6a26df3e2988bcdded204893
title: "挑戰 357：食物鏈"
challengeType: 28
dashedName: challenge-357
---

# --description--

給定一個包含 `[predator, prey]` 配對的陣列，傳回從頂級掠食者到最底層的食物鏈。

- 頂級掠食者是從未成為其他動物獵物的動物。
- 將串鏈以字串陣列形式傳回。

# --hints--

`getFoodChain([["cat", "mouse"]])` 應該傳回 `["cat", "mouse"]`。

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` 應該傳回 `["wolf", "deer", "grass"]`。

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` 應該傳回 `["hawk", "snake", "frog", "fly"]`。

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` 應該傳回 `["eagle", "fox", "rabbit", "grass"]`。

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` 應該傳回 `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`。

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
