---
id: 6a26df3e2988bcdded204896
title: "挑戰 360：口語化時間長度"
challengeType: 28
dashedName: challenge-360
---

# --description--

給定一個秒數，傳回口語化的英文時間長度。

- 將時間長度拆分成小時、分鐘和秒數。
- 跳過任何為零的數值。
- 適當使用單數或複數（`"1 hour"`、`"2 hours"`）。
- 如果有多個單位，最後兩個單位用 `"and"` 連接，倒數第二和倒數第三個單位用逗號分隔（`"1 hour, 2 minutes and 3 seconds"`）。

# --hints--

`getSpokenDuration(3723)` 應傳回 `"1 hour, 2 minutes and 3 seconds"`。

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` 應傳回 `"2 hours, 1 minute and 35 seconds"`。

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` 應傳回 `"2 hours, 22 minutes and 1 second"`。

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` 應傳回 `"7 minutes and 15 seconds"`。

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` 應傳回 `"4 hours and 55 seconds"`。

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` 應傳回 `"20 hours"`。

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` 應傳回 `"1 second"`。

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
