---
id: 6a26df95efa55a2524399745
title: "挑戰 364：兩桶之間"
challengeType: 28
dashedName: challenge-364
---

# --description--

給定兩桶油漆，每桶都有一個 RGB 顏色和一個滿度層級，傳回混合後的 RGB 顏色，格式為三個整數的陣列。

- 每桶是 JavaScript 的物件或 Python 的字典，包含一個 `color` 屬性（由三個整數 `[r, g, b]` 組成的陣列）和一個 `fullness` 屬性（範圍 0–100）。
- 混合顏色是根據滿度層級對兩個顏色的每個通道做加權平均，並將每個通道四捨五入到最接近的整數。

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` 應該傳回 `[125, 125, 125]`。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` 應該傳回 `[200, 200, 200]`。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` 應該傳回 `[100, 150, 200]`。

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` 應該傳回 `[114, 184, 170]`。

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` 應該傳回 `[68, 163, 123]`。

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
