---
id: 66d75dd0aa65a71600dc669b
title: Створіть програму для управління інвентарем
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

У цій лабораторній роботі ви створите програму для управління інвентарем, яка дозволить додавати, оновлювати, знаходити та видаляти товари з інвентарю. Ви будете використовувати масив об’єктів для представлення інвентарю, де кожен об’єкт матиме ключі name і quantity.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні оголосити порожній масив з іменем `inventory`, який зберігатиме об’єкти товарів з ключем `name` зі значенням рядка в нижньому регістрі та ключем `quantity` зі значенням цілого числа.
1. Ви повинні створити функцію з іменем `findProductIndex`, яка прийматиме назву товару як аргумент і повертатиме індекс відповідного об’єкта товару всередині масиву `inventory`. Функція завжди повинна використовувати назву товару в нижньому регістрі для пошуку. Якщо товар не знайдено, функція повинна повертати `-1`.
1. Ви повинні створити функцію з іменем `addProduct`, яка прийматиме об’єкт товару як аргумент.
1. Якщо товар уже присутній в інвентарі, функція `addProduct` повинна оновити його значення `quantity`, додавши кількість, передану у функцію, до поточного значення `quantity`, і вивести в консоль назву товару, за якою йде пробіл і `quantity updated`.
1. Якщо товар відсутній в інвентарі, функція `addProduct` повинна додати товар у масив `inventory` і вивести в консоль назву товару, за якою йде пробіл і `added to inventory`.
1. Ви повинні створити функцію з іменем `removeProduct`, яка прийматиме назву товару та кількість як аргументи.
1. Функція `removeProduct` повинна відняти передану кількість від відповідного об’єкта товару в інвентарі і вивести в консоль рядок `Remaining <product-name> pieces: <product-quantity>`, де `<product-name>` замінюється на назву товару, а `<product-quantity>` — на кількість товару.
1. Якщо кількість після віднімання дорівнює нулю, `removeProduct` повинна видалити об’єкт товару з інвентарю. Якщо кількість у інвентарі недостатня для віднімання, функція `removeProduct` повинна вивести в консоль рядок `Not enough <product-name> available, remaining pieces: <product-quantity>`.
1. Якщо товар для видалення відсутній в інвентарі, функція `removeProduct` повинна вивести в консоль `<product-name> not found`.

**Примітка:** Щоб уникнути конфліктів, залишайте лише ті виводи в консоль, які згадані в історіях користувача, під час запуску тестів.

# --hints--

Ви повинні оголосити порожній масив з іменем `inventory`.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

Ви повинні мати функцію з іменем `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` має повертати індекс об’єкта, у якого `name` дорівнює `"flour"` всередині масиву `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` має повертати індекс об’єкта, у якого `name` дорівнює `"flour"` всередині масиву `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` має повертати `-1`, коли не знайдено об’єкта, у якого `name` дорівнює `"flour"` всередині масиву `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

Ви повинні мати функцію з іменем `addProduct`.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` має додати `5` до значення `quantity` об’єкта, у якого `name` дорівнює `"flour"` у масиві `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` має вивести в консоль `flour quantity updated`, коли знайдено об’єкт, у якого `name` дорівнює `"flour"` у масиві `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` має додати `{name: "flour", quantity: 5}` до масиву `inventory`, коли не знайдено об’єкта, у якого `name` дорівнює `"flour"` в інвентарі.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` має вивести в консоль `flour added to inventory`, коли не знайдено об’єкта, у якого `name` дорівнює `"flour"` в інвентарі.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

Ви повинні створити функцію з іменем `removeProduct`.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` має вивести в консоль `flour not found`, коли не знайдено об’єкта, у якого `name` дорівнює `"flour"` у масиві `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` має відняти `5` від значення `quantity` об’єкта, у якого `name` дорівнює `"flour"` всередині масиву `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` має вивести в консоль `Remaining flour pieces: 15`, коли `inventory` дорівнює `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

Якщо `quantity` після віднімання дорівнює нулю, `removeProduct` має видалити об’єкт товару з інвентарю.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` має вивести в консоль `Not enough flour available, remaining pieces: 5`, коли `inventory` дорівнює `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
