---
id: 686ccc2c8b967e17ab18d593
title: 二分法の実装
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

二分法は、二分探索法とも呼ばれ、実数値関数の根を見つけるために二分探索を使います。平方根が存在する区間を徐々に狭めていき、指定された許容誤差内の値に収束させます。

例えば、許容誤差が `0.01` の場合、二分法は上限と下限の差が `0.01` 以下になるまで区間を半分にし続けます。

このラボでは、二分法を使って数値の平方根を求める関数を実装します。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `square_root_bisection` という名前の関数を3つのパラメータで定義してください。
    - 平方根を求めたい数値。
    - 許容誤差で、結果の許容できる誤差の範囲です。デフォルト値を設定してください。
    - 実行する最大反復回数。デフォルト値を設定してください。

1. `square_root_bisection` 関数は以下を行う必要があります。
    - 関数に渡された数値が負の場合、メッセージ `ValueError` を持つ `Square root of negative number is not defined in real numbers` を発生させてください。
    - `0` と `1` の場合、メッセージ `The square root of [number] is [number]` を出力し、その数値自体を平方根として返してください。
    - その他の正の数の場合、メッセージ `The square root of [square_target] is approximately [root]` とともに近似平方根を出力し、計算した平方根の値を返してください。
    - 許容誤差の条件を満たす値が見つからなければ、失敗メッセージ `Failed to converge within [maximum] iterations` を出力し、`None` を返してください。

**注意**: このラボではモジュールをインポートできません。

# --hints--

モジュールをインポートしないでください。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

`square_root_bisection` という名前の関数を用意してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

`square_root_bisection` 関数は3つのパラメータを持つ必要があります。

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

許容誤差と最大反復回数にデフォルト値を設定してください。

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

`square_root_bisection` 関数は、渡された数値が負の場合、メッセージ `ValueError` を持つ `Square root of negative number is not defined in real numbers` を発生させる必要があります。

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` は `0` を返す必要があります。

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` は `The square root of 0 is 0` を出力する必要があります。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` は `0.03162267660168379` と `0.031622876601683794` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` は `The square root of 0.001 is approximately X` を出力する必要があります。ここで `X` は `0.03162267660168379` と `0.031622876601683794` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` は `0.4999999` と `0.5000001` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` は `The square root of 0.25 is approximately X` を出力する必要があります。ここで `X` は `0.4999999` と `0.5000001` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` は `1` を返す必要があります。

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` は `The square root of 1 is 1` を出力する必要があります。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` は `8.999` と `9.001` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` は `The square root of 81 is approximately X` を出力する必要があります。ここで `X` は `8.999` と `9.001` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` は `14.999` と `15.001` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` は `The square root of 225 is approximately X` を出力する必要があります。ここで `X` は `14.999` と `15.001` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` は `14.99999` と `15.00001` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` は `The square root of 225 is approximately X` を出力する必要があります。ここで `X` は `14.99999` と `15.00001` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` は `14.9999999` と `15.0000001` の間の数値を返す必要があります。

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` は `The square root of 225 is approximately X` を出力する必要があります。ここで `X` は `14.9999999` と `15.0000001` の間の数値です。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` は `None` を返す必要があります。

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` は `Failed to converge within 10 iterations` を出力する必要があります。

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
