---
id: 684aaf9ec670c68d20efd0d0
title: Einen Benutzerkonfigurationsmanager erstellen
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

In diesem Labor erstellen Sie einen Benutzerkonfigurationsmanager, der es Benutzenden ermöglicht, ihre Einstellungen wie Thema, Sprache und Benachrichtigungen zu verwalten. Sie implementieren Funktionen zum Hinzufügen, Aktualisieren, Löschen und Anzeigen von Benutzereinstellungen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollen eine Funktion namens `add_setting` mit zwei Parametern definieren, die ein Wörterbuch von Einstellungen und ein Tuple mit einem Schlüssel-Wert-Paar repräsentieren.

1. Die `add_setting`-Funktion soll:
    
   - Den Schlüssel und den Wert in Kleinbuchstaben umwandeln.
   - Wenn die Schlüsseleinstellung existiert, `Setting '[key]' already exists! Cannot add a new setting with this name.` zurückgeben.
   - Wenn die Schlüsseleinstellung nicht existiert, das Schlüssel-Wert-Paar zum gegebenen Wörterbuch der Einstellungen hinzufügen und `Setting '[key]' added with value '[value]' successfully!` zurückgeben.
   - Die zurückgegebenen Nachrichten sollen den Schlüssel und Wert in Kleinbuchstaben enthalten.

1. Sie sollen eine Funktion namens `update_setting` mit zwei Parametern definieren, die ein Wörterbuch von Einstellungen und ein Tuple mit einem Schlüssel-Wert-Paar repräsentieren.

1. Die `update_setting`-Funktion soll:
       
   - Den Schlüssel und den Wert in Kleinbuchstaben umwandeln.
   - Wenn die Schlüsseleinstellung existiert, ihren Wert im gegebenen Wörterbuch der Einstellungen aktualisieren und `Setting '[key]' updated to '[value]' successfully!` zurückgeben.
   - Wenn die Schlüsseleinstellung nicht existiert, `Setting '[key]' does not exist! Cannot update a non-existing setting.` zurückgeben.
   - Die zurückgegebenen Nachrichten sollen den Schlüssel und Wert in Kleinbuchstaben enthalten.


1. Sie sollen eine Funktion namens `delete_setting` mit zwei Parametern definieren, die ein Wörterbuch von Einstellungen und einen Schlüssel repräsentieren.

1. Die `delete_setting`-Funktion soll:
    
   - Den übergebenen Schlüssel in Kleinbuchstaben umwandeln.
   - Wenn die Schlüsseleinstellung existiert, das Schlüssel-Wert-Paar aus dem gegebenen Wörterbuch der Einstellungen entfernen und `Setting '[key]' deleted successfully!` zurückgeben.
   - Wenn die Schlüsseleinstellung nicht existiert, `Setting not found!` zurückgeben.
   - Die zurückgegebenen Nachrichten sollen den Schlüssel in Kleinbuchstaben enthalten.

1. Sie sollen eine Funktion namens `view_settings` mit einem Parameter definieren, der ein Wörterbuch von Einstellungen repräsentiert.

1. Die `view_settings`-Funktion soll:
   
   - `No settings available.` zurückgeben, wenn das gegebene Wörterbuch der Einstellungen leer ist.
   - Wenn das Wörterbuch Einstellungen enthält, einen String zurückgeben, der die Einstellungen anzeigt. Der String soll mit `Current User Settings:` beginnen, gefolgt von den Schlüssel-Wert-Paaren, jeweils in einer neuen Zeile und mit dem Schlüssel großgeschrieben. Zum Beispiel sollte `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` Folgendes zurückgeben:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Zum Testen des Codes sollen Sie ein Wörterbuch namens `test_settings` erstellen, um einige Benutzerkonfigurationseinstellungen zu speichern.

# --hints--

Sie sollten ein Wörterbuch namens `test_settings` erstellen und einige Werte hinzufügen.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Sie sollten eine Funktion namens `add_setting` definieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

Die `add_setting`-Funktion sollte zwei Parameter haben.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` sollte den Schlüssel in Kleinbuchstaben umwandeln.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` sollte den Wert in Kleinbuchstaben umwandeln.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` sollte die Fehlermeldung `Setting 'theme' already exists! Cannot add a new setting with this name.` zurückgeben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` sollte ein neues Schlüssel-Wert-Paar hinzufügen und die Erfolgsmeldung `Setting 'volume' added with value 'high' successfully!` zurückgeben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` sollte das gegebene Schlüssel-Wert-Paar korrekt zum Wörterbuch hinzufügen.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Sie sollten eine Funktion namens `update_setting` definieren.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

Die `update_setting`-Funktion sollte zwei Parameter haben.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

Die `update_setting`-Funktion sollte den Schlüssel in Kleinbuchstaben umwandeln.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

Die `update_setting`-Funktion sollte den Wert in Kleinbuchstaben umwandeln.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` sollte einen bestehenden Schlüssel aktualisieren und die Erfolgsmeldung `Setting 'theme' updated to 'dark' successfully!` zurückgeben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` sollte die Fehlermeldung `Setting 'volume' does not exist! Cannot update a non-existing setting.` zurückgeben, wenn der Schlüssel nicht existiert.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` sollte das gegebene Schlüssel-Wert-Paar im Wörterbuch korrekt aktualisieren.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Sie sollten eine Funktion namens `delete_setting` definieren.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

Die `delete_setting`-Funktion sollte zwei Parameter haben.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` sollte den Schlüssel in Kleinbuchstaben umwandeln.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` sollte den bestehenden Schlüssel entfernen und die Erfolgsmeldung `Setting 'theme' deleted successfully!` zurückgeben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` sollte die Fehlermeldung `Setting not found!` zurückgeben, wenn der Schlüssel nicht existiert.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` sollte den gegebenen Schlüssel korrekt aus dem Wörterbuch entfernen.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Sie sollten eine Funktion namens `view_settings` definieren.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

Die `view_settings`-Funktion sollte einen Parameter haben.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` sollte die Nachricht `No settings available.` zurückgeben, wenn das gegebene Wörterbuch leer ist.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` sollte formatierte Einstellungen für ein nicht-leeres Wörterbuch zurückgeben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` sollte den ersten Buchstaben jedes Einstellungsnamens großschreiben.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` sollte die korrekten Ergebnisse anzeigen und mit einem Zeilenumbruch enden.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
