---
id: 66ec4c8e9878d8441956516f
title: Побудувати таблицю каталогу книг
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити `table` елемент, який перелічує інформацію про книги.
1. Ваша таблиця має містити заголовок таблиці з одним рядком.
1. Рядок у заголовку таблиці повинен містити чотири елементи заголовка таблиці з текстом `Title`, `Author`, `Genre` і `Publication Year` у такому порядку.
1. Ваша таблиця має містити тіло таблиці з принаймні п’ятьма рядками.
1. Кожен рядок у тілі таблиці повинен містити чотири елементи даних таблиці, які відображають Назву книги, Автора, Жанр і Рік видання.
1. Ваша таблиця має містити нижній колонтитул таблиці з одним рядком.
1. Рядок у нижньому колонтитулі таблиці повинен містити елемент даних таблиці, який охоплює чотири колонки і має текст `Total Books: N`, де `N` має бути замінений числом книг у вашій таблиці.

# --hints--

Ви повинні мати один `table` елемент.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Ви повинні мати один `thead` елемент всередині вашого `table` елемента.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

Ви повинні мати один `tr` елемент всередині вашого `thead` елемента.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

Ви повинні мати чотири `th` елементи в рядку вашого `thead` елемента.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

Ваші чотири `th` елементи повинні містити текст `Title`, `Author`, `Genre` і `Publication Year` у такому порядку.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

Ви повинні мати один `tbody` елемент всередині вашого `table` елемента.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

Ваш `tbody` елемент повинен містити принаймні п’ять рядків.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

Кожен рядок у вашому `tbody` елементі повинен мати рівно чотири дочірні `td` елементи.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

Кожен `td` елемент у тілі таблиці повинен містити текст з інформацією про книгу.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

Ви повинні мати один `tfoot` елемент всередині вашого `table` елемента.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

Ви повинні мати рівно один `tr` елемент у вашому `tfoot` елементі.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

Елемент `td` у рядку вашого `tfoot` елемента повинен мати атрибут `colspan`, встановлений у `4`.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

Елемент `td` у рядку вашого `tfoot` елемента повинен містити текст `Total Books: N`, де `N` — це число книг у вашій таблиці.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
