---
id: 66f55eac933ff64ce654ca74
title: Construir um Conversor de Markdown para HTML
challengeType: 25
dashedName: build-a-markdown-to-html-converter
demoType: onClick
---

# --description--

Markdown é uma linguagem de marcação usada para adicionar elementos de formatação a documentos de texto simples. Para este laboratório, todo o HTML e CSS foram fornecidos para você. Você usará JavaScript para completar o aplicativo Conversor de Markdown para HTML para que ele possa lidar com a conversão de construções básicas de Markdown em elementos HTML.

**Nota:** O resultado final não será um conversor completo de Markdown para HTML, mas você pode adicionar funcionalidades extras a ele se desejar.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `convertMarkdown` que não recebe parâmetros.
1. A função `convertMarkdown` deve usar expressões regulares para converter a entrada markdown de `#markdown-input` em HTML e deve retornar uma string contendo o código HTML.
1. A função `convertMarkdown` deve converter títulos de nível um, dois e três nos elementos correspondentes `h1`, `h2` e `h3`. Um título em markdown é indicado por tantos caracteres `#` quanto seu nível seguido de um espaço e do texto do título. Os caracteres `#` devem ser colocados no início da linha: pode haver espaços mas nenhum outro caractere antes deles.
1. A função `convertMarkdown` deve converter texto em **negrito** em elementos `strong`. O texto em negrito no markdown é indicado por um par de dois asteriscos duplos ou um par de dois underscores duplos envolvendo o texto.
1. A função `convertMarkdown` deve converter texto *itálico* em elementos `em`. Texto *itálico* em markdown é indicado por um par de asteriscos ou um par de underscores envolvendo o texto.
1. A função `convertMarkdown` deve converter imagens em elementos `img`. Uma imagem em markdown é indicada por `![texto-alt](fonte-da-imagem)`, onde `texto-alt` é o valor do atributo `alt` e `fonte-da-imagem` é o valor do atributo `src`.
1. A função `convertMarkdown` deve converter links em elementos de âncora. Um link em markdown é indicado por `[texto do link](URL)`, onde `texto do link` é o texto a ser incluído dentro das tags de âncora e `URL` é o valor do atributo `href`.
1. A função `convertMarkdown` deve converter aspas em elementos `blockquote`. Uma citação em markdown é indicada por um `>` seguido de um espaço e o texto da citação.  The `>` character should be placed at the beginning of the line: there can be spaces but no other characters before it.
1. Quando você insere texto dentro de `#markdown-input`, o código HTML bruto retornado por `convertMarkdown` deve ser exibido dentro de `#html-output`.
1. Quando você insere texto dentro de `#markdown-input`, o código HTML retornado por `convertMarkdown` deve ser renderizado dentro de `#preview`.

Nota: você deve trabalhar com o evento `input` para este projeto.

Aqui está uma tabela contendo todo o markdown que `convertMarkdown` deve ser capaz de processar e o HTML esperado após a conversão:

<table>
  <thead>
    <tr>
      <th>Markdown</th>
      <th>HTML</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td><code># heading 1</code></th>
      <td><code>&lth1&gtheading 1&lt/h1&gt</code></th>
    </tr>
    <tr>
      <td><code>## heading 2</code></th>
      <td><code>&lth2&gtheading 2&lt/h2&gt</code></th>
    </tr>
    <tr>
      <td><code>### heading 3</code></th>
      <td><code>&lth3&gtheading 3&lt/h3&gt</code></th>
    </tr>
    <tr>
      <td><code>**bold text**</code> or <code>__bold text__</code></th>
      <td><code>&ltstrong&gtbold text&lt/strong&gt</code></th>
    </tr>
    <tr>
      <td><code>*italic text*</code> or <code>_italic text_</code></th>
      <td><code>&ltem&gtitalic text&lt/em&gt</code></th>
    </tr>
    <tr>
      <td><code>![alt-text](image-source)</code></th>
      <td><code>&ltimg alt="alt-text" src="image-source"&gt</code></th>
    </tr>
    <tr>
      <td><code>[link text](URL)</code></th>
      <td><code>&lta href="URL"&gtlink text&lt/a&gt</code></th>
    </tr>
    <tr>
      <td><code>> quote</code></th>
      <td><code>&ltblockquote&gtquote&lt/blockquote&gt</code></th>
    </tr>
  </tbody>
</table>

**Nota:** Certifique-se de vincular seu arquivo JavaScript no seu HTML.

# --hints--

Você deve ter uma função chamada `convertMarkdown`.

```js
assert.isFunction(convertMarkdown);
```

Quando o valor de `#markdown-input` for `# título 1`, `convertMarkdown()` deve retornar `<h1>título 1</h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Quando o valor de `#markdown-input` for `# título 1`, `<h1>título 1</h1>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Quando o valor de `#markdown-input` for `# título 1`, um elemento `h1` com o texto `título 1` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Quando o valor de `#markdown-input` for `some text # title 1`, `convertMarkdown()` não deve converter `# title 1` em um elemento `h1`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text # title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH1 = testDiv.querySelector("h1")
assert.notExists(testH1);
assert.equal(testDiv.innerText, input.value)
```

Quando o valor de `#markdown-input` for `# título 1` seguido por `# título alternativo` em uma nova linha, `convertMarkdown()` deve retornar `<h1>título 1</h1><h1>título alternativo</h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1\n# alternate title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 1");
assert.equal(headings[1].innerText, "alternate title");
```

Quando o valor de `#markdown-input` for `## título 2`, `convertMarkdown()` deve retornar `<h2>título 2</h2>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 2");
```

Quando o valor de `#markdown-input` for `## título 2`, `<h2>título 2</h2>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

Quando o valor de `#markdown-input` for `## título 2`, um elemento `h2` com o texto `título 2` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

Quando o valor de `#markdown-input` for `some text ## title 2`, `convertMarkdown()` não deve converter `## title 2` em um elemento `h2`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH2 = testDiv.querySelector("h2")
assert.notExists(testH2);
assert.equal(testDiv.innerText, input.value)
```

Quando o valor de `#markdown-input` for `## título 2` seguido por `## título 2 alt` em uma nova linha, `convertMarkdown()` deve retornar `<h2>título 2</h2><h2>título 2 alt</h2>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2\n## title 2 alt";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 2");
assert.equal(headings[1].innerText, "title 2 alt");
```

Quando o valor de `#markdown-input` for `### título 3`, `convertMarkdown()` deve retornar `<h3>título 3</h3>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Quando o valor de `#markdown-input` for `### título 3`, `<h3>título 3</h3>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Quando o valor de `#markdown-input` for `### title 3`, um elemento `h3` com o texto `title 3` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Quando o valor de `#markdown-input` for `some text ### title 3`, `convertMarkdown()` não deve converter `### title 3` em um elemento `h3`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH3 = testDiv.querySelector("h3")
assert.notExists(testH3);
assert.equal(testDiv.innerText, input.value)
```

Quando o valor de `#markdown-input` for `### título 3` seguido de `### terceiro título` em uma nova linha, `convertMarkdown()` deve retornar `<h3>título 3</h3><h3>terceiro título</h3>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3\n### third title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 3");
assert.equal(headings[1].innerText, "third title");
```

Quando o valor de `#markdown-input` for `**este texto está em negrito**`, `convertMarkdown()` deve retornar `<strong>este texto está em negrito</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(strongs[0].innerText, "this is bold");
```

Quando o valor de `#markdown-input` for `**isto está em negrito**`, `<strong>isto está em negrito</strong>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Quando o valor de `#markdown-input` for `**isto está em negrito**`, um elemento `strong` com o texto `isto está em negrito` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Quando o valor de `#markdown-input` for `**este é negrito**` seguido por `**este também é negrito**` em uma nova linha, `convertMarkdown()` deve retornar `<strong>este é negrito</strong><strong>este também é negrito</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**\n**this is also bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

Quando o valor de `#markdown-input` for `__isto está em negrito__`, `<strong>isto está em negrito</strong>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Quando o valor de `#markdown-input` for `__isto está em negrito__`, um elemento `strong` com o texto `isto está em negrito` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Quando o valor de `#markdown-input` for `__este é negrito__` seguido por `__este também é negrito__` em uma nova linha, `convertMarkdown()` deve retornar `<strong>este é negrito</strong><strong>este também é negrito</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "__this is bold__\n__this is also bold__";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

Quando o valor de `#markdown-input` for `*isto está em itálico*`, `convertMarkdown()` deve retornar `<em>isto está em itálico</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `*isto está em itálico*`, `<em>isto está em itálico</em>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `*isto está em itálico*`, um elemento `em` com o texto `isto está em itálico` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `*este é itálico*` seguido por `*este também é itálico*` em uma nova linha, `convertMarkdown()` deve retornar `<em>este é itálico</em><em>este também é itálico</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*\n*this is also italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

Quando o valor de `#markdown-input` for `_isto está em itálico_`, `convertMarkdown()` deve retornar `<em>isto está em itálico</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `_isto está em itálico_`, `<em>isto está em itálico</em>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `_isto está em itálico_`, um elemento `em` com o texto `isto está em itálico` deve ser adicionado como filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Quando o valor de `#markdown-input` for `_isto está em itálico_` seguido por `_isto também está em itálico_` em uma nova linha, `convertMarkdown()` deve retornar `<em>isto está em itálico</em><em>isto também está em itálico</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_\n_this is also italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

Quando o valor de `#markdown-input` for `# **title 1**` ou `# __title 1__`, `convertMarkdown()` deve retornar `<h1><strong>title 1</strong></h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# **title 1**";
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = convertMarkdown();
let headings = testDiv1.querySelectorAll("h1");
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

input.value = "# __title 1__";
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = convertMarkdown();
headings = testDiv2.querySelectorAll("h1");
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Quando o valor de `#markdown-input` for `# **title 1**` ou `# __title 1__`, `<h1><strong>title 1</strong></h1>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = output.innerText;
let headings = testDiv1.querySelectorAll("h1")
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

output.innerText = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = output.innerText;
headings = testDiv2.querySelectorAll("h1")
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Quando o valor de `#markdown-input` for `# **title 1**` ou `# __title 1__`, você define o HTML interno de `#preview` para `<h1><strong>title 1</strong></h1>`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
let headings = preview.querySelectorAll("h1")
let testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

preview.innerHTML = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
headings = preview.querySelectorAll("h1")
testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Quando o valor de `#markdown-input` for `![alt-text](image-source)`, `convertMarkdown()` deve retornar `<img alt="alt-text" src="image-source">`.

```js
const input = document.querySelector("#markdown-input");
input.value = "![alt-text](image-source)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Quando o valor de `#markdown-input` for `![alt-text](image-source)`, `<img alt="alt-text" src="image-source">` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Quando o valor de `#markdown-input` for `![alt-text](image-source)`, `<img alt="alt-text" src="image-source">` deve ser anexado como um filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const imgs = preview.querySelectorAll("img");
assert.lengthOf(preview.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Quando o valor de `#markdown-input` for `![alt-text](image-source)` seguido por `![alt-text-2](image-source-2)` em uma nova linha, `convertMarkdown()` deve retornar `<img alt="alt-text" src="image-source"><img alt="alt-text-2" src="image-source-2">`.

```js
const input = document.querySelector("#markdown-input");
input.value = "![alt-text](image-source)\n![alt-text-2](image-source-2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(imgs, 2);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
assert.equal(imgs[1].alt, "alt-text-2");
assert.isTrue(imgs[1].src.endsWith("image-source-2"));
```

Quando o valor de `#markdown-input` for `[link text](URL)`, `convertMarkdown()` deve retornar `<a href="URL">link text</a>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Quando o valor de `#markdown-input` for `[link text](URL)`, `<a href="URL">link text</a>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Quando o valor de `#markdown-input` for `[link text](URL)`, `<a href="URL">link text</a>` deve ser adicionado como um filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const anchors = preview.querySelectorAll("a");
assert.lengthOf(preview.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Quando o valor de `#markdown-input` for `[link text](URL)` seguido por `[link text 2](URL2)` em uma nova linha, `convertMarkdown()` deve retornar `<a href="URL">link text</a><a href="URL2">link text 2</a>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)\n[link text 2](URL2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(anchors, 2);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
assert.isTrue(anchors[1].href.endsWith("URL2"));
assert.equal(anchors[1].innerText, "link text 2");
```

Quando o valor de `#markdown-input` for `> esta é uma citação`, `convertMarkdown()` deve retornar `<blockquote>esta é uma citação</blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Quando o valor de `#markdown-input` for `> isto é uma citação`, `<blockquote>isto é uma citação</blockquote>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Quando o valor de `#markdown-input` for `> this is a quote`, `<blockquote>this is a quote</blockquote>` deve ser anexado como um filho de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Quando o valor de `#markdown-input` for `> esta é uma citação` seguido por `> esta é outra citação` em uma nova linha, `convertMarkdown()` deve retornar `<blockquote>esta é uma citação</blockquote><blockquote>esta é outra citação</blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote\n> this is another quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(quotes, 2);
assert.equal(quotes[0].innerText, "this is a quote");
assert.equal(quotes[1].innerText, "this is another quote");
```


Quando o valor de `#markdown-input` for `some text > not a quote anymore`, `convertMarkdown()` não deve converter `> not a quote anymore` em um elemento `blockquote`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text > not a quote anymore";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote")
assert.notExists(testQuote);
assert.equal(testDiv.innerText, input.value)
```

Quando o valor de `#markdown-input` for `> **this is a *quote***`, `convertMarkdown()` deve retornar `<blockquote><strong>this is a <em>quote</em></strong></blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> **this is a *quote***";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Quando o valor de `#markdown-input` for `> **este é um *citação***`, `<blockquote><strong>este é um <em>citação</em></strong></blockquote>` deve ser exibido dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Quando o valor de `#markdown-input` for `> **this is a *quote***`, você deve definir o HTML interno de `#preview` como `<blockquote><strong>this is a <em>quote</em></strong></blockquote>`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
const strongs = preview.querySelectorAll("blockquote>strong");
const italics = preview.querySelectorAll("blockquote>strong>em");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Você deve ter apenas um elemento `script` no seu HTML.

```js
assert.lengthOf(document.querySelectorAll("script"), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <div id="container">
        <div class="container">
            <h2>Markdown Input:</h2>
            <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
        </div>
        <div class="container">
            <h2>Raw HTML Output:</h2>
            <div id="html-output"></div>
        </div>
        <div class="container">
            <h2>HTML Preview:</h2>
            <div id="preview"></div>
        </div>
    </div>
</body>

</html>
```

```css
* {
     box-sizing: border-box;
}
 body {
     font-family: Arial, sans-serif;
     padding: 20px;
}
 #markdown-input {
     width: 100%;
     height: 100px;
}
 #html-output, #preview {
     height: 100px;
     display: inline-block;
     width: 100%;
     border: 1px solid #ccc;
     padding: 10px;
     margin: auto;
     white-space: pre-wrap;
     background-color: #f9f9f9;
}
 @media (min-width: 600px) {
     #markdown-input, #html-output, #preview {
         height: 200px;
         margin: 0;
    }
     #container {
         display: flex;
         justify-content: space-evenly;
         gap: 10px;
    }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <div id="container">
        <div class="container">
            <h2>Markdown Input:</h2>
            <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
        </div>
        <div class="container">
            <h2>Raw HTML Output:</h2>
            <div id="html-output"></div>
        </div>
        <div class="container">
            <h2>HTML Preview:</h2>
            <div id="preview"></div>
        </div>
    </div>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
     box-sizing: border-box;
}
 body {
     font-family: Arial, sans-serif;
     padding: 20px;
}
 #markdown-input {
     width: 100%;
     height: 100px;
}
 #html-output, #preview {
     height: 100px;
     display: inline-block;
     width: 100%;
     border: 1px solid #ccc;
     padding: 10px;
     margin: auto;
     white-space: pre-wrap;
     background-color: #f9f9f9;
}
 @media (min-width: 600px) {
     #markdown-input, #html-output, #preview {
         height: 200px;
         margin: 0;
    }
     #container {
         display: flex;
         justify-content: space-evenly;
         gap: 10px;
    }
}
```

```js
const patternReplacementArray = [
    { h6: [/^[ \t]*######[ \t](.+$)/gm, "<h6>$1</h6>"] },
    { h5: [/^[ \t]*#####[ \t](.+$)/gm, "<h5>$1</h5>"] },
    { h4: [/^[ \t]*####[ \t](.+$)/gm, "<h4>$1</h4>"] },
    { h3: [/^[ \t]*###[ \t](.+$)/gm, "<h3>$1</h3>"] },
    { h2: [/^[ \t]*##[ \t](.+$)/gm, "<h2>$1</h2>"] },
    { h1: [/^[ \t]*#[ \t](.+$)/gm, "<h1>$1</h1>"] },
    { hr: [/(?<=\n|^)---(?=\n|$)/, "<hr>"] },
    { strong: [/((?:\*|_){2})(.*?)\1/g, "<strong>$2</strong>"] },
    { em: [/(\*|_)(.*?)\1/g, "<em>$2</em>"] },
    { img: [/!\[\s*(.*?)\s*\]\(\s*(.*?)\s*\)/g, '<img alt="$1" src="$2">'] },
    { a: [/\[(.*?)\]\(\s*(.*?)\s*\)/g, '<a href="$2">$1</a>'] },
    { blockquote1: [/^[ \t]*> ([^\n]+?)$/gm, "<blockquote>$1</blockquote>\n"] },
    { blockquote2: [/(?<=<blockquote>)(.+?)<\/blockquote>\n\n^([^<\n]+?)$/gm, "$1 $2</blockquote>\n"] },
    { ul: [/(?:^[-*+]\s+.*?$\n)+\s*?/gm, "<ul>\n$&</ul>"] },
    { ol: [/(?:^\d*\.\s+?\w.*?$\n)+\s*?/gm, "<ol>\n$&</ol>"] },
    { uli: [/^[-*+]\s+(.*?$)\n/gm, "<li>$1</li>\n"] },
    { oli: [/^\d*\.\s+?(.*?$)\n/gm, "<li>$1</li>\n"] },
    { p1: [/^([^<>\s].*?)(?=\n\s*$)/gsm, "<p>$1</p>"] },
    { p2: [/^\w*?(?:<strong>|<em>).+(?:<\/strong>|<\/em>)\s*\w*?[ \t]*$/gm, "<p>$&</p>"] }
]

const markdownInput = document.getElementById('markdown-input');
const convertMarkdown = () => {
    let markdown = markdownInput.value;
    let matchUList = false;
    let matchOList = false;
    patternReplacementArray.forEach((item) => {
        const [key] = Object.keys(item);
        const [pattern, replacement] = Object.values(item)[0];
        if (key === "blockquote2") {
            while (markdown.match(pattern)) {
                markdown = markdown.replace(pattern, replacement);
            }
        } else if (key === "ul") {
            if (markdown.match(pattern)) {
                matchUList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "ol") {
            if (markdown.match(pattern)) {
                matchOList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "uli" && !matchUList) {
            return;
        } else if (key === "oli" && !matchOList) {
            return;
        } else {
            markdown = markdown.replace(pattern, replacement);
        }
    })
    return markdown;
}

markdownInput.addEventListener('input', () => {
    const html = convertMarkdown();
    const out = document.getElementById('html-output');
    const preview = document.getElementById('preview');
    out.textContent = html;
    preview.innerHTML = html;
});
```
