---
id: a5229172f011153519423690
title: Construa um Calculador de Soma Ímpar de Fibonacci
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

Neste laboratório você irá construir um calculador de soma de Fibonacci ímpar que calcula a soma de todos os números ímpares de Fibonacci que são menores ou iguais a um inteiro positivo dado.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário**

1. Você deve ter uma função `sumFibs` que aceita um número como argumento.
1. A função `sumFibs` deve retornar a soma de todos os números ímpares de Fibonacci que são menores ou iguais ao número fornecido.
1. A sequência de Fibonacci começa com `0` e `1` e cada número subsequente é a soma dos dois anteriores.
1. Apenas os números ímpares de Fibonacci devem ser adicionados à soma.

# --hints--

Você deve ter uma função `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` deve retornar um número.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` deve retornar `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` deve retornar `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` deve retornar `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` deve retornar `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` deve retornar `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
