---
id: 689f7996426c1534fa5bea66
title: Crea un traduttore di punteggi per il golf
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

Nel gioco del golf, ogni buca ha un `par`, che indica la media del numero di `strokes` che un golfista dovrebbe fare per mandare la pallina in buca e completare il gioco. A seconda di quanto i tuoi `par` siano sopra o sotto `strokes`, c'è un soprannome diverso.

In questo laboratorio, scriverai una funzione che converte il `par` e il `strokes` nel loro soprannome.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `golfScore`.
1. `golfScore` dovrebbe prendere due argomenti numerici, che sono il par del campo e il numero di colpi effettuati.
1. `golfScore` dovrebbe restituire una stringa.
1. `golfScore` dovrebbe restituire `"Hole-in-one!"` se `strokes` è `1`.
1. `golfScore` dovrebbe restituire `"Eagle"` se `strokes` è minore o uguale a `par` meno `2`.
1. `golfScore` dovrebbe restituire `"Birdie"` se `strokes` è uguale a `par` meno `1`.
1. `golfScore` dovrebbe restituire `"Par"` se `strokes` è uguale a `par`.
1. `golfScore` dovrebbe restituire `"Bogey"` se `strokes` è uguale a `par` più `1`.
1. `golfScore` dovrebbe restituire `"Double Bogey"` se `strokes` è uguale a `par` più `2`.
1. `golfScore` dovrebbe restituire `"Go Home!"` se `strokes` è maggiore o uguale a `par` più `3`.

# --hints--

Dovresti creare una funzione chiamata `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` dovrebbe prendere due parametri.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` dovrebbe restituire una stringa

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` dovrebbe restituire la stringa `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` dovrebbe restituire la stringa `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` dovrebbe restituire la stringa `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` dovrebbe restituire la stringa `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` dovrebbe restituire la stringa `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` dovrebbe restituire la stringa `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` dovrebbe restituire la stringa `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` dovrebbe restituire la stringa `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` dovrebbe restituire la stringa `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` dovrebbe restituire la stringa `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` dovrebbe restituire la stringa `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` dovrebbe restituire la stringa `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` dovrebbe restituire la stringa `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` dovrebbe restituire la stringa `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` dovrebbe restituire la stringa `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` dovrebbe restituire la stringa `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` dovrebbe restituire la stringa `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` dovrebbe restituire la stringa `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` dovrebbe restituire la stringa `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` dovrebbe restituire la stringa `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` dovrebbe restituire la stringa `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
