---
id: 6852f0c3fcedb964ab959921
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Testen Sie nun, ob das Attribut `read` korrekt zu Ihrer `Email` Klasse hinzugefügt wurde. Da Sie bereits ein E-Mail-Objekt aus den vorherigen Schritten haben, geben Sie das Attribut `read` aus, um zu sehen, dass es jetzt standardmäßig `False` ist.

# --hints--

Sie sollten `email_obj.read` ausgeben, um das gelesene Attribut zu testen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
