---
id: 6852ef95ccc27660275692cd
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Ihre `Email`-Klasse muss Informationen über jede E-Mail speichern, wenn sie erstellt wird.

Entfernen Sie innerhalb Ihrer `Email`-Klasse das vorhandene `pass`-Schlüsselwort und erstellen Sie die `__init__`-Methode, die `self`, `sender`, `receiver`, `subject` und `body` als Parameter entgegennimmt.

# --hints--

Sie sollten eine `__init__`-Methode innerhalb der `Email`-Klasse definieren.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Ihre `__init__`-Methode sollte die Parameter `self`, `sender`, `receiver`, `subject` und `body` haben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
