---
id: 6852f06fdde1b4634a263f03
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Testen Sie nun die Klasse `Email`, indem Sie ein E-Mail-Objekt erstellen und überprüfen, ob es die Informationen korrekt speichert. Sie geben ein paar Attribute aus, um zu verifizieren, dass alles funktioniert.

Erstellen Sie ein E-Mail-Objekt namens `email_obj` mit `alice@example.com` als Absender, `bob@example.com` als Empfänger, `Hello` als Betreff und `Hi Bob!` als Inhalt. Geben Sie dann die Attribute `sender` und `subject` jeweils in eigenen print-Anweisungen aus, um zu überprüfen, ob sie korrekt gespeichert sind.

# --hints--

Sie sollten ein E-Mail-Objekt namens `email_obj` mit den angegebenen Parametern erstellen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Sie sollten `email_obj.sender` ausgeben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Sie sollten `email_obj.subject` ausgeben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
--fcc-editable-region--

--fcc-editable-region--
```
