---
id: 6852ea0c3cadb54e897acdac
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie einen *E-Mail-Simulator*, der das Senden, Empfangen und Verwalten von E-Mails zwischen verschiedenen Nutzenden simuliert. Sie lernen Klassen, Objekte und wie man Code objektorientiert organisiert.

Beginnen Sie damit, eine Klasse namens `Email` mit dem Schlüsselwort `class` zu erstellen.

# --hints--

Sie sollten eine Klasse namens `Email` erstellen.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
