---
id: 6852f11d412c29665985a915
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Da Benutzer jetzt E-Mails senden können, benötigen sie eine Möglichkeit, empfangene E-Mails zu speichern. Jeder Benutzer sollte einen Posteingang haben, um seine E-Mails zu sammeln.

Fügen Sie der `inbox`-Klasse in der `User`-Methode ein `__init__`-Attribut hinzu und initialisieren Sie es als leere Liste `[]`. Darin werden alle E-Mails gespeichert, die der Benutzer erhält.

# --hints--

Sie sollten `self.inbox = []` in der `User`-Klasse in der `__init__`-Methode hinzufügen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
