---
id: 6852f01116912861c7d1b0e6
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Innerhalb der `__init__`-Methode müssen Sie die Parameterwerte den Attributen des Objekts zuweisen, damit jede E-Mail ihre Informationen speichern kann.

Entfernen Sie innerhalb Ihrer `__init__`-Methode das Schlüsselwort `pass` und weisen Sie die Parameter `self.sender`, `self.receiver`, `self.subject` und `self.body` zu.

# --hints--

Sie sollten den Parameter `sender` dem `self.sender` zuweisen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Sie sollten den Parameter `receiver` dem `self.receiver` zuweisen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Sie sollten den Parameter `subject` dem `self.subject` zuweisen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Sie sollten den Parameter `body` dem `self.body` zuweisen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
