---
id: 6853f2da9b8ed891e9fdc5fc
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Jetzt zeigen wir die Absender- und Empfängerinformationen an. Um anzuzeigen, wer die E-Mail gesendet und empfangen hat, fügen Sie zwei ausgabe-Anweisungen zur `display_full_email`-Methode in folgendem Format hinzu:

- `From: sender`, wobei `sender` durch den Namen des Absenders ersetzt wird  
- `To: receiver`, wobei `receiver` durch den Namen des Empfängers ersetzt wird

# --hints--

Sie sollten f-Strings verwenden, um den Namen des Absenders mit `From: {self.sender.name}` auszugeben.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

Sie sollten f-Strings verwenden, um den Namen des Empfängers mit `To: {self.receiver.name}` auszugeben.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
