---
id: 6853f4ded9f84896ff182dd8
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Nach der Überprüfung, ob der Posteingang leer ist, geben Sie die Nachricht `\nYour Emails:` aus. Danach durchlaufen Sie alle E-Mails im Posteingang mit einer `for`-Schleife unter Verwendung von Enumeration. Dies ist die Syntax für Enumeration mit einer Startnummer: `enumerate(iterable, start=start_number)`.

Verwenden Sie Enumeration, um die E-Mails beginnend bei `1` zu nummerieren. Verwenden Sie `i` (den Index) und `email` (das E-Mail-Objekt) als Iterationsvariablen.

Innerhalb der Schleife geben Sie einen f-String mit dem Iterationsindex gefolgt von einem `.` aus, dann ein Leerzeichen und die String-Darstellung des E-Mail-Objekts in folgendem Format:

```md
i. string_representation
```

Dabei ist `i` der Index und `string_representation` die Darstellung des E-Mail-Objekts.

# --hints--

Sie sollten die Kopfzeile `\nYour Emails:` außerhalb der `if`-Anweisung ausgeben.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Sie sollten eine `for`-Schleife haben, die `enumerate(self.emails, start=1)` durchläuft.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Ihre `for`-Schleife sollte `i` und `email` verwenden, um `enumerate(self.emails, start=1)` zu durchlaufen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Innerhalb der Schleife sollten Sie jede E-Mail-Zusammenfassung mit einem nummerierten Präfix ausgeben, gefolgt von dem formatierten String aus der `__str__`-Methode. Vergessen Sie nicht, nach der Nummer ein `.` hinzuzufügen.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
        
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--
```
