---
id: 6853f4584f508e95cc7a9f4d
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Jetzt erstellen Sie eine Methode, um alle E-Mails im Posteingang aufzulisten. Diese Methode soll den Fall behandeln, dass der Posteingang leer ist, und eine nummerierte Liste der E-Mails anzeigen, wenn E-Mails vorhanden sind.

Fügen Sie Ihrer `list_emails`-Klasse eine Methode namens `Inbox` hinzu, die `self` als Parameter erhält. Erstellen Sie zuerst eine `if`-Anweisung, um zu prüfen, ob der Posteingang leer ist, indem Sie die `emails`-Liste testen. Wenn sie leer ist, geben Sie die Nachricht `Your inbox is empty.\n` aus. Fügen Sie eine `return`-Anweisung hinzu, um die Methode zu beenden.

# --hints--

Sie sollten eine Methode namens `list_emails` in der `Inbox`-Klasse definieren, die `self` als Parameter erhält.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Sie sollten prüfen, ob `self.emails` leer ist, indem Sie `if not self.emails` verwenden.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Sie sollten die Nachricht `Your inbox is empty.\n` ausgeben, wenn keine E-Mails vorhanden sind.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Sie sollten am Ende der `return`-Anweisung eine `if`-Anweisung haben.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
