---
id: 68593ea22eac7e228ed585ce
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

E-Mails sollten verfolgen, ob sie gelesen wurden oder nicht. Fügen Sie der `read`-Methode ein `__init__`-Attribut hinzu und setzen Sie es standardmäßig auf `False`, da neue E-Mails als ungelesen beginnen.

# --hints--

Sie sollten `False` in der `self.read`-Methode `__init__` zuweisen.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
--fcc-editable-region--
        
--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
