---
id: 685941422c30cd3dc41b1d0a
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Fügen wir nun eine Kopfzeile zur E-Mail-Anzeige hinzu. Geben Sie in der `display_full_email`-Methode nach dem Aufrufen von `mark_as_read()` `\n--- Email ---` aus, um die E-Mail-Anzeige mit einer klaren Kopfzeile zu beginnen.

# --hints--

Sie sollten die Kopfzeile `\n--- Email ---` in der `display_full_email`-Methode ausgeben.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print('\\\\n--- Email ---')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
--fcc-editable-region--
        
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
