---
id: 68594df7779a164e6a61e8ae
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Der Posteingang benötigt eine Methode, um eine bestimmte E-Mail zu lesen. Wenn ein Benutzer eine E-Mail lesen möchte, gibt er die Nummer der E-Mail an (beginnend mit dem Index 0), die er sehen möchte, und die Methode zeigt den vollständigen Inhalt der E-Mail an.

Fügen Sie Ihrer `read_email`-Klasse eine Methode namens `Inbox` hinzu, die einen `index`-Parameter entgegennimmt. Überprüfen Sie zuerst, ob der Posteingang leer ist, und geben Sie die Nachricht `Inbox is empty.\n` aus, falls dies der Fall ist. Fügen Sie danach eine `return`-Anweisung hinzu, um die Methode zu beenden.

# --hints--

Sie sollten eine Methode namens `read_email` in der `Inbox`-Klasse definieren, die `self` und `index` als Parameter entgegennimmt.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

Sie sollten prüfen, ob `self.emails` leer ist, indem Sie `if not self.emails` verwenden.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

Sie sollten `Inbox is empty.\n` ausgeben, wenn keine E-Mails vorhanden sind, und dann zurückkehren.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
