---
id: 6859400f8f7f872e50e1238c
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Die `User`-Klasse benötigt eine Methode, um E-Mails an andere Benutzer zu senden. Beim Senden einer E-Mail erstellen Sie ein neues `Email`-Objekt und fügen es dem Posteingang des Empfängers hinzu.

Erstellen Sie eine Methode namens `send_email` in Ihrer `User`-Klasse. Diese Methode soll Parameter für `receiver`, `subject` und `body` entgegennehmen.

Erstellen Sie innerhalb der Methode ein neues `Email`-Objekt mit den folgenden Parameterwerten und weisen Sie es einer Variablen namens `email` zu:

- `sender`: `self` (der aktuelle Benutzer)
- `receiver`: `receiver` (der Benutzer, der die E-Mail erhält)
- `subject`: `subject` (der Betreff der E-Mail)
- `body`: `body` (der Inhalt der E-Mail)

# --hints--

Sie sollten eine `send_email`-Methode in der `User`-Klasse definieren.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

Die `send_email`-Methode sollte die Parameter `self`, `receiver`, `subject` und `body` haben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

Die Methode sollte ein `Email`-Objekt namens `email` mit den korrekten Parametern erstellen.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        
--fcc-editable-region--
    
--fcc-editable-region--
```
