---
id: 685cdabd4cded3b340193f72
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Fügen Sie nun die Betreffzeile zur E-Mail-Anzeige hinzu. Fügen Sie in der `display_full_email`-Methode eine Ausgabeanweisung hinzu, um den Betreff der E-Mail in folgendem Format anzuzeigen:

- `Subject: subject`, wobei `subject` durch den Betreff der E-Mail ersetzt wird

# --hints--

Sie sollten f-Strings verwenden, um den Betreff mit `Subject: {self.subject}` auszugeben.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
