---
id: 685c1eba966c6396850736ea
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Während Benutzer E-Mails senden und Posteingänge haben können, ist eine eigene Klasse erforderlich, um Posteingangsoperationen effizient zu verwalten.

Die `Inbox`-Klasse wird eine Liste von E-Mails speichern und Methoden bereitstellen, um neue E-Mails hinzuzufügen, alle E-Mails aufzulisten und sie zu verwalten.

Erstellen Sie eine neue Klasse namens `Inbox` mit einer `__init__`-Methode, die eine leere Liste namens `emails` initialisiert.

# --hints--

Sie sollten eine `Inbox`-Klasse erstellen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

Die `Inbox`-Klasse sollte eine `__init__`-Methode haben, die `self` als Parameter entgegennimmt.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Sie sollten `self.emails = []` in der `Inbox`-Klasse `__init__`-Methode initialisieren.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
