---
id: 685c1e32254e5993a5c8823e
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Jetzt, wo Sie eine erste Einrichtung der `Email` Klasse haben, ist es Zeit, Benutzer zu erstellen, die E-Mails senden und empfangen können. Jeder Benutzer wird einen Namen haben und E-Mail-Operationen ausführen können.

Erstellen Sie eine `User` Klasse mit einer `__init__` Methode, die `self` und `name` als Parameter entgegennimmt. Innerhalb der Klasse weisen Sie den `name` Parameter `self.name` zu.

# --hints--

Sie sollten eine `User` Klasse erstellen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

Die `User` Klasse sollte eine `__init__` Methode haben, die `self` und `name` als Parameter entgegennimmt.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Sie sollten den `name` Parameter `self.name` zuweisen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
