---
id: 68875554c6c32de788afc64e
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Benutzende benötigen eine Möglichkeit, ihre E-Mails richtig zu lesen.

Fügen Sie der `Email`-Klasse eine `display_full_email`-Methode hinzu, die nur `self` als Parameter erhält. Rufen Sie innerhalb dieser Methode die `mark_as_read()`-Methode auf, um die E-Mail als gelesen zu markieren, wenn sie jemand ansieht.

Diese Methode wird verwendet, um den vollständigen Inhalt der E-Mail ansprechend formatiert anzuzeigen.

# --hints--

Sie sollten eine `display_full_email`-Methode in der `Email`-Klasse definieren.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

Die `display_full_email`-Methode sollte nur `self` als Parameter erhalten.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

Die Methode sollte `self.mark_as_read()` aufrufen, um die E-Mail als gelesen zu markieren.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
