---
id: 688754d6d59578df1313cb3c
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Jetzt ist es Zeit, die Klassen User und Inbox zu verbinden, damit E-Mails tatsächlich zugestellt werden können!

Benutzer benötigen richtige Inbox-Objekte anstelle von einfachen Listen, und die `send_email`-Methode sollte E-Mails an den Posteingang des Empfängers zustellen.

Aktualisieren Sie die `User`-Klasse `__init__` so, dass jeder Benutzer ein echtes Inbox-Objekt erhält, anstatt der zuvor erstellten leeren Liste.

# --hints--

Sie sollten die User-Klasse aktualisieren, um `self.inbox = Inbox()` zu initialisieren.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
