---
id: 6887550fe36f64e27074475a
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Lassen Sie Alice eine E-Mail an Bob senden, um zu prüfen, ob die E-Mail-Zustellung korrekt funktioniert.

Verwenden Sie die `send_email`-Methode von Alice, um Bob eine E-Mail mit dem Betreff `"Hello"` und dem Inhalt `"Hi Bob, how are you?"` zu senden.

# --hints--

Sie sollten `alice.send_email()` mit den richtigen Parametern aufrufen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

alice = User("Alice")
bob = User("Bob")
--fcc-editable-region--

--fcc-editable-region--
```
