---
id: 68879d6de6398893be8c4b7b
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Fügen Sie nun den E-Mail-Text hinzu, um den Hauptinhalt zu vervollständigen. Fügen Sie in der `display_full_email`-Methode eine weitere print-Anweisung im Format `Body: body` hinzu (wobei `body` den Inhalt der E-Mail darstellt), um den Inhalt der E-Mail anzuzeigen.

# --hints--

Sie sollten den E-Mail-Text mit `f'Body: {self.body}'` ausgeben.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Body: {self.body}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
--fcc-editable-region--
        
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

```
