---
id: 688754f3494b17e0b352d67d
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Jetzt ist es Zeit, das komplette E-Mail-System zu testen! Erstellen Sie zwei Benutzer und sehen Sie die E-Mail-Funktionalität in Aktion.

Erstellen Sie zwei User-Objekte: `alice` mit dem Namen `"Alice"` und `bob` mit dem Namen `"Bob"`. Dies zeigt, wie Benutzer erstellt werden können und miteinander interagieren.

# --hints--

Sie sollten ein User-Objekt namens `alice` mit dem Namen `"Alice"` erstellen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

Sie sollten ein User-Objekt namens `bob` mit dem Namen `"Bob"` erstellen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
