---
id: 68b930ee0e7b8fb81895e59a
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Jetzt fügen Sie eine Methode hinzu, um eine E-Mail als gelesen zu markieren. Wenn jemand eine E-Mail öffnet, möchten Sie ihren Status von ungelesen auf gelesen ändern. Diese Methode wird einfach sein – sie muss nur das `read`-Attribut auf `True` setzen.

Fügen Sie Ihrer `mark_as_read`-Klasse eine Methode namens `Email` hinzu. Diese Methode soll nur `self` als Parameter erhalten und das `read`-Attribut auf `True` setzen.

# --hints--

Sie sollten eine Methode namens `mark_as_read` in der `Email`-Klasse definieren.
Diese sollte nur `self` als Parameter erhalten.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

Die Methode sollte `self.read` auf `True` setzen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
