---
id: 68b936bf4f47d6c29a669f99
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Rufen Sie weiterhin innerhalb der `send_email`-Methode die `receive_email`-Methode des `receiver` `inbox` auf, um die E-Mail zum Posteingang des Empfängers hinzuzufügen.

# --hints--

Sie sollten `receiver.inbox.receive_email(email)` zur `send_email`-Methode hinzufügen.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_stmt("receiver.inbox.receive_email(email)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
        
--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)
```
