---
id: 68b96d6a290d52d5e57a7485
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Entfernen Sie die `email_obj` und die folgenden ausgabe-Anweisungen.

# --hints--

Sie sollten die `email_obj`-Erstellungszeile nicht in Ihrem Code haben.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

Sie sollten `print(email_obj.sender)` nicht in Ihrem Code haben.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

Sie sollten `print(email_obj.subject)` nicht in Ihrem Code haben.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

Sie sollten `print(email_obj.read)` nicht in Ihrem Code haben.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

Sie sollten `email_obj.mark_as_read()` nicht in Ihrem Code haben.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
