---
id: 68b9350183ff05bd106c79e4
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Testen Sie nun die `mark_as_read`-Methode, die Sie in der `Email`-Klasse erstellt haben. Verwenden Sie die Methode für das vorhandene `Email`-Objekt `email_obj`, um den Status zu ändern. Nachdem Sie die E-Mail als gelesen markiert haben, geben Sie das `read`-Attribut von `email_obj` aus, um die Änderung zu bestätigen.

# --hints--

Sie sollten die `mark_as_read`-Methode auf dem `email_obj` aufrufen, um die E-Mail als gelesen zu markieren.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Sie sollten das `read`-Attribut von `email_obj` ausgeben, nachdem Sie die `mark_as_read`-Methode aufgerufen haben.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
--fcc-editable-region--

--fcc-editable-region--
```
