---
id: 68b9366fae8b13c1278761b1
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Ihr Posteingang benötigt eine Möglichkeit, neue E-Mails zu empfangen. Wenn jemand einem Benutzer eine E-Mail sendet, sollte diese zu seinem Posteingang hinzugefügt werden.

Fügen Sie Ihrer `receive_email`-Klasse eine Methode namens `Inbox` hinzu, die `self` und ein E-Mail-Objekt `email` als Parameter entgegennimmt. Fügen Sie im Methodenrumpf die `email` mit der `emails`-Methode zur `append`-Liste hinzu.

# --hints--

Sie sollten in der `receive_email`-Klasse eine Methode namens `Inbox` definieren, die `self` und `email` als Parameter entgegennimmt.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Sie sollten die `email` an `self.emails` anhängen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
