---
id: 68be6af1f43cc44541debe41
title: Schritt 37
challengeType: 20
dashedName: step-37
---

# --description--

Sie sollten den Fall behandeln, in dem der Benutzer versucht, eine E-Mail mit einem ungültigen Index zu löschen, genauso wie Sie es in der `read_email`-Methode gemacht haben.

Innerhalb der `delete_email`-Methode:

- Wandeln Sie den 1-basierten `index`-Parameter in einen 0-basierten Index um, indem Sie 1 subtrahieren und ihn in einer Variablen namens `actual_index` speichern.
- Erstellen Sie eine if-Anweisung, um zu prüfen, ob der `actual_index` kleiner als 0 oder größer oder gleich der Länge der `self.emails`-Liste ist. Falls ja, geben Sie die Nachricht `Invalid email number.\n` aus und verwenden Sie `return`, um die Methode zu verlassen.

# --hints--


Sie sollten 1 vom `index`-Parameter subtrahieren und in einer Variablen namens `actual_index` speichern.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

Sie sollten eine if-Anweisung erstellen, um zu prüfen, ob der `actual_index` kleiner als 0 oder größer oder gleich der Länge der `self.emails`-Liste ist.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})
```

Innerhalb der `if`-Anweisung sollten Sie die Nachricht `Invalid email number.\n` ausgeben und `return` verwenden, um die Methode zu verlassen.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
            
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--

```
