---
id: 68e43e884ef09d4d2ab05631
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, das Ergebnis der Verwendung der `replace()`-Methode zu sehen.

Beginnen Sie damit, eine `console.log()` mit dem Wert `"After using the replace() method:"` hinzuzufügen.

Fügen Sie dann unter dieser `console.log()` eine weitere `console.log()` mit der `replacedString`-Variablen hinzu.

Schauen Sie sich die Konsole an, und Sie sollten sehen, dass der neue Satz `"I love dogs."` anstelle des ursprünglichen `"I love cats."` sagt.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten `"After using the replace() method:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Sie sollten die `replacedString`-Variable in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
