---
id: 68e41084bb275539f72de9fe
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop Ă¼ben Sie den Umgang mit den `replace()`, `replaceAll()` und `repeat()` Methoden, indem Sie eine String-Transformer-App erstellen.

Beginnen Sie damit, eine Variable namens `originalString` zu erstellen und ihr den Wert von `"I love cats."` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `originalString` haben.

```js
assert.exists(originalString);
```

Ihre `originalString` Variable sollte ein String sein.

```js
assert.isString(originalString);
```

Sie sollten der `"I love cats."` Variable den String `originalString` zuweisen.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
