---
id: 68e4417b7317835208145323
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Erstellen Sie als Nächstes eine neue Variable namens `dogsOnlySentence` und weisen Sie ihr `exampleSentence.replace("cats", "dogs")` zu.

Protokollieren Sie dann unterhalb dieser Variable `"Replacing all occurrences of cats with dogs:"` in der Konsole. Fügen Sie unterhalb von `console.log()` eine weitere `console.log()` für die `dogsOnlySentence` Variable hinzu.

Sie sollten nun den Text `"I love dogs and cats are so much fun!"` in der Konsole sehen. Das ist jedoch nicht das gewünschte Ergebnis. Im nächsten Schritt lernen Sie, wie Sie das beheben, damit der Text `"I love dogs and dogs are so much fun!"` lautet.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine `dogsOnlySentence` Variable haben.

```js
assert.exists(dogsOnlySentence);
```

Ihre `dogsOnlySentence` Variable sollte ein String sein.

```js
assert.isString(dogsOnlySentence);
```

Sie sollten `exampleSentence.replace("cats", "dogs")` Ihrer `dogsOnlySentence` Variable zuweisen.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Sie sollten `"Replacing all occurrences of cats with dogs:"` in der Konsole protokollieren.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Sie sollten die `dogsOnlySentence` Variable in der Konsole protokollieren.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
