---
id: 68e44abea678a65b877864dc
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Es wäre schön, dieses wiederholte Wort in einem Satz zu sehen. Erstellen Sie eine neue Variable namens `newSentence` und weisen Sie ihr das Ergebnis der Verkettung des Wortes `"I "` mit der Variable `repeatedLove` zu, gefolgt von der Verkettung des Strings `" learning."`. Geben Sie dann die Variable `newSentence` in der Konsole aus. Sie können entweder den `+`-Operator oder Template-Literale für die String-Verkettung verwenden.

Jetzt sollten Sie den Text `"I love love love  learning."` sehen.

*HINWEIS*: Nach dem letzten `"love"` befindet sich ein zusätzliches Leerzeichen, das korrekt ist. Im letzten Schritt werden Sie dieses zusätzliche Leerzeichen entfernen.

# --hints--

Sie sollten eine Variable namens `newSentence` haben.

```js
assert.exists(newSentence);
```

Ihre `newSentence`-Variable sollte ein String sein.

```js
assert.isString(newSentence);
```

Sie sollten das Wort `"I "` mit der Variable `repeatedLove` verketten, gefolgt von der Verkettung des Strings `" learning."`. Dieses Ergebnis sollte der Variable `newSentence` zugewiesen werden. Achten Sie auf die Leerzeichen.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Sie sollten NICHT den Literalstring `"I love love love  learning."` der Variable `newSentence` zuweisen. Bitte lesen Sie die Anweisungen noch einmal für den korrekten Ansatz.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
