---
id: 68e44519c769bf545bb3f22a
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Um die gewünschte Ausgabe von `"I love dogs and dogs are so much fun!"` zu erhalten, müssen Sie die `replaceAll()`-Methode anstelle der `replace()`-Methode verwenden. Diese Methode gibt für alle Übereinstimmungen mit dem Teilstring einen neuen String zurück.

Aktualisieren Sie die `exampleSentence.replace("cats", "dogs")`, damit sie stattdessen die `replaceAll()`-Methode verwendet. Nun sollten Sie den korrekten Text in der Konsole sehen.

# --hints--

Sie sollten die `exampleSentence.replace("cats", "dogs")`-Codezeile aktualisieren, um die `replaceAll()`-Methode zu verwenden.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
