---
id: 68e42b62d7093345c6975f22
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Fügen Sie unter Ihrer `originalString`-Variable eine `console.log()` mit dem String `"Original string:"` hinzu.

Dann fügen Sie unter dieser `console.log()` eine weitere `console.log()` mit der `originalString`-Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"Original string:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[0], "Original string:");
```

Sie sollten die `originalString`-Variable in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
