---
id: 68e4465fba0bd256868b7deb
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Jetzt ist es Zeit, mit einer neuen String-Methode zu arbeiten.

Beginnen Sie damit, eine Variable namens `learningSentence` zu erstellen und den String `"I love learning!"` zuzuweisen.

Geben Sie dann unter dieser Variable `"Original learning sentence:"` in der Konsole aus. Fügen Sie unter `console.log()` eine weitere `console.log()` für die `learningSentence` Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine `learningSentence` Variable haben.

```js
assert.exists(learningSentence);
```

Ihre `learningSentence` Variable sollte ein String sein.

```js
assert.isString(learningSentence);
```

Sie sollten der `"I love learning!"` Variable den String `learningSentence` zuweisen.

```js
assert.equal(learningSentence, "I love learning!");
```

Sie sollten `"Original learning sentence:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Sie sollten die `learningSentence` Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
