---
id: 68e44c305f1f935d949d76bb
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Wie im vorherigen Schritt erwähnt, gibt es ein zusätzliches Leerzeichen nach dem letzten `"love"` im Satz `"I love love love  learning."`. Um dieses zusätzliche Leerzeichen zu entfernen, können Sie die `trimEnd()`-Methode verwenden.

Am Ende der `"love ".repeat(3)`-Methode verketten Sie die `trimEnd()`-Methode. Sie können Methoden so verketten:

```js
.firstMethod().secondMethod()
```

Wenn Sie nun die Konsole überprüfen, sollten Sie sehen, dass das zusätzliche Leerzeichen entfernt wurde.

Und mit dieser letzten Änderung ist Ihr Workshop abgeschlossen!

# --hints--

Sie sollten die `trimEnd()`-Methode an die `"love ".repeat(3)`-Methode anhängen.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
