---
id: 68e448ec4ac4205968c0a805
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Es wäre schön, wenn das Wort `"love"` im Satz `"I love learning!"` dreimal wiederholt würde.

Wie Sie sich aus einer früheren Lektion erinnern, können Sie einen String mit der `repeat()`-Methode eine bestimmte Anzahl von Malen wiederholen. Hier ist ein Beispiel:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Da Strings unveränderlich sind, wird diese Methode den ursprünglichen String nicht verändern. Sie gibt einen neuen String mit dem wiederholten Inhalt zurück.

Erstellen Sie eine Variable namens `repeatedLove` und weisen Sie ihr `"love ".repeat(3)` zu. Geben Sie dann die `repeatedLove`-Variable in der Konsole aus, damit Sie das Ergebnis sehen können.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Variable namens `repeatedLove` haben.

```js
assert.exists(repeatedLove);
```

Ihre `repeatedLove`-Variable sollte ein String sein.

```js
assert.isString(repeatedLove);
```

Sie sollten `"love ".repeat(3)` der `repeatedLove`-Variable zuweisen.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Sie sollten die `repeatedLove`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
