---
id: 68e43f94365a964f26c88f78
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt ist es Zeit, mit einer neuen String-Methode zu arbeiten.

Beginnen Sie damit, eine neue Variable namens `exampleSentence` zu erstellen und ihr den String `"I love cats and cats are so much fun!"` zuzuweisen.

Geben Sie dann unterhalb dieser Variable `"Original sentence:"` in der Konsole aus. Fügen Sie darunter `console.log()` eine weitere `console.log()` mit der `exampleSentence` Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine `exampleSentence` Variable haben.

```js
assert.exists(exampleSentence);
```

Ihre `exampleSentence` Variable sollte ein String sein.

```js
assert.isString(exampleSentence);
```

Sie sollten der `"I love cats and cats are so much fun!"` Variable den String `exampleSentence` zuweisen.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Sie sollten `"Original sentence:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Sie sollten die `exampleSentence` Variable in der Konsole ausgeben.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
