---
id: 68e42de6396f5c49d90d1951
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

In einer vorherigen Lektion haben Sie die `replace()`-Methode kennengelernt, mit der ein Teilstring in einem String gefunden und durch einen anderen Wert ersetzt wird.

Hier ist ein Beispiel:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Denken Sie daran, dass Strings unveränderlich sind, was bedeutet, dass der ursprüngliche `text` in diesem Beispiel nicht verändert wird. Stattdessen wird ein neuer String erstellt.

Erstellen Sie eine neue Variable namens `replacedString` und weisen Sie ihr `originalString.replace("cats", "dogs")` zu.

# --hints--

Sie sollten eine Variable namens `replacedString` haben.

```js
assert.exists(replacedString);
```

Ihre `replacedString`-Variable sollte ein String sein.

```js
assert.isString(replacedString);
```

Sie sollten `originalString.replace("cats", "dogs")` der `replacedString`-Variable zuweisen.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
