---
id: 680900675ae3d54ee19590c9
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

All'interno del tuo elemento `select`, devi creare delle opzioni tra cui l'utente può scegliere la propria fonte di energia. Il `select` dovrebbe avere sette opzioni elencate qui sotto:

| Fonte di energia             |
|-----------------------------|
| `Select one`                   |
| `Bitten by a strange creature`                   |
| `Radioactive exposure`                   |
| `Science experiment`                   |
| `Alien heritage `                   |
| `Ancient artifact discovery`                   |
| `Other`                   |

La prima opzione (`Select one`) dovrebbe avere una stringa vuota come suo attributo `value` (cioè `value=''`). Questo funge da segnaposto nell'interfaccia utente.

# --hints--

Dovresti creare il primo elemento `option` con il testo `Select one` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

Dovresti dare al primo elemento `option` un attributo `value` con una stringa vuota.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.getAttribute('value'), '');
```

Dovresti creare un secondo elemento `option` con il testo `Bitten by a strange creature` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

Dovresti creare un terzo elemento `option` con il testo `Radioactive exposure` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

Dovresti creare un quarto elemento `option` con il testo `Science experiment` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

Dovresti creare un quinto elemento `option` con il testo `Alien heritage` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

Dovresti creare un sesto elemento `option` con il testo `Ancient artifact discovery` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

Dovresti creare un settimo elemento `option` con il testo `Other` all'interno del tuo elemento `select`.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>
              
            </select>
            --fcc-editable-region--
          </label>
      </form>
    </div>
  )
};
```
