---
id: 686ccc2c8b967e17ab18d593
title: реалізувати метод ділення навпіл
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

Метод ділення навпіл, також відомий як метод бінарного пошуку, використовує бінарний пошук для знаходження коренів дійсної функції. Він працює шляхом звуження інтервалу, в якому лежить квадратний корінь, доки не збіжиться до значення в межах заданої точності.

Наприклад, якщо точність дорівнює `0.01`, метод ділення навпіл буде постійно ділити інтервал навпіл, доки різниця між верхньою та нижньою межами не стане меншою або рівною `0.01`.

У цій лабораторній роботі ви реалізуєте функцію, яка використовує метод ділення навпіл для знаходження квадратного кореня числа.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з назвою `square_root_bisection` з трьома параметрами:
    - Число, для якого потрібно знайти квадратний корінь.
    - Точність — прийнятна похибка для результату. Ви повинні встановити значення точності за замовчуванням.
    - Максимальна кількість ітерацій для виконання. Ви повинні встановити значення за замовчуванням.

1. Функція `square_root_bisection` повинна:
    - Викликати `ValueError` з повідомленням `Square root of negative number is not defined in real numbers`, якщо передане число від’ємне.
    - Для чисел `0` та `1` вивести повідомлення: `The square root of [number] is [number]` і повернути саме число як квадратний корінь.
    - Для будь-якого іншого додатного числа вивести приблизний квадратний корінь з повідомленням: `The square root of [square_target] is approximately [root]` і повернути обчислене значення кореня.
    - Якщо жодне значення не відповідає умові точності, вивести повідомлення про невдачу: `Failed to converge within [maximum] iterations` і повернути `None`.

**Примітка**: Ви не можете імпортувати жоден модуль для цієї лабораторної роботи.

# --hints--

Вам не слід імпортувати жоден модуль.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

Ви повинні мати функцію з назвою `square_root_bisection`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

Ваша функція `square_root_bisection` повинна мати три параметри.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

Ви повинні встановити значення за замовчуванням для точності та максимальної кількості ітерацій.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

Ваша функція `square_root_bisection` повинна викликати `ValueError` з повідомленням `Square root of negative number is not defined in real numbers`, коли передане число від’ємне.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` повинна повертати `0`.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` повинна виводити `The square root of 0 is 0`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` повинна повертати число між `0.03162267660168379` та `0.031622876601683794`.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` повинна виводити `The square root of 0.001 is approximately X`, де `X` — число між `0.03162267660168379` та `0.031622876601683794`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` повинна повертати число між `0.4999999` та `0.5000001`.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` повинна виводити `The square root of 0.25 is approximately X`, де `X` — число між `0.4999999` та `0.5000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` повинна повертати `1`.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` повинна виводити `The square root of 1 is 1`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` повинна повертати число між `8.999` та `9.001`.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` повинна виводити `The square root of 81 is approximately X`, де `X` — число між `8.999` та `9.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` повинна повертати число між `14.999` та `15.001`.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` повинна виводити `The square root of 225 is approximately X`, де `X` — число між `14.999` та `15.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` повинна повертати число між `14.99999` та `15.00001`.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` повинна виводити `The square root of 225 is approximately X`, де `X` — число між `14.99999` та `15.00001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` повинна повертати число між `14.9999999` та `15.0000001`.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` повинна виводити `The square root of 225 is approximately X`, де `X` — число між `14.9999999` та `15.0000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` повинна повертати `None`.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` повинна виводити `Failed to converge within 10 iterations`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
