---
id: 6718d2d59337c822ecb697ff
title: Tengeneza programu ya usimamizi wa akaunti ya benki
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua darasa la vitu linaloitwa `BankAccount` lenye kijenzi. Kijenzi kinapaswa kuweka `balance` ya msingi kuwa `0` na kuanzisha safu ya data tupu iitwayo `transactions` kuhifadhi rekodi za miamala kama vitu.

2. Kila muamala ulihifadhiwa katika safu ya data ya `transactions` unapaswa kuwa kitu chenye sifa mbili: `type` na `amount`. Sifa ya `type` inapaswa kuwa `deposit` au `withdraw`, na sifa ya `amount` inapaswa kuwa kiasi kilichowekwa au kuchukuliwa.

3. Unapaswa kufafanua njia iitwayo `deposit` inayopokea kiasi cha amana kama kigezo. Wakati kiasi cha amana ni kikubwa kuliko `0`, inapaswa:

   - Ongeza kitu kipya kwenye safu ya data ya miamala chenye aina ya `deposit` na kiasi kilichowekwa.
   - Sasisha salio.   
   - Rudisha `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. Ikiwa kiasi ni kidogo au sawa na `0`, njia ya `deposit` inapaswa kurudisha `"Deposit amount must be greater than zero."`.

5. Unapaswa kufafanua njia iitwayo `withdraw` inayopokea kiasi kama kigezo. Njia hii inapaswa kusasisha salio la sasa kulingana na uondoaji. Wakati kiasi kinachotakiwa kuchukuliwa ni kikubwa kuliko `0` na kidogo au sawa na salio la sasa, inapaswa:

   - Ongeza kitu kipya kwenye safu ya data ya miamala chenye aina ya `withdraw` na kiasi kilichochukuliwa.
   - Sasisha salio.
   - Rudisha `"Successfully withdrew $[amount]. New balance: $[balance]"`. 

6. Ikiwa kiasi kinachotakiwa kuchukuliwa ni kidogo au sawa na `0` au kikubwa kuliko salio la sasa, njia ya `withdraw` inapaswa kurudisha `"Insufficient balance or invalid amount."`.

7. Unapaswa kufafanua njia iitwayo `checkBalance` inayorudisha salio la sasa kwa muundo wa `"Current balance: $[balance]"`.

8. Unapaswa kufafanua njia iitwayo `listAllDeposits` inayorudia kupitia safu ya data ya miamala na kurudisha amana zote kwa muundo wa `"Deposits: amount,amount,..."`.

9. Unapaswa kufafanua njia iitwayo `listAllWithdrawals` inayorudia kupitia safu ya data ya miamala na kurudisha uondoaji wote kwa muundo wa `"Withdrawals: amount,amount,..."`.

10. Unapaswa kuunda mfano mpya wa `BankAccount` uitwao `myAccount`.

11. Akaunti yako ya benki ya `myAccount` inapaswa kuwa na miamala angalau mitano.

12. Akaunti yako ya benki ya `myAccount` inapaswa kuwa na amana angalau mbili.

13. Akaunti yako ya benki ya `myAccount` inapaswa kuwa na uondoaji angalau mbili.

14. Akaunti yako ya benki ya `myAccount` inapaswa kuwa na salio kubwa kuliko `$100`.

# --hints--

Unapaswa kufafanua darasa la vitu linaloitwa `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

Kitu cha `BankAccount` kinapaswa kuwa na `balance` ya `0` na safu ya data tupu `transactions` kuhifadhi rekodi za miamala.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

Unapaswa kuwa na njia ya `deposit` inayopokea kiasi cha amana kama kigezo.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

Unapaswa kuwa na njia ya `withdraw` inayopokea kiasi cha uondoaji kama kigezo.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

Unapaswa kuwa na njia ya `checkBalance` inayokagua salio la sasa.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

Unapaswa kuwa na njia ya `listAllDeposits` inayoorodhesha amana zote.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

Unapaswa kuwa na njia ya `listAllWithdrawals` inayoorodhesha uondoaji wote.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)` inapaswa kurudisha `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` inapaswa kurudisha `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` inapaswa kurudisha `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

Wakati salio la akaunti ni `100`, `BankAccount.withdraw(150)` inapaswa kurudisha `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` inapaswa kurudisha `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` inapaswa kurudisha `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

Wakati salio la akaunti ni `200`, `BankAccount.withdraw(150)` inapaswa kurudisha `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

Wakati salio la akaunti ni `200`, `BankAccount.checkBalance()` inapaswa kurudisha `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

Unapoweka amana `10`, `35`, `90`, njia ya `listAllDeposits` inapaswa kurudisha `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

Unapochukua amana `20`, `50`, `100`, njia ya `listAllWithdrawals` inapaswa kurudisha `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

Unapaswa kuwa na mfano wa `BankAccount` uitwao `myAccount`.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

Akaunti yako ya benki ya `myAccount` inapaswa kuwa na miamala angalau mitano.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

Akaunti yako ya benki ya `myAccount` inapaswa kuwa na amana angalau mbili.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

Akaunti yako ya benki ya `myAccount` inapaswa kuwa na uondoaji angalau mbili.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```


Akaunti yako ya benki ya `myAccount` inapaswa kuwa na salio kubwa kuliko `$100`.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
