---
id: 66b205e6eacba4c4e54ea434
title: Crie um mascarador de e-mails
challengeType: 14
dashedName: build-an-email-masker
---

# --description--

Neste laboratório, você vai mascarar a parte do nome de usuário de um endereço de e-mail com asteriscos. Mascaramento é um termo utilizado para ocultar ou substituir informações sensíveis com asteriscos ou outros caracteres.

Por exemplo, se o endereço de e-mail fosse `myEmail@email.com`, o endereço de e-mail mascarado seria `m*******l@email.com`.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Crie uma função chamada `maskEmail` que recebe `email` como argumento.
2. Dentro da função, você deve mascarar o `email` e anexar o nome de domínio a ele.  Lembre-se de que você pode usar métodos como `slice`, `repeat`, `indexOf`, `slice` ou até `replace` como auxílio.
3. Fora da função, declare uma variável chamada `email` para armazenar o endereço de e-mail que você deseja mascarar.
4. Chame a função `maskEmail` com a variável `email` e produza a saída para o console.
5. `maskEmail("apple.pie@example.com")` deve retornar `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` deve retornar `"f**********p@example.com"`.

# --hints--

Você deve definir uma função chamada `maskEmail`.

```js
assert.isFunction(maskEmail);
```

A função `maskEmail` deve receber uma string, `email`, como argumento.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Fora da função, você deve ter uma variável `email`.

```js
assert.isDefined(email);
```

Você deve atribuir um endereço de e-mail válido à variável `email`.

```js
assert.match(email, /^[\w\.-]+@[a-zA-Z\d\.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` deve retornar `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` deve retornar `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

A função `maskEmail` deve produzir o resultado correto.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
```

Você deve registrar a chamada da função `maskEmail` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
