---
id: 66db57ad34c7089b9b41bfd6
title: Einen Lightbox-Viewer erstellen
challengeType: 25
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

Eine Lightbox wird auf Websites verwendet, um multimediale Inhalte ansprechender darzustellen, indem ein Popup- oder Modal-Fenster über dem Hauptinhalt der Seite angezeigt wird.

In diesem Labor erstellen Sie eine Lightbox-Galerie, die Vollbildbilder anzeigt, wenn auf eine Miniaturansicht geklickt wird. Für jedes Bild sind zwei Versionen vorhanden: eine Miniaturansicht und ein Vollbildbild. Das Vollbildbild ist dasselbe wie die Miniaturansicht, jedoch ohne das `-thumbnail`-Suffix.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `div` mit einer Klasse von `gallery` innerhalb Ihres `body` haben.
1. Innerhalb des `.gallery`-Elements sollten Sie drei Bild-Miniaturen haben, jeweils mit der Klasse `gallery-item`. Verwenden Sie für die Miniaturbilder die folgenden Links:
   
    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg`

1. Sie sollten ein `div` mit einer Klasse von `lightbox` innerhalb Ihres `body` haben.
1. Sie sollten ein `button` mit einem `id` von `close-btn` innerhalb Ihres `.lightbox`-Elements haben. Sie können `&times;` als Text verwenden, wenn Sie möchten.
1. Sie sollten ein `img` mit einem `id` von `lightbox-image` innerhalb Ihres `.lightbox`-Elements haben.
1. Ihr `.lightbox`-Element sollte eine feste Position haben, damit die Vorschau über den aktuellen Bildern geöffnet wird.
1. Ihr `.lightbox`-Element sollte den gesamten Viewport abdecken, indem Höhe und Breite auf 100 % des Containers gesetzt werden. Sie sollten sicherstellen, dass das `.lightbox`-Element in der oberen linken Ecke des Containers beginnt.
1. `.lightbox` sollte eine Hintergrundfarbe haben. Anfangs sollte seine `display`-Eigenschaft auf `none` gesetzt sein, um es zu verbergen.
1. Wenn Sie auf eines Ihrer `.gallery-item`-Elemente klicken, sollte die `.lightbox`-Element-`display`-Eigenschaft auf `flex` gesetzt werden, damit das `.lightbox`-Element und die beiden Elemente darin sichtbar werden.
1. Wenn Sie auf eines Ihrer `.gallery-item`-Elemente klicken, sollte die `#lightbox-image`-Element-`src` auf eine Vollbildversion des angeklickten Bildes gesetzt werden, indem `-thumbnail` aus dem `src`-Attribut des Bildes entfernt wird. Die Vollbildbilder befinden sich unter den folgenden Links:

    - `https://cdn.freecodecamp.org/curriculum/labs/stonehenge.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/storm.jpg`
    - `https://cdn.freecodecamp.org/curriculum/labs/trees.jpg`
  
1. Wenn Ihr `.lightbox`-Element sichtbar ist und Sie auf das `#close-btn`- oder das `.lightbox`-Element klicken, sollte die `.lightbox`-Element-`display`-Eigenschaft wieder auf `none` gesetzt werden.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet und die JavaScript-Datei in Ihrem HTML verlinken.

# --hints--

Sie sollten ein `div` mit der Klasse `gallery` innerhalb Ihres `body`-Elements haben.

```js
assert.equal(document.querySelector("body .gallery")?.tagName, "DIV");
```

Innerhalb des `.gallery`-Elements sollten Sie drei `img`-Elemente mit der Klasse `gallery-item` haben.

```js
assert.lengthOf(document.querySelectorAll(".gallery .gallery-item"), 3)
```

Innerhalb des `.gallery`-Elements sollten Sie ein `img`-Element mit der `src`, die auf `https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg` gesetzt ist, haben.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Innerhalb des `.gallery`-Elements sollten Sie ein `img`-Element mit der `src`, die auf `https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg` gesetzt ist, haben.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Innerhalb des `.gallery`-Elements sollten Sie ein `img`-Element mit der `src`, die auf `https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg` gesetzt ist, haben.

```js
const images = document.querySelectorAll(".gallery .gallery-item");
let srcCount = 0; 
const targetSrc = "https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg";

for (let image of images) {
  if (image.src === targetSrc) { 
    srcCount++; 
  }
}

assert.strictEqual(srcCount, 1);
```

Sie sollten ein `div`-Element mit der Klasse `lightbox` innerhalb Ihres `body`-Elements haben.

```js
const lightboxEl = document.querySelector("body .lightbox");
assert.exists(lightboxEl);
assert.equal(lightboxEl?.tagName, "DIV");
```

Innerhalb Ihres `.lightbox`-Elements sollten Sie ein `button`-Element mit der `id`, die auf `close-btn` gesetzt ist, haben.

```js
assert.equal(document.querySelector(".lightbox #close-btn")?.tagName, "BUTTON");
```

Innerhalb Ihres `.lightbox`-Elements sollten Sie ein `img`-Element mit der `id`, die auf `lightbox-image` gesetzt ist, haben.

```js
assert.equal(document.querySelector(".lightbox #lightbox-image")?.tagName, "IMG");
```

Ihr `.lightbox`-Element sollte eine feste Positionierung haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.position, "fixed");
```

Ihr `.lightbox`-Element sollte den gesamten Viewport abdecken.

```js
const style = new __helpers.CSSHelp(document).getStyle(".lightbox");

const width = style?.width;
const height = style?.height;

assert.oneOf(
  width,
  ["100%", "100vw"]
);

assert.oneOf(
  height,
  ["100%", "100vh"]
);
```

Ihr `.lightbox`-Element sollte mit der oberen linken Ecke des Containers ausgerichtet sein.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.top, "0px");
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.left, "0px");
```

Ihr `.lightbox`-Element sollte eine Hintergrundfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".lightbox")?.backgroundColor);
```

Ihr `.lightbox`-Element sollte anfangs verborgen sein.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(".lightbox")?.display, "none");
```

Wenn Sie auf eines Ihrer `.gallery-item`-Elemente klicken, sollte die `.lightbox`-Element-`display`-Eigenschaft auf `flex` gesetzt werden, damit `.lightbox` und die beiden Elemente darin sichtbar werden.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

assert.strictEqual(getComputedDisplay(lightbox), "none");

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");
```

Wenn Sie auf eines Ihrer `.gallery-item`-Elemente klicken, sollte die `#lightbox-image`-Element-`src` auf eine Vollbildversion des angeklickten Bildes gesetzt werden, indem `-thumbnail` aus dem `src`-Attribut des Bildes entfernt wird.

```js
const galleryItems = document.querySelectorAll(".gallery-item");
const lightboxImage = document.getElementById("lightbox-image");

assert.isAbove(galleryItems.length, 0);

function getFullSizeImageUrl(thumbnailUrl) {
  return thumbnailUrl.replace("-thumbnail", "");
}

galleryItems.forEach((item) => {
  const thumbnailUrl = item.src; 
  const expectedFullSizeUrl = getFullSizeImageUrl(thumbnailUrl);

  item.dispatchEvent(new MouseEvent("click", { bubbles: true }));

  assert.strictEqual(lightboxImage.src, expectedFullSizeUrl);
});
```

Wenn Ihr `.lightbox`-Element sichtbar ist und Sie auf die `#close-btn`-Schaltfläche klicken, sollte die `.lightbox`-Element-`display` wieder auf `none` gesetzt werden.

```js
const lightbox = document.querySelector(".lightbox");
const closeBtn = document.getElementById("close-btn"); 
const galleryItem = document.querySelector(".gallery-item");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

closeBtn.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

Wenn Ihr `.lightbox`-Element sichtbar ist und Sie auf das `.lightbox`-Element klicken, sollte die `.lightbox`-Element-`display` wieder auf `none` gesetzt werden.

```js
const lightbox = document.querySelector(".lightbox");

function getComputedDisplay(element) {
  return window.getComputedStyle(element).display;
}

const galleryItem = document.querySelector(".gallery-item");
galleryItem.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "flex");

lightbox.dispatchEvent(new Event("click", { bubbles: true }));

assert.strictEqual(getComputedDisplay(lightbox), "none");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
  </head>

  <body>

  </body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Lightbox Viewer</title>
    <link rel="stylesheet" href="styles.css">
  </head>

  <body>
    <h1 class="title">Museum Gallery</h1>

    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/stonehenge-thumbnail.jpg" alt="Stonehenge landscape with large stone structures." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/trees-thumbnail.jpg" alt="A peaceful forest with tall green trees." class="gallery-item">
      <img src="https://cdn.freecodecamp.org/curriculum/labs/storm-thumbnail.jpg" alt="A dramatic stormy sky with dark clouds and lightning." class="gallery-item">
    </div>
    <div id="lightbox" class="lightbox">
      <button id="close-btn" class="close-btn">&times;</button>
      <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  background-color: #f0f0f0;
  color: #333;
}

.title {
  margin-top: 100px;
  text-align: center;
  background-color: rgb(211, 221, 224);  
}

.gallery {
  display: flex;
  flex-wrap: wrap;
  gap: 10px;
  justify-content: center;
  padding: 20px;
}

.gallery-item {
  width: 250px;
  height: 300px;
  object-fit: cover;
  cursor: pointer;
  transition: transform 0.2s;
}

.gallery-item:hover {
  transform: scale(1.05);
}

.lightbox {
  display: none;
  position: fixed;
  top: 0;
  left: 0;
  width: 100%;
  height: 100%;
  background-color: rgba(0, 0, 0, 0.8);
  justify-content: center;
  align-items: center;
  z-index: 1000;
}

.lightbox img {
  max-width: 90%;
  max-height: 90%;
}

.close-btn {
  position: absolute;
  top: 20px;
  right: 10px;
  font-size: 2em;
  color: white;
  cursor: pointer;
  border: none;
  background-color: transparent;
}
```

```js
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");

function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; 
}

function closeLightbox() {
  lightbox.style.display = "none"; 
}

const closeButton = document.getElementById("close-btn");

closeButton.addEventListener("click", closeLightbox);

lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); 
  }
});

const galleryItems = document.querySelectorAll(".gallery-item");

galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); 
  });
});
```
