---
id: a9bd25c716030ec90084d8a1
title: 청키 몽키 알고리즘 구현하기
challengeType: 26
dashedName: implement-the-chunky-monkey-algorithm
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `chunkArrayInGroups`라는 이름의 함수를 작성하세요. 이 함수는 첫 번째 인수로 배열을, 두 번째 인수로 숫자를 받습니다. 함수는 배열을 두 번째 인수와 같은 길이의 작은 배열들로 나누고, 이들을 2차원 배열로 반환해야 합니다.

# --hints--

`chunkArrayInGroups(["a", "b", "c", "d"], 2)`는 `[["a", "b"], ["c", "d"]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups(['a', 'b', 'c', 'd'], 2), [
  ['a', 'b'],
  ['c', 'd']
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3)`는 `[[0, 1, 2], [3, 4, 5]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3), [
  [0, 1, 2],
  [3, 4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2)`는 `[[0, 1], [2, 3], [4, 5]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2), [
  [0, 1],
  [2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4)`는 `[[0, 1, 2, 3], [4, 5]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4), [
  [0, 1, 2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3)`는 `[[0, 1, 2], [3, 4, 5], [6]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3), [
  [0, 1, 2],
  [3, 4, 5],
  [6]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4)`는 `[[0, 1, 2, 3], [4, 5, 6, 7], [8]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4), [
  [0, 1, 2, 3],
  [4, 5, 6, 7],
  [8]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2)`는 `[[0, 1], [2, 3], [4, 5], [6, 7], [8]]`를 반환해야 합니다.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2), [
  [0, 1],
  [2, 3],
  [4, 5],
  [6, 7],
  [8]
]);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function chunkArrayInGroups(arr, size) {
  let out = [];

  for (let i = 0; i < arr.length; i += size) {
    out.push(arr.slice(i, i + size));
  }

  return out;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```
