---
id: 68c34ace730ca43e5fa7b8d6
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

이제 표에 또 다른 행을 추가할 시간입니다.

먼저 `tr` 요소를 하나 더 추가하세요. 그 `tr` 요소 안에 `th` 요소를 추가하고, `scope` 속성을 `"row"`로 설정한 뒤 `10:00 AM` 텍스트 내용을 넣으세요.

그다음, 다음 텍스트 내용을 가진 세 개의 `td` 요소를 추가하세요:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

`tr` 요소 안에 두 번째 `tbody` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

두 번째 `th` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

`th` 요소는 `10:00 AM` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

`th` 요소는 `scope` 속성을 `"row"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

두 번째 `tr` 요소는 세 개의 `td` 요소를 가져야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

첫 번째 `td` 요소는 `Accessibility Deep Dive` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

두 번째 `td` 요소는 `CSS for Beginners` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

세 번째 `td` 요소는 `Inclusive Design Principles` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
