---
id: 68c34f4aeb2b1a4170506751
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

다음으로, 표에 다섯 번째 행을 추가하세요. `tr` 요소를 하나 더 추가하는 것부터 시작하세요. 그 `tr` 요소 안에 `th` 요소를 추가하고, `scope` 속성을 `"row"`로 설정하며 텍스트 내용은 `12:30 PM`로 하세요. 그런 다음 그 `th` 요소 아래에 텍스트 내용이 `td`인 `Lunch Break` 요소를 추가하세요. `td` 요소에는 세 개의 트랙 전체에 걸치도록 `colspan` 속성을 `3`로 설정해야 합니다.

# --hints--

`tr` 요소 안에 다섯 번째 `tbody` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

다섯 번째 `th` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

`th` 요소의 텍스트 내용은 `12:30 PM`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

`th` 요소에는 `scope` 속성이 `"row"`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

다섯 번째 `td` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

`td` 요소의 텍스트 내용은 `Lunch Break`이어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

`td` 요소에는 `colspan` 속성이 `3`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
