---
id: 68c34329d430583a6cd5f7f1
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`tr` 요소 안에 `th` 요소 네 개를 추가하세요. 첫 번째 `th` 요소에는 `Time` 텍스트가 있어야 하고, 두 번째에는 `Track A` 텍스트가, 세 번째에는 `Track B` 텍스트가, 네 번째에는 `Track C` 텍스트가 있어야 합니다.

# --hints--

`tr` 요소에는 `th` 요소 네 개가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

첫 번째 `th` 요소에는 `Time` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

두 번째 `th` 요소에는 `Track A` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

세 번째 `th` 요소에는 `Track B` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

네 번째 `th` 요소에는 `Track C` 텍스트가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

--fcc-editable-region--
    <thead>
      <tr>
        
      </tr>
    </thead>
--fcc-editable-region--
  </table>
</body>
</html>
```
