---
id: 68c346d03709343bff132707
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

다음 몇 단계에서는 표의 본문을 작성할 것입니다.

`tbody` 요소 아래에 `thead` 요소를 추가하는 것부터 시작하세요. 그런 다음 `tbody` 안에 `tr` 요소를 추가하세요.

# --hints--

`tbody` 요소 아래에 `thead` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("tbody"));
```

`tr` 요소 안에 `tbody` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
