---
id: 68c34e201b347c40acda8111
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

이제 표에 또 다른 행을 추가할 시간입니다.

`tr` 요소를 하나 더 추가하는 것부터 시작하세요. 그 `tr` 요소 안에 `th` 요소를 추가하고, `scope` 속성을 `"row"`로 설정한 뒤 `11:30 AM` 텍스트 내용을 넣으세요.

그다음, 다음 텍스트 내용을 가진 세 개의 `td` 요소를 추가하세요:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

`tr` 요소 안에 네 번째 `tbody` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

네 번째 `th` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

`th` 요소는 `11:30 AM` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

`th` 요소는 `scope` 속성을 `"row"`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

네 번째 `tr` 요소는 세 개의 `td` 요소를 가져야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

첫 번째 `td` 요소는 `AR/VR in Education` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

두 번째 `td` 요소는 `JavaScript Fundamentals` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

세 번째 `td` 요소는 `Design Systems at Scale` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
