---
id: 68c34a35c91b053dbae2a8c6
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`scope` 속성의 또 다른 값은 `row`으로, 헤더 셀이 해당 행 전체의 헤더임을 나타냅니다.

`th` 요소 안에 `scope` 속성을 추가하고 값으로 `row`를 넣으세요.

# --hints--

`th` 요소에 `scope` 속성을 추가하고 값으로 `row`를 넣으세요.

```js
assert.equal(document.querySelector("tbody tr th")?.getAttribute("scope"), "row");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        <th>9:00 AM</th>
        --fcc-editable-region--
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
