---
id: 68c34d57ed874e4004f7b5d4
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

지금 `td` 요소는 텍스트 내용이 `Break`인 상태로 한 열만 차지하고 있습니다. 하지만 세 개의 열 전체를 차지하면 좋겠습니다.

이전 워크숍과 수업에서 기억하듯이, `colspan` 속성을 사용하면 표 셀을 여러 열에 걸쳐 확장할 수 있습니다.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

`colspan` 요소에 `td` 속성을 추가하고 그 값을 `3`로 설정하세요.

# --hints--

텍스트 내용이 `td`인 `Break` 요소에 `colspan` 속성을 `3`로 설정해야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
