---
id: 68c3420cb0176439a95af1b8
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

다음 단계는 표 머리글을 추가하는 것입니다.

`thead` 요소 아래에 `caption` 요소를 추가하세요. `thead` 안에 `tr` 요소를 추가하세요.

# --hints--

`thead` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("table thead"));
```

`tr` 안에 `thead` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
