---
id: 68c341332d98fe38fae733af
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 표를 만들어 나갈 시간입니다.

`table` 요소 아래에 `h1` 요소를 추가하세요. `table` 요소 안에 `caption` 요소를 추가하고 텍스트 `Schedule by Track and Time`를 넣으세요.

# --hints--

`table` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("table"));
```

`table` 요소 안에 `caption` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("table caption"));
```

`caption` 요소 안에 텍스트 `Schedule by Track and Time`가 있어야 합니다.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
