---
id: 68c343692285523aef0fc6bf
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`scope` 속성은 헤더 셀이 행, 열 또는 행이나 열 그룹의 헤더인지 지정하는 데 사용됩니다. 예시는 다음과 같습니다.

```html
<th scope="col">Example Header</th>
```

이것은 화면 읽기 프로그램이 헤더 셀과 데이터 셀 간의 관계를 이해하는 데 도움이 됩니다.

모든 `th` 요소에 `scope` 속성을 추가하고 값은 `col`로 설정하세요.

# --hints--

각 `th` 요소에는 값이 `scope`인 `"col"` 속성이 있어야 합니다.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
