---
id: 68c34c608cf9e53f5d36b1d5
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

다음으로, 표에 세 번째 행을 추가하세요. 또 다른 `tr` 요소를 추가하는 것부터 시작하세요. 그 `tr` 요소 안에 `th` 요소를 추가하고, `scope` 속성을 `"row"`로 설정한 뒤 `11:00 AM` 텍스트 내용을 넣으세요. 그런 다음 그 `th` 요소 아래에 `td` 요소를 추가하고 `Break` 텍스트 내용을 넣으세요.

# --hints--

`tr` 요소 안에 세 번째 `tbody` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

세 번째 `th` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

`th` 요소는 `11:00 AM` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

`th` 요소는 `scope`로 설정된 `"row"` 속성을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

세 번째 `td` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

`td` 요소는 `Break` 텍스트 내용을 가져야 합니다.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
