---
id: 68c33e502cd97f2fd52cb63b
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 기술 컨퍼런스 일정을 구성하면서 접근 가능한 표 요소를 만드는 방법을 배우세요.

먼저 텍스트가 `h1`인 `Tech Conference 2025 Schedule` 요소를 추가하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

`h1` 요소에는 `Tech Conference 2025 Schedule` 텍스트가 있어야 합니다.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
