---
id: 68c3482278a14e3cf225e897
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`tr` 요소 안에 `th` 요소를 추가하고 `9:00 AM` 텍스트를 넣으세요. 그런 다음 그 `th` 요소 아래에 `td` 요소 세 개를 추가하고 각각 `Keynote: Tech Future`, `Intro to Web Dev`, `UX for All` 텍스트를 넣으세요.

# --hints--

`th` 요소 안에 `tr` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("tbody tr th"));
```

`th` 요소에는 `9:00 AM` 텍스트 내용이 있어야 합니다.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

`td` 요소 아래에 `th` 요소가 총 세 개 있어야 합니다.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

첫 번째 `td` 요소에는 `Keynote: Tech Future` 텍스트 내용이 있어야 합니다.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

두 번째 `td` 요소에는 `Intro to Web Dev` 텍스트 내용이 있어야 합니다.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

세 번째 `td` 요소에는 `UX for All` 텍스트 내용이 있어야 합니다.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
