---
id: 5dc1798ff86c76b9248c6eb3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie unter dem `h1`-Element ein `h2`-Element mit diesem Text hinzu:

`Cat Photos`

# --hints--

Ihr `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte ein schließendes Tag haben. Schließende Tags haben diese Syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Sie sollten nur ein `h1`-Element haben. Entfernen Sie das zusätzliche.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Der Text Ihres `h1` Elements sollte `CatPhotoApp` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

Ihre `h2`-Tags sollten in Kleinbuchstaben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
assert.notMatch(code, /<\/?H2>/);
```

Ihr `h2`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Ihr `h2`-Element sollte ein schließendes Tag haben. Schließende Tags haben direkt nach dem `/`-Zeichen ein `<`.

```js
assert.match(code, /<\/h2\>/);
```

Der Text Ihres `h2`-Elements sollte `Cat Photos` sein. Platzieren Sie nur den Text `Cat Photos` zwischen den öffnenden und schließenden `h2`-Tags.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

Ihr `h2`-Element sollte unter dem `h1`-Element sein. Das `h1`-Element hat eine höhere Bedeutung und muss über dem `h2`-Element stehen.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
