---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop arbeiten Sie weiter mit grundlegenden HTML-Elementen wie Überschriften, Absätzen und Listen, indem Sie eine Katzenfoto-App erstellen.

Beginnen Sie den Workshop, indem Sie ein `h1`-Element mit dem Text `CatPhotoApp` hinzufügen.

# --hints--

Der Text `CatPhotoApp` sollte im Code vorhanden sein. Überprüfen Sie gegebenenfalls Ihre Rechtschreibung.

```js
assert.match(code, /catphotoapp/i);
```

Ihr `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Ihre `h1`-Tags sollten in Kleinbuchstaben geschrieben sein. Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben.

```js
assert.notMatch(code, /<\/?H1>/);
```

Ihr `h1`-Element sollte ein schließendes Tag haben. Schließende Tags haben diese Syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Der Text des `h1`-Elements sollte `CatPhotoApp` sein. Sie haben entweder den Text weggelassen, einen Tippfehler oder er steht nicht zwischen dem öffnenden und schließenden Tag des `h1`-Elements.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
