---
id: 5dc23f9bf86c76b9248c6eba
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Sie können Bilder zu Ihrer Website hinzufügen, indem Sie das `img`-Element verwenden. `img`-Elemente haben ein öffnendes Tag ohne schließendes Tag. Ein Element ohne schließendes Tag wird als <dfn>void element</dfn> bezeichnet.

Fügen Sie ein `img`-Element unter dem `p`-Element hinzu. Zu diesem Zeitpunkt wird im Browser kein Bild angezeigt.

# --hints--

Ihr `img`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Ihr `img`-Element sollte kein schließendes Tag `</img>` haben.

```js
assert.notMatch(code, /<\/img\>/);
```

Sie sollten nur ein `img`-Element haben. Entfernen Sie alle zusätzlichen.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Ihr `img`-Element sollte unter dem `p`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
