---
id: 5dc24165f86c76b9248c6ebc
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Alle `img`-Elemente sollten ein `alt`-Attribut haben. Der Text des `alt`-Attributs wird von Screenreadern verwendet, um die Barrierefreiheit zu verbessern, und wird angezeigt, falls das Bild nicht geladen werden kann.

Hier ist ein Beispiel für ein `img`-Element mit einem `alt`-Attribut:

```html
<img src="cat.jpg" alt="A cat">
```

Fügen Sie im `img`-Element ein `alt`-Attribut mit diesem Text hinzu:

`A cute orange cat lying on its back`

# --hints--

Ihr Code sollte ein `img`-Element enthalten. Sie haben das `img`-Element in einem früheren Schritt entfernt.

```js
assert.exists(document.querySelector('img'));
```

Ihr `img`-Element hat kein `alt`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen sind.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Der Wert des `alt`-Attributs Ihres `img`-Elements ist etwas anderes als `A cute orange cat lying on its back`. Stellen Sie sicher, dass der Wert des `alt`-Attributs in Anführungszeichen steht.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
