---
id: 5dc23991f86c76b9248c6eb8
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Im vorherigen Schritt haben Sie die `h1`, `h2`, den Kommentar und die `p`-Elemente innerhalb des `main`-Elements platziert. Das nennt man *Verschachtelung*. Verschachtelte Elemente sollten zwei Leerzeichen weiter rechts vom Element, in dem sie verschachtelt sind, stehen. Dieser Abstand heiĂźt EinrĂĽckung und dient dazu, HTML leichter lesbar zu machen.

Hier ist ein Beispiel fĂĽr Verschachtelung und EinrĂĽckung:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Das `h1`-Element, das `h2`-Element und der Kommentar sind im folgenden Code zwei Leerzeichen weiter eingerĂĽckt als das `main`-Element. Verwenden Sie die Leertaste auf Ihrer Tastatur, um vor dem `p`-Element zwei weitere Leerzeichen hinzuzufĂĽgen, damit es ebenfalls richtig eingerĂĽckt ist.

# --hints--

Ihr Code sollte ein `h2`-Element mit dem Text `Cat Photos` enthalten. MĂ¶glicherweise haben Sie es versehentlich gelĂ¶scht, es fehlen sowohl Ă¶ffnendes als auch schlieĂźendes Tag, oder der Text wurde geĂ¤ndert.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

Sie sollten die `ul`- oder `li`-Elemente aus dem Beispiel nicht hinzufĂĽgen.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

Sie sollten die EinrĂĽckung in der Zeile mit Ihrem `h2`-Element nicht Ă¤ndern. Das Ă¶ffnende Tag sollte 6 Leerzeichen vom Zeilenanfang entfernt beginnen. Sie kĂ¶nnen den Schritt neu starten, um die ursprĂĽngliche EinrĂĽckung wiederherzustellen.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Ihr Code sollte einen Kommentar enthalten. Sie haben den Kommentar aus einem frĂĽheren Schritt entfernt.

```js
assert.match(code, /<!--.*-->/);
```

Der Text des Kommentars sollte `TODO: Add link to cat photos` sein. Ă„ndern Sie weder den Text noch die EinrĂĽckung des Kommentars.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

Sie sollten die EinrĂĽckung in der Zeile mit Ihrem Kommentar-Element nicht Ă¤ndern. Das Ă¶ffnende Tag sollte 6 Leerzeichen vom Zeilenanfang entfernt beginnen. Sie kĂ¶nnen den Schritt neu starten, um die ursprĂĽngliche EinrĂĽckung wiederherzustellen.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Ihr Code sollte ein `p`-Element enthalten. Sie haben das `p`-Element aus einem frĂĽheren Schritt entfernt.

```js
assert.exists(document.querySelector('p'));
```

Der Text des `p`-Elements sollte `Everyone loves cute cats online!` sein. Ă„ndern Sie weder den Text, die EinrĂĽckung noch die Zeichensetzung des `p`-Elements.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Das Ă¶ffnende `p`-Tag sollte die gleiche EinrĂĽckung haben wie Ihre `h2`- und Kommentar-Elemente. Das Ă¶ffnende Tag sollte 6 Leerzeichen vom Zeilenanfang entfernt beginnen.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
