---
id: 5dc24614f86c76b9248c6ebd
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Sie können mit dem Ankerelement (`a`) auf eine andere Seite verknüpfen.

Hier ist ein Beispiel, das auf `https://www.freecodecamp.org` verknüpft:

```html
<a href="https://www.freecodecamp.org"></a>
```

Fügen Sie nach dem Absatz ein Ankerelement hinzu, das auf `https://freecatphotoapp.com` verknüpft. Zu diesem Zeitpunkt wird der Link in der Vorschau noch nicht angezeigt.

# --hints--

Ihr Ankerelement (`a`) sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Ihr Ankerelement (`a`) sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/a\>/);
```

Ihr Ankerelement (`a`) sollte unter dem `p`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Ihr Ankerelement (`a`) hat kein `href`-Attribut. Prüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen ist und/oder vor allen Attributnamen Leerzeichen stehen.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

Ihr Ankerelement (`a`) sollte auf `https://freecatphotoapp.com` verknüpfen. Sie haben entweder die URL weggelassen oder einen Tippfehler.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Obwohl Sie das `a`-Attribut des Ankerelements (`href`) auf den richtigen Link gesetzt haben, wird empfohlen, den Wert eines Attributs immer in Anführungszeichen zu setzen.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
