---
id: 5dc2385ff86c76b9248c6eb7
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 hat ein paar Elemente, die verschiedene Inhaltsbereiche beschreiben. Diese Elemente machen dein HTML einfacher zu lesen und helfen mit Suchmaschinenoptimierung (SEO - Search Engine Optimization) und Barrierefreiheit.

Das `main`-Element wird verwendet, um den Hauptinhalt des Hauptteils eines HTML-Dokuments darzustellen. Der Inhalt des `main`-Elements sollte für das Dokument einzigartig sein und nicht in anderen Teilen des Dokuments wiederholt werden.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Kennzeichne den Hauptteil deiner Seite, indem du einen öffnenden `<main>` Tag vor dem `h1`-Element und einen schließenden `</main>` Tag nach dem `p`-Element hinzufügst.

# --hints--

Dein `main`-Element sollte einen öffnenden Tag haben. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Dein `main`-Element sollte einen schließenden Tag haben. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/main\>/);
```

Der öffnende Tag deines `main`-Elements sollte unter dem öffnenden `body` Tag sein. Du hast die Reihenfolge vertauscht.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Der öffnende Tag deines `main`-Elements sollte über dem `h1`-Element sein. Sie stehen in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Der schließende Tag deines `main`-Elements sollte unter dem `p` Element sein. Du hast die Reihenfolge vertauscht.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
