---
id: 5dfa2407b521be39a3de7be1
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Um Links in einem neuen Tab zu öffnen, können Sie das `target`-Attribut im Anker-Element (`a`) verwenden.

Das `target`-Attribut gibt an, wo das verlinkte Dokument geöffnet wird. `target="_blank"` öffnet das verlinkte Dokument in einem neuen Tab oder Fenster.

Hier ist die grundlegende Syntax für ein `a`-Element mit einem `target`-Attribut:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Fügen Sie dem öffnenden Tag des `cat photos`-Anker-Elements (`a`) ein `target`-Attribut mit dem Wert `_blank` hinzu, damit der Link in einem neuen Tab geöffnet wird.

# --hints--

Ihr `p`-Element sollte ein verschachteltes Anker-Element (`a`) mit dem Text `cat photos` enthalten. Möglicherweise haben Sie es gelöscht oder einen Tippfehler.

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

Ihr Anker-Element (`a`) hat kein `target`-Attribut. Überprüfen Sie, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder vor allen Attributnamen Leerzeichen vorhanden sind.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Der Wert des `target`-Attributs sollte `_blank` sein. Sie haben entweder den Wert weggelassen oder einen Tippfehler. Denken Sie daran, dass Attributwerte in Anführungszeichen stehen sollten.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
