---
id: 5dfa3589eacea3f48c6300ae
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Fügen Sie innerhalb des zweiten `section`-Elements ein neues `h2`-Element mit dem Text `Cat Lists` hinzu.

# --hints--

Ihr `section` Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Ihr `h2`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Ihr `h2`-Element sollte ein schließendes Tag haben. Schließende Tags haben direkt nach dem `/`-Zeichen ein `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Ihr zweites `h2`-Element sollte direkt über dem schließenden Tag des zweiten `section`-Elements stehen. Es befindet sich nicht an der richtigen Stelle.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

Das zweite `h2`-Element sollte den Text `Cat Lists` enthalten. Sie haben entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
