---
id: 5dfb6a35eacea3f48c6300b4
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Ein `figcaption`-Element wird verwendet, um eine Bildunterschrift hinzuzufügen, die das Bild beschreibt, das im `figure`-Element enthalten ist.

Hier ist ein Beispiel für ein `figcaption`-Element mit der Bildunterschrift `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

Fügen Sie nach dem im `figure`-Element verschachtelten Bild ein `figcaption`-Element mit folgendem Text hinzu:

`Cats love lasagna.`

# --hints--

Ihr `figcaption`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('figcaption'));
```

Ihr `figcaption`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/figcaption\>/);
```

Das `figcaption`-Element sollte im `figure`-Element verschachtelt sein.

```js
assert.exists(document.querySelector('figure > figcaption'));
```

Das Lasagna-`img`-Element sollte im `figure`-Element verschachtelt sein.

```js
assert.exists(document.querySelector('figure > img'))
assert.equal(
    document.querySelector('figure > img')?.getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Das `figcaption`-Element, das im `figure`-Element verschachtelt ist, sollte unter dem `img`-Element stehen. Sie haben sie in der falschen Reihenfolge.

```js
assert.equal(
  document.querySelector('figcaption')?.previousElementSibling.nodeName, 'IMG'
);
```

Der Text Ihres `figcaption`-Elements sollte `Cats love lasagna.` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.match(
  document.querySelector('figcaption')?.innerText?.trim(), /Cats love lasagna.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          
--fcc-editable-region--          
        </figure>
      </section>
    </main>
  </body>
</html>
```
