---
id: 5ef9b03c81a63668521804d3
title: Schritt 30
challengeType: 0
dashedName: step-30
---

# --description--

Füge nach der geordneten Liste ein weiteres `figure`-Element hinzu.

# --hints--

Dein `figure`-Element sollte ein öffnendes Tag haben. Opening tags have this syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

Dein `figure` sollte ein schließendes Tag haben. Closing tags have a `/` just after the `<` character.

```js
assert.lengthOf(code.match(/<\/figure>/g), 2);
```

Genau über dem abschließenden Tag des zweiten `section`-Elements sollte ein `figure`-Element stehen.

```js
assert.equal(document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
