---
id: 5ef9b03c81a63668521804e8
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--

Verschachteln Sie ein `p`-Element mit dem Text `No Copyright - freeCodeCamp.org` innerhalb des `footer`-Elements.

# --hints--

Sie haben entweder das `footer`-Element gelöscht oder es fehlt ein öffnendes Tag oder schließendes Tag.

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

Ihr `footer`-Element sollte ein `p`-Element enthalten. Stellen Sie sicher, dass Sie ein öffnendes Tag und ein schließendes Tag für das `p`-Element hinzufügen.

```js
assert.exists(document.querySelector('footer > p'));
```

Ihr `footer`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 3);
```

Der Text Ihres `p`-Elements sollte `No Copyright - freeCodeCamp.org` sein. Sie haben entweder den Text ausgelassen, einen Tippfehler oder er befindet sich nicht zwischen dem öffnenden und schließenden Tag des `p`-Elements.

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0]?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /^No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      
--fcc-editable-region--      
    </footer>
  </body>
</html>
```
