---
id: 5ef9b03c81a63668521804ea
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Beachten Sie, dass alles, was Sie bisher zur Seite hinzugefügt haben, sich innerhalb des `body`-Elements befindet. Alle Inhalts-Elemente der Seite, die auf der Seite gerendert werden sollen, kommen in das `body`-Element. Andere wichtige Informationen gehören jedoch in das `head`-Element.

Das `head`-Element wird verwendet, um Metadaten über das Dokument zu enthalten, wie zum Beispiel dessen Titel, Links zu Stylesheets und Skripten. Metadaten sind Informationen über die Seite, die nicht direkt auf der Seite angezeigt werden.

Fügen Sie ein `head`-Element oberhalb des `body`-Elements hinzu.

# --hints--

Sie haben entweder das `body`-Element gelöscht oder es fehlt ein öffnendes Tag oder ein schließendes Tag.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Ihr `head`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Ihr `head`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /\<\/head\>/);
```

Ihr `head`-Element sollte oberhalb des öffnenden `body`-Elements stehen. Sie haben es an einer anderen Stelle eingefügt.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
