---
id: 5ef9b03c81a63668521804d0
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

Um ein bestimmtes Wort oder eine Phrase hervorzuheben, können Sie das `em`-Element verwenden.

Heben Sie das Wort `love` im `figcaption`-Element hervor, indem Sie es in ein Hervorhebungs-`em`-Element einschließen.

# --hints--

Ihr Hervorhebungs-`em`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

Ihr Hervorhebungs-`em`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /<\/em\>/);
```

Sie haben entweder das `figcaption`-Element gelöscht oder es fehlt ein öffnendes oder schließendes Tag.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

Ihr Hervorhebungs-`em`-Element sollte den Text `love` umschließen. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

Der Text von `figcaption` sollte `Cats love lasagna` sein. Überprüfen Sie Tippfehler und ob die notwendigen Leerzeichen um die öffnenden und schließenden Tags des `em`-Elements vorhanden sind.

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
