---
id: 5ef9b03c81a63668521804d1
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

Fügen Sie nach dem `figure`-Element ein weiteres `h3`-Element mit dem Text hinzu:

`Top 3 things cats hate:`

# --hints--

Es sollte ein `h3`-Element direkt über dem schließenden Tag des zweiten `section`-Elements stehen. Stellen Sie sicher, dass es ein öffnendes und schließendes Tag hat.

```js
assert.equal(
  document.querySelectorAll('main > section')[1]?.lastElementChild.nodeName,
    'H3'
);
assert.lengthOf(code.match(/<\/h3\>/g),  2);
```

Das neue `h3`-Element sollte den Text `Top 3 things cats hate:` enthalten. Achten Sie darauf, den Doppelpunkt am Ende des Texts einzuschließen.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild.innerText?.trim().toLowerCase()
    .replace(/\s+/g, ' '), 'top 3 things cats hate:'
);
```

Über dem neuen `figure`-Element sollte ein `h3` stehen. Möglicherweise haben Sie versehentlich das `figure`-Element gelöscht.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  ?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'H3');
assert.equal(secondSectionLastElemNode.previousElementSibling.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
