---
id: 5ef9b03c81a63668521804e9
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

Verwandeln Sie den vorhandenen `freeCodeCamp.org`-Text in einen Link, indem Sie ihn in ein Ankerelement (`a`) einschließen. Das `href`-Attribut sollte auf `https://www.freecodecamp.org` gesetzt werden.

# --hints--

Ihr Ankerelement (`a`) sollte innerhalb des `p`-Elements im `footer` verschachtelt sein. Stellen Sie sicher, dass Sie ein öffnendes Tag und ein schließendes Tag für das Ankerelement (`a`) hinzufügen.

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Ihr Ankerelement (`a`) sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Ihr Ankerelement (`a`) sollte ein `href`-Attribut mit dem Wert `https://www.freecodecamp.org` haben. Möglicherweise haben Sie das Attribut/den Wert weggelassen oder einen Tippfehler.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Der Linktext sollte `freeCodeCamp.org` sein. Sie haben entweder den Text weggelassen oder einen Tippfehler.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Nachdem Sie das Ankerelement (`a`) verschachtelt haben, sollte der einzige sichtbare Inhalt des `p`-Elements im Browser `No Copyright - freeCodeCamp.org` sein. Überprüfen Sie den Text, die Leerzeichen oder die Zeichensetzung sowohl des `p` als auch des verschachtelten Ankerelements.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
