---
id: 5ef9b03c81a63668521804eb
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

Das `title`-Element bestimmt, was Browser in der Titelleiste oder im Tab für die Seite anzeigen.

Fügen Sie ein `title`-Element innerhalb des `head`-Elements mit dem folgenden Text hinzu:

`CatPhotoApp`

# --hints--

Sie haben entweder das `head`-Element gelöscht oder es fehlt ein öffnendes Tag oder schließendes Tag.

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

Ihr `title`-Element sollte im `head`-Element verschachtelt sein. Stellen Sie sicher, dass Sie ein öffnendes Tag und ein schließendes Tag für das `title`-Element hinzufügen.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

Ihr `title`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert.match(code, /\<\/title\>/);
```

Der Text Ihres `title`-Elements sollte `CatPhotoApp` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
