---
id: 5f07c98cdb9413cbd4b16750
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Es ist an der Zeit, einen neuen Abschnitt hinzuzufügen. Füge ein zweites `section`-Element unter dem vorhandenen `section`-Element hinzu.

# --hints--

Dein `section`Element sollte ein öffnendes Tag haben. Opening tags have this syntax: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Du solltest nur ein öffnendes `section`-Tag haben. Entferne bitte alle zusätzlichen.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Dein `section`-Element sollte ein abschließendes Tag haben. Closing tags have a `/` just after the `<` character.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Du solltest nur ein abschließendes `section`-Tag hinzufügen. Entferne bitte alles Zusätzliche.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Das zweite `section`-Element sollte nicht im ersten `section`-Element eingebettet sein.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Beide `section`-Elemente sollten sich zwischen dem öffnenden und abschließenden Tag deines `main`-Elements befinden.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```
