---
id: 69b83e35f19ba26ba1fa517a
title: Tengeneza Mfululizo wa Taa za Barabarani
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

Katika maabara hii, utaiga mizunguko ya taa za barabarani inayoweza kusanidiwa na kurekodi kasoro.

Utafanya kazi na vitu vya usanidi vinavyoelezea hatua za taa ya barabarani. Kila kitu cha usanidi kina sifa zifuatazo:

- `fault`: bendera ya thamani ya kweli au siyo kweli inayosababisha kusimamisha mapema wakati ni `true`.
- `phases`: safu ya data ya vitu vya hatua.

Kila kitu cha hatua ndani ya `phases` kina sifa zifuatazo:

- `color`: mfuatano wa herufi unaowakilisha rangi ya taa (`"green"`, `"yellow"`, au `"red"`).
- `duration`: nambari kamili chanya inayowakilisha muda wa hatua kwa sekunde.

Unaweza kurejelea vitu vya `config1`, `config2`, `config3`, na `config4` vilivyotolewa kama mifano ya vitu vya usanidi vinavyowezekana.

**Kumbuka**: Usiongeze taarifa zozote za ziada za `console.log()`, kwani zinaweza kusababisha mitihani kushindwa.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

# Hadithi za Mtumiaji

1. Unapaswa kuwa na kitendakazi kinachoitwa `runSequence` chenye vigezo viwili: `config` na `cycles`, ambapo `config` inawakilisha kitu cha usanidi na `cycles` inawakilisha idadi kubwa ya mara mfululizo unaweza kuendeshwa kabla ya kusimama.

1. Unapaswa kutekeleza `runSequence(config, cycles)` kwa kutumia mzunguko wa `for` au `while` kupita kupitia kila hatua katika idadi iliyotolewa ya mizunguko.

1. Kitendakazi cha `runSequence` kinapaswa:
    - Rekodi `No phases found` na urudi mara moja ikiwa `config.phases` ni tupu.
    - Rekodi `Faulted phase!` na simamisha uigaji mapema ikiwa `config.fault` imewekwa kuwa `true`.
    - Rekodi `Invalid phase detected` ikiwa `duration <= 0`.
    - Rekodi `Switching to [color] for [duration] s` kwa kila hatua halali. Badilisha `[color]` na `[duration]` kwa sifa zinazolingana za kitu cha hatua.
    - Kwa mfano, `runSequence(config1, 1)` inapaswa kurekodi:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Unapaswa kuwa na kitendakazi kinachoitwa `generateTimeline` chenye vigezo viwili: `config` na `cycles`.

1. Kitendakazi cha `generateTimeline` kinapaswa:
    - Rekodi muda uliopita wa jumla baada ya kila hatua katika mizunguko ndani ya safu ya data, ukiongeza `duration` ya kila hatua kwenye jumla inayorudiwa unaporudia.
    - Shughulikia hatua zote zilizo na kasoro na zisizo halali bila uthibitisho, hata kama `config.fault` ni `true` au `duration <= 0`.
    - Rudisha safu ya data ya alama za muda za jumla.
    - Kwa mfano, `generateTimeline(config1, 1)` inapaswa kurudisha safu ya data `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

Hupaswi kufanya mabadiliko katika vitu vya `config1`, `config2`, `config3`, na `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Unapaswa kuwa na kitendakazi kinachoitwa `runSequence`.

```js
assert.isFunction(runSequence);
```

Kitendakazi chako cha `runSequence` kinapaswa kupokea vigezo viwili, `config` na `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` inapaswa kurekodi `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` kwa mpangilio.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` inapaswa kurekodi `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` kwa mpangilio.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` inapaswa kurekodi `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` kwa mpangilio.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` inapaswa kurekodi `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` inapaswa kurekodi `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Unapaswa kuwa na kitendakazi kinachoitwa `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

Kitendakazi chako cha `generateTimeline` kinapaswa kupokea vigezo viwili, `config` na `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` inapaswa kurudisha `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` inapaswa kurudisha `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` inapaswa kurudisha `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` inapaswa kurudisha `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` inapaswa kurudisha `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
