---
id: 6391d1a4f7ac71efd0621380
title: 레시피 페이지 프로젝트 만들기
challengeType: 25
dashedName: top-build-a-recipe-project
---

# --description--

웹사이트는 몇 가지 레시피로 링크된 메인 인덱스 페이지로 구성됩니다. 완성 시점에는 웹사이트가 아주 예쁘게 보이지 않을 수 있습니다.

**사용자 스토리:**

1. 레시피 페이지에는 `DOCTYPE` 태그가 포함되어야 합니다.
1. 레시피 페이지에는 `html` 요소가 있어야 하며, 그 자식으로 `head`와 `body` 요소가 있어야 합니다.
1. `title` 요소 안에 `head` 요소가 있어야 하며, 그 안에 `The Odin Recipes` 텍스트가 있어야 합니다.
1. `h1` 요소가 보이고 그 안에 `Creamy Chocolate Fudge` 텍스트가 있어야 합니다.
1. `alt` 속성이 있는 관련 이미지가 보여야 합니다.
1. 이미지 아래에 `h2` 요소가 있고 그 안에 `Description` 텍스트가 있어야 합니다.
1. `Description` 아래에 레시피를 설명하는 두 개의 단락이 보여야 합니다.
1. `h2` 요소가 있고 그 안에 `Ingredients` 텍스트가 있어야 합니다.
1. `Ingredients` 제목 아래에는 레시피에 필요한 재료가 들어 있는 순서 없는 목록이 있어야 합니다.
1. 재료 목록 아래에 `Steps`라는 제목이 하나 더 있어야 합니다.
1. 레시피를 완성하는 데 필요한 몇 가지 단계가 들어 있는 순서 있는 목록이 보여야 합니다.
1. 단계 아래에는 `h2` 요소가 있고 그 안에 `More Recipes` 텍스트가 있어야 합니다.
1. 순서 없는 목록 안에 몇 개의 항목이 있고, 그 안에 앵커 요소가 포함된 다른 레시피로 연결되는 링크가 몇 개 보여야 합니다.
1. 이 앵커 요소들은 `href` 속성을 가지고 있어야 하며, 그 값은 `#`로 설정되어야 합니다.

# --hints--

`DOCTYPE` 태그가 있어야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+?html\s*?>/gi));
```

`html` 요소가 있어야 하며, 그 안에 `head`와 `body` 요소가 있어야 합니다.

```js
const html = document.querySelectorAll('html')[0];
const head = document.querySelectorAll('html > head')[0];
const body = document.querySelectorAll('html > body')[0];

assert(html && head && body);
```

`title` 요소 안에 `head` 요소가 있어야 하며, 그 안에 `The Odin Recipes` 텍스트가 포함되어야 합니다.

```js
assert(document.querySelectorAll('HEAD > TITLE')[0]?.innerText == 'The Odin Recipes');
```

`h1` 요소 안에 `body` 요소가 있어야 하며, 그 안에 `Creamy Chocolate Fudge` 텍스트가 포함되어야 합니다.

```js
assert(document.querySelectorAll('BODY > H1')[0]?.innerText == 'Creamy Chocolate Fudge');
```

`alt` 속성이 있는 이미지가 있어야 합니다.

```js
const img = document.querySelectorAll('IMG')[0];

assert(img && img.alt !='' && img.src != '')
```

`h2` 요소가 있어야 하며, 그 안에 `Description` 텍스트가 포함되어야 합니다.

```js
const h2 = document.querySelectorAll('H2')[0];

assert(h2?.innerText == 'Description');
```

레시피를 설명하는 `p` 요소가 최소 두 개 있어야 합니다.

```js
const paragraphs = document.querySelectorAll('P');

assert(paragraphs?.length > 1);
```

`h2` 요소가 있어야 하며, 그 안에 `Ingredients` 텍스트가 포함되어야 합니다.

```js
const h2 = document.querySelectorAll('H2')[1];

assert(h2?.innerText == 'Ingredients');
```

재료를 포함하는 `ul` 요소가 중첩된 `li` 요소가 있어야 합니다.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "Ingredients") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);
        } else {
            assert(false);
        }
    }
});
```

`h2` 요소가 있어야 하며, 그 안에 `Steps` 텍스트가 포함되어야 합니다.

```js
const h2 = document.querySelectorAll('H2')[2];

assert(h2?.innerText == 'Steps');
```

레시피 단계가 들어 있는 `ol` 요소가 포함된 `li` 요소가 있어야 합니다.

```js
const orderedList = document.querySelectorAll('OL')[0];
const listItems = document.querySelectorAll('OL > LI');

assert(orderedList && listItems && listItems.length > 1);
```

`h2` 요소가 있어야 하며, 그 안에 `More Recipes` 텍스트가 포함되어야 합니다.

```js
const h2 = document.querySelectorAll('H2')[3];

assert(h2?.innerText == 'More Recipes');
```

다른 레시피로 링크하는 `ul` 요소가 포함된 `li` 요소가 중첩된 `a` 요소가 있어야 합니다.

```js
const headers = document.querySelectorAll("h2");

headers.forEach(header => {
    if (header.textContent.trim() === "More Recipes") {
        const next = header.nextElementSibling;

        if(next.tagName === "UL") {
            const listItems = next.querySelectorAll("LI");

            assert(listItems.length > 1);

            listItems.forEach(listItem => {
                const anchor = listItem.querySelector("A");

                assert(anchor);
            });
        } else {
            assert(false);
        }
    }
});
```

레시피로 링크하는 앵커 요소들은 `href` 속성을 가지고 있어야 하며, 그 값은 `#`로 설정되어야 합니다.

```js
const anchorTags = document.querySelectorAll("a");

const allAnchorsHaveHrefHash = [...anchorTags].every(function(anchorTag) {
  return anchorTag.hasAttribute("href") && anchorTag.getAttribute("href") === "#";
});

assert(allAnchorsHaveHrefHash && anchorTags.length > 0); 
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>The Odin Recipes</title>
  </head>
  <body>
    <h1>Creamy Chocolate Fudge</h1>
    <img src="https://cdn.freecodecamp.org/curriculum/odin-project/build-a-recipe-page/build-a-recipe-page-01.jpg" alt="A delicious chocolate fudge dessert">
    <h2>Description</h2>
    <p>This recipe is for a rich and creamy chocolate fudge that is sure to satisfy your sweet tooth. It's perfect for a special occasion or as a tasty treat for any time of the year.</p>
    <p>This recipe is easy to follow and only requires a few simple ingredients. With just a few steps, you'll be able to create a delicious dessert that everyone will love.</p>
    <h2>Ingredients</h2>
    <ul>
      <li>1 cup sugar</li>
      <li>1/2 cup unsalted butter</li>
      <li>1/4 cup milk</li>
      <li>1/4 cup cocoa powder</li>
      <li>1/4 cup chocolate chips</li>
      <li>1/4 tsp salt</li>
      <li>1 tsp vanilla extract</li>
    </ul>
    <h2>Steps</h2>
    <ol>
      <li>In a medium saucepan, melt the butter over medium heat.</li>
      <li>Add the sugar, milk, cocoa powder, and salt to the saucepan and stir until well combined.</li>
      <li>Bring the mixture to a boil, stirring constantly, and then reduce the heat to low and simmer for 5 minutes.</li>
      <li>Remove the saucepan from the heat and stir in the chocolate chips and vanilla extract until the chocolate is melted and the mixture is smooth.</li>
      <li>Pour the fudge into a greased 8-inch square pan and let it cool completely before cutting into squares.</li>
    </ol>
    <h2>More Recipes</h2>
    <ul>
      <li><a href="#">Peanut Butter Cookies</a></li>
      <li><a href="#">Lemon Bars</a></li>
      <li><a href="#">Chocolate Chip Pancakes</a></li>
    </ul>
  </body>
</html>
```

```css

```
