---
id: 680b3ef395479b0e449ecb6e
title: Tekeleza algoriti ya selection sort
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

Selection sort ni algoriti maarufu ya kupanga ambayo hufundishwa katika kozi nyingi za sayansi ya kompyuta.

Algoriti hii hufanya kazi kwa kutafuta mara kwa mara kipengele kidogo zaidi kutoka sehemu isiyopangwa ya orodha na kukibadilisha na kipengele cha kwanza kisichopangwa. Huanza kwa kuchagua thamani ya kiwango cha chini katika orodha nzima na kuibadilisha na kipengele cha kwanza. Kisha huenda kwenye nafasi ya pili, kutafuta thamani ndogo zaidi katika vipengele vilivyobaki visivyopangwa, na kuibadilisha na kipengele cha pili. Mchakato huu unaendelea, ukipitia orodha kipengele kimoja kwa wakati, hadi orodha nzima ipangwe.

Selection sort husababisha ugumu wa wakati wa mduara wa mraba (quadratic) katika hali bora, wastani, na mbaya kabisa. Ugumu wa nafasi utakuwa wa kawaida `O(1)` kwa sababu upangaji unafanywa mahali hapo na kiasi cha kumbukumbu kinachotumika ni cha kawaida bila kujali ukubwa wa orodha.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua kitendakazi kinachoitwa `selection_sort`.
1. Kitendakazi chako cha `selection_sort` kinapaswa kuwa na kigezo kimoja kinachowakilisha orodha ya vitu.
1. Kitendakazi chako cha `selection_sort` kinapaswa kuchukua orodha na kupanga vitu mahali hapo kutoka kidogo hadi kikubwa.
1. Kitendakazi chako cha `selection_sort` kinapaswa kubadilisha orodha ya ingizo mahali hapo, na kuirudisha mara itakapopangwa.
1. Kitendakazi chako cha `selection_sort` kinapaswa kufuata algoriti ya selection sort, kubadilisha kipengele kidogo zaidi kutoka sehemu isiyopangwa ya orodha na kipengele cha kwanza kisichopangwa.
1. Kitendakazi chako cha `selection_sort` hakipaswi kufanya mabadilishano yasiyo ya lazima wakati kipengele kidogo tayari kiko katika nafasi sahihi.
1. Kitendakazi chako cha `selection_sort` hakipaswi kutumia njia iliyo jengwa ndani ya `sort()` wala kitendakazi cha `sorted()`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `selection_sort`.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

Kitendakazi chako cha `selection_sort` kinapaswa kuwa na kigezo kimoja.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

Hupaswi kuleta moduli yoyote au kutumia njia za kupanga zilizojengwa ndani katika msimbo wako.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

`selection_sort` inapaswa kurudisha orodha ile ile kama orodha ya ingizo.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

`selection_sort` inapaswa kubadilisha orodha ya ingizo mahali hapo. Hupaswi kutumia njia yoyote ya kuongeza au kuondoa vitu kutoka orodha.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

Kitendakazi chako cha `selection_sort` kinapaswa kufuata algoriti ya selection sort, kubadilisha thamani ya kiwango cha chini katika sehemu isiyopangwa ya orodha na kipengele cha kwanza kisichopangwa. Epuka mabadilishano yasiyo ya lazima wakati thamani ya kiwango cha chini tayari iko katika nafasi sahihi.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` inapaswa kurudisha `[1, 2, 33, 67, 89, 245]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` inapaswa kurudisha `[3, 5, 12, 15, 16, 23, 72, 99, 567]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` inapaswa kurudisha `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]`.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

Kitendakazi chako cha `selection_sort` kinapaswa kupanga kwa usahihi orodha yoyote ya nambari.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
