---
id: a6b0bb188d873cb2c8729495
title: 实现一个超文本标记语言实体转换器
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

本实验是关于将字串中的特殊字符转换为对应的超文本标记语言实体。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `convertHTML` 函数，它接受一个字串作为参数。
1. `convertHTML` 函数应通过将参数字串中的特殊字符转换为其对应的超文本标记语言实体来返回一个新建字串。

   - `&` 应转换为 `&amp;`。
   - `<` 应转换为 `&lt;`。
   - `>` 应转换为 `&gt;`。
   - `&quot;` 应转换为 `&quot;`。
   - `'` 应转换为 `&apos;`。

# --hints--

你应该有一个 `convertHTML` 函数。

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` 应该返回字串 `Dolce &amp; Gabbana`。

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` 应该返回字串 `Hamburgers &lt; Pizza &lt; Tacos`。

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` 应该返回字串 `Sixty &gt; twelve`。

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` 应该返回字串 `Stuff in &quot;quotation marks&quot;`。

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` 应该返回字串 `Schindler&apos;s List`。

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` 应该返回字串 `&lt;&gt;`。

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` 应该返回字符串 `abc`。

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
