---
id: 6551f6f9ec6db270eb83f3db
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Unaweza kuona `11` ikichapishwa kwenye terminal kwa sababu `'Hello World'` ina herufi 11.

Kitendakazi kingine muhimu kilicho ndani ni `type()`, ambacho hurudisha aina ya data ya kigezo. Badilisha wito wako wa `print()` ili kuchapisha aina ya data ya `text`.

# --hints--

Unapaswa kuita kitendakazi cha `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Unapaswa kupitisha `text` kwa kitendakazi cha `type()` kwa kuiweka kati ya mabano ya kawaida.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Unapaswa kupitisha `type(text)` kwa kitendakazi cha `print()` kwa kuiweka kati ya mabano ya kawaida. Zingatia kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Hupaswi kuwa na `print(len(text))` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
