---
id: 6551ecf597038c36637f3759
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Vigezo ni sehemu muhimu ya Python na lugha yoyote ya programu. Kigezo ni jina linalorejea au kuashiria kitu. Unaweza kutangaza kigezo kwa kuandika jina la kigezo upande wa kushoto wa opereta ya mgawo `=` na kuelezwa thamani ya kuweka kwa kigezo hicho upande wa kulia wa opereta ya mgawo:

```py
variable_name = value
```

Tengeneza kigezo kinachoitwa `number` na weka thamani `5` kwa kigezo chako kipya.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `number`. Zingatia kuweka jina la kigezo mwanzoni mwa mstari.

```js
assert.match(code, /^number\s*=/m)
```

Unapaswa kuweka thamani `5` kwa kigezo chako cha `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mstari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
