---
id: 6551f5756c20146fc94f8675
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Unaweza kupata idadi ya herufi katika mfuatano wa herufi kwa kutumia kitendakazi kilicho ndani ya `len()`.

Badilisha wito wako wa `print()` kwa kupitisha `len(text)` badala ya `text[-1]`.

# --hints--


Unapaswa kuita kitendakazi cha `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Unapaswa kupitisha kigezo `text` kwa kitendakazi cha `len()` kwa kuingiza kati ya mabano ya kawaida.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Unapaswa kupitisha `len(text)` kwa kitendakazi cha `print()` kwa kuingiza kati ya mabano ya kawaida. Zingatia kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Hupaswi kuwa na `print(text[-1])` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Msimbo wako una sintaksia isiyo sahihi na/au ulinganifu usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
