---
id: 6551f3bfed58796ebb1268dc
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kila herufi ya mfuatano wa herufi inaweza kurejeewa kwa kutumia kielezo cha nambari. Hesabu ya kielezo huanza kwa sifuri. Hivyo herufi ya kwanza ya mfuatano wa herufi ina kielezo cha `0`. Kwa mfano, katika mfuatano wa herufi `'Hello World'`, `'H'` iko kwenye kielezo cha `0`, `'e'` iko kwenye kielezo cha `1`, na kadhalika.

Kila herufi ya mfuatano wa herufi inaweza kupatikana kwa kutumia uandishi wa mabano. Unahitaji kuandika jina la mabadiliko ukifuatiwa na mabano ya mraba na kuongeza kielezo cha herufi ndani ya mabano hayo:

```py
text = 'Hello World'
r = text[8]
```

Sasa, badala ya kuchapisha `text`, chapisha tu herufi iliyoko kwenye kielezo cha `6`.

# --hints--

Unapaswa bado kuita kitendakazi cha `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Unapaswa kupitisha `text[6]` kwa kitendakazi cha `print()` kwa kuiweka ndani ya mabano ya kawaida. Tazama kwa makini kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uandikaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
