---
id: 6551f1e8737dbe6e0cb08ea4
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Sasa chapisha mabadiliko yako mpya.

# --hints--

Haupaswi kuondoa wito wako wa `print(type(text))` uliopo.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Unapaswa kuita kitendakazi cha `print()` ukipitisha kigezo cha `shift`. Zingatia kuweka wito la kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
