---
id: 6551f22653e6ce6e4f1d1370
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Badilisha wito lako la `print(shift)` ili uchapishe aina ya data ya mabadiliko yako ya `shift`.

# --hints--

Unapaswa kuendelea kutumia wito wako wa `print(type(text))` uliopo.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Unapaswa kupitisha `shift` kwa kitendakazi cha `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Unapaswa kupitisha `type(shift)` kwa kitendakazi chako cha `print()`. Zingatia kuweka wito la kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Hupaswi kuwa na `print(shift)` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Msimbo wako una sintaksia isiyo sahihi na/au uandikaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
