---
id: 6551f4259d66026ee28cb205
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Unaweza pia kupata herufi za mfuatano wa herufi kuanzia mwisho wa mfuatano huo. Herufi ya mwisho ina kielezo cha `-1`, ya pili kutoka mwisho `-2` na kadhalika.

Sasa badilisha wito wako wa `print()` uliopo ili kuchapisha herufi ya mwisho katika mfuatano wako wa herufi.

# --hints--

Unapaswa bado kuita kitendakazi cha `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Unapaswa kupitisha `text[-1]` kwa kitendakazi cha `print()` kwa kuiweka kati ya mabano ya kawaida. Zingatia kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
