---
id: 6551fb531ce7f773071a0b84
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Kama unavyoona, matokeo ya kuchapisha `type(text)` ni `<class 'str'>`, ambayo ina maana kwamba kigezo chako ni mfuatano wa herufi, kinachoonyeshwa kama `str`.

Sasa nenda kwenye mstari mpya na tengeneza kigezo kingine kinachoitwa `shift` na weka thamani `3` kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `shift`. Zingatia kuweka jina la kigezo mwanzoni mwa mstari.

```js
assert.match(code, /^shift\s*=/m)
```

Unapaswa kuweka thamani `3` kwa kigezo chako cha `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
