---
id: 6551ef504d91936d2d4e54f8
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Hoja ni kitu au usemi wa kihesabu unaopitishwa kwa kitendakazi — huongezwa kati ya mabano ya kawaida ya kufungua na kufunga — wakati kitendakazi kinapoitwa:

```py
greet = 'Hello!'
print(greet)
```

Msimbo katika mfano hapo juu ungetachapisha mfuatano wa herufi `'Hello!'`, ambao ni thamani ya kigezo `greet` kilichopitishwa kwa `print()` kama hoja.

Chapisha kigezo chako cha `text` kwenye skrini kwa kupitisha kigezo cha `text` kama hoja kwa kitendakazi cha `print()`.

# --hints--

Unapaswa kupitisha `text` kwa kitendakazi cha `print()` kwa kuongeza jina la kigezo hiki ndani ya mabano ya kawaida ya kufungua na kufunga. Zingatia kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
