---
id: 65520e6f2b9678799977f24d
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Ondoa wito la mwisho la `print()`. Kisha, badala ya `text[0]`, pitisha `text[0].lower()` kama hoja kwa wito lako la `.find()` na uone matokeo.

# --hints--

Haupaswi kuwa na `print(text.lower())` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Bado unapaswa kuchapisha tofauti yako ya `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Unapaswa kusahihisha wito lako la `alphabet.find(text[0])` ili litumie `text[0].lower()` kama hoja. Zingatia kuweka wito wa njia mwanzoni mwa mstari.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
