---
id: 655251308f31958d06cdf267
title: Hatua 36
challengeType: 20
dashedName: step-36
---

# --description--

Kwa sasa, herufi iliyofichwa inasasishwa kila mara katika kila mzunguko. Itakuwa bora kuhifadhi mfuatano wa herufi uliofichwa katika kigezo kipya. Kabla ya mzunguko wako wa `for`, tangaza kigezo kinachoitwa `encrypted_text` na weka thamani ya mfuatano wa herufi tupu (`''`) kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `encrypted_text` kabla ya mzunguko wako wa `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi tupu kwa kigezo chako cha `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
