---
id: 655261b2e1f2c197093f3993
title: Hatua 41
challengeType: 20
dashedName: step-41
---

# --description--

Kwa sasa, nafasi huandikwa kama `'c'`. Ili kudumisha nafasi asilia katika ujumbe wa kawaida, utahitaji usemi wa `if` wa masharti. Huu unaundwa na neno kuu la `if`, sharti, na nukta mbili `:`.

```py
if x != 0:
    print(x)
```

Katika mfano ulio juu, sharti la usemi wa `if` ni `x != 0`. Msimbo wa `print(x)`, ndani ya mwili wa usemi wa `if`, hufanyika tu wakati sharti linapothibitishwa kuwa `True` (katika mfano huu, maana yake `x` ni tofauti na sifuri).

Juu kabisa ya mzunguko wako wa for, badilisha `print(char == ' ')` na usemi wa `if`. Sharti la usemi huu wa `if` linapaswa kuthibitishwa kuwa `True` ikiwa `char` ni nafasi tupu na `False` vinginevyo. Ndani ya mwili wa `if`, chapisha mfuatano wa herufi `'space!'`. Kumbuka kuingiza mstari huu kwa nafasi sahihi.

# --hints--

Haupaswi kuwa na `print(char == ' ')` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Unapaswa kubadilisha `print(char == ' ')` na usemi wa `if` unaochochea wakati `char == ' '`. Usitumie mabano ya kawaida kufunika sharti la `if` na kumbuka kujumuisha nukta mbili ya mwisho.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Unapaswa kuchapisha mfuatano wa herufi `'space!'` ndani ya usemi wako mpya wa `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Msimbo wako una sintaksia isiyo sahihi na/au kuingizwa kwa mistari isiyo sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
