---
id: 6552008c0d9d9075cbec9772
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Utaitumia njia ya `.find()` kupata nafasi katika alfabeti ya kila herufi katika ujumbe wako. Njia ni kama kitendakazi, lakini ni mali ya kitu.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Juu, njia ya `.find()` inafanywa *kwenye* `sentence` (mfuatano wa herufi wa kutafutwa), na `'r'` (herufi ya kutafuta) hupitishwa kama hoja. Wito wa `sentence.find('r')` utarudisha `4`, ambayo ni kielezo cha tukio la kwanza la `'r'` ndani ya `sentence`.

Mwisho wa msimbo wako, fanya wito wa `.find()` kwenye `alphabet` na pitisha `'z'` kama hoja kwa njia hiyo.

# --hints--

Unapaswa kufanya wito wa njia ya `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Unapaswa kufanya wito wa njia ya `.find()` kwenye kigezo cha `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Unapaswa kufanya wito wa njia ya `.find()` kwenye kigezo cha `alphabet` na kupitisha `'z'` kwa njia hiyo. Zingatia kuweka wito wa njia mwanzoni mwa mstari.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
