---
id: 65522255d5b9cd80f335c6f2
title: Hatua 29
challengeType: 20
dashedName: step-29
---

# --description--

Mwisho wa mwili wa mzunguko wako, tangaza kigezo kinachoitwa `new_index` na weka thamani ya `index + shift` kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `new_index` ndani ya mzunguko wako wa `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Unapaswa kuweka `index + shift` kwa kigezo chako kipya mwishoni mwa mwili wa mzunguko wako wa `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
