---
id: 65521fc818947e800bffe48a
title: Hatua 26
challengeType: 20
dashedName: step-26
---

# --description--

Ndani ya mzunguko wa `for`, kabla ya kuchapisha herufi ya sasa, tangaza kigezo kinachoitwa `index` na weka thamani inayorejeshwa na `alphabet.find(char)` katika kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kipya kinachoitwa `index` mwanzoni mwa mzunguko wako wa `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Unapaswa kuweka `alphabet.find(char)` katika kigezo chako kipya cha `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
