---
id: 6552487e689f6e895f658717
title: Hatua 33
challengeType: 20
dashedName: step-33
---

# --description--

Sasa unahitaji kuunda kigezo cha `new_char` mwishoni mwa mwili wa mzunguko wako. Weka thamani yake kuwa `alphabet[new_index]`.

# --hints--

Unapaswa kuunda kigezo cha `new_char` ndani ya mzunguko wa `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Unapaswa kuweka kigezo chako cha `new_char` kuwa `alphabet[new_index]` mwishoni mwa mwili wa mzunguko wako.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
