---
id: 655208d59b131e7816f18c96
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Kitendakazi cha `print()` kinakupa matokeo tu kwenye konsoli, lakini vitendakazi na njia zinaweza kuwa na thamani ya return ambayo unaweza kutumia katika msimbo wako.

Sasa weka thamani ya `alphabet.find(text[0])` katika kigezo kinachoitwa `index`. Kwa njia hii, `index` itahifadhi thamani iliyorejeshwa na `alphabet.find(text[0])`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `index`. Tazama kwa makini kuweka jina la kigezo mwanzoni mwa mstari.

```js
assert.match(code, /^index\s*=/m)
```

Unapaswa kuweka thamani ya return inayotolewa na `alphabet.find(text[0])` katika kigezo chako cha `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
