---
id: 65525e359ca28d938baa82c5
title: Hatua 39
challengeType: 20
dashedName: step-39
---

# --description--

Unaweza kupata athari sawa na `a = a + b` kwa kutumia opereta ya kuongeza:

```py
a += b
```

Opereta ya kuongeza inakuwezesha kuongeza thamani kwa kigezo kisha ukaweka mgawo wa thamani kwa kigezo hicho.

Tumia opereta ya `+=` kuongeza thamani na kuiweka mgawo wa thamani kwa wakati mmoja kwa `encrypted_text`.

# --hints--

Unapaswa kutumia opereta ya kuongeza kuongeza `alphabet[new_index]` kwa thamani ya sasa ya `encrypted_text` ndani ya mwili wa mzunguko wako.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
