---
id: 65520f8282faf57a0db4f7fe
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Kama unavyoona kutoka kwenye matokeo, `'h'` iko kwenye kielezo `7` katika mfuatano wa herufi `alphabet`. Sasa unahitaji kupata herufi iliyoko kwenye kielezo `7` pamoja na thamani ya `shift`. Kwa hilo, unaweza kutumia kiendeshaji cha kuongeza, `+`, kwa njia ile ile unavyotumia kwa kuongeza kihesabu.

Badilisha tofauti yako ya `shifted` ili ihifadhi thamani ya `alphabet` kwenye kielezo `index + shift`.

# --hints--

Unapaswa kuweka thamani ya `alphabet[index + shift]` kwenye tofauti yako ya `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au mkao usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
