---
id: 65524d2a1a253b8bb5197ae2
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` inarudisha tena `-1` kwa herufi kubwa, na pia kwa nafasi ya uandishi. Utashughulikia nafasi hiyo baadaye.

Kwa sasa, badala ya kurudia juu ya `text`, badilisha mzunguko wa `for` urudie juu ya `text.lower()`.

# --hints--

Unapaswa kubadilisha mzunguko wako wa `for` urudie juu ya `text.lower()` badala ya `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
