---
id: 65521203d7165c7b84b22ad4
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Mzunguko hukuruhusu kupita kwa mpangilio mfululizo wa vipengele na kutekeleza vitendo kwa kila kimoja.

Katika kesi hii, utatumia mzunguko wa `for`. Hapa ni jinsi unavyoweza kurudia juu ya `text`:

```py
for i in text:
```

`for` ni neno kuu linaloashiria aina ya mzunguko. `i` ni kigezo kinachochukua thamani ya vipengele katika `text` kwa mfululizo. Sentensi inamalizika kwa nukta mbili, `:`.

Chini ya mstari ambapo ulitangaza `alphabet`, andika mzunguko wa `for` kurudia juu ya `text`. Tumia `i` kama kigezo cha mzunguko.

Kufanya hivyo, kuna kosa kwenye terminali. Utajifunza kuhusu hilo katika hatua inayofuata.

# --hints--

Unapaswa kutumia neno kuu la `for` kuunda mzunguko. Hakikisha kuweka neno kuu la `for` mwanzoni mwa mstari na acha nafasi ya wazi baada ya neno kuu.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Unapaswa kuandika kigezo cha `i` baada ya neno kuu la `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Unapaswa kuandika neno kuu la `in` baada ya `for i `. Hakikisha unaacha nafasi karibu na neno kuu la `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Unapaswa kuandika `text` baada ya `for i in `. Usisahau kuongeza `:` ya mwisho.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Mzunguko wako wa `for` unapaswa kuwekwa chini ya mstari wa msimbo `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
