---
id: 655209a4a27dd37873c4cac3
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` hurejesha kielezo cha herufi inayolingana ndani ya mfuatano wa herufi. Ikiwa herufi haipatikani, hurejesha `-1`. Kama unavyoona, herufi ya kwanza katika `text`, herufi kubwa `'H'`, haipatikani, kwa kuwa `alphabet` ina herufi ndogo tu.

Unaweza kubadilisha mfuatano wa herufi kuwa sawa na herufi ndogo kwa kutumia njia ya `.lower()`. Ongeza wito mwingine wa `print()` kuchapisha `text.lower()` na kuona matokeo.

# --hints--

Unapaswa bado kuwa na `print(index)` katika msimbo wako. Zingatia kuweka wito la kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Unapaswa kuchapisha `text.lower()`. Zingatia kuweka wito la kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
