---
id: 655220953ba90d80514d7ee2
title: Hatua 27
challengeType: 20
dashedName: step-27
---

# --description--

Kwa sasa, kitendakazi cha `print()` kinapokea hoja moja `char`, lakini kinaweza kupokea hoja nyingi, zikitenganishwa kwa koma.

Ongeza hoja ya pili kwa `print(char)` ili ichapishe herufi na kielezo chake ndani ya alfabeti.

# --hints--

Unapaswa kuongeza `index` kama hoja ya pili kwa wito lako la `print(char)` lililopo. Usisahau kutenganisha hoja kwa koma.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
