---
id: 6553a572f7a65718f1e42e18
title: Hatua 49
challengeType: 20
dashedName: step-49
---

# --description--

Kitendakazi ni sehemu ya msimbo inayoweza kutumika tena. Tayari umejifunza baadhi ya vitendakazi vilivyojengwa ndani, kama `print()`, `find()` na `len()`. Lakini pia unaweza kutangaza vitendakazi maalum kama hivi:

```py
def function_name():
    <code>
```

Tamko la kitendakazi huanza na neno kuu la `def` linalofuatiwa na jina la kitendakazi — jina halali la kigezo — na mabano ya kawaida. Tamko linaisha kwa alama ya nukta mbili.

Mara moja baada ya kigezo chako cha `shift`, tangaza kitendakazi kinachoitwa `caesar` na punguza mistari yote inayofuata ili kumpa kitendakazi chako mwili.

# --hints--

Unapaswa kutumia neno kuu la `def` kutangaza kitendakazi kipya.

```js
assert.match(code, /^def\s+/m)
```

Unapaswa kuandika `caesar` kama jina la kitendakazi baada ya neno kuu la `def`. Kumbuka kuongeza nafasi baada ya `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Unapaswa kuongeza mabano ya kawaida baada ya jina la kitendakazi. Usisahau nukta mbili ya mwisho.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Unapaswa kupunguza mistari yote baada ya `shift = 3` ili iwe mwili wa kitendakazi chako kipya.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
