---
id: 6553ed69ece88d29594748aa
title: Hatua 52
challengeType: 20
dashedName: step-52
---

# --description--

Ili kutekeleza, kitendakazi kinahitaji kuitwa (au kuamshwa) kwa kuongeza jozi ya mabano ya kawaida baada ya jina lake, kama hii:

```py
function_name()
```

Mwisho wa msimbo wako, ita wito wa kitendakazi chako cha `caesar`. Zingatia ulinganifu wa mistari.

# --hints--

Unapaswa kuita kitendakazi chako cha `caesar`. Hakikisha unaandika wito wa kitendakazi mwanzoni mwa mstari.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
