---
id: 6553980e0527fa115c705646
title: Hatua 42
challengeType: 20
dashedName: step-42
---

# --description--

Sasa, badala ya kuchapisha `'space!'`, tumia opereta ya kuongeza kuongeza nafasi (ambayo kwa sasa imehifadhiwa katika `char`) kwenye thamani ya sasa ya `encrypted_text`.

# --hints--

Haupaswi kuwa na `print('space!')` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Unapaswa kutumia opereta ya `+=` kuongeza `char` kwenye thamani ya sasa ya `encrypted_text` ndani ya tamko lako jipya la `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
