---
id: 6553efd6ada3f42aa2d75448
title: Hatua 53
challengeType: 20
dashedName: step-53
---

# --description--

Sasa unapaswa kuona matokeo tena. Ingawa njia hii inafanya kazi, haiboresha sana matumizi ya msimbo tena na tena. Kufanya wito kwa kitendakazi chako mara kwa mara kutatoa matokeo sawa. Hata hivyo, kitendakazi kinaweza kutangazwa na *vigezo* ili kuleta utofauti na kubadilika:

```py
def function_name(param_1, param_2):
    <code>
```

Vigezo ni vigezo vya kubadilika ambavyo unaweza kutumia ndani ya kitendakazi chako. Kitendakazi kinaweza kutangazwa na idadi tofauti ya vigezo. Katika mfano uliotangulia, `param_1` na `param_2` ni vigezo.

Badilisha tamko la kitendakazi chako ili kipokee vigezo viwili vinavyoitwa `message` na `offset`.

Baada ya hapo, utaona kosa linalojitokeza kwenye terminal. Utaona jinsi ya kulitatua katika hatua zinazofuata.

# --hints--

Kitendakazi chako cha `caesar` kinapaswa kupokea `message` na `offset` kama vigezo. Kumbuka kutenganisha vigezo kwa koma.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
