---
id: 6553f3fc92741c2bf8ded140
title: Hatua 54
challengeType: 20
dashedName: step-54
---

# --description--

Ndani ya mwili wa kitendakazi chako, badilisha majina ya vigezo `text` na `shift` kuwa `message` na `offset`, mtawalia.

# --hints--

Unapaswa kubadilisha majina yote ya `text` kuwa `message`.

```js
assert.match(code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*:/);
assert.match(code, /print\s*\(\s*("|')(plain text|(plain )?message):\1\s*,\s*message\s*\)/);
```

Unapaswa kubadilisha jina la `shift` kuwa `offset`.

```js
assert.match(code, /new_index\s*=\s*\(\s*index\s*\+\s*offset\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/);
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
--fcc-editable-region--
```
