---
id: 6553f4f66099802c6ae94613
title: Hatua 55
challengeType: 20
dashedName: step-55
---

# --description--

Kwa sasa, msimbo wako unaleta `TypeError`, kwa sababu kitendakazi cha `caesar` kimefafanuliwa na vigezo viwili (`message` na `offset`), kwa hiyo kinatarajia kuitwa na hoja mbili.

Kuitwa kwa `caesar()` bila hoja zinazohitajika kunasimamisha utekelezaji wa msimbo.

Toa thamani za `message` na `offset` kwa kupitisha `text` na `shift` kama hoja katika wito wa kitendakazi cha `caesar`.

# --hints--

Unapaswa kupitisha `text` na `shift` kama hoja katika wito wa kitendakazi chako kwa kuziweka ndani ya mabano ya kawaida. Usisahau kutenganisha hoja kwa koma.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
