---
id: 6553a755879b131a445e664c
title: Hatua 44
challengeType: 20
dashedName: step-44
---

# --description--

Jaribu kuweka thamani ya mfuatano wa herufi `'Hello Zaira'` katika tofauti yako ya `text` na uone kinachotokea kwenye terminali.

Utaona hitilafu ya `string index out of range`. Usijali, utagundua jinsi ya kuirekebisha hivi karibuni!

# --hints--

Unapaswa kuwa na tofauti ya `text`.

```js
assert.match(code, /^text\s*=/m)
```

Unapaswa kuweka thamani ya mfuatano wa herufi `'Hello Zaira'` katika tofauti yako ya `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
