---
id: 6553a7d8d05cbb1ae335a665
title: Hatua 45
challengeType: 20
dashedName: step-45
---

# --description--

Wakati mzunguko unafikia herufi `Z`, jumla ya `index + shift` inazidi kielezo cha mwisho cha mfuatano wa herufi `alphabet`. Kwa hiyo, `alphabet[new_index]` inajaribu kutumia kielezo kisicho halali, jambo ambalo husababisha `IndexError` kutupwa.

Unaweza kuona kuwa matokeo kwenye terminal yanakoma katika nafasi mara moja kabla ya `Z`, `print` ya mwisho kabla ya kosa kutupwa.

Katika kesi hii, operator wa modulo (`%`) unaweza kutumika kurudisha mabaki ya mgawanyo kati ya nambari mbili. Kwa mfano: `5 % 2` ni sawa na `1`, kwa sababu 5 kugawanywa na 2 kuna gawio la 2 na mabaki ya 1.

Zungusha `index + shift` na mabano ya kawaida, kisha tumia modulo kwenye usemi wa kihesabu na `26`, ambayo ni urefu wa `alphabet`.

# --hints--

Unapaswa kuwa na `new_index = (index + shift) % 26` katika tamko lako la `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Msimbo wako una sintaksia isiyo halali na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
