---
id: 6553f6086add4b2cbb99fd78
title: Hatua 56
challengeType: 20
dashedName: step-56
---

# --description--

Mwisho wa msimbo wako, baada ya wito wako wa `caesar(text, shift)` uliopo, fanya wito wa kitendakazi chako tena.

Mara hii, pitia tofauti ya `text` na nambari kamili `13` kama hoja.

# --hints--

Unapaswa kufanya wito wa kitendakazi chako tena, mara hii ukipitia `text` na `13` kama hoja.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Unapaswa kuhifadhi wito wa kitendakazi uliopo.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
