---
id: 65549f90cf78131c96ebcf28
title: Hatua 61
challengeType: 20
dashedName: step-61
---

# --description--

Wakati wa kuandika msimbo, kusomeka ni muhimu sana. Maoni hutumika kama vidokezo vyenye ufanisi, vinavyoelezea mantiki nyuma ya msimbo wako. Hii ni muhimu unaporudi kwenye mradi baada ya muda na pia husaidia wenzako kuelewa msimbo huo.

Katika Python, unaweza kuandika maoni kwa kutumia `#`. Kila kitu kinachofuata `#` hakitatekelezwa.

Kabla ya tamko lako la `if`, ongeza maoni yanayosema `Append space to the message`.

# --hints--

Unapaswa kuongeza maoni yanayosema `Append space to the message`. Usisahau kuandika alama ya `#` mwanzoni.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
