---
id: 655491bd5b98b813fa5bedca
title: Hatua 58
challengeType: 20
dashedName: step-58
---

# --description--

Sasa badilisha tamko la kitendakazi chako: badilisha jina la kitendakazi kuwa `vigenere` na kigezo cha pili kuwa `key`.

# --hints--

Unapaswa kubadilisha jina la kitendakazi chako kuwa `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Kitendakazi chako cha `vigenere` kinapaswa kuchukua `message` na `key` kama vigezo.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
