---
id: 6554d25dc5ceaa354307a77e
title: Hatua 70
challengeType: 20
dashedName: step-70
---

# --description--

Usimbaji na ufumbuzi ni michakato ya kinyume na kitendakazi chako kinaweza kufanya yote mawili kwa marekebisho machache.

Ongeza kigezo cha tatu kinachoitwa `direction` kwenye ufafanuzi wa kitendakazi chako. Pia, toa maoni kwenye mistari miwili ya mwisho ya msimbo ili kuepuka makosa kwenye konsoli.

# --hints--

Unapaswa kubadilisha mistari miwili ya mwisho kwenye msimbo wako kuwa maoni. Weka `#` mwanzoni mwa kila mstari.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

Unapaswa kuongeza `direction` kama kigezo cha tatu cha kitendakazi chako.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
