---
id: 6554930320d70414e7b6acc6
title: Hatua 59
challengeType: 20
dashedName: step-59
---

# --description--

Futa variable yako ya `shift`. Kisha, tangaza variable mpya iitwayo `custom_key` na weka thamani ya mfuatano wa herufi `'python'` kwa variable hii.

# --hints--

Unapaswa kufuta variable ya `shift` na thamani yake.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Unapaswa kutangaza variable iitwayo `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `'python'` kwa variable yako ya `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
