---
id: 6554e0adc7bb193cbfdb36d5
title: Hatua 73
challengeType: 20
dashedName: step-73
---

# --description--

Kagua kama kitendakazi kinaweza kufungua mfuatano wa herufi kurudi kwenye maandishi ya kawaida.

Tangaza kigezo kingine kinachoitwa `decryption` na weka thamani ya `vigenere(encryption, custom_key, -1)`.

# --hints--

Unapaswa kuita `vigenere` ukipitisha `encryption`, `custom_key` na `-1` kama hoja.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

Unapaswa kutangaza kigezo cha `decryption`.

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

Kigezo chako cha `decryption` kinapaswa kuwa na thamani ya `vigenere(encryption, custom_key, -1)`.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
