---
id: 65549561463f0016876e852c
title: Hatua 60
challengeType: 20
dashedName: step-60
---

# --description--

Kwa kuwa ufunguo wako ni mfupi kuliko maandishi utakayohitaji kuyaficha, utahitaji kurudia ili kuzalisha maandishi yote yaliyofichwa. Mwanzoni mwa mwili wa kitendakazi chako, tangaza kigezo cha `key_index` na uweke thamani yake kuwa sifuri.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `key_index` mwanzoni mwa mwili wa kitendakazi chako.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

Unapaswa kuweka thamani ya `0` kwa kigezo chako cha `key_index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
