---
id: 6554a88d5af937226f4a9121
title: Hatua 65
challengeType: 20
dashedName: step-65
---

# --description--

Njia ya `.index()` ni sawa kabisa na njia ya `.find()` lakini hutupa ubaguzi wa `ValueError` ikiwa haiwezi kupata sehemu ndogo ya mfuatano wa herufi.

Ubaguzi wa `ValueError` ni ubaguzi uliyojengwa ndani unaotolewa wakati hoja yenye aina sahihi lakini thamani isiyofaa inapopita kwa kitendakazi.

Baada ya kuongeza `key_index`, tangaza kigezo kinachoitwa `offset`. Tafuta kielezo ambacho `key_char` ina katika `alphabet` na weka thamani yake katika `offset`. Tumia `.index()` kupata kielezo hicho.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

Kigezo chako cha `offset` kinapaswa kuhifadhi thamani ya `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

Msimbo wako una sintaksia isiyo sahihi na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
