---
id: 65551a628bcb7e121e32d04b
title: Hatua 79
challengeType: 20
dashedName: step-79
---

# --description--

Kwa sasa, alama za uandishi, herufi maalum au nambari hazitafsiriwi/kutafsiriwa kwa usahihi.

Kagua hili kwa kuongeza alama ya mshangao mwishoni mwa mfuatano wa herufi wa `text`.

# --hints--

Unapaswa kuwa na kigezo cha `text`.

```js
({ test: () => assert(__userGlobals.has("text")) })
```

Kigezo chako cha `text` kinapaswa kuwa sawa na mfuatano wa herufi `'Hello Zaira!'`.

```js
({ test: () => assert.equal(__userGlobals.get("text"), "Hello Zaira!") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'
--fcc-editable-region--
def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
