---
id: 655619327c7b364166f8dd6f
title: Hatua 46
challengeType: 20
dashedName: step-46
---

# --description--

Ikiwa unataka kuingiza herufi zaidi katika mfuatano wa herufi `alphabet`, kama vile nambari au herufi maalum, utagundua kuwa ni lazima ubadilishe kwa mkono operand ya kulia ya operesheni ya modulo.

Badilisha `26` na `len(alphabet)` ili kuepuka tatizo hili.

# --hints--

Unapaswa kubadilisha thamani ya `new_index` kwa kubadilisha `26` na `len(alphabet)`.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
