---
id: 655a2a7210094920069b117c
title: Hatua 40
challengeType: 20
dashedName: step-40
---

# --description--

Opereta za kulinganisha zinakuwezesha kulinganisha vitu viwili kulingana na thamani zao. Unaweza kutumia opereta ya kulinganisha kwa kuiweka kati ya vitu unavyotaka kulinganisha.
Zinarudisha thamani ya *thamani ya kweli au siyo kweli* — yaani `True` au `False` — kulingana na ukweli wa usemi wa kihesabu.

Python ina opereta zifuatazo za kulinganisha:

<table>
  <thead>
    <tr>
      <th>Opereta</th>
      <th>Maana</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Usawa</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Si sawa</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Kubwa kuliko</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Ndogo kuliko</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Kubwa kuliko au sawa na</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Ndogo kuliko au sawa na</td>
    </tr>
  </tbody>
</table>

Mwanzoni mwa mwili wa mzunguko wako, chapisha matokeo ya kulinganisha `char` na nafasi (`' '`). Tumia opereta ya usawa `==` kwa hilo.

# --hints--

Unapaswa kulinganisha `char` na nafasi kwa kutumia opereta ya usawa ndani ya mzunguko wako wa `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Unapaswa kuchapisha matokeo ya kulinganisha `char` na nafasi ndani ya mzunguko wako wa `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Unapaswa kuchapisha matokeo ya kulinganisha `char` na nafasi mwanzoni mwa mzunguko wako.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
