---
id: 6568b7a40c7ba0ccbc4b4425
title: Hatua 32
challengeType: 20
dashedName: step-32
---

# --description--

Kama unavyoona, kila herufi ya mfuatano wa herufi uliowekwa thamani upya inachapishwa ndani ya mzunguko.

Rudi kwenye mfuatano wa herufi wa awali kwa kufuta mgawo wa thamani wa `text`.

# --hints--

Haupaswi kuwa na `text = 'Albatross'` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

Bado unapaswa kuwa na kigezo cha `text` chenye thamani ya `'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
