---
id: 65687d2f8c7ee27b0446cef3
title: Hatua 24
challengeType: 20
dashedName: step-24
---

# --description--

Msimbo unaotekelezwa kila mara — uliowekwa baada ya `:` — ni mwili wa mzunguko. Msimbo huu lazima uandikwe kwa nafasi. Katika Python, inashauriwa kutumia nafasi 4 kwa kila kiwango cha nafasi. Kiwango hiki cha nafasi ni sehemu ya msimbo.

```py
for i in text:
    print(i)
```

Python hutegemea nafasi kuonyesha sehemu za msimbo. Nukta ya wima mwishoni mwa mstari ni ishara kwamba sehemu mpya ya msimbo itafuata ikiwa imeandikwa kwa nafasi.

Hivyo, wakati hakuna sehemu ya msimbo iliyopangwa kwa nafasi baada ya nukta ya mwisho, utekelezaji wa msimbo unasimama na `IndentationError` hutupwa. Msimbo huu hautaonyesha matokeo na badala yake utatupa `IndentationError`:

```py
for i in text:
print(i)
```

Toa mwili kwa mzunguko wako wa `for` kwa kuongeza wito wa `print(i)`. Kumbuka kuandika mwili wa mzunguko kwa nafasi.

# --hints--

Unapaswa kuongeza `print(i)` kwa mwili wa mzunguko wako wa `for`. Lipa kipaumbele nafasi inayotumika.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au nafasi isiyo sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
