---
id: 656886ec4e4d9b96ded3c565
title: Hatua 25
challengeType: 20
dashedName: step-25
---

# --description--

Kigezo cha kurudia kinaweza kuwa na jina lolote halali, lakini ni rahisi kuliita jina lenye maana.

Badilisha jina la kigezo chako cha `i` kuwa `char`.

# --hints--

Unapaswa kubadilisha `i` na `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Msimbo wako una sintaksia isiyo halali na/au uelekezaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
