---
id: 65688a50e6c998a21d8e41d3
title: Hatua 30
challengeType: 20
dashedName: step-30
---

# --description--

Mfuatano wa herufi hauwezi kubadilika, ambayo inamaanisha hayawezi kubadilishwa mara tu yanapotengenezwa. Kwa mfano, unaweza kufikiri kuwa msimbo ufuatao unabadilisha thamani ya `my_string` kuwa mfuatano wa herufi `'train'`, lakini hii si halali:

```py
my_string = 'brain'
my_string[0] = 't'
```

Thibitisha kwa kutumia uandishi wa mabano kupata herufi ya kwanza katika `text` na jaribu kuibadilisha kuwa herufi unayochagua. Utaona matokeo yamefifia na kosa linaonekana.

# --hints--

Unapaswa bado kuwa na `text = 'Hello World'` katika msimbo wako.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Unapaswa kupata herufi ya kwanza katika `text` kwa kutumia `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Unapaswa kutumia kielezi `=` kuweka thamani ya herufi unayochagua kwa `text[0]`. Usisahau kuifunga herufi hiyo kwa alama za nukuu moja au alama za nukuu maradufu.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
