---
id: 65688b5a1655a7a6caede847
title: Hatua 31
challengeType: 20
dashedName: step-31
---

# --description--

Unapojaribu kubadilisha herufi binafsi za mfuatano wa herufi kama ulivyofanya katika hatua iliyopita, unapata `TypeError`, ambayo hutokea wakati kitu cha aina isiyofaa kinatumika katika msimbo wako.

Kama unavyoona kutoka kwenye ujumbe wa kosa, mfuatano wa herufi hauwezi kupewa mgawo wa kipengee, kwa sababu ni thabiti. Hata hivyo, kigezo kinaweza kupewa mfuatano mwingine wa herufi:

```py
message = 'Hello World'
message = 'Hello there!'
```

Futa mstari wa `text[0]` na weka thamani `text` mfuatano wa herufi `'Albatross'`.

# --hints--

Haupaswi kuwa na `text[0] = 'C'` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Unapaswa kuweka thamani `text` mfuatano wa herufi `'Albatross'`. Usibadilishe `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
