---
id: 65952b70924edb449faf1950
title: Hatua 19
challengeType: 20
dashedName: step-19
---

# --description--

Tangaza kigezo kipya kinachoitwa `shifted`. Tumia uandishi wa mabano kupata thamani ya `alphabet` kwa kielezo `index` na weka thamani hiyo katika kigezo chako kipya.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `shifted`. Zingatia kuweka jina la kigezo mwanzoni mwa mstari.

```js
assert.match(code, /^shifted\s*=/m)
```

Unapaswa kuweka thamani ya `alphabet[index]` katika kigezo chako cha `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uandishi usio sahihi wa nafasi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
