---
id: 659526f3e5424b3eac4b6aa7
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Ifuatayo, chapisha thamani ya mabadiliko ya `index` kwenye konsoli.

# --hints--

Unapaswa kuchapisha mabadiliko yako ya `index`. Zingatia kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
