---
id: 659530ca823d6c47f61daa8a
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Chapisha tofauti yako ya `shifted`.

# --hints--

Unapaswa kuchapisha tofauti yako ya `shifted`. Tazama kwa makini kuweka wito wa kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa mistari usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
