---
id: 65951b637ddec63611c482c0
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Unaweza kutumia kitendakazi kilicho ndani `print()` ili kuchapisha matokeo ya msimbo wako kwenye terminali.

Vitendakazi ni sehemu za msimbo zinazoweza kutumika tena ambazo unaweza kuitisha, au kuanzisha, ili kuendesha msimbo wake unapotakiwa. Ili kuitisha kitendakazi, unahitaji tu kuandika jozi ya mabano ya kawaida kando ya jina lake. Utajifunza zaidi kuhusu vitendakazi hivi karibuni.

Kwa sasa, nenda kwenye mstari mpya na ongeza wito tupu kwa kitendakazi `print()`. Huatakiwi kuona matokeo yoyote bado.

# --hints--

Unapaswa kuwa na `print()` katika msimbo wako. Zingatia kuweka wito la kitendakazi mwanzoni mwa mstari.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Msimbo wako una sintaksia isiyo sahihi na/au uingizaji wa nafasi usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
