---
id: 659db3e3670d3e712be82593
title: Hatua 22
challengeType: 20
dashedName: step-22
---

# --description--

Kurudia mchakato wa kutafuta herufi ndani ya alfabeti na kubaini herufi iliyohamishwa kwa kila tabia katika `text` kunaweza kuwa kazi ya kuchosha. Kwa bahati nzuri, unaweza kuifanya iwe rahisi kwa kutumia mzunguko.

Kwa sasa, futa mistari yote ya msimbo chini ya tamko la kigezo cha `alphabet`.

# --hints--

Unapaswa bado kuwa na `text = 'Hello World'` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Unapaswa bado kuwa na `shift = 3` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Unapaswa bado kuwa na `alphabet = 'abcdefghijklmnopqrstuvwxyz'` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Unapaswa kufuta kigezo cha `index` na thamani yake.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Haupasi kuwa na `print(index)` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Unapaswa kufuta kigezo cha `shifted` na thamani yake.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

Haupasi kuwa na `print(shifted)` katika msimbo wako.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Msimbo wako una sintaksia isiyo sahihi na/au uandikaji usio sahihi.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
