---
id: 6624e0ec0bf47017eed84b1f
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Aina ya kwanza ya usimbaji (cipher) utakayojenga huitwa *Caesar* cipher. Hasa, utachukua kila herufi katika ujumbe wako, utafute nafasi yake katika alfabeti, uchukue herufi iliyopo baada ya nafasi 3 katika alfabeti, na kubadilisha herufi ya awali na herufi mpya.

Ili kutekeleza hili, utatumia njia ya `.find()` iliyojadiliwa katika hatua iliyopita. Badilisha wito wako wa `.find()` uliopo kwa kuipatia `text[0]` kama hoja badala ya `'z'`.

# --hints--

Unapaswa kubadilisha wito wako wa `alphabet.find('z')` uliopo kwa kumpatia `text[0]` njia hiyo.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
