---
id: 66605a64a9822f01ec5ae8ce
title: Hatua 50
challengeType: 20
dashedName: step-50
---

# --description--

Katika Python, *wigo* wa kigezo huamua wapi kigezo hicho kinaweza kufikiliwa:

- Vigezo vilivyoelezwa nje ya kitendakazi vina *wigo wa kimataifa* na vinaweza kufikiliwa kutoka sehemu yoyote ya msimbo wako.

- Vigezo vilivyoelezwa ndani ya kitendakazi ni *vya ndani* kwa kitendakazi hicho na haviwezi kufikiliwa nje yake.

Ili kuona hili likifanyika, jaribu kuchapisha kigezo cha `alphabet` mwishoni mwa msimbo wako. Hii itasababisha hitilafu ya `NameError`.

Unapaswa kuona ujumbe wa kosa unaoonyesha kuwa `alphabet` haijaelezwa. Hii ni kwa sababu `alphabet` imeelezwa ndani ya kitendakazi cha `caesar` na haipatikani nje yake.

# --hints--

Unapaswa kujaribu kuchapisha kigezo cha `alphabet` nje ya kitendakazi cha caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
