---
id: 5dc17dc8f86c76b9248c6eb5
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

コメントを使うと、ブラウザーの表示に影響を及ぼさずにメッセージを残すことができます。 また、コメントを使ってコードを非アクティブにすることも可能です。 HTML のコメントは `<!--` で始まり、テキストの行を囲み、`-->` で終了します。

例えば `TODO: Remove h1` というコメントは次のようになります:

```html
<!-- TODO: Remove h1 -->
```

Add a comment above the `p` element with this text:

`TODO: Add link to cat photos`

# --hints--

コメントは `<!--` で始める必要があります。 コメントの開始を定義する文字が 1 つ以上欠けています。

```js
assert.match(code, /<!--/);
```

コメントは `-->` で終わる必要があります。  コメントの終了を定義する文字が 1 つ以上欠けています。

```js
assert.match(code, /-->/);
```

コード内に余分なコメントを開始 / 終了する文字があってはいけません。 余分な `<!--` または `-->` がブラウザーに表示されています。

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

コメントが `p` 要素の上にあるようにしてください。 順番に誤りがあるようです。

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

コメントには `TODO: Add link to cat photos` というテキストを含めてください。

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
