---
id: 5dc17d3bf86c76b9248c6eb4
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 要素を `h2` 要素の下に作成し、それに次のテキストを設定してください:

`Everyone loves cute cats online!`

# --hints--

Your `p` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

`p` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
assert.match(code, /<\/p\>/);
```

`p` 要素のテキストは `Everyone loves cute cats online!` でなければなりません。テキストに誤字脱字があります。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Your `p` element should be below the `h2` element. 順番に誤りがあるようです。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
