---
id: 5dc174fcf86c76b9248c6eb2
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

The first element you will learn about is the `h1` element. The `h1` element is a heading element and is used for the main heading of a webpage.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

The text `CatPhotoApp` should be present in the code. You may want to check your spelling.

```js
assert.match(code, /catphotoapp/i);
```

Your `h1` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

`h1` 要素のテキストは `CatPhotoApp` でなければなりません。 テキストに誤字脱字があるか、テキストが `h1` 要素の開始タグと終了タグの間に入っていない可能性があります。

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

ウェブページを変更するブラウザ拡張機能が使用されているようです。 拡張機能をオフにしてください。

```js
if(__checkForBrowserExtensions){
  assert.isAtMost(document.querySelectorAll('script').length, 2);
  assert.equal(document.querySelectorAll('style').length, 1);
  assert.equal(document.querySelectorAll('link').length, 0);
}
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
