---
id: 5dc1798ff86c76b9248c6eb3
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 要素の下に、次のテキストを持つ `h2` 要素を追加してください:

`Cat Photos`

# --hints--

`h1` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素には終了タグが必要です。 Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

`h1` 要素が 1 つだけ必要です。 余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`h1` 要素のテキストは 'CatPhotoApp' でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```

`h2` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('h2'));
```

`h2` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
assert.match(code, /<\/h2\>/);
```

Your `h2` element's text should be `Cat Photos`. `Cat Photos` というテキストだけを、`h2` の開始タグと終了タグの間に置いてください。

```js
assert.equal(document.querySelector('h2')?.innerText.toLowerCase(), 'cat photos');
```

`h2` 要素は `h1` 要素の下に置く必要があります。 `h1` 要素の方が重要度が高いので、`h2` 要素より上になければなりません。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
