---
id: 5dc23f9bf86c76b9248c6eba
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`img` 要素を使うと、ウェブサイトに画像を追加することができます。 `img` 要素は開始タグを持ちますが、終了タグを持ちません。 終了タグを持たない要素は<dfn>空要素</dfn>として知られています。

`p` 要素の下に、`img` 要素を追加してください。 この時点では、ブラウザーに画像は表示されません。

# --hints--

`img` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

`img` 要素は終了タグ `</img>` を持たないようにしてください。

```js
assert.notMatch(code, /<\/img\>/);
```

`img` 要素は 1 つだけにしてください。 余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

`img` 要素は `p` 要素の下に置く必要があります。 順番に誤りがあるようです。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```
