---
id: 5ddbd81294d8ddc1510a8e56
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

リンクのテキストは、アンカー (`a`) 要素の開始タグと終了タグの間に配置します。

下記は `click here to go to freeCodeCamp.org` というテキストのリンクの例です。

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

アンカー要素に、`cat photos` というアンカーテキストを追加してください。 これがリンクのテキストになります。

# --hints--

Your anchor (`a`) element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Your anchor (`a`) element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/a\>/);
```

アンカー (`a`) 要素のテキストは `cat photos` でなければなりません。 リンクテキストがアンカー (`a`) 要素の開始タグと終了タグの間に置かれていることを確認してください。

```js
assert.equal(
  document.querySelector('a')?.innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
