---
id: 5dfa37b9eacea3f48c6300b0
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

順序なしリストを作成するには、`ul` 要素を使用します。

テキスト `Things cats love:` を持つ `h3` 要素の後に、順序なしリスト (unordered list、`ul`) 要素を追加してください。 この時点では何も表示されません。

# --hints--

`ul` 要素には開始タグが必要です。 Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('ul'));
```

`ul` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
assert.match(code, /<\/ul>/);
```

`ul` 要素は、2 つ目の `section` 要素の終了タグの上にある必要があります。

```js
const secondSectionLastElemNode =
  document.querySelectorAll('main > section')[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName, 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
