---
id: 5dfa2407b521be39a3de7be1
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

リンクを新しいタブで開くには、アンカー (`a`) 要素に `target` 属性を設定します。

`target` 属性は、リンクされたドキュメントをどこで開くかを指定します。 `target="_blank"` を指定すると、リンクされたドキュメントは新しいタブまたはウィンドウで開かれます。

`target` 属性を持つ `a` 要素の基本的な構文は次の通りです:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

リンクが新しいタブで開くようにするために、アンカー (`a`) 要素の開始タグに `target` 属性を追加し、値 `_blank` を設定してください。

# --hints--

`p` 要素内にはアンカー (`a`) 要素とテキスト `cat photos` がネストされているはずです。 削除されたか、誤字脱字がある可能性があります。

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

アンカー (`a`) 要素に `target` 属性がありません。 Check that there is a space after the opening tag's name and/or there are spaces before all attribute names.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

`target` 属性の値は `_blank` である必要があります。 You have either omitted the value or have a typo. Remember that attribute values should be surrounded with quotation marks.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
