---
id: 5dfa3589eacea3f48c6300ae
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

2 つ目の `section` 要素内に、新しい `h2` 要素を追加して、テキストを `Cat Lists` としてください。

# --hints--

Your `section` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Your `h2` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Your `h2` element should have a closing tag. 終了タグは `<` の直後に `/` があるタグです。

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Your second `h2` element should be right above the second `section` element's closing tag. It is not in the correct position.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

The second `h2` element should have the text `Cat Lists`. テキストに誤字脱字があります。

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
