---
id: 5dfb655eeacea3f48c6300b3
title: ステップ 25
challengeType: 0
dashedName: step-25
---

# --description--

`figure` 要素はそれ自体で完結する 1 つのコンテンツを表します。これを使って画像をキャプションと関連付けることができます。

Nest the image you just added within a `figure` element.

# --hints--

Your `figure` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert.exists(document.querySelector('figure'));
```

Your `figure` element should have a closing tag. 終了タグは `<` の直後に `/` があるタグです。

```js
assert.match(code, /<\/figure\>/);
```

There should be an `figure` element right above the second `section` element's closing tag.

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

The lasagna `img` element should be nested in the `figure` element.

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
