---
id: 5ef9b03c81a63668521804d4
title: ステップ 34
challengeType: 0
dashedName: step-34
---

# --description--

`strong` 要素は、あるテキストの重要性や緊急性が高いことを示すために使われます。

In the `figcaption` you just added, indicate that `hate` is of strong importance by wrapping it in a `strong` element.

# --hints--

Your `strong` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('strong'));
```

Your `strong` element should have a closing tag. 終了タグは `<` の直後に `/` があるタグです。

```js
assert.match(code, /<\/strong\>/);
```

`strong` 要素が、`Cats hate other cats.` というテキスト内の単語 `hate` を囲む必要があります。 テキストが設定されていないか、誤字脱字があります。

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    ?.querySelector('strong')
    ?.innerText?.trim().toLowerCase(), 'hate'
);
```

`figcaption` のテキストは `Cats hate other cats.` でなければなりません。 誤字がないか、`strong` 要素の開始タグと終了タグの前後に適切なスペースがあるか確認してください。

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
   secondFigCaption?.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
