---
id: 5ef9b03c81a63668521804ec
title: ステップ 40
challengeType: 0
dashedName: step-40
---

# --description--

ページのコンテンツ全体が `html` 要素内にネストされていることに注意しましょう。 `html` 要素は HTML ページのルート要素であり、ページ内のすべてのコンテンツがこの中に含まれます。

また、`html` 要素に `lang` 属性を追加することで、ページの言語を指定することができます。

このページの言語が英語であることを指定するために、`html` の開始タグに `lang` 属性と値 `en` を追加してください。

# --hints--

`html` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

`html` 要素は値が `en` に設定されている `lang` 属性を持つ必要があります。 属性または値が設定されていないか、誤字脱字があります。

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
