---
id: 5ef9b03c81a63668521804ea
title: ステップ 38
challengeType: 0
dashedName: step-38
---

# --description--

ここまでページに追加してきた物がすべて `body` 要素内にあることに注目しましょう。 ページ上に表示されるべきコンテンツの要素は、すべて `body` 要素の中に入れます。 ですが、その他の重要な情報を `head` 要素に入れます。

`head` 要素は、ドキュメントについてのメタデータを格納するために使われます。例えばドキュメントのタイトル、スタイルシートへのリンク、スクリプトなどです。 メタデータとは、ページ上には表示されない、そのページ自体に関する情報です。

`body` 要素の上に `head` 要素を追加してください。

# --hints--

`body` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

`head` 要素には開始タグが必要です。 Opening tags have the following syntax: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

`head` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /\<\/head\>/);
```

`head` 要素は `body` 要素の開始タグの上に置く必要があります。 どこか違う場所に配置されているようです。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
