---
id: 66731cd027ef3acb155669f5
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

The first element you will learn about is the `h1` element. The `h1` element is a heading element and is used for the main heading of a webpage.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

コード内には `CatPhotoApp` というテキストが必要です。 スペルが正しいか確認してください。

```js
assert.match(code, /catphotoapp/i);
```

`h1` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素には終了タグが必要です。 Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

`h1` 要素のテキストは `CatPhotoApp` でなければなりません。 テキストに誤字脱字があるか、テキストが `h1` 要素の開始タグと終了タグの間に入っていない可能性があります。

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

ウェブページを変更するブラウザ拡張機能が使用されているようです。 拡張機能をオフにしてください。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
