---
id: 6690bb0b4062b114997f708a
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 has some elements that identify different content areas. These elements make your HTML easier to read and help with Search Engine Optimization (SEO) and accessibility.

`main` 要素は、HTML ドキュメント本文のメインとなるコンテンツを表すために使われます。 `main` 要素内のコンテンツはそのドキュメント特有のコンテンツであり、また、他の部分と重複しない内容であるべきです。

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identify the main section of this page by adding a `<main>` opening tag before the `h1` element, and a `</main>` closing tag after the `p` element.

# --hints--

Your `main` element should have an opening tag. 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('main'));
```

Your `main` element should have a closing tag. 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /<\/main\>/);
```

Your `main` element's opening tag should be below the `body` element's opening tag. 順番に誤りがあるようです。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Your `main` element's opening tag should be above the `h1` element. 順番が誤っているようです。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Your `main` element's closing tag should be below the `p` element. 順番に誤りがあるようです。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /see more cat photos in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```
