---
id: 6690be42942f421648b119eb
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

`p` 要素の下に、`img` 要素を追加してください。 この時点では、ブラウザーに画像は表示されません。

# --hints--

`img` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('img'));
```

`img` 要素は終了タグ `</img>` を持たないようにしてください。

```js
assert.notMatch(code, /<\/img\>/);
```

`img` 要素は 1 つだけにしてください。 余分なものは削除してください。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

`img` 要素は `p` 要素の下に置く必要があります。 順番に誤りがあるようです。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


