---
id: 6690be902f051416f668962b
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>attributes</dfn> are special words used inside the opening tag of an element to control the element's behavior. The `src` attribute in an `img` element specifies the image's URL (where the image is located).

下記は `src` 属性に freeCodeCamp ロゴ の URL を指定した `img` 要素の例です。

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

既存の `img` 要素に、`src` 属性を追加してこの URL を指定してください:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Your code should have an `img` element. You may have removed the `img` element or you have not surrounded the `src` attribute's value with quotes.

```js
assert.exists(document.querySelector('img'));
```

`img` 要素には `src` 属性が必要です。 You have either omitted the attribute or have a typo. Make sure there is a space between the element name and the attribute name.

```js
assert.exists(document.querySelector('img').src);
```

Your `img` element's `src` attribute should be set to '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. You have either omitted the URL or have a typo. The case of the URL is important.

```js
assert.equal(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Although you have set the `img` element's `src` to the correct URL, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

`img` 要素を `>` または `/>` で閉じてください。

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```


