---
id: 669842cee55b683aa13ee758
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

Within the second `section` element, add a new `h2` element with the text `Cat Lists`.

# --hints--

`section` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

`h2` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

`h2` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

2 つ目の `h2` 要素は、2 つ目の `section` 要素の終了タグのすぐ上にある必要があります。 位置が正しくないようです。

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

2 つ目の `h2` 要素には `Cat Lists` というテキストが必要です。 テキストに誤字脱字があります。

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


