---
id: 66984ec64a3652522e81c678
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

Turn the existing `freeCodeCamp.org` text into a link by enclosing it in an anchor (`a`) element. The `href` attribute should be set to `https://www.freecodecamp.org`.

# --hints--

アンカー (`a`) 要素は、`footer` 要素内にある `p` 要素の中にネストされている必要があります。 アンカー (`a`) 要素の開始タグと終了タグが追加されていることを確認してください。

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

アンカー (`a`) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 3);
```

アンカー (`a`) 要素は値が `https://www.freecodecamp.org` に設定されている `href` 属性を持つ必要があります。 属性または値が設定されていないか、誤字脱字があります。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

リンクのテキストは `freeCodeCamp.org` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' '),
    'freecodecamp.org'
);
```

アンカー (`a`) 要素をネストした後は、ブラウザーに表示される `p` 要素の中身は `No Copyright - freeCodeCamp.org` だけになるはずです。 `p` 要素およびネストされたアンカー要素の、テキスト、スペース、句読点を再確認してください。

```js
const pText = document.querySelectorAll('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```


