---
id: 66984e0cb4febb50c3794edf
title: ステップ 33
challengeType: 0
dashedName: step-33
---

# --description--

The `footer` element is used to define a footer for a document or section. A footer typically contains information about the author of the document, copyright data, links to terms of use, contact information, and more.

`main` 要素の後に `footer` 要素を追加してください。

# --hints--

`main` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

`footer` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('footer'));
```

`footer` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /<\/footer\>/);
```

`footer` 要素は `main` 要素の終了タグの下に置く必要があります。 どこか違う場所に配置されているようです。

```js
assert.equal(document.querySelector('main').nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```


