---
id: 66984fe8efbd8c53a78c9b62
title: ステップ 37
challengeType: 0
dashedName: step-37
---

# --description--

The `title` element determines what browsers show in the title bar or tab for the page.

下記のテキストを使用して、`head` 要素内に `title` 要素を追加してください:

`CatPhotoApp`

# --hints--

`head` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

`title` 要素は `head` 要素の中にネストされている必要があります。 `title` 要素の開始タグと終了タグが追加されていることを確認してください。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

`title` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /\<\/title\>/);
```

`title` 要素のテキストは `CatPhotoApp` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


