---
id: 66984905e95af5432faa28b1
title: ステップ 30
challengeType: 0
dashedName: step-30
---

# --description--

To improve the accessibility of the image you added, add an `alt` attribute with the text:

`Five cats looking around a field.`

# --hints--

`figure` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

`figure` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
assert.lengthOf(code.match(/<\/figure>/g), 2);
```

最後の `section` 要素の終了タグのすぐ上に `figure` 要素が必要です。

```js
assert.equal(document.querySelectorAll('main > section')[1].lastElementChild.nodeName, 'FIGURE');
```

5 匹の猫の `img` 要素は `figure` 要素内にネストされている必要があります。

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.equal(
    catsImg?.getAttribute('src')?.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg'
);
```

5 匹の猫の `img` 要素には `alt` 属性があり、値が `Five cats looking around a field.` に設定されている必要があります。

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.match(
  catsImg
    .getAttribute('alt')
    .replace(/\s+/g, ' '),
    /^Five cats looking around a field\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
--fcc-editable-region--
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg">
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```


