---
id: 669840900adeba35e4ebb917
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

A link's text must be placed between the opening and closing tags of an anchor (`a`) element.

下記は `click here to go to freeCodeCamp.org` というテキストのリンクの例です。

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Add the anchor text `link to cat pictures` to the anchor element. This will become the link's text.

# --hints--

アンカー (`a`) 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('a'));
```

アンカー (`a`) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /<\/a\>/);
```

Your anchor (`a`) element's text should be `link to cat pictures`. Make sure to put the link text between the anchor (`a`) element's opening tag and closing tag.

```js
assert.equal(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```


