---
id: 66984e8ab54bee517207be42
title: ステップ 34
challengeType: 0
dashedName: step-34
---

# --description--

Nest a `p` element with the text `No Copyright - freeCodeCamp.org` within the `footer` element.

# --hints--

`footer` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

`footer` 要素内に `p` 要素が 1 つ必要です。 `p` 要素の開始タグと終了タグが追加されていることを確認してください。

```js
assert.exists(document.querySelector('footer > p'));
```

`footer` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があるタグです。

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 2);
```

`p` 要素のテキストは `No Copyright - freeCodeCamp.org` でなければなりません。 テキストに誤字脱字があるか、テキストが `p` 要素の開始タグと終了タグの間に入っていない可能性があります。

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0].innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```


