---
id: 66984218720e29393b1adad2
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

Before adding any new content, you should make use of a `section` element to separate the cat photos content from the future content.

`section` 要素は、ドキュメント内のチャプター、ヘッダー、フッターなどのセクションを定義するために使われます。 これは、SEO とアクセシビリティに役立つセマンティック要素です。

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Take your `h2`, comment, `p`, and anchor (`a`) elements and nest them in a `section` element.

# --hints--

Your `section` element should have an opening tag. 開始タグは `<elementName>` のような形式の構文です。

```js
assert.exists(document.querySelector('section'));
```

Your `section` element should have a closing tag. 終了タグは `<` の直後に `/` があります。

```js
assert.match(code, /<\/section\s*>/i);
```

The entire `section` element should be between the opening and closing tags of the `main` element.

```js
assert.equal(document.querySelector('section').parentNode.nodeName, 'MAIN');
```

The existing `h2`, comment, `p` element, and anchor (`a`) element should be between the opening and closing tags of the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 3)
```

The `h1` element should not be nested in the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```


