---
id: 6ac0ae563f6de70930f47ef7
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

코드를 작성할 때 기존 값을 재사용하거나 참조하고 싶을 때가 많습니다. 변수를 사용하면 그렇게 할 수 있습니다.

변수에 값을 할당하려면 변수 이름 다음에 대입 연산자(`=`)를 쓰고, 그 뒤에 할당할 값을 적으세요:

```py
sentence = 'I am learning Python'
```

강조된 줄로 가서 `sentence`라는 변수를 만들고 문자열 `I am learning Python`을 할당하세요.

# --hints--

코드에 `sentence = 'I am learning Python'`가 있어야 합니다. 줄의 맨 앞에서 시작했는지 확인하세요.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
