---
id: 6ac0aaba914d81b32567c076
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`print`은 터미널에 출력을 보내는 데 사용할 수 있는 내장 함수입니다. 보시다시피, 텍스트 `Hello`가 이제 터미널에 보입니다.

이제 강조 표시된 줄로 가서 텍스트 `freeCodeCamp`를 터미널에 출력하세요:

```py
print('freeCodeCamp')
```

줄의 맨 처음에서 시작했는지 확인하세요.

# --hints--

코드에 `print('freeCodeCamp')`가 있어야 합니다. 줄의 맨 처음에서 시작했는지 확인하세요.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
