---
id: 6ac0a4463662d9738c393e21
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이전 단계에서, 작은따옴표나 큰따옴표로 감싼 문자들의 수열인 **문자열**을 만들었어요. 프로그래밍에서 문자열은 텍스트를 나타내는 데 사용됩니다.

**터미널**은 코드의 출력을 볼 수 있는 곳입니다.

출력을 볼 수 있도록 기존 코드를 다음과 같이 바꾸세요:

```py
print('Hello')
```

줄의 맨 앞에서 시작했는지 꼭 확인하세요.

# --hints--

코드에 `print('Hello')`가 있어야 합니다. 줄의 맨 앞에서 시작했는지 꼭 확인하세요.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
