---
id: 6ac09f8420b4b67e075d2cd2
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서 첫 번째 Python 프로그램을 만들게 됩니다.

여기는 Python 코드를 작성하는 코드 편집기입니다. 1번째 줄이 강조되어 있습니다. 그곳으로 가서 입력하세요:

```py
'Hello'
```

줄의 처음부터 시작하고 따옴표를 잊지 마세요.

완료하면 "check your code" 버튼을 클릭해 올바른지 확인하세요.

# --hints--

`'Hello'`를 입력해야 합니다. 줄의 처음부터 시작하고 따옴표를 잊지 마세요.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
