---
id: 6ac0b35f3b29a27f7e207295
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이전 단계에서 변수 `sentence`를 만들었어요. 이제 강조 표시된 줄에 `print(sentence)`를 입력해 터미널에 그 값을 출력하세요.

이로써 첫 번째 Python 프로그램을 완성했어요.

# --hints--

코드에 `print(sentence)`가 있어야 해요. 줄 맨 앞에서 시작했는지 확인하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
