---
id: 63ee3ff1381756f9716727f2
title: Zoezi la msingi la CSS D
challengeType: 25
dashedName: css-foundations-exercise-d
---

# --description--

Kwa zoezi hili, tumekupa faili la HTML lililokamilika, hivyo utalazimika tu kuhariri faili la CSS. Kwa zoezi hili, ni muhimu zaidi kuelewa jinsi uunganishaji wa mfululizo wa vichaguzi unavyofanya kazi kuliko jinsi ya kuongeza sifa za kipengele.

1. Unapaswa kuona `width` ya `300px` kwenye darasa la vitu la `avatar` na `proportioned`.
1. Unapaswa kuipa urefu ili ibaki na uwiano wa mraba wa awali (usitengeneze thamani ya urefu kwa pikseli moja kwa moja!).
1. Unapaswa kuipa vipengele vyenye madarasa ya `avatar` na `distorted` `width` ya `200px`.
1. Unapaswa kuipa `height` mara mbili ya upana wake.

# --hints--

Unapaswa kuwa na `width` ya `300px` kwenye darasa la vitu la `avatar` na `proportioned`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`);
assert(style?.width === '300px');
```

Unapaswa kuwa na urefu wa `auto` kwenye darasa la vitu la `avatar` na `proportioned`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`)
assert(style?.height === 'auto');
```

Unapaswa kutumia kichaguzi cha uunganishaji wa mfululizo kwenye darasa la vitu la `avatar` na `proportioned`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.proportioned`) || new __helpers.CSSHelp(document).getStyle(`.proportioned.avatar`);
assert(style);
```

Unapaswa kuwa na `width` ya `200px` kwenye darasa la vitu la `avatar` na `distorted`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style?.width === '200px');
```

Unapaswa kutumia kichaguzi cha uunganishaji wa mfululizo kwenye darasa la vitu la `avatar` na `distorted`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style);
```

Unapaswa kuwa na `height` mara mbili ya upana kwenye darasa la vitu la `avatar` na `distorted`.

```js
const style = new __helpers.CSSHelp(document).getStyle(`.avatar.distorted`) || new __helpers.CSSHelp(document).getStyle(`.distorted.avatar`);
assert(style?.height === '400px');
```

# --seed--

## --seed-contents--

```css

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Chaining Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <!-- Use the classes BELOW this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
    <!-- Use the classes ABOVE this line -->
    <div>
      <img class="original proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="original distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Chaining Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <!-- Use the classes BELOW this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
    </div>
    <!-- Use the classes ABOVE this line -->
    <div>
      <img class="avatar proportioned" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-01.jpg" alt="Two kittens on a blanket.">
      <img class="avatar distorted" src="https://cdn.freecodecamp.org/curriculum/odin-project/css-foundations/css-foundations-02.jpg" alt="A cute kitten in the woods.">
  </body>
</html>
```

```css
.avatar.proportioned {
  height: auto;
  width: 300px;
}

.avatar.distorted {
  height: 400px;
  width: 200px;
}
```
