---
id: 63ee3fe9381756f9716727f1
title: Zoezi la msingi la CSS C
challengeType: 25
dashedName: css-foundations-exercise-c
---

# --description--

Tujenge kidogo kutoka kwenye zoezi lililopita, ambapo uliweka madarasa mengi kwenye kipengele kimoja ili kutumia sheria mbili tofauti kwa kipengele hicho.

1. Unapaswa kuona mandhari ya rangi nyeusi na maandishi meupe kwenye kipengele cha kwanza cha `button`.
1. Unapaswa kuona mandhari ya rangi ya manjano kwenye kipengele cha pili cha `button`.
1. Unapaswa kuweka ukubwa wa fonti wa `28px` kwa kutumia orodha ya vichaguzi.
1. Unapaswa kuwa na orodha ya fonti zenye `Helvetica` na `Times New Roman` na `sans-serif` kama njia mbadala katika orodha ya vichaguzi.
1. Unapaswa kuona jina la darasa la vitu la kipekee kwenye kila kipengele.
1. Unapaswa kuwa na orodha ya vichaguzi kwa mitindo ambayo vipengele vyote viwili vinashirikiana nayo.

# --hints--

Unapaswa kuwa na mandhari ya rangi ya `black` kwenye kipengele cha kwanza.

```js
const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.backgroundColor, 'black');

```

Kipengele chako cha kwanza kinapaswa kuwa na rangi ya maandishi ya `white`.

```js

const classes = document.querySelectorAll('button')?.[0].classList;

const style = new __helpers.CSSHelp(document).getStyle(`.${classes[0]}`);

assert.equal(style?.color, 'white');

```

Unapaswa kuweka `font-size` kwa vipengele vyote viwili kuwa `28px` kwa kutumia orodha ya vichaguzi.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.equal(eitherOr()?.fontSize, '28px');
```

Unapaswa kuwa na jina la darasa la vitu la kipekee kwenye kila kipengele.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

for(let i = 0; i < elementOneClasses.length; i++){
    assert(![...elementTwoClasses].includes(elementOneClasses[i]));
}

```

Unapaswa kuwa na orodha ya vichaguzi kwa mitindo ambayo vipengele vyote viwili vinashirikiana nayo.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}
assert.exists(eitherOr());
```

Unapaswa kuwa na orodha ya fonti zenye `Helvetica` na `Times New Roman` na `sans-serif` kama njia mbadala katika orodha ya vichaguzi.

```js
const classOne = document.querySelectorAll('button')?.[0].classList?.[0];
const classTwo = document.querySelectorAll('button')?.[1].classList?.[0];

function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle(`.${classOne}, .${classTwo}`) ?? a.getStyle(`.${classTwo}, .${classOne}`);
}

assert.equal(eitherOr()?.fontFamily, 'Helvetica, "Times New Roman", sans-serif');
```

Kipengele kila kimoja kinapaswa kuwa na darasa la vitu moja tu.

```js
const elementOneClasses = document.querySelectorAll('button')?.[0].classList;
const elementTwoClasses = document.querySelectorAll('button')?.[1].classList;

assert(elementOneClasses.length === 1 && elementTwoClasses.length === 1);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button>Click Me!</button>
    <button>No, Click Me!</button>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
  <html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Grouping Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <button class="inverted">Click Me!</button>
    <button class="fancy">No, Click Me!</button>
  </body>
</html>
```

```css
.inverted,
.fancy {
  font-family: Helvetica, "Times New Roman", sans-serif;
  font-size: 28px;
}

.inverted {
  background-color: black;
  color: white;
}

.fancy {
  background-color: yellow;
}
```
