---
id: 63ee3f71381756f9716727ef
title: Zoezi la msingi la CSS A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**Lengo:**  
Katika zoezi hili, utafanya mazoezi ya kuongeza CSS kwenye faili ya HTML kwa kutumia njia zote tatu: CSS ya nje, CSS ya ndani, na CSS ndani ya mstari. Unapaswa kutumia tu vichaguzi vya aina kwa zoezi hili wakati wa kuongeza mitindo kupitia njia za nje na za ndani. Pia unapaswa kutumia maneno muhimu kwa rangi (mfano "blue") badala ya kutumia thamani za `RGB` au `HEX`.

## Hadithi za mtumizi

- Unapaswa kuona kipengele cha `div` chenye maandishi.
  - Kinapaswa kuwa na mandhari ya nyuma ya `red`, maandishi ya `white`, ukubwa wa fonti wa `32px`, maandishi yakiwa yamepangiliwa katikati na `bold`.
  - CSS ya kipengele cha `div` inapaswa kuongezwa kwa njia ya nje, na kwa kutumia kichaguzi cha aina.
- Unapaswa kuona kipengele cha `p` chenye maandishi.
  - Kinapaswa kuwa na mandhari ya nyuma ya `green`, maandishi ya `white`, na ukubwa wa fonti wa `18px`.
  - CSS ya kipengele cha `p` inapaswa kuongezwa ndani, na kwa kutumia kichaguzi cha aina.
- Unapaswa kuona kipengele cha `button` chenye maandishi.
  - Kipengele cha `button` kinapaswa kuwa na mandhari ya nyuma ya `orange` na ukubwa wa fonti wa `18px`.
  - CSS ya kipengele cha `button` inapaswa kuongezwa kwa kutumia mitindo ndani ya mstari.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `div` chenye maandishi.

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

Unapaswa kuwa na karatasi ya mitindo ya nje yenye mitindo ya kipengele cha `div`.

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

Kipengele chako cha `div` hakipaswi kuwa na CSS yake imeongezwa kwa kutumia mitindo ya ndani au ndani ya mstari.

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

Kipengele chako cha `div` kinapaswa kuwa na `background-color` ya `red` na `color` ya `white`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

Kipengele chako cha `div` kinapaswa kuwa na `font-weight` imewekwa kuwa `bold`, `font-size` imewekwa kuwa `32px`, na `text-align` imewekwa kuwa `center`.

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

Unapaswa kuwa na kipengele kimoja cha `p` na kinapaswa kuwa na maandishi.

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

Kipengele chako cha `p` kinapaswa kuwa na mitindo yake imeongezwa ndani kwa kutumia kipengele cha `style`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

Kipengele chako cha `p` kinapaswa kuwa na `font-size` ya `18px` na kuwa na `color` imewekwa kuwa `white`.

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

Unapaswa kuwa na kipengele kimoja cha `button` chenye maandishi.

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

Kipengele chako cha `button` kinapaswa kuwa na mtindo ndani ya mstari.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```
 
Kipengele chako cha `button` kinapaswa kuwa na `background-color` yake imewekwa kuwa `orange`.

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

Kipengele chako cha `button` kinapaswa kuwa na `font-size` yake imewekwa kuwa `18px`.

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
