---
id: 63ee3fe4381756f9716727f0
title: Zoezi la msingi la CSS B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**Lengo:** Kuna vipengele kadhaa katika faili la HTML lililotolewa, ambavyo utalazimika kuongeza sifa ya kipengele ya darasa la vitu au ID. Kisha utalazimika kuongeza kanuni katika faili la CSS lililotolewa kwa kutumia sintaksia sahihi ya kichaguzi.

## Hadithi za mtumizi

1. Unapaswa kuona rangi ya nyuma ya `yellow` kwa vipengele vyote vyenye nambari zisizoshirikiwa katika orodha.

1. Unapaswa kuwa na kichaguzi cha `class` kinachotumika kwa vipengele vyote vyenye nambari zisizoshirikiwa katika orodha.

1. Unapaswa kuona kwamba kipengele cha pili katika orodha kina maandishi ya `blue` na `font-size` ya `36px`.

1. `font-size` na rangi ya maandishi kwenye kipengele cha pili inapaswa kuwekwa kwa kutumia sifa ya kipengele ya `id`.

1. Unapaswa kuona kwamba kipengele cha tatu katika orodha kina `font-size` ya `24px`.

1. `font-size` kwenye kipengele cha tatu inapaswa kuwekwa kwa kutumia sifa ya kipengele ya `class`.

1. Unapaswa kuona kwamba kipengele cha nne katika orodha kina rangi ya nyuma ya `red`, `font-size` ya `24px`, na `font-weight` ya `bold`.

1. `font-size` ya kipengele cha nne inapaswa kuwekwa kwa sifa ya kipengele ya `class`. `font-weight` na rangi inapaswa kuwekwa kwa sifa ya kipengele ya `id`.

# --hints--

Kipengele kila kimoja chenye nambari ishirini na moja kinapaswa kuwa na sifa ya kipengele ya `class`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

Vipengele vyako vyenye nambari zisizoshirikiwa vinapaswa kuwa na sifa ya kipengele ya `background-color` ya `yellow`.

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

Kipengele chako cha pili kinapaswa kuwa na maandishi ya `blue` na `font-size` ya `36px`.

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

Kipengele chako cha tatu kinapaswa kuwa na maandishi na `font-size` ya `24px`.

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

Kipengele cha nne kinapaswa kuwa na `font-size` ya `24px`.

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

Kipengele cha nne kinapaswa kuwa na `red` `background-color`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

Kipengele cha nne kinapaswa kuwa na `font-weight` ya `bold`.

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
