---
id: 61537bb9b1a29430ac15ad38
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comece este projeto vinculando seu arquivo `styles.css` à página.

# --hints--

Você deve ter um elemento `link` dentro do seu elemento `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Seu elemento `link` deve ter `rel="stylesheet"` e `href="styles.css"`.

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
