---
id: 61537bb9b1a29430ac15ad38
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comece este projeto vinculando seu arquivo `styles.css` à página.

# --hints--

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
