---
id: 614385513d91ae5c251c2052
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge ein `title`-Element mit dem Text `Magazine` hinzu, ein `link`-Element für das `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`-Font-Stylesheet, einen `link` für das `https://use.fontawesome.com/releases/v5.8.2/css/all.css`-FontAwesome-Stylesheet und einen `link` für dein `./styles.css`-Stylesheet.

Dein Font-Stylesheet lädt drei verschiedene Fonts: `Anton`, `Baskervville` und `Raleway`.

# --hints--

Dein Code sollte drei `link`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Dein`link`-Element sollte innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Deine drei `link`-Elemente sollten ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines deiner Link-Elemente sollte `href` auf `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` eingestellt haben.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Eines deiner Link-Elemente sollte `href` auf `https://use.fontawesome.com/releases/v5.8.2/css/all.css` eingestellt haben.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Eines deiner `link`-Elemente sollte `href`- auf den Wert `styles.css` gesetzt haben.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

Dein Code sollte über ein `title`-Element verfügen.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

Dein`title`-Element sollte innerhalb des `head`-Elements liegen.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Dein Projekt sollte den Titel `Magazine` tragen. Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
