---
id: a3566b1109230028080c9345
title: Projetar um Algoritmo para Somar Todos os Números
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

Neste laboratório, você precisará projetar um algoritmo para somar todos os números.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `sumAll` que aceita um array de dois números.
1. `sumAll([n, m])` deve retornar a soma de `n` e `m` mais a soma de todos os números entre eles. O número menor nem sempre virá primeiro. Por exemplo, `sumAll([4,1])` deve retornar `10` porque a soma de todos os números entre `1` e `4` (ambos inclusivos) é `10`.

# --hints--

Você deve ter uma função chamada `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` deve ser um número.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` deve retornar `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` deve retornar `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` deve retornar `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` deve retornar `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
