---
id: a0b5010f579e69b815e7c5d6
title: Tengeneza kitendakazi cha kubadilisha neno kwa akili
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

Katika maabara hii, utaunda kitendakazi kinachofanya utafutaji na kubadilisha katika sentensi kwa kutumia hoja zilizotolewa kisha kurudisha sentensi mpya.

**Kumbuka:** Hifadhi herufi ya kwanza ya neno la awali unaporibadilisha. Kwa mfano, ikiwa unakusudia kubadilisha neno `Book` na neno `dog`, linapaswa kubadilishwa kama `Dog`

**Lengo**: Timiza hadithi za mtumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `myReplace`.
1. Kitendakazi `myReplace` kinapaswa kuchukua hoja tatu: mfuatano wa herufi, neno la kubadilishwa, na neno la kubadilisha nalo.
1. Kitendakazi `myReplace` kinapaswa kurudisha mfuatano wa herufi mpya ulio na kubadilishwa.
1. Unapaswa kuhifadhi herufi ya kwanza ya neno la awali unaporibadilisha.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` inapaswa kuchukua hoja tatu.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` inapaswa kurudisha mfuatano wa herufi `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` inapaswa kurudisha mfuatano wa herufi `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` inapaswa kurudisha mfuatano wa herufi `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` inapaswa kurudisha mfuatano wa herufi `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` inapaswa kurudisha mfuatano wa herufi `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` inapaswa kurudisha mfuatano wa herufi `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
