---
id: 612e89d254fe5d3df7d6693d
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

html ボックスモデルをリセットしたので、内部の要素にもこれを適用する必要があります。そのためには、すべての他の要素の `box-sizing` プロパティを継承に設定し、対象の要素が親要素と同じ値を持つようにします。

また、セレクターに続く特別なキーワードである疑似要素も対象にする必要があります。使用する疑似要素は `::before` と `::after` です。

`::before` セレクターは選択した要素の最初の子要素である疑似要素を作成し、`::after` セレクターは選択した要素の最後の子要素である疑似要素を作成します。これらの疑似要素は装飾的なコンテンツを作成するためによく使われ、後ほどこのプロジェクトで見ていきます。

まずは、`*` を使ってすべての要素を対象にする CSS セレクターを作成し、疑似要素 `::before` と `::after` も含めてください。`box-sizing` プロパティを `inherit` に設定しましょう。

# --hints--

`html` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

`html` セレクターには `box-sizing` プロパティが `border-box` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

`*, ::before, ::after` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

`*, ::before, ::after` セレクターには `box-sizing` プロパティが `inherit` に設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
