---
id: 612e804c54d5e7308d7ebe56
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`.keys` 要素の中に、クラスが `div` の `key` 要素を7つ追加してください。

# --hints--

7つの新しい `div` 要素を作成してください。

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

7つの新しい `div` 要素は、`.keys` 要素の中に入っている必要があります。

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

7つの新しい `div` 要素はすべて、`class` が `key` に設定されている必要があります。

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
