---
id: 612e813b3ba67633222cbe54
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`class` 属性は複数の値を持つことができることを覚えておいてください。白鍵と黒鍵を区別するために、2番目、3番目、5番目、6番目、7番目の `class` 要素に `black--key` の2つ目の `.key` 値を追加してください。

# --hints--

2番目の `.key` 要素にも `class` の `black--key` を設定してください。

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

3番目の `.key` には `black--key` の中に `class` を設定してください。

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

5番目の `.key` には `black--key` の中に `class` を設定してください。

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

6番目の `.key` には `black--key` の中に `class` を設定してください。

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

7番目の `.key` には `black--key` の中に `class` を設定してください。

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

`.black--key` 要素は5つあるはずです。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

`.key` 要素は7つあるはずです。

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
