---
id: 612e96fc87fe8e44f69f7ec5
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

キーのスタイルを設定する時間です。`#piano`ルールの下に、`.keys`セレクターで`class`要素を指定してください。新しいルールに`background-color`プロパティとして`#040404`、`width`プロパティとして`949px`、`height`プロパティとして`180px`を設定してください。

# --hints--

`.keys`セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.keys'));
```

`.keys`セレクターには`background-color`プロパティを`#040404`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor, 'rgb(4, 4, 4)');
```

`.keys`セレクターには`width`プロパティを`949px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.width, '949px');
```

`.keys`セレクターには`height`プロパティを`180px`に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.keys')?.height, '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
