---
id: 612e7d1c29fb872d6384379c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

既存の`div`の中に2つ目の`div`をネストし、`class`を`keys`に設定してください。

# --hints--

2つ目の`div`要素を作成してください。

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

新しい`div`要素は既存の`div`要素の中に入っている必要があります。

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

新しい`div`要素の`class`は`keys`に設定してください。

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
