---
id: 68e45a2b4ab9b2f48b1cc6df
title: Projete uma Página de Layout de Planos de Preços
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário**

1. Sua página deve ter um elemento `h1` com o texto `Pricing Plans`.
2. Sua página deve ter um elemento `div` com a classe `pricing-container` abaixo do elemento `h1`.
    - O seletor `.pricing-container` deve ter uma propriedade `display` com o valor `flex` e uma propriedade `flex-wrap` com o valor `wrap`.
3. Dentro do elemento `.pricing-container`, você deve ter três elementos `div` com a classe `pricing-card` para representar os planos de preços.
    - Um dos elementos `.pricing-card` deve ter a classe `basic-plan` além da classe `pricing-card`.
    - Um dos elementos `.pricing-card` deve ter a classe `pro-plan` além da classe `pricing-card`.
    - Um dos elementos `.pricing-card` deve ter a classe `premium-plan` além da classe `pricing-card`.
4. Seu elemento `.basic-plan` deve ter um elemento `h2` com o texto `Basic`.
5. Seu elemento `.basic-plan` deve ter um elemento `p` com o texto `$9/month`.
6. Seu elemento `.pro-plan` deve ter um elemento `h2` com o texto `Pro`.
7. Seu elemento `.pro-plan` deve ter um elemento `p` com o texto `$19/month`.
8. Seu elemento `.premium-plan` deve ter um elemento `h2` com o texto `Premium`.
9. Seu elemento `.premium-plan` deve ter um elemento `p` com o texto `$29/month`.
10. Cada um dos seus elementos `.pricing-card` deve:
    - Use Flexbox para empilhar seu conteúdo em uma coluna e justificar o espaço entre os filhos usando `space-between`.
    - Defina a propriedade `flex` para `0 0 200px` para dar uma largura consistente e evitar que ela cresça ou encolha no layout.
    - Defina a propriedade `border` para `2px solid black` para ver como os diferentes cartões ocupam espaço.
11. O elemento `.basic-plan` deve aparecer primeiro no layout. Você deve usar a propriedade `order` para isso.
12. O elemento `.pro-plan` deve aparecer em segundo no layout. Você deve usar a propriedade `order` e definir sua propriedade `flex-grow` como `2` para que ocupe mais espaço do que os outros planos.
13. O elemento `.premium-plan` deve vir por último no layout. Você deve usar a propriedade `order` para isso.

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Sua página deve ter um elemento `h1` com o texto `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Você deve ter um elemento `div` abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Seu elemento `div` deve ter a classe `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Você deve ter um seletor `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Seu seletor `pricing-container` deve ter a propriedade `display` com o valor `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

Seu seletor `.pricing-container` deve ter a propriedade `flex-wrap` com o valor `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Você deve ter três elementos `div` com a classe `pricing-card` dentro do elemento `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Um dos seus elementos `.pricing-card` deve ter a classe `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Um dos seus elementos `.pricing-card` deve ter a classe `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Um dos seus elementos `.pricing-card` deve ter a classe `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Você deve ter um seletor `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

Seu seletor `.pricing-card` deve ter a propriedade `display` com o valor `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

Seu seletor `.pricing-card` deve ter a propriedade `flex-direction` com o valor `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Seu seletor `.pricing-card` deve ter a propriedade `justify-content` com o valor `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Seu seletor `.pricing-card` deve ter a propriedade `flex` com o valor `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

Seu seletor `.pricing-card` deve ter a propriedade `border` com um valor de `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Seu seletor `.basic-plan` deve ter a propriedade `order` com o valor `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

Seu elemento `.basic-plan` deve ter um elemento `h2` com o texto `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Seu elemento `.basic-plan` deve ter um elemento `p` com o texto `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

Seu elemento `.pro-plan` deve ter um elemento `h2` com o texto `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

Seu elemento `.pro-plan` deve ter um elemento `p` com o texto `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

Seu elemento `.pro-plan` deve ter a propriedade `order` com o valor `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

Seu elemento `.pro-plan` deve ter a propriedade `flex-grow` com um valor de `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

Seu elemento `.premium-plan` deve ter um elemento `h2` com o texto `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

Seu elemento `.premium-plan` deve ter um elemento `p` com o texto `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

Seu elemento `premium-plan` deve ser o último elemento no layout e ter a propriedade `order` e o valor `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="cross">✗</span> Advanced analytics</li>
          <li><span class="cross">✗</span> Custom integrations</li>
          <li><span class="cross">✗</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">✓</span> Access to our core features</li>
          <li><span class="check">✓</span> Email support</li>
          <li><span class="check">✓</span> Advanced analytics</li>
          <li><span class="check">✓</span> Custom integrations</li>
          <li><span class="check">✓</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```

