---
id: 66a9b8f14b963916a3baa732
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Per aggiungere un sottotitolo a una `table`, puoi usare l'elemento <dfn>table caption</dfn>.

Ecco un esempio che usa l'elemento `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

All'interno del tuo elemento `table`, annida un elemento `caption` con il testo `Calculus Final Exam Grades`.

# --hints--

Dovresti avere un tag di apertura `caption`.

```js
assert.match(code, /<caption>/i);
```

Dovresti avere un tag di chiusura `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Il tuo elemento `caption` dovrebbe essere all'interno del tuo elemento `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Il tuo `caption` dovrebbe contenere il testo `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
