---
id: 66a9c006e4f58819396605b3
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Ora che hai completato la sezione head, è il momento di aggiungere il <dfn>corpo della tabella</dfn>, `tbody`. Il corpo della tabella rappresenterà tutti i nomi degli studenti e i loro voti.

Aggiungi un elemento corpo della tabella alla tua tabella.

# --hints--

Dovresti avere un tag di apertura `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Dovresti avere un tag di chiusura `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Il tuo tag di apertura `tbody` dovrebbe venire prima del tag di chiusura `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
