---
id: 66a9c14e3b34c719e34bae20
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Per aggiungere i dati degli studenti alla tabella, dovrai usare gli elementi riga della tabella e dati della tabella.

L'elemento <dfn>dati della tabella</dfn>, `td`, serve a creare una cella nella tabella.

Ecco un esempio di utilizzo dell'elemento `td` per una tabella di giocatori sportivi:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

All'interno del tuo elemento corpo della tabella, aggiungi un elemento riga della tabella, `tr`, con tre elementi dati della tabella, `td`.

Il primo elemento dati della tabella dovrebbe contenere il cognome di `Davis`.

Il secondo elemento dati della tabella dovrebbe contenere il nome di `Alex`.

Il terzo elemento dati della tabella dovrebbe contenere il voto di `54`.

# --hints--

Dovresti avere un elemento `tr` all'interno del tuo elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Dovresti avere un elemento `td` con il testo `Davis` all'interno dell'elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Dovresti avere un elemento `td` con il testo `Alex` all'interno dell'elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Dovresti avere un elemento `td` con il testo `54` all'interno dell'elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Dovresti avere tutti e tre gli elementi `td` all'interno dell'elemento `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
