---
id: 66a9ca099e3ecb1be10a2696
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora è il momento di aggiungere altri due studenti alla tabella.

Seguendo lo stesso pattern dei passaggi precedenti, aggiungi una quarta riga studente nella tabella. Usa i seguenti dati per gli elementi della tabella:

- Cognome: `Thompson`
- Nome: `Jane`
- Voto: `77`

Per la quinta riga studente della tabella, usa i seguenti dati per gli elementi della tabella:

- Cognome: `Williams`
- Nome: `Natalie`
- Voto: `83`

# --hints--

Dovresti avere un quarto `tr` all'interno del tuo elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Dovresti avere un elemento `td` con il testo `Thompson` all'interno del quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Dovresti avere un elemento `td` con il testo `Jane` all'interno del quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Dovresti avere un elemento `td` con il testo `77` all'interno del quarto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Dovresti avere un quinto `tr` all'interno del tuo elemento `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Dovresti avere un elemento `td` con il testo `Williams` all'interno del quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Dovresti avere un elemento `td` con il testo `Natalie` all'interno del quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Dovresti avere un elemento `td` con il testo `83` all'interno del quinto elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
