---
id: 66a9bcf00f13a418368a272e
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

L'elemento intestazione della tabella consiste in un elemento <dfn>riga della tabella</dfn>, `tr`, che contiene gli elementi <dfn>cella di intestazione della tabella</dfn>, `th`.

Ecco un esempio che usa gli elementi `tr` e `th` per una tabella sportiva:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

All'interno del tuo elemento `thead`, aggiungi un elemento `tr`.

All'interno del tuo elemento `tr`, aggiungi tre elementi `th`.

Il primo elemento `th` dovrebbe contenere il testo `Last Name`. Il secondo elemento `th` dovrebbe contenere il testo `First Name`. Il terzo elemento `th` dovrebbe contenere il testo `Grade`.

# --hints--

Dovresti avere un tag di apertura `tr`.

```js
assert.match(code, /<tr>/i);
```

Dovresti avere un tag di chiusura `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Il tuo elemento `tr` dovrebbe essere all'interno del tuo elemento `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Il tuo primo elemento `th` dovrebbe avere il testo `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Il tuo secondo elemento `th` dovrebbe avere il testo `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Il tuo terzo elemento `th` dovrebbe avere il testo `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Tutti i tuoi elementi `th` dovrebbero essere all'interno del tuo elemento `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
