---
id: 66a9b48a211a73155621d0df
title: Passaggio 1
challengeType: 0
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso delle tabelle HTML costruendo una tabella per l'esame finale di un gruppo di studenti.

Nelle lezioni precedenti, hai imparato a lavorare con l'elemento `table` per rappresentare dati tabellari.

All'interno del tuo elemento `body`, annida un elemento `table`.

# --hints--

Dovresti avere un tag di apertura `table`.

```js
assert.match(code, /<table>/i);
```

Dovresti avere un tag di chiusura `table`.

```js
assert.match(code, /<\/table>/i);
```

Il tuo elemento `table` dovrebbe essere all'interno del tuo elemento `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
