---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Ora è il momento di aggiungere altri due studenti alla tabella.

Seguendo lo stesso pattern del passaggio precedente, aggiungi una seconda riga per lo studente nella tabella. Usa i seguenti dati per gli elementi della tabella:

- Cognome: `Doe`
- Nome: `Samantha`
- Voto: `92`

Per la terza riga della tabella degli studenti, usa i seguenti dati per gli elementi della tabella:

- Cognome: `Rodriguez`
- Nome: `Marcus`
- Voto: `88`

# --hints--

Dovresti avere un secondo `tr` all’interno del tuo elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Dovresti avere un elemento `td` con il testo `Doe` all’interno del secondo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Dovresti avere un elemento `td` con il testo `Samantha` all’interno del secondo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Dovresti avere un elemento `td` con il testo `92` all’interno del secondo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Dovresti avere un terzo `tr` all’interno del tuo elemento `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Dovresti avere un elemento `td` con il testo `Rodriguez` all’interno del terzo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Dovresti avere un elemento `td` con il testo `Marcus` all’interno del terzo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Dovresti avere un elemento `td` con il testo `88` all’interno del terzo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
