---
id: 66a9cb26ec6bd41cf6c82bc5
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

L'ultima sezione da aggiungere alla tabella sarà l'elemento <dfn>table foot</dfn>, `tfoot`. L'elemento table foot verrà usato per mostrare la media dei voti di tutti gli studenti nella tabella.

Aggiungi un elemento `tfoot` alla tabella.

All'interno dell'elemento `tfoot`, aggiungi un elemento `tr`.

All'interno dell'elemento `tr`, aggiungi due elementi `td`.

Il primo elemento `td` dovrebbe contenere il testo `Average Grade`. Il secondo elemento `td` dovrebbe contenere il testo `78.8`.

# --hints--

Dovresti avere un tag di apertura `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Dovresti avere un tag di chiusura `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Dovresti avere un elemento `tr` dentro il tuo elemento `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Il tuo primo elemento `td` dovrebbe contenere il testo `Average Grade`. Assicurati che sia dentro l'elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Il tuo secondo elemento `td` dovrebbe contenere il testo `78.8`. Assicurati che sia dentro l'elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Il tuo elemento `tr` dovrebbe avere due elementi `td` dentro di esso.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
