---
id: 66a9bb8578459a175432e7d0
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Per la prima sezione della tabella, dovrai raggruppare il contenuto dell'intestazione che rappresenta le etichette delle colonne per il nome, il cognome e il voto finale dell'esame dello studente.

L'elemento <dfn>table head</dfn>, `thead`, serve a raggruppare il contenuto dell'intestazione in una tabella.

Ecco un esempio che usa l'elemento `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Sotto il tuo elemento `caption`, aggiungi un elemento table head.

# --hints--

Dovresti avere un tag di apertura `thead`.

```js
assert.match(code, /<thead>/i);
```

Dovresti avere un tag di chiusura `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Il tuo elemento `thead` dovrebbe essere all'interno del tuo elemento `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
