---
id: 6626b8dcf5057f896f948440
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora che l'insegnante ha tutte le informazioni di cui ha bisogno, vuole poter inviare un messaggio allo studente con i risultati.

Completa la funzione `studentMsg` con `totalScores` e `studentScore` come parametri. La funzione dovrebbe restituire una stringa che rappresenta un messaggio per lo studente.

Se lo studente ha superato il corso, la stringa dovrebbe seguire questo formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Se lo studente non ha superato il corso, la stringa dovrebbe seguire questo formato:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Sostituisci `average-goes-here` con la media dei punteggi totali.  
Sostituisci `grade-goes-here` con il voto dello studente.

**Suggerimenti**

- Usa la funzione `getAverage` per ottenere la media della classe.  
- Usa la funzione `getGrade` per ottenere il voto dello studente.  
- Usa la concatenazione di stringhe (`+`) per costruire il messaggio.  
- Fai attenzione alla punteggiatura e agli spazi nel messaggio.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` dovrebbe restituire il seguente messaggio: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` dovrebbe restituire il seguente messaggio: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

La tua funzione `studentMsg` dovrebbe restituire il messaggio corretto in base al punteggio dello studente e alla media della classe.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
