---
id: 662693f82c91a66be46c881b
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

Un insegnante ha finito di correggere i test dei suoi studenti e ha bisogno del tuo aiuto per calcolare la media dei voti della classe.

Completa la funzione `getAverage` che prende in input un array di voti e restituisce la media dei voti.

La media si calcola sommando tutti i voti e dividendo per il numero totale di voti.

```md
average = sum of all scores / total number of scores
```

Sono state fornite alcune chiamate alla funzione per permetterti di testare il tuo codice.

**Suggerimenti**

- Puoi usare un loop per iterare su `scores` array e sommare tutti i voti.
- Puoi usare la proprietà `length` per ottenere il numero totale di voti.

# --hints--

La tua funzione `getAverage` dovrebbe restituire un numero.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` dovrebbe restituire `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` dovrebbe restituire `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` dovrebbe restituire `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

La tua funzione `getAverage` dovrebbe restituire la media dei voti.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
