---
id: 6626b4c58c027d86478ff5eb
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

L'insegnante è davvero soddisfatto del programma che hai creato finora. Ma ora vuole un modo semplice per verificare se uno studente ha una votazione sufficiente. Una votazione sufficiente è qualsiasi cosa che non sia un `"F"`.

Completa la funzione `hasPassingGrade` che prende come parametro il punteggio di uno studente. La tua funzione dovrebbe restituire `true` se lo studente ha una votazione sufficiente e `false` se non ce l'ha.

**Consigli**

- Usa la funzione `getGrade` per ottenere il voto dello studente. Poi verifica se il voto è sufficiente o no.

# --hints--

La tua funzione `hasPassingGrade` dovrebbe restituire un valore booleano.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

La tua funzione `hasPassingGrade` dovrebbe restituire `true` se il voto è un `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

La tua funzione `hasPassingGrade` dovrebbe restituire `false` se il voto è un `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

La tua funzione `hasPassingGrade` dovrebbe restituire `true` per tutti i voti sufficienti.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
