---
id: 6626a060c4006f793e10cb33
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora l'insegnante ha bisogno del tuo aiuto per convertire il punteggio dello studente in un voto lettera.

Completa la funzione `getGrade` che prende un numero `score` come parametro. La tua funzione dovrebbe restituire una stringa che rappresenta un voto lettera basato sul punteggio.

Ecco i punteggi e i loro voti lettera corrispondenti:

| Intervallo di punteggio | Voto   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Suggerimenti**

- Ricorda che hai imparato le istruzioni condizionali (`if`, `else if` e `else`).
- Ricorda che hai imparato gli operatori di confronto (`>`, `<`, `>=`, `<=`, `===`).
- Ricorda che hai imparato l'operatore logico AND (`&&`).

# --hints--

La tua funzione `getGrade` dovrebbe restituire `"A++"` se il punteggio è `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

La tua funzione `getGrade` dovrebbe restituire `"A"` se il punteggio è `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

La tua funzione `getGrade` dovrebbe restituire `"B"` se il punteggio è compreso tra `80` e `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

La tua funzione `getGrade` dovrebbe restituire `"C"` se il punteggio è `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

La tua funzione `getGrade` dovrebbe restituire `"D"` se il punteggio è compreso tra `60` e `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

La tua funzione `getGrade` dovrebbe restituire `"F"` se il punteggio è `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

La tua funzione `getGrade` dovrebbe restituire `"F"` se il punteggio è compreso tra `0` e `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
