---
id: 68cabc534d2e741f33a3fe1d
title: Costruisci un tracker per le statistiche di un personaggio di gioco
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

In questo laboratorio costruirai un tracker per le statistiche di un personaggio di gioco. Il programma ti permetterà di creare un personaggio con attributi specifici, aggiornare quegli attributi e recuperare le statistiche attuali del personaggio.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Crea una classe chiamata `GameCharacter` che rappresenta un personaggio di gioco e gestisce le sue statistiche.
2. Quando viene istanziato, un nuovo oggetto `GameCharacter` dovrebbe avere i seguenti attributi:
    - `_name` impostato sulla stringa fornita al momento dell'istanza.
    - `_health` impostato su `100`.
    - `_mana` impostato su `50`.
    - `_level` impostato su `1`.
3. Crea una proprietà `name` per l'accesso in sola lettura al nome del personaggio.
4. Per la proprietà `health`:
     - Definisci un getter che restituisce la salute attuale.
     - Definisci un setter che impedisce che la salute venga impostata sotto `0` e limita la salute a `100`.
5. Per la proprietà `mana`:
     - Definisci un getter che restituisce il mana attuale.
     - Definisci un setter che impedisce che il mana venga impostato sotto `0` e limita il mana a `50`.
6. Crea un getter per `level` che restituisce il livello attuale del personaggio.
7. Definisci un metodo chiamato `level_up` che:
    - Aumenta il livello del personaggio di 1.
    - Reimposta la salute a `100` e il mana a `50` usando i rispettivi setter delle proprietà.
    - Stampa un messaggio nella forma `<name> leveled up to <level>!` (dove `<name>` e `<level>` devono essere sostituiti con il nome del personaggio e il nuovo livello, rispettivamente).
8. Definisci un metodo `__str__` che restituisce una stringa formattata che include:
    - Il nome del personaggio.
    - Il livello del personaggio.
    - La salute attuale del personaggio.
    - Il mana attuale del personaggio.
   Per esempio, un personaggio chiamato `Kratos`, subito dopo l'istanza, dovrebbe essere rappresentato come segue:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Esempio di utilizzo

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Dovresti avere una classe `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Dovresti avere un metodo `__init__` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

Il metodo `__init__` dovrebbe avere due parametri, con il primo che è `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Dovresti impostare `self._name` sulla stringa fornita al momento dell'istanza dentro il metodo `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Dovresti assegnare `100` a `self._health` dentro il metodo `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Dovresti assegnare `50` a `self._mana` dentro il metodo `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Dovresti assegnare `1` a `self._level` dentro il metodo `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Dovresti avere un metodo `name` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

Il metodo `name` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Dovresti restituire `self._name` nel metodo `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

Il metodo `name` dovrebbe avere un decoratore `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Dovresti avere un metodo `health` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

Il metodo `health` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Dovresti restituire `self._health` nel metodo `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

Il metodo `health` dovrebbe avere un decoratore `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Dovresti creare un metodo setter `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

Il metodo setter `health` dovrebbe avere un decoratore `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

Il metodo setter `health` dovrebbe avere due parametri, con il primo che è `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Dovresti impostare `self._health` a `0` se il valore dato al setter `health` è inferiore a `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Dovresti impostare `self._health` a `100` se il valore dato al setter `health` è superiore a `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Dovresti impostare `self._health` al valore dato al setter `health` quando il valore è tra `0` e `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

Dovresti avere un metodo `mana` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

Il metodo `mana` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Dovresti restituire `self._mana` nel metodo `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

Il metodo `mana` dovrebbe avere un decoratore `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Dovresti creare un metodo setter `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

Il metodo setter `mana` dovrebbe avere un decoratore `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

Il metodo setter `mana` dovrebbe avere due parametri, con il primo che è `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Dovresti impostare `self._mana` a `0` se il valore dato al setter `mana` è inferiore a `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Dovresti impostare `self._mana` a `50` se il valore dato al setter `mana` è superiore a `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Dovresti impostare `self._mana` al valore dato al setter `mana` se il valore è tra `0` e `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Dovresti avere un metodo `level` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

Il metodo `level` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Dovresti restituire `self._level` nel metodo `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

Il metodo `level` dovrebbe avere un decoratore `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Dovresti avere un metodo `level_up` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

Il metodo `level_up` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Dovresti aumentare `self._level` di `1` nel metodo `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Dovresti impostare `self.health` a `100` nel metodo `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Dovresti impostare `self.mana` a `50` nel metodo `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

Il metodo `level_up` dovrebbe stampare `<name> leveled up to <level>!` (dove `<name>` e `<level>` devono essere sostituiti con il nome del personaggio e il nuovo livello, rispettivamente).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Dovresti avere un metodo `__str__` nella classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

Il metodo `__str__` dovrebbe avere un solo parametro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Il tuo metodo `__str__` dovrebbe restituire una stringa con le statistiche del personaggio usando il formato fornito.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
