---
id: 6894c1b715b8d18a18ef3203
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, farai pratica con le classi e gli oggetti di Python costruendo un inventario di strumenti musicali.

Per definire una classe in Python, usa la parola chiave `class` seguita dal nome della classe e dai due punti.

```python
class ClassName:
    # class contents go here
```

Ricorda che in Python i nomi delle classi usano solitamente il PascalCase, cioè ogni parola inizia con una lettera maiuscola senza spazi tra le parole.

Inizia creando una classe chiamata `MusicalInstrument`. Aggiungi `pass` come corpo della classe per ora.

# --hints--

Devi creare una classe chiamata `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
