---
id: 6895e235935106f88281b648
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Puoi consultare gli attributi delle tue istanze usando la notazione a punto:

```py
instance_name.attribute_name
```

Questo ti permette di recuperare i valori che sono stati archiviati durante l'inizializzazione.

Prova a stampare gli attributi `name` e `instrument_type` di entrambe le tue istanze. Potrai vedere i valori che hai assegnato quando le hai create.

# --hints--

Dovresti stampare l'attributo `name` di `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

Dovresti stampare l'attributo `instrument_type` di `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

Dovresti stampare l'attributo `name` di `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

Dovresti stampare l'attributo `instrument_type` di `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
