---
id: 6895e356fd877efd6e3657ee
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--


All'interno della classe `MusicalInstrument`, definisci un altro metodo chiamato `play` con solo il parametro `self`. Il parametro `self` permette ai metodi di consultare gli attributi dell'istanza specifica.

# --hints--

Dovresti definire un metodo chiamato `play` con un parametro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
