---
id: 6895e5337b1914050b0af7bd
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora, chiama il metodo `get_fact()` sull'istanza `instrument_1` e stampa il risultato.

Dato che `get_fact()` restituisce una stringa invece di stamparla direttamente, devi usare `print()` per mostrare la stringa restituita.

# --hints--

Dovresti chiamare `get_fact()` su `instrument_1` e stampare il risultato.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
