---
id: 6895e0452b3370f26f365a46
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Nel metodo `__init__` devi assegnare i parametri agli attributi dell’istanza. Gli attributi dell’istanza sono variabili che appartengono a una specifica istanza di una classe.

All’interno del tuo metodo `__init__`, assegna il parametro `name` a `self.name` e il parametro `instrument_type` a `self.instrument_type`.

# --hints--

Dovresti assegnare il parametro `name` a `self.name` all’interno del metodo `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

Dovresti assegnare il parametro `instrument_type` a `self.instrument_type` all’interno del metodo `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
