---
id: 6895dffc64e97df0e62e9ef9
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

I metodi sono funzioni definite all'interno di una classe che possono eseguire azioni utilizzando i dati della classe. Ecco come creare i metodi:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

In Python, il metodo `__init__` è un metodo speciale noto come metodo inizializzatore, che viene chiamato quando crei una nuova istanza della classe.

All'interno della tua classe `MusicalInstrument`, aggiungi un metodo `__init__` con tre parametri: `self`, `name` e `instrument_type`.

Il parametro `self` è un modo standard per riferirsi all'istanza della classe ed è richiesto come primo parametro per tutti i metodi di istanza in una classe.

# --hints--

Dovresti definire un metodo `__init__` all'interno della tua classe `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Il tuo metodo `__init__` dovrebbe avere tre parametri: `self`, `name` e `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
