---
id: 6895e4940e39950288e406c2
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Successivamente, aggiungi un altro metodo alla tua classe chiamato `get_fact` che restituisce un fatto sull'elemento.

Questo metodo dovrebbe avere solo il parametro `self` e restituire una f-string che dice `The [instrument name] is part of the [instrument type] family of instruments.`.

A differenza del metodo `play` che stampa un messaggio direttamente, questo metodo dovrebbe restituire una stringa che può essere usata altrove nel tuo codice.

# --hints--

Dovresti definire un metodo chiamato `get_fact` con un parametro `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

Il tuo metodo `get_fact` dovrebbe restituire una f-string contenente il nome e il tipo dello strumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
