---
id: 6895e3f91c069e0068a42582
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

All'interno del tuo metodo `play`, usa `print()` per mostrare un messaggio riguardo al suonare lo strumento.

Usa una f-string per includere il nome dello strumento nel messaggio: `The [instrument name] is fun to play!`. `[instrument name]` dovrebbe essere sostituito con il nome reale dello strumento.

# --hints--

Il tuo metodo `play` dovrebbe usare una f-string per mostrare il messaggio `The [instrument name] is fun to play!`. Usa `self.name` per consultare il nome dello strumento.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
