---
id: 6895e197180c9af6988e1e19
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Crea una seconda istanza chiamata `instrument_2` per un `Trumpet` che è uno strumento `brass`.

Nota come puoi creare più istanze dallo stesso modello di classe. Ogni istanza archivia i propri valori degli attributi in modo indipendente dalle altre istanze.

# --hints--

Dovresti creare un'istanza chiamata `instrument_2` della classe `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

Dovresti passare `Trumpet` e `brass` come argomenti quando crei `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
