---
id: 6895e0bdb0348bf46952ab67
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora è il momento di creare istanze della tua classe `MusicalInstrument`.

Per creare un'istanza di una classe, chiami il nome della classe come una funzione, passando gli argomenti richiesti per il metodo `__init__` (eccetto `self`, che Python gestisce automaticamente).

```python
my_instance = ClassName(arg1, arg2)
```

Quando crei questa istanza, Python chiama il tuo metodo `__init__` e archivia i valori come attributi dell'istanza che appartengono solo a questa specifica istanza.

Fuori dalla definizione della classe, crea un'istanza chiamata `instrument_1` per un `Oboe` che è uno strumento `woodwind`.

# --hints--

Dovresti creare un'istanza della classe `MusicalInstrument` chiamata `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

Dovresti passare `Oboe` e `woodwind` come argomenti quando crei `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
