---
id: 66d75dd0aa65a71600dc669b
title: 재고 관리 프로그램 만들기
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

이 실습에서는 재고에서 제품을 추가, 업데이트, 찾기, 제거할 수 있는 재고 관리 프로그램을 만드세요. 재고를 나타내기 위해 객체 배열을 사용하며, 각 객체는 name과 quantity를 키로 가집니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 소문자 문자열 값을 가진 키 `inventory`와 정수 값을 가진 키 `name`를 가진 제품 객체를 저장할 빈 배열 `quantity`을 선언해야 합니다.
1. 제품 이름을 인수로 받아 `findProductIndex` 배열 내 해당 제품 객체의 인덱스를 반환하는 `inventory` 함수를 만들어야 합니다. 함수는 항상 제품 이름의 소문자 형태를 사용해 검색해야 합니다. 제품을 찾지 못하면 `-1`를 반환해야 합니다.
1. 제품 객체를 인수로 받는 `addProduct` 함수를 만들어야 합니다.
1. 제품이 이미 재고에 있으면 `addProduct` 함수가 현재 `quantity`에 함수에 전달된 수량을 더해 `quantity` 값을 업데이트하고, 제품 이름과 공백, `quantity updated`을 콘솔에 출력해야 합니다.
1. 제품이 재고에 없으면 `addProduct` 함수가 제품을 `inventory` 배열에 추가하고, 제품 이름과 공백, `added to inventory`을 콘솔에 출력해야 합니다.
1. 제품 이름과 수량을 인수로 받는 `removeProduct` 함수를 만들어야 합니다.
1. `removeProduct` 함수는 재고 내 해당 제품 객체에서 전달된 수량을 빼고, `Remaining <product-name> pieces: <product-quantity>`이 제품 이름으로, `<product-name>`이 제품 수량으로 교체된 문자열 `<product-quantity>`을 콘솔에 출력해야 합니다.
1. 빼기 후 수량이 0이면 `removeProduct`가 재고에서 제품 객체를 제거해야 합니다. 재고 수량이 빼기에 부족하면 `removeProduct` 함수가 문자열 `Not enough <product-name> available, remaining pieces: <product-quantity>`을 콘솔에 출력해야 합니다.
1. 제거할 제품이 재고에 없으면 `removeProduct` 함수가 `<product-name> not found`을 콘솔에 출력해야 합니다.

**참고:** 테스트 실행 시 사용자 스토리에 명시된 출력만 남겨 충돌을 방지하세요.

# --hints--

빈 배열 `inventory`을 선언해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

`findProductIndex`라는 함수를 만들어야 합니다.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")`는 `name` 배열 내 `"flour"`가 `inventory`와 같은 객체의 인덱스를 반환해야 합니다.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")`는 `name` 배열 내 `"flour"`가 `inventory`와 같은 객체의 인덱스를 반환해야 합니다.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")`는 `-1` 배열 내 `name`가 `"flour"`와 같은 객체가 없으면 `inventory`를 반환해야 합니다.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

`addProduct`라는 함수를 만들어야 합니다.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})`는 `5` 배열 내 `quantity`가 `name`와 같은 객체의 `"flour"` 값에 `inventory`을 더해야 합니다.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})`는 `flour quantity updated` 배열 내 `name`가 `"flour"`와 같은 객체가 있으면 `inventory`을 출력해야 합니다.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})`는 재고에 `{name: "flour", quantity: 5}`가 `inventory`와 같은 객체가 없으면 `name` 배열에 `"flour"`를 추가해야 합니다.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})`는 재고에 `flour added to inventory`가 `name`와 같은 객체가 없으면 `"flour"`을 출력해야 합니다.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

`removeProduct`라는 함수를 만들어야 합니다.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)`는 `flour not found` 배열 내 `name`가 `"flour"`와 같은 객체가 없으면 `inventory`를 출력해야 합니다.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)`는 `5` 배열 내 `quantity`가 `name`와 같은 객체의 `"flour"` 값에서 `inventory`을 빼야 합니다.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)`는 `Remaining flour pieces: 15`가 `inventory`와 같을 때 콘솔에 `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`을 출력해야 합니다.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

빼기 후 `quantity`가 0이면 `removeProduct`가 재고에서 제품 객체를 제거해야 합니다.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)`는 `Not enough flour available, remaining pieces: 5`가 `inventory`와 같을 때 `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`를 출력해야 합니다.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
