---
id: 61537c9eecea6a335db6da79
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Debajo de su elemento `.header`, cree un nuevo elemento `div` y asígnele una `clase` de `gallery`. Este `div` actuará como un contenedor para las imágenes de la galería.

Dentro de ese elemento `.gallery`, cree nueve elementos `img`.

# --hints--

Debe crear un elemento `div` en su elemento `body`.

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

Su nuevo elemento `div` debe tener una `clase` con el valor establecido como `gallery`.

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

Su nuevo elemento `div` debe ir después de su elemento `.header`.

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Su elemento `.gallery` debe tener nueve elementos `img`.

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
