---
id: 615f171d05def3218035dc85
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Elimina el margen de tu elemento body, a `font-family` dale el valor `sans-serif`, y agrega un `background-color` con el valor `#f5f6f7`.


# --hints--

Debe tener un selector `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

El selector de `body` debe tener una propiedad de `margin` establecida en 0 como valor.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginTop, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginBottom, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginLeft, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginRight, '0px');
```

El selector de `body` debe tener una propiedad de `font-family` establecida en `sans-serif` como valor.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily, 'sans-serif');
```

El selector de `body` debe tener una propiedad de `background-color` establecida en `#f5f6f7` como valor.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}

```
