---
id: 6494da0daf5df5197963671d
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Observa cómo el borde azul de la imagen se extiende más allá del borde rojo de la galería. Esto se debe a la forma en que los navegadores calculan el tamaño de los elementos contenedores.

La propiedad `box-sizing` se utiliza para establecer este comportamiento. Por defecto, se utiliza el modelo `content-box`. Con este modelo, cuando un elemento tiene un ancho específico, ese ancho se calcula solo en función del contenido del elemento. Los valores de padding y border se suman al ancho total, por lo que el elemento crece para acomodar estos valores.

Intente establecer `box-sizing` en `content-box` explícitamente, con el selector global `*`. En este punto, no verá ningún cambio, porque está usando el valor predeterminado.

# --hints--

Debe tener un selector `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

El selector `*` debe tener una propiedad `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

El selector `*` debe tener su propiedad `box-sizing` establecida en `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
