---
id: 68e6aff5325410118e76e7be
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على العمل مع ARIA من خلال بناء وحدة تحكم صوتية يمكن الوصول إليها.

لبدء العمل، أضف عنصر `h1` يحتوي على النص `Audio Controls`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يكون عنصر `h1` داخل عنصر `body` الخاص بك.

```js
assert.exists(document.querySelector('body > h1'));
```

يجب أن يحتوي عنصر `h1` على النص `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

--fcc-editable-region--
    
--fcc-editable-region--

  </body>
</html>
```
