---
id: 68eac57fadcc662ce917e75e
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

بعد ذلك، أضف عنصر `button` مع تعيين الخاصية `type` إلى القيمة `button`.

يجب أن يكون نص عنصر الـ `button` هو `Play`.

# --hints--

يجب أن يكون لديك عنصر `button`.

```js
assert.exists(document.querySelector('button'));
```

يجب أن يحتوي عنصر الـ `button` على خاصية `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

يجب أن تكون الخاصية `type` في عنصر الـ `button` مضبوطة على `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

يجب أن يحتوي عنصر الـ `button` على النص `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

--fcc-editable-region--
    
--fcc-editable-region--

  </body>
</html>
```
