---
id: 68eb7f2397da251a40e04aa2
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بعد ذلك، أضف عنصر `div` ليحتوي على جميع العناصر الخاصة بضبط مستوى الصوت.

داخل عنصر `div`، قم بتسكين عنصر `span` مع خاصية `id` مضبوطة على `volume-label`.
أيضًا، يجب أن يكون نص عنصر `span` هو `Volume`.

سيتم الرجوع إلى خاصية `id` لاحقًا بواسطة خصائص ARIA حتى تتمكن تقنيات المساعدة من التعرف على هذا النص كمسمى لضبط مستوى الصوت.

# --hints--

يجب أن يكون لديك عنصر `div` أسفل عنصر `button` الخاص بك.

```js
assert.exists(document.querySelector('button + div'));
```

يجب أن يكون لديك عنصر `span` داخل عنصر `div` الخاص بك.

```js
assert.exists(document.querySelector('button + div > span'));
```

يجب أن يحتوي عنصر `span` الخاص بك على خاصية `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

يجب أن تحتوي خاصية `id` في عنصر `span` الخاص بك على القيمة `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

يجب أن يحتوي عنصر `span` الخاص بك على النص `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
