---
id: 68ec0a48c9d6db3594d768c3
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

حان الوقت الآن لإضافة عنصر `span` آخر.

يجب أن يحتوي عنصر `span` هذا على خاصية `id` مضبوطة على `volume-description`.
أيضًا، يجب أن يكون نص هذا العنصر `span` هو `Adjust the sound level`.

مماثل لعنصر الـ `span` السابق، سيُشار لاحقًا إلى الـ `id` الخاص بهذا الـ `span` بواسطة خصائص ARIA حتى تتمكن تقنيات المساعدة من تحديد هذا النص كالوصف لعنصر التحكم في مستوى الصوت.

# --hints--

يجب أن يكون لديك عنصران `span` داخل عنصر `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

يجب أن يحتوي عنصر `span` الثاني على خاصية `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

يجب أن تحتوي خاصية `id` في عنصر `span` الثاني على القيمة `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

يجب أن يحتوي عنصر `span` الثاني على النص `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
