---
id: 68ed3a96e0ebcc36565b39e5
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

بعد ذلك، أضف عنصر `input` مع خاصية `type` مضبوطة على `range` لإنشاء شريط تمرير الصوت. اضبط خاصية `min` على `0`، وخاصية `max` على `100`، وخاصية `value` على `50` لتحديد مستوى الصوت الافتراضي.

# --hints--

يجب أن يكون لديك عنصر `input` أسفل عنصر `span` الثاني.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

يجب أن يحتوي عنصر `input` على خاصية `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

يجب أن تحتوي خاصية `type` في عنصر `input` على القيمة `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

يجب أن يحتوي عنصر `input` على خاصية `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

يجب أن تحتوي خاصية `min` في عنصر `input` على القيمة `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

يجب أن يحتوي عنصر `input` على خاصية `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

يجب أن تحتوي خاصية `max` في عنصر `input` على القيمة `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

يجب أن يحتوي عنصر `input` على خاصية `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

يجب أن تحتوي خاصية `value` في عنصر `input` على القيمة `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>

  </body>
</html>
```
