---
id: 68ed4204ab7edc66b6cb2383
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

كما تتذكر من الدروس السابقة، تُستخدم الخاصية `aria-labelledby` عندما يكون هناك نص موجود على الصفحة يمكن استخدامه كمسمى. في هذه الحالة، أضفت الخاصية `id` إلى عناصر `span` الخاصة بك والتي ستُستخدم هنا الآن.

أضف خاصية `aria-labelledby` إلى عنصر الـ `input` بقيمة `volume-label volume-description`.

# --hints--

يجب أن يحتوي عنصر الـ `input` على خاصية `aria-labelledby`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

يجب أن يحتوي عنصر الـ `input` على خاصية `aria-labelledby` بالقيمة `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
      <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
