---
id: 68ed488821845990c911a223
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

جهاز التحكم الصوتي القابل للوصول لديك يكاد يكتمل، لكن هناك شيء أخير يجب إضافته.

أضف عنصر `button` مع خاصية `type` مضبوطة على `button` ونص `Mute` أسفل عنصر `div`.

وبهذا التغيير، تكون قد أكملت هذه الورشة!

# --hints--

يجب أن يكون لديك عنصر `button` أسفل عنصر `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

يجب أن يحتوي عنصر `button` على خاصية `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

يجب أن تحتوي خاصية `type` في عنصر `button` على القيمة `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

يجب أن يحتوي عنصر `button` على النص `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
        aria-labelledby="volume-label volume-description">
    </div>

  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
        aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
