---
id: 61537a8054753e2f1f2a1574
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`head` 요소 안에 `meta` 태그를 추가하고 `name` 속성을 `viewport`으로, `content` 속성을 `width=device-width, initial-scale=1`로 설정하세요.

또한 `meta` 태그를 추가하고 `charset` 속성을 `UTF-8`으로 설정하세요.

# --hints--

`meta` 요소가 두 개 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

한 `meta` 요소는 `name`이 `viewport`로, `content`이 `width=device-width, initial-scale=1.0`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

다른 `meta` 요소는 `charset` 속성이 `UTF-8`로 설정되어야 합니다.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
