---
id: 6153a3952facd25a83fe8083
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

`::after` 가상 요소는 선택한 요소의 마지막 자식인 요소를 만듭니다. 마지막 이미지 뒤에 빈 요소를 추가할 때 사용할 수 있습니다. 이미지와 같은 `width` 을 주면 갤러리가 두 열 레이아웃일 때 마지막 이미지를 왼쪽으로 밀어냅니다. 현재는 플렉스 컨테이너에 `justify-content: center` 를 설정해서 가운데에 있습니다.

예시:

```CSS
.container::after {
  content: "";
  width: 860px;
}
```

`::after` 요소에 `.gallery` 가상 요소를 사용해 새 선택자를 만드세요. 빈 문자열 `content` 을 값으로 하는 `""` 속성과 `350px` 속성에 대해 `width` 을 추가하세요.

# --hints--

`.gallery::after` 선택자에는 `""` 속성에 빈 문자열 `content` 을 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.content === "\"\"");
```

`.gallery::after` 선택자에는 `350px` 속성에 `width` 을 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.width === "350px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
