---
id: 61539f32a206bd53ec116465
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

일부 이미지가 왜곡된 것을 확인하세요. 이는 이미지마다 종횡비가 다르기 때문입니다. 각 종횡비를 개별적으로 설정하는 대신, `object-fit` 속성을 사용해 이미지가 어떻게 동작할지 결정할 수 있습니다.

`.gallery img` 선택자에 `object-fit` 속성을 주고 값을 `cover`으로 설정하세요. 이렇게 하면 이미지가 종횡비를 유지하면서 `img` 컨테이너를 채우도록 지시하여 맞추기 위해 잘리게 됩니다.

# --hints--

`.gallery img` 선택자에 `object-fit` 속성을 주고 값을 `cover`으로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.objectFit === 'cover');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

--fcc-editable-region--
.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
--fcc-editable-region--
```
