---
id: 61537d86bdc3dd343688fceb
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

다음으로, 문서 내 순서에 따라 각 `img`에 `src` 속성을 지정하세요. 첫 번째 `img`는 `src` 값을 가진 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`을 가져야 합니다. 나머지는 모두 같지만 `1`를 문서 내 `img`의 번호로 바꾸세요.

# --hints--

모든 아홉 개의 `img` 요소는 `src` 속성을 가져야 합니다.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

첫 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

두 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

세 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

네 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

다섯 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

여섯 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

일곱 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

여덟 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

아홉 번째 `img` 요소는 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` 속성 값으로 `src`을 설정하세요.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
