---
id: 61537bb9b1a29430ac15ad38
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`head` 요소 안에 텍스트가 `title`로 설정된 `Photo Gallery` 요소를 추가하고, `link` 파일을 페이지에 추가하기 위해 `styles.css` 요소를 추가하세요.

# --hints--

`link` 요소에는 값이 `href`인 `styles.css` 속성이 있어야 합니다.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

코드에는 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

프로젝트의 제목은 `Photo Gallery`이어야 합니다.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

제목의 대소문자와 철자가 정확해야 한다는 점을 기억하세요.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
