---
id: 61537485c4f2a624f18d7794
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

표준 HTML 보일러플레이트로 시작하세요. `DOCTYPE` 선언, `html` 요소, `head` 요소, `body` 요소를 추가하세요.

여는 `lang` 태그에 `<html>` 속성을 추가하고 값으로 `en`을 설정하세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에는 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

`html`를 문서 유형으로 정의해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

`DOCTYPE` 선언은 타입 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 요소는 여는 태그가 있어야 합니다.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

여는 `<html>` 태그에는 `lang` 속성이 있고 값은 `en`이어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

여는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<head\s*>/i));
```

닫는 `head` 태그가 있어야 합니다.

```js
assert(code.match(/<\/head\s*>/i));
```

여는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<body\s*>/i));
```

닫는 `body` 태그가 있어야 합니다.

```js
assert(code.match(/<\/body\s*>/i));
```

`head`와 `body` 요소는 형제 관계여야 합니다.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 요소는 `html` 요소 안에 있어야 합니다.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
