---
id: 6494da0daf5df5197963671d
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

파란색 이미지 테두리가 빨간 갤러리 테두리를 넘어 확장되는 것을 확인하세요. 이는 브라우저가 컨테이너 요소의 크기를 계산하는 방식 때문입니다.

`box-sizing` 속성은 이 동작을 설정하는 데 사용됩니다. 기본적으로 `content-box` 모델이 사용됩니다. 이 모델에서는 요소에 특정 너비가 있으면 그 너비가 요소의 콘텐츠만을 기준으로 계산됩니다. 패딩과 테두리 값은 전체 너비에 추가되어 요소가 이 값을 수용하도록 커집니다.

글로벌 `box-sizing` 선택자를 사용해 `content-box`를 `*`으로 명시적으로 설정해 보세요. 이때 기본값을 사용하고 있으므로 아무 변화도 보이지 않을 것입니다.

# --hints--

`*` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` 선택자에는 `box-sizing` 속성이 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

`*` 선택자의 `box-sizing` 속성은 `content-box`으로 설정되어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
