---
id: 64dd136e4e8b0894f9c70d86
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`box-sizing` 방식을 알아냈으니, 추가한 CSS를 정리해서 변경 사항을 확인하세요.

`.gallery`와 `img` 선택자 및 그 안의 모든 규칙을 제거하세요.

# --hints--

`.gallery` 선택자가 없어야 합니다.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

`img` 선택자가 없어야 합니다.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('img'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
 * {
  box-sizing: border-box;
}

--fcc-editable-region--
.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
--fcc-editable-region--
```
