---
id: 6852ef95ccc27660275692cd
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Darasa lako la `Email` linahitaji kuhifadhi taarifa kuhusu kila barua pepe inapoanzishwa.

Ndani ya darasa lako la `Email`, ondoa neno la `pass` lililopo na unda njia ya `__init__` inayopokea kigezo cha `self`, `sender`, `receiver`, `subject`, na `body`.

# --hints--

Unapaswa kufafanua njia ya `__init__` ndani ya darasa la `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Njia yako ya `__init__` inapaswa kuwa na vigezo: `self`, `sender`, `receiver`, `subject`, na `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
--fcc-editable-region--
    pass
--fcc-editable-region--
```
