---
id: 6852f0c3fcedb964ab959921
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa jaribu kama sifa ya kipengele ya `read` imeongezwa kwa usahihi katika darasa la vitu la `Email`. Kwa kuwa tayari una kitu cha barua pepe kutoka hatua zilizopita, chapisha sifa ya kipengele ya `read` kuona kuwa sasa ni `False` kwa chaguo-msingi.

# --hints--

Unapaswa kuchapisha `email_obj.read` ili kujaribu sifa ya kusoma.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
