---
id: 6852ea0c3cadb54e897acdac
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda *Msimulizi wa Barua Pepe* unaosimuliza kutuma, kupokea, na kusimamia barua pepe kati ya watumizi tofauti. Utajifunza kuhusu madarasa ya vitu, vitu, na jinsi ya kupanga msimbo kwa njia ya kuzingatia vitu.

Anza kwa kuunda darasa la vitu linaloitwa `Email` kwa kutumia neno kuu la `class`.

# --hints--

Unapaswa kuunda darasa la vitu linaloitwa `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
