---
id: 6852f11d412c29665985a915
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa watumizi wanaweza kutuma barua pepe, wanahitaji njia ya kuhifadhi barua pepe wanazopokea. Kila mtumizi anapaswa kuwa na kisanduku cha barua pepe cha kukusanya barua pepe zao.

Ongeza sifa ya kipengele `inbox` kwenye njia `User` ya darasa la vitu `__init__` na uanze kuihifadhi kama orodha tupu `[]`. Hii itahifadhi barua pepe zote ambazo mtumizi anazipokea.

# --hints--

Unapaswa kuongeza `self.inbox = []` kwenye njia `User` ya darasa la vitu `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
