---
id: 6852f01116912861c7d1b0e6
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ndani ya njia ya `__init__`, unahitaji kuweka thamani za vigezo kwenye sifa za kipengele cha kitu ili kila barua pepe ihifadhi taarifa zake.

Ndani ya njia yako ya `__init__`, ondoa neno kuu la `pass` na weka thamani za vigezo kwenye `self.sender`, `self.receiver`, `self.subject`, na `self.body`.

# --hints--

Unapaswa kuweka thamani ya kigezo cha `sender` kwenye `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Unapaswa kuweka thamani ya kigezo cha `receiver` kwenye `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Unapaswa kuweka thamani ya kigezo cha `subject` kwenye `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Unapaswa kuweka thamani ya kigezo cha `body` kwenye `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
