---
id: 6853f4584f508e95cc7a9f4d
title: Hatua 30
challengeType: 20
dashedName: step-30
---

# --description--

Sasa utaunda njia ya kuorodhesha barua pepe zote zilizo kwenye kikasha cha barua. Njia hii inapaswa kushughulikia hali ambapo kikasha cha barua ni tupu na kuonyesha orodha yenye nambari za barua pepe pale ambapo barua zipo.

Ongeza njia iitwayo `list_emails` katika darasa la vitu la `Inbox` ambayo inachukua `self` kama kigezo. Kwanza, tengeneza usemi wa `if` ili kuangalia kama kikasha cha barua ni tupu kwa kupima orodha ya `emails`. Ikiwa ni tupu, chapisha ujumbe wa `Your inbox is empty.\n`. Ongeza usemi wa `return` kuondoka katika njia hiyo.

# --hints--

Unapaswa kufafanua njia iitwayo `list_emails` katika darasa la vitu la `Inbox` ambayo inachukua `self` kama kigezo.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Unapaswa kuangalia kama `self.emails` ni tupu kwa kutumia `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Unapaswa kuchapisha ujumbe wa `Your inbox is empty.\n` ikiwa hakuna barua pepe.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Unapaswa kuwa na usemi wa `return` mwishoni mwa usemi wa `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
