---
id: 6853f4ded9f84896ff182dd8
title: Hatua 31
challengeType: 20
dashedName: step-31
---

# --description--

Baada ya ukaguzi wa kisanduku cha barua tupu, chapisha ujumbe `\nYour Emails:`. Baadaye, rudia juu ya barua pepe zote katika kisanduku cha barua kwa kutumia mzunguko wa `for` wenye kuhesabu. Hii ni sintaksia ya kuhesabu kuanzia nambari fulani: `enumerate(iterable, start=start_number)`.

Tumia kuhesabu kuzipa nambari barua pepe kuanzia `1`. Tumia `i` (kielezo) na `email` (kielezo cha barua pepe) kama vigezo vya mzunguko.

Ndani ya mzunguko, chapisha f-string yenye kielezo cha mzunguko ikifuatiwa na `.`, kisha nafasi na uwakilishi wa mfuatano wa herufi wa kielezo cha barua pepe kwa muundo huu:

```md
i. string_representation
```

Ambapo `i` ni kielezo na `string_representation` ni uwakilishi wa kielezo cha barua pepe.

# --hints--

Unapaswa kuchapisha kichwa `\nYour Emails:` nje ya tamko la `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

Unapaswa kuwa na mzunguko wa `for` unaorudia juu ya `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Mzunguko wako wa `for` unapaswa kutumia `i` na `email` kurudia juu ya `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Ndani ya mzunguko, unapaswa kuchapisha muhtasari wa kila barua pepe ukiwa na kiambishi awali chenye nambari ikifuatiwa na mfuatano wa herufi uliopangwa kutoka kwa njia ya `__str__`. Usisahau kuongeza `.` baada ya nambari.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
