---
id: 6859400f8f7f872e50e1238c
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Darasa la `User` linahitaji njia ya kutuma barua pepe kwa watumizi wengine. Unapotuma barua pepe, utaunda kitu kipya cha `Email` na kuiongeza kwenye sanduku la mpokeaji.

Unda njia iitwayo `send_email` katika darasa lako la `User`. Njia hii inapaswa kuchukua vigezo vya `receiver`, `subject`, na `body`.

Ndani ya njia hiyo, unda kitu kipya cha `Email` kwa vigezo vifuatavyo na uweke thamani yake katika kigezo kinachoitwa `email`:

- `sender`: `self` (mtumizi wa sasa)
- `receiver`: `receiver` (mtumizi anayepokea barua pepe)
- `subject`: `subject` (mada ya barua pepe)
- `body`: `body` (maudhui ya barua pepe)

# --hints--

Unapaswa kufafanua njia ya `send_email` katika darasa la `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

Njia ya `send_email` inapaswa kuwa na vigezo: `self`, `receiver`, `subject`, na `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

Njia hiyo inapaswa kuunda kitu cha `Email` kinachoitwa `email` kwa vigezo sahihi.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
