---
id: 685c1e32254e5993a5c8823e
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Sasa umeanzisha darasa la `Email`, ni wakati wa kuunda watumizi wanaoweza kutuma na kupokea barua pepe. Kila mtumizi atakuwa na jina na ataweza kufanya shughuli za barua pepe.

Unda darasa la `User` lenye njia ya `__init__` inayopokea kigezo cha `self` na `name`. Ndani ya darasa hilo, weka thamani ya kigezo cha `name` kwa `self.name`.

# --hints--

Unapaswa kuunda darasa la `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

Darasa la `User` linapaswa kuwa na njia ya `__init__` inayopokea kigezo cha `self` na `name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Unapaswa kuweka thamani ya kigezo cha `name` kwa `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
