---
id: 68875554c6c32de788afc64e
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Watumizi wanahitaji njia ya kusoma barua pepe zao ipasavyo.

Ongeza njia ya `display_full_email` katika darasa la `Email` ambayo inachukua `self` kama kigezo pekee. Ndani ya njia hii, fanya wito wa njia ya `mark_as_read()` ili kuashiria barua pepe kama imesomwa mtu anapoiangalia.

Njia hii itatumika kuonyesha maudhui kamili ya barua pepe kwa muundo mzuri.

# --hints--

Unapaswa kufafanua njia ya `display_full_email` katika darasa la `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

Njia ya `display_full_email` inapaswa kuchukua `self` kama kigezo pekee.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

Njia hii inapaswa kufanya wito wa `self.mark_as_read()` kuashiria barua pepe kama imesomwa.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
