---
id: 688754d6d59578df1313cb3c
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Sasa ni wakati wa kuunganisha madarasa ya User na Inbox ili barua pepe ziweze kufikishwa kweli!

Watumizi wanahitaji kuwa na vitu halisi vya Inbox badala ya orodha rahisi, na njia ya `send_email` inapaswa kufikisha barua pepe kwenye kisanduku cha mpokeaji.

Sahihisha darasa la `User` `__init__` kwa njia ambayo kila mtumizi anapata kitu halisi cha Inbox badala ya orodha tupu iliyotengenezwa awali.

# --hints--

Unapaswa kusahihisha darasa la User kuanzisha `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
--fcc-editable-region--
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
