---
id: 68b96d6a290d52d5e57a7485
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Ondoa `email_obj` na kauli za chapisha zinazofuata.

# --hints--

Haupaswi kuwa na mstari wa kuunda `email_obj` katika msimbo wako.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

Haupaswi kuwa na `print(email_obj.sender)` katika msimbo wako.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

Haupaswi kuwa na `print(email_obj.subject)` katika msimbo wako.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

Haupaswi kuwa na `print(email_obj.read)` katika msimbo wako.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

Haupaswi kuwa na `email_obj.mark_as_read()` katika msimbo wako.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
