---
id: 68b930ee0e7b8fb81895e59a
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa utaongeza njia ya kuweka barua pepe kama imesomwa. Wakati mtu anapofungua barua pepe, unataka kubadilisha hali yake kutoka isijasomwa kuwa imesomwa. Njia hii itakuwa rahisi - inahitaji tu kuweka sifa ya kipengele `read` kuwa `True`.

Ongeza njia iitwayo `mark_as_read` katika darasa la vitu lako la `Email`. Njia hii inapaswa kuchukua `self` kama kigezo pekee na kuweka sifa ya kipengele `read` kuwa `True`.

# --hints--

Unapaswa kufafanua njia iitwayo `mark_as_read` katika darasa la vitu la `Email`.
inayopaswa kuchukua `self` kama kigezo pekee.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

Njia hiyo inapaswa kuweka `self.read` kuwa `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False
--fcc-editable-region--
    
--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
