---
id: 68b9350183ff05bd106c79e4
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Sasa, jaribu njia ya `mark_as_read` uliyounda katika darasa la vitu la `Email`. Tumia njia hiyo kwenye kipengele cha `Email` kilichopo `email_obj` kubadilisha hali. Baada ya kuweka barua pepe kama imesomwa, chapisha sifa ya kipengele ya `read` ya `email_obj` kuthibitisha mabadiliko.

# --hints--

Unapaswa kuita njia ya `mark_as_read` kwenye `email_obj` kuweka barua pepe kama imesomwa.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Unapaswa kuchapisha sifa ya kipengele ya `read` ya `email_obj` baada ya kuita njia ya `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
--fcc-editable-region--

--fcc-editable-region--
```
