---
id: 68be6af1f43cc44541debe41
title: Hatua 37
challengeType: 20
dashedName: step-37
---

# --description--

Unapaswa kushughulikia hali ambapo mtumizi anajaribu kufuta barua pepe kwa kutumia kielezo kisicho halali kama ulivyofanya katika njia ya `read_email`.

Ndani ya njia ya `delete_email`:

- Badilisha kigezo cha `index` kilichoanzia kwa 1 kuwa kielezo kinachoanzia kwa 0 kwa kutoa 1 na kuhifadhi katika kigezo kinachoitwa `actual_index`.
- Tengeneza taarifa ya if ili kuangalia kama `actual_index` ni kidogo kuliko 0 au kikubwa au sawa na urefu wa orodha ya `self.emails`. Ikiwa ni hivyo, chapisha ujumbe wa `Invalid email number.\n` na tumia `return` kutoka katika njia hiyo.

# --hints--


Unapaswa kutoa 1 kutoka kwa kigezo cha `index` na kuhifadhi katika kigezo kinachoitwa `actual_index`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

Unapaswa kutengeneza taarifa ya if ili kukagua kama `actual_index` ni kidogo kuliko sifuri au ni kubwa au sawa na urefu wa orodha ya `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_conditions()[0].is_equivalent("actual_index < 0 or actual_index >= len(self.emails)")`))
    }
})
```

Ndani ya taarifa ya `if`, unapaswa kuchapisha ujumbe wa `Invalid email number.\n` na kutumia `return` kutoka katika njia hiyo.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[0].is_equivalent("print('Invalid email number.\\\\n')")`))
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[1].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
            
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
--fcc-editable-region--
        
--fcc-editable-region--

```
