---
id: 68be6a4447e61c4395ef2359
title: Hatua 36
challengeType: 20
dashedName: step-36
---

# --description--

Sanduku lako la barua pia linahitaji njia ya kufuta barua pepe.

Neno kuu la Python `del` linaweza kutumika kufuta vitu kutoka kwenye orodha kwa kutumia kielezo chao.

Ongeza njia iitwayo `delete_email` katika darasa lako la `Inbox`. Kama `read_email`, inapaswa:

- Kuchukua kigezo cha `index`, kukagua kama sanduku la barua ni tupu, na kuchapisha ujumbe wa `Inbox is empty.\n` ikiwa ni hivyo. Tumia `return` kutoka njia hiyo.

# --hints--

Unapaswa kufafanua njia iitwayo `delete_email` katika darasa la `Inbox` ambayo inachukua `self` na `index` kama vigezo.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

Unapaswa kwanza kukagua kama sanduku la barua ni tupu kwa kutumia `if not self.emails:`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

Ikiwa sanduku la barua ni tupu, unapaswa kuchapisha ujumbe wa `Inbox is empty.\n` na kutumia `return` kutoka njia hiyo.

```js
({
    test: () => 
    {
    assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\\nreturn")`))
}
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
        
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
--fcc-editable-region--
    
--fcc-editable-region--
```
