---
id: 66a37f37ef5823a313de8c26
title: بناء مدار القمر
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

في هذا المختبر، ستُنشئ حركة بسيطة لمدار القمر حول الأرض باستخدام HTML وCSS. ستكون الأرض في مركز النظام، والقمر سيدور حولها.

**الهدف:** حقق قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تُحاذي كل العناصر إلى مركز الصفحة عن طريق تعيين الارتفاع إلى `100%` من إطار العرض وتعيين تخطيط flexbox في عنصر `body`.
   
2. يجب أن يكون لديك عنصر `div` يحمل الصنف `space`.

3. داخل عنصر `.space`، يجب أن يكون هناك عنصران `div` آخران يحملان الصنفين `earth` و`orbit` على التوالي وبالترتيب.

4. داخل عنصر `.orbit`، يجب أن يكون هناك عنصر `div` آخر يحمل الصنف `moon`.
   
5. يجب أن يكون عنصر `div` الذي يحمل `class` باسم `space` في مركز الصفحة وأن يكون له عرض وارتفاع بقيمة `200px`.

6. يجب أن يستخدم عنصر `div` الذي يحمل `class` باسم `space` تموضعًا من نوع `relative`.

7. يجب أن يستخدم عنصر `.earth` تموضعًا من نوع `absolute`. ضع مركز عنصر `.earth` في منتصف عنصر الأب على المحورين الرأسي (top) والأفقي (left). بعد ذلك، حرك عنصر `.earth` للخلف بمقدار نصف عرضه وارتفاعه ليتمركز داخل عنصر الأب `.space`.

8. يجب أن يكون لعنصر `.earth` عرض وارتفاع بقيمة `100px`.

9. يجب أن يكون لعنصر `.orbit` عرض وارتفاع بقيمة `200px`.

10. يجب أن يكون عنصر `.orbit` متموضعًا باستخدام `absolute`. يجب أن يكون الزاوية السفلية اليمنى له في مركز عنصر `.space` باستخدام خاصية `transform` التي تحركه بمقدار `-50%` على المحورين الرأسي والأفقي.

11. يجب أن يكون مسار مدار القمر حول الأرض دائريًا.

12. يجب أن يكون عنصر `.moon` متموضعًا باستخدام `absolute` وأن يكون له عرض وارتفاع بقيمة `30px`. يجب أن يتموضع عنصر `.moon` في أعلى عنصر `.orbit` ويكون متمركزًا أفقيًا.

13. يجب تعديل التموضع الأفقي لعنصر `.moon` بتحريكه إلى اليسار بمقدار نصف عرضه.

14. يجب أن يحتوي محدِّد `.earth` على لون خلفية و`border-radius` بقيمة `50%`.

15. يجب أن يحتوي محدِّد `.moon` على لون خلفية و`border-radius` بقيمة `50%`.

16. يجب تعريف حركة `@keyframes orbit` التي تدور عنصر `.orbit` 360 درجة حول مركزه. يجب تطبيق هذه الحركة على عنصر `.orbit` لمدة `5` ثوانٍ، مع دالة توقيت خطية وتكرار لا نهائي.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML وتطبيق CSS.

# --hints--

يجب أن تُحاذى محتويات عنصر `body` في مركز الصفحة عن طريق تعيين الارتفاع إلى `100%` من ارتفاع إطار العرض، و`display` إلى `flex`، و`justify-content` و`align-items` إلى `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

يجب أن يكون لديك عنصر رئيسي `div` يحمل الصنف `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

يجب أن يكون لعنصر `.space` عرض وارتفاع بقيمة `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

يجب أن يكون لعنصر `.space` خاصية `position` بقيمة `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

يجب أن يكون لديك `div` بصنف `earth` داخل `div` له `class` بقيمة `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

يجب أن يكون عنصر `.earth` هو الابن الأول لعنصر `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

يجب أن يكون لعنصر `.earth` عرض وارتفاع بقيمة `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

يجب أن يستخدم عنصر `.earth` تموضعًا من نوع `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

يجب أن يكون لعنصر `.earth` خاصيتي `top` و`left` بقيمة `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

يجب تحريك عنصر `.earth` للخلف وللأعلى بمقدار `50%` عن طريق تعيين خاصية `transform` إلى `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

يجب أن يكون لديك عنصر `div` يحمل الصنف `orbit` داخل عنصر `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

يجب أن يأتي عنصر `.orbit` بعد عنصر `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

يجب أن يكون لعنصر `.orbit` عرض وارتفاع بقيمة `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

يجب أن يستخدم عنصر `.orbit` تموضعًا من نوع `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

يجب أن يتموضع عنصر `.orbit` في مركز عنصر `.space` باستخدام خاصية `transform` التي تحركه بمقدار `-50%` على المحورين الرأسي والأفقي.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

يجب أن يكون لديك عنصر `div` يحمل الصنف `moon` داخل عنصر `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

يجب أن يكون لعنصر `.moon` عرض وارتفاع بقيمة `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

يجب أن يستخدم عنصر `.moon` تموضعًا من نوع `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

يجب أن يكون الحافة العلوية لعنصر `.moon` محاذية للحافة العلوية لكتلته الحاوية.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

يجب تحريك الحافة اليسرى لعنصر `.moon` إلى مركز الكتلة الحاوية (`50%` من الجانب الأيسر).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

يجب تعديل عنصر `.moon` أفقيًا بترجمته إلى `-50%` عبر المحور x.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

يجب أن يكون لمحدد `.earth` لون خلفية.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

يجب أن يحتوي محدِّد `.earth` على `border-radius` بقيمة `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

يجب أن يكون لمحدد `.moon` لون خلفية.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

يجب أن يحتوي محدِّد `.moon` على `border-radius` بقيمة `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

يجب أن يكون لديك قاعدة `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

يجب أن تُسمى قاعدة `@keyframes` الجديدة باسم `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

يجب أن تحتوي قاعدة `@keyframes orbit` على محدد `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

يجب أن تحتوي قاعدة `@keyframes orbit` على محدد `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

يجب أن يحتوي محدد `0%` على خاصية `transform` مضبوطة إلى `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

يجب أن يأتي محدد `100%` بعد محدد `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

يجب أن يحتوي محدد `100%` على خاصية `transform` مضبوطة إلى `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

يجب أن يحتوي محدد `.orbit` على خاصية `animation` مضبوطة إلى `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

