---
id: 69a92f0b9cc04bb0d5327bb5
title: تنفيذ مكدس
challengeType: 26
dashedName: implement-a-stack
---

# --description--

في هذا المختبر، ستنفذ هيكل بيانات المكدس باستخدام الدوال. المكدس هو هيكل بيانات من نوع آخر يدخل أولاً يخرج أخيراً (LIFO)، حيث تُضاف العناصر وتُزال من الأعلى.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تجتاز لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك دالة `initStack` تُرجع كائنًا يحتوي على خاصية `collection` مضبوطة على مصفوفة فارغة.
1. يجب أن يكون لديك دالة `push` تضيف عنصرًا إلى أعلى المكدس.
1. يجب أن يكون لديك دالة `pop` تزيل وتُرجع العنصر الأعلى من المكدس، أو `undefined` إذا لم يكن هناك عنصر.
1. يجب أن يكون لديك دالة `peek` تُرجع العنصر الأعلى من المكدس دون إزالته، أو `undefined` إذا لم يكن هناك عنصر.
1. يجب أن يكون لديك دالة `isEmpty` تُرجع `true` إذا كان المكدس لا يحتوي على عناصر، و`false` خلاف ذلك.
1. يجب أن يكون لديك دالة `clear` تزيل جميع العناصر من المكدس.

**ملاحظة:** تعتمد معظم الاختبارات على دالتي `initStack` و`push`. نفذهما أولاً، إذ تتطلب اختبارات `pop` و`peek` و`isEmpty` و`clear` إضافة عناصر إلى المكدس.

# --hints--

يجب أن يكون لديك دالة `push`.

```js
assert.isFunction(push);
```

يجب أن تضيف دالة `push` عنصرًا إلى أعلى المكدس.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

يجب أن يكون لديك دالة `pop`.

```js
assert.isFunction(pop);
```

يجب أن تزيل دالة `pop` العنصر الأعلى من المكدس وتُرجعه.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

يجب أن تُرجع دالة `pop` القيمة `undefined` إذا كان المكدس فارغًا.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

يجب أن تُرجع دالة `pop` القيم الخاطئة بشكل صحيح.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

يجب أن يكون لديك دالة `peek`.

```js
assert.isFunction(peek);
```

يجب أن تُرجع دالة `peek` العنصر الأعلى من المكدس دون إزالته.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

يجب أن تُرجع دالة `peek` القيمة `undefined` إذا كان المكدس فارغًا.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

يجب أن تُرجع دالة `peek` القيم الخاطئة بشكل صحيح.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

يجب أن يكون لديك دالة `isEmpty`.

```js
assert.isFunction(isEmpty);
```

يجب أن تُرجع دالة `isEmpty` القيمة `true` للمكدس الفارغ.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

يجب أن تُرجع دالة `isEmpty` القيمة `false` للمكدس غير الفارغ.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

يجب أن تُرجع دالة `isEmpty` القيمة `false` عندما يكون العنصر الأعلى قيمة خاطئة.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

يجب أن يكون لديك دالة `clear`.

```js
assert.isFunction(clear);
```

يجب أن تزيل دالة `clear` جميع العناصر من المكدس.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
