---
id: 68e97fe79367ad7b5dd6c9cd
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Katika hatua hii ongeza kipengele cha `h1` chenye maandishi ya `XYZ Bookstore`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Maandishi ya kipengele chako cha `h1` yanapaswa kuwa `XYZ Bookstore`. Hakikisha upitie makosa ya tahajia au herufi kubwa.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele cha `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
