---
id: 68ebdbacdd3fa474132cc975
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Chini ya kipengele cha `h1`, ongeza kipengele cha `p` chenye maandishi haya: `Browse our collection of amazing books!`.

# --hints--

Unapaswa kuwa na kipengele cha `p`.

```js
assert.exists(document.querySelector('p'));
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
