---
id: 68eca6364c5616d393389a30
title: Hatua 18
challengeType: 0
dashedName: step-18
---

# --description--

Kumbuka, kipengele cha HTML kinaonekana hivi:

```html
<element attribute="value">
    inner text
</element>
```

Chini ya kipengele chenye darasa la `card-container`, ongeza kipengele kipya cha `p` chenye maandishi haya:

```md
Review your selections and continue to checkout.
```

Chini ya kipengele cha `p`, tengeneza kipengele cha `div` chenye sifa ya `class` kuwekwa kuwa `btn-container`. Kontena hili litaunganisha vipengele vyako vya kitufe cha urambazaji.

# --hints--

Kipengele kipya cha `p` kinapaswa kuwa chini ya kipengele chenye darasa la `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Maandishi ya kipengele chako kipya cha `p` yanapaswa kuwa `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Unapaswa kuwa na kipengele cha `div` chenye sifa ya `class` kuwa `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Kipengele chako cha `div` chenye sifa ya `class` kuwa `btn-container` kinapaswa kuja baada ya kipengele kipya cha `p`.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
