---
id: 68eca6e5a6759ed4ea0034dc
title: Hatua 19
challengeType: 0
dashedName: step-19
---

# --description--

Ndani ya kipengele chenye darasa la `btn-container`, ongeza vipengele viwili vya `button`:

Kitufe cha kwanza:

- Id: `view-cart-btn`
- Darasa la vitu: `btn`
- Maandishi: `View Cart`

Kitufe cha pili:

- Id: `checkout-btn`
- Darasa la vitu: `btn`
- Maandishi: `Checkout`


Hongera! Umefanikiwa kujenga muundo wa ukurasa wa duka la vitabu ukitumia divs, madarasa ya vitu, na ids kupanga maudhui yako.

# --hints--

Unapaswa kuwa na vipengele viwili vya `button` vilivyopangwa ndani ya kipengele cha `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Kitufe chako cha kwanza ndani ya kipengele chenye darasa la `btn-container` kinapaswa kuwa na sifa ya kipengele ya `id` yenye thamani ya `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Kitufe chako cha kwanza ndani ya kipengele chenye darasa la `btn-container` kinapaswa kuwa na sifa ya kipengele ya `class` yenye thamani ya `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Maandishi ya kitufe chako cha kwanza yanapaswa kuwa `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Kitufe chako cha pili ndani ya kipengele chenye darasa la `btn-container` kinapaswa kuwa na sifa ya kipengele ya `id` yenye thamani ya `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Kitufe chako cha pili ndani ya kipengele chenye `class="btn-container"` kinapaswa kuwa na sifa ya kipengele ya `class` yenye thamani ya `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Maandishi ya kitufe chako cha pili yanapaswa kuwa `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
  <p>Review your selections and continue to checkout.</p>
  
  <div class="btn-container">
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
