---
id: 68ec6d9a315221aa31e54816
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Sifa ya kipengele `class` hutumika kutambua kipengele kimoja au zaidi kwa ajili ya mtindo. Tofauti na sifa ya kipengele `id`, majina ya darasa la vitu hayahitaji kuwa ya kipekee: vipengele vingi vinaweza kushiriki darasa la vitu moja.

Hapa kuna mfano:

```html
<p class="example">example paragraph</p>
```

Ongeza sifa ya kipengele `class` kwenye kipengele chako cha `div` na weka thamani yake kuwa `card-container`.

# --hints--

Kipengele chako cha `div` kinapaswa kuwa na sifa ya kipengele `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya kipengele `class` yenye thamani ya `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
