---
id: 68eca2a795b333ca5fee30a8
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Sasa tengeneza kadi ya kitabu ya pili. Ongeza kipengele kingine cha `div` chenye sifa ya `class` ikiwa na thamani ya `card`. Angalia jinsi unavyoweza kutumia tena jina la darasa la vitu kwa vipengele vingi ili kutumia mtindo unaolingana.

# --hints--

Unapaswa kuwa na vipengele viwili vya `div` vyenye sifa ya `class` yenye thamani ya `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

Kipengele cha pili chenye darasa la `card` kinapaswa kuwa ndani ya kipengele chenye darasa la `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Vipengele viwili vyenye darasa la `card` vinapaswa kuwa ndugu wa daraja moja.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
