---
id: 68ec9332a9b5b2b32487bd00
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sifa ya kipengele `id` huongeza kitambulisho cha kipekee kwa kipengele cha HTML. Kila `id` inapaswa kuwa ya kipekee ndani ya ukurasa na inapaswa kutumika mara moja tu.

Thamani za `id` haziwezi kuwa na nafasi na zinapaswa kuwa na herufi, nambari, alama za mstari chini, na alama za dash tu.

Hapa kuna mfano:

```html
<p id="para">example paragraph</p>
```

Ongeza sifa ya kipengele `id` kwa kipengele chako chenye darasa la `card` na weka thamani yake kuwa `sally-adventure-book`.

# --hints--

Kipengele chako chenye darasa la `card` kinapaswa kuwa na sifa ya kipengele `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Kipengele chako chenye darasa la `card` kinapaswa kuwa na `id` yenye thamani ya `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
