---
id: 68ec6cd5b7e8f5a8f7319e32
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Kipengele cha `div` kinatumika kama kontena kuunganisha vipengele vingine vya HTML. Utatumia zaidi kipengele cha `div` wakati unapotaka kuunganisha vipengele vya HTML ambavyo vitashiriki seti ya mitindo ya CSS.

Chini ya kipengele cha `p`, ongeza kipengele cha `div`. `div` hii itakuwa kontena la kadi zako za vitabu.

**Kumbuka**: Warsha hii haitekelezi CSS. Madarasa na vipengele vilivyopangwa ni muhimu kwa mitindo ya CSS, lakini katika warsha hii vinatumika tu kwa muundo na kuunganisha maudhui. Utajifunza jinsi mitindo inavyofanya kazi katika moduli inayofuata.

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert.exists(document.querySelector('div'));
```

Kipengele chako cha `div` kinapaswa kuwa chini ya kipengele chako cha `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
