---
id: 68ec98b38d83a3c28dd30efe
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Ndani ya kipengele cha kwanza chenye darasa la vitu la `card`, ongeza kipengele cha `h2` chenye maandishi `Sally's SciFi Adventure`.

# --hints--

Unapaswa kuwa na kipengele cha `h2` kilichopangwa ndani ya kipengele chenye darasa la vitu la `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

Maandishi ya kipengele chako cha `h2` yanapaswa kuwa `Sally's SciFi Adventure`. 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
