---
id: 68eca11e99e3c5c894ca9d69
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Kipengele cha `button` kinatumika kuunda vitufe vinavyoweza kubofyanwa kwenye ukurasa wa mtandao. Vitufe ni vipengele vinavyoshirikisha mtumizi ambavyo anaweza kubofya ili kufanya vitendo.

Ongeza kipengele cha `button` ndani ya kipengele chenye `class` ya `card`, mpe kitufe sifa ya `class` iliyowekwa kuwa `btn`, na maandishi ya `Buy Now`.

# --hints--

Unapaswa kuwa na kipengele cha `button` kilichopangwa ndani ya kipengele chenye darasa la `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Kipengele chako cha `button` kinapaswa kuwa na `class` yenye thamani ya `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Maandishi ya kipengele chako cha `button` yanapaswa kuwa `Buy Now`. 

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
