---
id: 68eca5412b5dd9d06b3d6404
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Chini ya kipengele cha `h2` katika kadi ya pili, ongeza kipengele cha `p` chenye maandishi haya:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

Unapaswa kuwa na kipengele cha `p` kilichopangwa ndani ya kipengele cha pili chenye darasa la `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

Maandishi ya kipengele cha `p` katika kadi yako ya pili yanapaswa kuwa `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

Kipengele cha `p` katika kadi yako ya pili kinapaswa kuwa chini ya kipengele chako cha `h2`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
