---
id: 68eca5c0065c82d256d29ca3
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Ndani ya kadi ya pili, ongeza kipengele cha `button` chenye sifa ya `class` iliyowekwa kuwa `btn` na maandishi ya `Buy Now`.

Vipengele vyote viwili vya `button` sasa vinashiriki `class` moja, ambayo ina maana vinaweza kupambwa kwa mtindo unaoendana pamoja.

# --hints--

Unapaswa kuwa na kipengele cha `button` kilichopangwa ndani ya kipengele cha pili chenye darasa la `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

Kipengele cha `button` cha kadi yako ya pili kinapaswa kuwa na sifa ya `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

Kipengele cha `button` cha kadi yako ya pili kinapaswa kuwa na `class` yenye thamani ya `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Maandishi ya kipengele cha `button` cha kadi yako ya pili yanapaswa kuwa `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
