---
id: 68ec99e478211dc578699944
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Chini ya kipengele cha `h2` katika kipengele cha kwanza chenye darasa la `card`, ongeza kipengele cha `p` chenye maandishi yafuatayo:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Unapaswa kuwa na kipengele cha `p` kilichopangwa ndani ya kipengele chenye darasa la `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
