---
id: 69134b6cb43aa5f254950a10
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utaunda ukurasa wa duka la vitabu kwa kuunda kadi za vitabu zinazoonyesha taarifa kuhusu vitabu tofauti. Utajifunza kupanga maudhui kwa kutumia vipengele vya `div`, madarasa, na vitambulisho (IDs).

Anza ukurasa wako wa duka la vitabu kwa kuunda muundo wa msingi wa HTML.

Ongeza tamko la `<!DOCTYPE html>` na vipengele vya `html` na `head`.

Ongeza sifa ya kipengele ya `lang` kwa kipengele cha `html` na uiweke kuwa `"en"`.

# --hints--

Msimbo wako unapaswa kujumuisha tamko la `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Msimbo wako unapaswa kujumuisha lebo za kufungua na kufunga za `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Kipengele chako cha `html` kinapaswa kuwa na sifa ya kipengele ya `lang` iliyowekwa kuwa `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Msimbo wako unapaswa kujumuisha lebo za kufungua na kufunga za `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Kipengele chako cha `head` kinapaswa kupangwa ndani ya kipengele cha `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
