---
id: 69134d14150d51f8066a3246
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa, boresha muundo wa hati yako ya HTML ili kuhakikisha ukurasa wako umeandikwa kwa usahihi.

Ndani ya kipengele cha `head`, ongeza kipengele cha `<meta charset="UTF-8">`.

Mwishowe, ongeza kipengele cha `body` chini ya sehemu ya `head`. Hapa ndipo maudhui yote yanayoonekana ya ukurasa yatakapo wekwa.

# --hints--

Msimbo wako unapaswa kujumuisha kipengele cha `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Msimbo wako unapaswa kujumuisha lebo za kufungua na kufunga za `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Kipengele chako cha `body` kinapaswa kuwa chini ya kipengele cha `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
