---
id: 691f85085787a8c4ca589fbb
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza kipengele cha `title` ndani ya kipengele cha `head`.

Weka kichwa cha ukurasa kuwa `XYZ Bookstore Page`.

# --hints--

Msimbo wako unapaswa kujumuisha lebo za kufungua na kufunga za `<title>`.

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Kichwa chako kinapaswa kuwa `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Kipengele chako cha `title` kinapaswa kupangwa ndani ya kipengele cha `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
