---
id: 6853e58d52ec64349a5760ca
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop implementerai l'algoritmo del percorso più breve. Scriverai una funzione Python che calcola il percorso più breve tra i nodi di un grafo e restituisce anche il percorso seguito.

Per esempio, dato un grafo in cui le città sono collegate da strade con distanze diverse, l'algoritmo troverà la rotta più breve da una città all'altra. Se vuoi viaggiare dalla città A alla città D, l'algoritmo potrebbe trovare che andare A ⇨ B ⇨ C ⇨ D (totale: 15 km) è più breve che andare direttamente A ⇨ D (20 km).

Per iniziare, definisci una variabile chiamata `INF` e assegnale il valore `float('inf')`, che rappresenta l'infinito positivo. Più avanti la userai per indicare una distanza infinita tra due nodi.

# --hints--

Dovresti definire la variabile `INF` all'inizio del tuo file.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Dovresti assegnare `float('inf')` alla variabile `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
