---
id: 68625ad34aa3f58d7f3b9b07
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Devi archiviare il numero di nodi nel grafo. Per questo ti serve una variabile che corrisponda alla lunghezza della matrice di adiacenza.

All'interno della funzione `shortest_path`, crea una variabile `n` e impostala sulla lunghezza della `matrix`.

# --hints--

Dovresti inizializzare `n` con la lunghezza della `matrix`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
