---
id: 686251fb49ebc489ba60ccb5
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Dovrai creare una lista 2D per rappresentare la matrice di adiacenza del grafo. Questa matrice servirà a rappresentare i pesi degli archi tra i nodi nel grafo.

Crea una variabile chiamata `adj_matrix` e assegnale una lista 2D che rappresenti il grafo con i seguenti pesi:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Dovresti creare una variabile chiamata `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` dovrebbe essere una lista 2D contenente i valori forniti.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
