---
id: 686253f1b051998ad4904e3e
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora creerai la funzione principale che accetta tre parametri: la matrice di adiacenza, il nodo di partenza e un nodo obiettivo opzionale.

Crea una funzione chiamata `shortest_path` che prende tre parametri: `matrix`, `start_node` e `target_node`. Assegna `None` come valore predefinito per `target_node`.

Il parametro `target_node` è impostato su `None` per impostazione predefinita, indicando che se non viene specificato un nodo obiettivo, la funzione dovrebbe calcolare i percorsi più brevi dal nodo di partenza a tutti gli altri nodi del grafo.

Aggiungi per ora una istruzione `pass` all'interno del corpo della funzione.

# --hints--

Dovresti creare una funzione chiamata `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

La funzione dovrebbe prendere tre parametri: `matrix`, `start_node` e `target_node` con `target_node` che ha un valore predefinito di `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
