---
id: 686b78bc13d518e1267448e4
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Oltre a tenere traccia delle distanze, devi anche tenere traccia dei percorsi effettivi seguiti per raggiungere ogni nodo.

Creerai una lista in cui ogni voce archivia il percorso seguito per raggiungere quel nodo. Inizialmente, il percorso di ogni nodo conterrà solo se stesso.

Le list comprehension offrono un modo conciso per creare liste. Per esempio:

```py
[x * 2 for x in range(3)]
```

Crea una variabile chiamata `paths` e inizializzala usando una list comprehension che crea una lista contenente `[node_no]` per ogni `node_no` in `range(n)`.

# --hints--

Dovresti avere una variabile chiamata `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Dovresti inizializzare `paths` usando una list comprehension.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

La tua list comprehension dovrebbe usare `node_no` per iterare su `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

La tua list comprehension dovrebbe valutare `[node_no]` per ogni `node_no` in `range(n)`.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
