---
id: 686b73007fc582db0132c361
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Devi tenere traccia della distanza più breve conosciuta dal nodo di partenza a ogni altro nodo nel grafo.

Per farlo, crea una variabile chiamata `distances` e inizializzala come una lista contenente un solo elemento: `INF`.  


# --hints--

Dovresti creare una variabile chiamata `distances`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

La lista `distances` dovrebbe essere inizializzata a `[INF]`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
