---
id: 68760464eb9c2e79cc913fac
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

In questo passaggio, aggiungerai un loop che verrà eseguito una volta per ogni nodo nel grafo. Questo loop permetterà all'algoritmo di aggiornare distanze e percorsi su più passaggi.

Crea un loop `for` che si esegue `n` volte. Usa `_` come variabile del loop dato che non devi usare il valore dell'iterazione.

All'interno del loop, devi prepararti a selezionare il prossimo nodo da processare creando due variabili:

- una per contenere la distanza più piccola trovata finora nell'iterazione corrente  
- e un'altra per archiviare l'indice del nodo che ha questa distanza più piccola.

Crea le variabili `min_distance` e `current` e impostale rispettivamente su `INF` e `-1`.

# --hints--

Dovresti avere un loop `for` all'interno della funzione `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Il tuo loop `for` dovrebbe iterare su `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Il tuo for loop dovrebbe usare `_` come variabile di iterazione.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Dovresti creare una variabile chiamata `min_distance` e impostarla su `INF` all'interno del for loop.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Dovresti creare una variabile chiamata `current` e impostarla su `-1` all'interno del for loop.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
