---
id: 6876089f7a500b7c782e5b83
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Devi decidere se il nodo corrente è una scelta migliore rispetto a quello che hai già trovato (se ce n’è uno). Per farlo, aggiungerai un’istruzione condizionale all’interno del loop.

La condizione deve fare due cose:

- Verificare se il nodo non è stato ancora visitato.
- Confrontare la distanza nota con l’attuale `min_distance`.

All’interno del loop interno `for`, aggiungi un’istruzione `if` che controlli se `node_no` non è stato visitato **e** se `distances[node_no]` è minore di `min_distance`.

Aggiungi `pass` come segnaposto all’interno del blocco condizionale.

# --hints--

Dovresti aggiungere un’istruzione `if` all’interno del tuo loop interno `for`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

La tua istruzione `if` dovrebbe verificare se il nodo è non visitato (`visited[node_no]` è falsy) e `distances[node_no]` è minore di `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
