---
id: 68777816e95d4936f129819f
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Prima di provare ad aggiornare la distanza verso un vicino, devi verificare che il vicino sia sia raggiungibile sia non visitato. Poi calcolerai quale sarebbe la distanza totale per raggiungere quel vicino passando dal nodo corrente.

All'interno del loop `for`, aggiungi un'istruzione `if` che controlla:

- Che `distance` non sia uguale a `INF` (dove questo significa che esiste un arco tra i nodi)
- Che il vicino `node_no` non sia ancora stato visitato

All'interno del blocco `if`, crea una variabile chiamata `new_distance` e assegnale la somma di `distances[current]` (la distanza più breve verso il nodo corrente) e `distance` (la distanza dal nodo corrente al vicino).

# --hints--

Dovresti avere un'istruzione `if` che verifica se `distance` non è uguale a `INF` e se il nodo vicino non è stato visitato.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Dovresti creare una variabile chiamata `new_distance` all'interno del blocco `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Dovresti assegnare a `new_distance` la somma di `distances[current]` e `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
