---
id: 68e4c0ff9edcc92771d7549a
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Al momento, la lista `distances` ha solo un elemento. Tuttavia, ti serve un valore di distanza per ogni nodo nel grafo.

In Python, puoi moltiplicare una lista per un intero per ripeterla. Per esempio, `[0] * 3` crea `[0, 0, 0]`.

Aggiorna la lista `distances` per contenere `n` copie di `INF` moltiplicando `[INF]` per `n`.

# --hints--

Dovresti aggiornare `distances` per contenere `n` copie di `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
