---
id: 689f7996426c1534fa5bea66
title: 创建一个高尔夫得分翻译器
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

在高尔夫比赛中，每个洞都有一个 `par`，指职业高尔夫球手为完成比赛而将球击入洞中的平均 `strokes`。 根据你实际 `strokes` 与 `par` 的差距，会有不同的昵称来描述成绩。

在本实验中，你将编写一个函数，它将 `par` 和 `strokes` 转换为相应的昵称。 

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应创建一个名为 `golfScore` 的函数。
1. `golfScore` 应包含两个数字参数，即球场的标准杆数和击出的杆数。
1. `golfScore` 应返回一个字符串。
1. 如果 `strokes` 是 `1`，`golfScore` 应返回 `"Hole-in-one!"`。
1. 如果 `strokes` 小于等于 `par` 减 `2`，`golfScore` 应返回 `"Eagle"`。
1. 如果 `strokes` 等于 `par` 减 `1`，`golfScore` 应返回 `"Birdie"`。
1. 如果 `strokes` 等于 `par`，`golfScore` 应返回 `"Par"`。
1. 如果 `strokes` 等于 `par` 加 `1`，`golfScore` 应返回 `"Bogey"`。
1. 如果 `strokes` 等于 `par` 加 `2`，`golfScore` 应返回 `"Double Bogey"`。
1. 如果 `strokes` 大于等于 `par` 加 `3`，`golfScore` 应返回 `"Go Home!"`。

# --hints--

你应创建一个名为 `golfScore` 的函数。

```js
assert.isFunction(golfScore);
```

`golfScore` 应包含两个参数。

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` 应返回一个字符串

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` 应返回字符串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` 应该返回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` 应返回字符串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` 应该返回字串 `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` 应返回字符串 `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` 应返回字符串 `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` 应该返回字串 `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` 应返回字符串 `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` 应该返回字串 `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` 应该返回字串 `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` 应返回字符串 `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` 应返回字符串 `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` 应该返回字串 `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` 应返回字符串 `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` 应该返回字串 `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` 应该返回字串 `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` 应返回字符串 `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` 应该返回字串 `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` 应返回字串 `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` 应返回字串 `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` 应返回字符串 `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
