---
id: 66c07238b01053abaf812065
title: بناء آلة حاسبة للمضروب
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

المضروب هو حاصل ضرب عدد صحيح وكل الأعداد الصحيحة التي تحته. على سبيل المثال، المضروب لـ `5` هو `5 * 4 * 3 * 2 * 1 = 120`. في هذا المختبر، ستنشئ آلة حاسبة للمضروب تأخذ رقمًا من المستخدم وتحسب المضروب لذلك الرقم.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب إعلان متغير `num` وتعيين رقم من اختيارك له. يجب أن يكون الرقم المُعيَّن بين `1` و `20` شاملة.
1. أنشئ دالة باسم `factorialCalculator` تأخذ رقمًا كمعلمة وتُرجع المضروب لذلك الرقم.
1. داخل الدالة، أعلن متغير `result` وعيّن له القيمة `1`. باستخدام حلقة، استعرض كل الأعداد من `1` إلى الرقم المُدخَل (شاملة) ولكل رقم، اضرب متغير `result` في الرقم الحالي وعيّن الناتج إلى متغير `result`. يمكنك اختيار استخدام حلقة `for` أو `while` أو `do...while` هنا.
1. يجب استدعاء دالة `factorialCalculator` مع `num` كمعلمة وتعيين النتيجة إلى المتغير `factorial`.
1. يجب تخزين الناتج النهائي بصيغة `Factorial of [num] is [factorial]` وتعيينه إلى المتغير `resultMsg`.
1. يجب إخراج قيمة `resultMsg` إلى وحدة التحكم.

# --hints--

يجب أن يكون لديك متغير `num`.

```js
assert.isDefined(num);
```

يجب تعيين قيمة للمتغير `num`

```js
assert.isNotNull(num);
```

يجب أن تكون قيمة `num` بين `1` و `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

يجب أن تكون قيمة `num` رقمًا.

```js
assert.isNumber(num);
```

يجب أن يكون لديك دالة باسم `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

يجب أن تأخذ دالة `factorialCalculator` رقمًا كمعلمة.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

يجب أن يُرجع المضروب لـ `5` القيمة `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

يجب أن يُرجع المضروب لـ `7` القيمة `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

يجب استدعاء دالة `factorialCalculator` مع المتغير `num` كمعلمة.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

يجب تعريف متغير `factorial` وتعيين نتيجة دالة `factorialCalculator` إليه.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

يجب أن تُنتج دالة `factorialCalculator` النتيجة الصحيحة.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

يجب أن يحتوي متغير `resultMsg` على سلسلة نصية بصيغة `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

يجب إخراج قيمة `resultMsg` إلى وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
