---
id: 6690e10ebe2181212abc9652
title: Visitenkarte gestalten
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `div` mit einem `class`-Attribut mit dem Wert `business-card` haben, das alle anderen Elemente enthält.  
1. Innerhalb des `.business-card`-Elements sollte ein `img`-Element mit einem `class`-Attribut mit dem Wert `profile-image` vorhanden sein. Sie können die Bildquelle auf `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` setzen, wenn Sie möchten. Es sollte ein `alt` mit einer aussagekräftigen Beschreibung vorhanden sein.  
1. Innerhalb des `.business-card`-Elements, nach dem `img`-Element, sollten drei `p`-Elemente mit einem `class`-Attribut vorhanden sein, die jeweils die Werte `full-name`, `designation` und `company` haben.  
1. Das erste `p`-Element sollte Ihren Namen enthalten.  
1. Das zweite `p`-Element sollte Ihre Berufsbezeichnung enthalten.  
1. Das dritte `p`-Element sollte den Namen Ihres Unternehmens enthalten.  
1. Unter dem dritten `hr`-Element sollte ein `p`-Element vorhanden sein.  
1. Nach dem `hr`-Element sollten zwei `p`-Elemente folgen. Im ersten `p`-Element sollte eine E-Mail-Adresse als Text stehen. Im zweiten `p`-Element sollte der Text eine Telefonnummer sein.  
1. Nach den zwei `p`-Elementen sollte ein `a`-Element mit der Klasse `portfolio-link`, dem Text `Portfolio` und einem Link zu einer gültigen URL vorhanden sein.  
1. Nach dem `hr`-Element sollte ein `a`-Element mit dem Portfolio-Link vorhanden sein.  
1. Sie sollten ein weiteres `div`-Element mit einem `class`-Attribut mit dem Wert `social-media` haben. Innerhalb dieses Elements sollte ein `h2`-Element mit dem Text `Connect with me` sein.  
1. Innerhalb der `.social-media`-Elemente sollten drei `a`-Elemente sein. Die `a`-Elemente sollten die Texte `Twitter`, `LinkedIn` und `GitHub` enthalten, jeweils mit Links zu gültigen Websites.  
1. Sie sollten das Stylesheet `styles.css` mit dem `link`-Tag im `head`-Abschnitt der HTML-Datei verknüpfen.  
1. Der Hintergrund Ihrer Seite sollte `rosybrown` sein. Die Schriftart insgesamt sollte `Arial` mit einem Fallback von `sans-serif` sein.  
1. Der `.business-card`-Selektor sollte Eigenschaften haben, um das Element `300px` breit zu machen und eine Hintergrundfarbe Ihrer Wahl zu setzen. Außerdem sollten Sie einen `padding` von `20px` rundherum und einen oberen `margin` von `100px` setzen. Der Text sollte zentriert ausgerichtet sein und die Schriftgröße `16px` betragen. Die linken und rechten Außenabstände sollten auf `auto` gesetzt sein.  
1. Der `.profile-image`-Selektor sollte eine `max-width`-Eigenschaft mit dem Wert `100%` haben.  
1. Alle `p`-Elemente sollten einen oberen und unteren `margin` von `5px` haben.  
1. Alle Links auf der Seite sollten nicht unterstrichen sein.

# --hints--

Sie sollten ein Haupt-`div` mit einem `class`-Attribut mit dem Wert `business-card` haben.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

Innerhalb des `.business-card`-Elements sollte ein `img`-Element mit einem `class`-Attribut mit dem Wert `profile-image` sein.

```js
assert.exists(document.querySelector('img.profile-image'));
```

Die Bildquelle sollte auf ein gültiges Bild gesetzt werden.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

Das `alt`-Attribut des Bildes sollte eine aussagekräftige Beschreibung enthalten.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

Innerhalb des `.business-card`-Elements, nach dem `img`-Element, sollte ein `p`-Element mit einem `class`-Attribut mit dem Wert `full-name` sein.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Nach dem `.full-name`-Element sollte ein `p`-Element mit einem `class`-Attribut mit dem Wert `designation` folgen.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Nach dem `.designation`-Element sollte ein `p`-Element mit einem `class`-Attribut mit dem Wert `company` sein.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Das erste `p`-Element, das mit der Klasse `full-name` versehen ist, sollte Ihren Namen enthalten.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Das zweite `p`-Element, das ein `class`-Attribut mit dem Wert `designation` hat, sollte Ihre Berufsbezeichnung enthalten.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Das dritte `p`-Element, das ein `class`-Attribut mit dem Wert `company` hat, sollte den Namen Ihres Unternehmens enthalten.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Nach dem dritten `p`-Element, das ein `class`-Attribut mit dem Wert `company` hat, sollte ein `hr`-Element folgen.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Nach dem ersten `hr`-Element sollte ein `p`-Element mit einer E-Mail-Adresse als Text folgen.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Nach dem E-Mail-`p`-Element sollte ein weiteres `p`-Element mit einer Telefonnummer als Text folgen.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Nach dem Telefonnummer-`p`-Element sollte ein `a`-Element mit der Klasse `portfolio-link` folgen.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

Das `.portfolio-link`-Element sollte den Text `Portfolio` enthalten.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Ihr `.portfolio-link`-Element sollte eine gültige `href` haben.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Nach dem `a`-Element mit Ihrem Portfolio sollte ein weiteres `hr`-Element als Trennlinie folgen.

```js
assert.exists(document.querySelector('a + hr'));
```

Nach dem zweiten `hr`-Element sollte ein `div`-Element mit einem `class`-Attribut mit dem Wert `social-media` folgen.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

Innerhalb des `.social-media`-Elements sollte ein `h2`-Element mit dem Text `Connect with me` sein.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

Innerhalb des `.social-media`-Elements sollten drei `a`-Elemente sein.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Das erste `a`-Element sollte den Text `Twitter` enthalten.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

Das `href` des ersten `a`-Elements sollte auf einen gültigen Link zeigen.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Das zweite `a`-Element sollte den Text `LinkedIn` enthalten.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

Das `href` des zweiten `a`-Elements sollte auf einen gültigen Link zeigen.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Das dritte `a`-Element sollte den Text `GitHub` enthalten.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

Das `href` des dritten `a`-Elements sollte auf einen gültigen Link zeigen.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Sie sollten die Hintergrundfarbe der Seite im `body`-Selektor auf `rosybrown` setzen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Sie sollten die Schriftart der Seite im `Arial`-Element auf `body` mit einem Fallback von `sans-serif` setzen.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Ihr `.business-card`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `300px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Ihr `.business-card`-Selektor sollte eine `background-color`-Eigenschaft mit einem gültigen Farbwert haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Ihr `.business-card`-Selektor sollte eine `padding`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Ihr `.business-card`-Selektor sollte einen oberen Außenabstand von `100px` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Der Text im `.business-card`-Element sollte zentriert ausgerichtet sein.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

Die Schriftgröße des Textes im `.business-card`-Element sollte `16px` betragen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

Die linken und rechten Außenabstände des `.business-card`-Elements sollten auf `auto` gesetzt sein.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Ihr `.profile-image`-Selektor sollte eine `max-width`-Eigenschaft mit dem Wert `100%` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

Ihre `p`-Elemente sollten einen oberen und unteren `margin` von `5px` haben.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

Die Links auf der Seite sollten nicht unterstrichen sein.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
