---
id: 6735a7370e0ae93a4577c689
title: Diseña una estructura de periódico
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

En este laboratorio, diseñarás una estructura de periódico usando CSS Grid, incluyendo conceptos como filas de la cuadrícula y columnas de la cuadrícula.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio. 

**Historias de usuario:**

1. Tu página debe contener un elemento `main` con la clase `newspaper-layout`.
1. La `.newspaper-layout` debe incluir un `header` con la clase `title` que contenga un `h1` para mostrar el nombre del periódico.
1. La `.newspaper-layout` debe incluir un `article` con la clase `feature-article` para el artículo principal de noticias.
1. El `.feature-article` debe incluir un elemento `h2` para el título del artículo seguido de un elemento `p` para el contenido del artículo.
1. Debes agregar tres elementos `article` más para artículos más pequeños, con las clases `small-article1`, `small-article2` y `small-article3`.
1. Cada uno de los artículos más pequeños debe incluir un elemento `h3` para el título del artículo seguido de un elemento `p` para el contenido del artículo.
1. La `.newspaper-layout` debe incluir un elemento `article` con la clase `secondary-article` para una sección adicional de noticias.
1. El `.secondary-article` debe incluir un elemento `h2` para el título del artículo seguido de un elemento `p` para el contenido del artículo.
1. La `.newspaper-layout` debe incluir un `figure` con la clase `cover-image` para mostrar una imagen que represente el contenido del periódico.
1. Los elementos con las clases `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` y `small-article3` deben tener una propiedad `grid-area` establecida con el mismo nombre de clase.
1. Your `.newspaper-layout` should use CSS Grid with a `grid-template-areas` property to define the arrangement of sections:
   - El `.title` debe abarcar toda la fila superior.
   - La `.feature-article` debe abarcar dos columnas en la segunda fila, con la `.cover-image` en la tercera columna.
   - La `.secondary-article` debe abarcar dos columnas en la tercera fila, con la `.cover-image` en la tercera columna.
   - Los tres artículos pequeños deberían aparecer en la cuarta fila.
1. La `.newspaper-layout` debe tener una propiedad `grid-template-columns` que divida el espacio en tres columnas iguales.
1. Debes establecer la propiedad `grid-template-rows` de `.newspaper-layout` en `auto` para la primera fila y dividir el espacio restante en partes iguales para las otras filas.
1. Debes agregar un espacio entre los elementos de la cuadrícula.
1. Asegúrate de que la imagen dentro de `.cover-image` quepa en el espacio designado estableciendo su `max-width` en `100%`.

**Nota:** Asegúrate de vincular tu hoja de estilos en tu HTML y aplica tu CSS.

# --hints--

Debes tener un elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Tu elemento `main` debe tener una `class` de `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Debes tener un elemento `header` con la `class` de `title` dentro de tu elemento `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Tu elemento `header` debe contener un elemento `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Tu elemento `.title` debe contener algo de texto.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Debes tener un elemento `article` para cada artículo con las clases `feature-article`, `small-article1`, `small-article2`, `small-article3` y `secondary-article`, respectivamente.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

El `.feature-article` debe incluir un elemento `h2` seguido de un elemento `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` y `.small-article3` deben incluir un elemento `h3` seguido de un elemento `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

El `.secondary-article` debe incluir un elemento `h2` seguido de un elemento `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Debes tener un elemento `figure` con la clase `cover-image` dentro de tu `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Tu `.cover-image` debe contener un elemento `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

La propiedad `display` del elemento `.newspaper-layout` debe establecerse en `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Tu elemento `.newspaper-layout` debe tener una propiedad `grid-template-areas` con `title` que abarque las tres columnas de la primera fila.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

La segunda fila de la plantilla de la cuadrícula debe contener un `feature-article` que abarque dos columnas y un `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

La tercera fila de la plantilla de la cuadrícula debe contener un `secondary-article` que abarque dos columnas y un `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

La cuarta fila de la estructura de la cuadrícula debe contener `small-article1`, `small-article2` y `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Tu selector `.newspaper-layout` debe usar la propiedad `grid-template-columns` para dividir el espacio en tres partes iguales. Recuerda que puedes usar la función `repeat` o definir manualmente cada unidad fraccional.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Tu selector `.newspaper-layout` debe tener una propiedad `grid-template-rows` configurada en `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Debes agregar un espacio entre los elementos de la cuadrícula.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` debe tener un `grid-area` de `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` debe tener un `grid-area` de `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` debe tener un `grid-area` de `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` debe tener un `grid-area` de `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` debe tener un `grid-area` de `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` debe tener un `grid-area` de `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` debe tener un `grid-area` de `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

El `img` dentro de `.cover-image` debe tener un `max-width` de `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
