---
id: 680a41ef5bfdc059cfd9d4b5
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

En conferencias anteriores, aprendiste sobre diferentes tipos de datos que puedes almacenar en una variable. Acabas de asignar un valor entero. Ahora necesitas asignar una cadena, que es una secuencia de caracteres encerrados entre comillas simples o dobles:

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

Declara otra variable llamada `alphabet` y asigna la cadena `abcdefghijklmnopqrstuvwxyz` a esta variable.

# --hints--

Debes declarar una variable llamada `alphabet`. Presta atención para colocar el nombre de la variable al principio de la línea.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Debes asignar la cadena `abcdefghijklmnopqrstuvwxyz` a la `alphabet` variable. Recuerda usar comillas simples o dobles para encerrar la cadena.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
