---
id: 6818eb3f98c0bca5af7c8d03
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Ahora tu c贸digo es reusable. Sin embargo, la funci贸n `caesar` imprime un mensaje en la terminal y devuelve por defecto `None`. Para probarlo, imprime `encrypted_text` al final de tu c贸digo.

# --hints--

Debes imprimir `encrypted_text` fuera del cuerpo de la funci贸n.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)

--fcc-editable-region--
```
