---
id: 6818fcb250f34e62cab32c39
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Podría haber notado que, aunque el mensaje está cifrado, los caracteres en mayúscula no se han modificado. Esto ocurre porque la tabla de traducción no incluye letras mayúsculas.

Para solucionarlo, necesitarás modificar las cadenas pasadas a su llamada `str.maketrans()` utilizando el método `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Actualice su llamada `str.maketrans()` concatenando a cada argumento la versión en mayúsculas del propio argumento.

# --hints--

Debe modificar los argumentos de su llamada `str.maketrans()` concatenando a cada argumento la versión en mayúsculas del propio argumento.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
