---
id: 6818e9c7c8dc5694e8f64fc3
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora, prueba la función `caesar` llamándola con la cadena `freeCodeCamp` y el número `3` como argumentos. Asignar la llamada de función a una variable llamada `encrypted_text`.

# --hints--

Debe tener una variable llamada `encrypted_text` en el ámbito global.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Debe asignar `caesar('freeCodeCamp', 3)` a la variable `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--


--fcc-editable-region--
```
