---
id: 6818ed1adc249bb91d373729
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

En conferencias anteriores, aprendiste que la declaración `return` se utiliza para devolver un valor de una función, para que puedas utilizarlo en otro lugar de tu código:

```py
def spam():
    return 'Spam!'
```

Elimina la llamada `print(encrypted_text)` de tu función. Luego, borra la declaración de la variable `encrypted_text` y devuelve directamente `text.translate(translation_table)` de tu función.

# --hints--

No debes tener `print(encrypted_text)` dentro de la función `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

No debes tener una variable `encrypted_text` dentro de la función `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Tu función `caesar` debe devolver `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)

```
