---
id: 6819be9055828d65c813eec5
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

La función `isinstance()` devuelve `True` si su primer argumento es una instancia del segundo argumento, y `False` en caso contrario:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Reemplace la condición actual de su declaración `if` con una llamada a `isinstance()`. Pase `shift` como el primer argumento, y `int` como el segundo argumento.

# --hints--

Su declaración `if` debe usar `isinstance(shift, int)` como su condición.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    if True:
        return 'Shift must be an integer value.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
