---
id: 681a44f4c3235f7d8f428545
title: Paso 25
challengeType: 20
dashedName: step-25
---

# --description--

Ahora vas a probar la función `decrypt`. Reemplaza el valor asignado a `encrypted_text` con la siguiente cadena, que representa un mensaje para descifrar: `Pbhentr vf sbhaq va hayvxryl cynprf.`.

Luego, declara una variable llamada `decrypted_text` y asígnale una llamada a `decrypt` con `encrypted_text` como su primer argumento y un desplazamiento de `13` como el segundo argumento.

Finalmente, muestra el `decrypted_text` en la terminal. Con eso, el cifrado César está completo.

# --hints--

Debes asignar `Pbhentr vf sbhaq va hayvxryl cynprf.` a `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

Debes tener una variable llamada `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

Debes asignar `decrypt(encrypted_text, 13)` a `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

Debes imprimir `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)
--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
