---
id: 681a2a6cf95e82ca30866364
title: Paso 23
challengeType: 20
dashedName: step-23
---

# --description--

Declara dos funciones llamadas `encrypt` y `decrypt`, ambas con parámetros `text` y `shift`.

Usarás `encrypt` para el proceso de encriptación, y `decrypt` para la desencriptación, etiquetando las dos acciones con un nombre descriptivo.

Devuelve una llamada `caesar` pasando `text` y `shift` desde ambas funciones nuevas, pero asegúrate de pasar también `False` como el tercer argumento para la función `decrypt`.

# --hints--

Debes tener una función llamada `encrypt` con dos parámetros, `text` y `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Tu función `encrypt` debe devolver `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Debes tener una función llamada `decrypt` con dos parámetros, `text` y `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Tu función `decrypt` debe devolver `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
