---
id: af7588ade1100bde429baf20
title: 欠けている文字検出器を作る
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `fearNotLetter` という名前の関数が必要です。
1. `fearNotLetter` 関数は1つの引数を受け取る必要があります。それはアルファベット順に並んだ文字の範囲を表す文字列で、1文字欠けている場合があります。
1. 関数は渡された文字範囲の中で欠けている文字を見つけて返す必要があります。
1. 範囲内のすべての文字が揃っている場合、関数は `undefined` を返す必要があります。

# --hints--

`fearNotLetter` 関数が必要です。

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` は文字列 `d` を返す必要があります。

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` は文字列 `i` を返す必要があります。

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` は文字列 `u` を返す必要があります。

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` は文字列 `e` を返す必要があります。

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` は `undefined` を返す必要があります。

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
