---
id: 61fd78621573aa5e8b512f5e
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

3番目の`tr`要素に、テキスト`th`を持つ`Savings Funds set aside for emergencies.`要素を追加してください。そのテキストのうち`Savings `を除く部分を、`span`属性が`class`に設定された`description`要素で囲んでください。

その後、次のテキストを持つ3つの`td`要素を順に追加してください：`$500`、`$650`、`$728`。3番目の`td`要素には、`class`属性を`current`に設定してください。

# --hints--

3番目の`tr`には`th`要素が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isNotNull(tableRow?.querySelector('th'));
```

`th`要素にはテキスト`Savings Funds set aside for emergencies.`が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Savings Funds set aside for emergencies.'
);
```

テキスト`Funds set aside for emergencies.`は`span`要素で囲んでください。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Funds set aside for emergencies.'
);
```

`span`要素には`class`属性を`description`に設定してください。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

3つの`td`要素が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

1番目の`td`要素にはテキスト`$500`が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$500');
```

2番目の`td`要素にはテキスト`$650`が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$650');
```

3番目の`td`要素にはテキスト`$728`が必要です。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$728');
```

3番目の`td`要素には`class`を`current`に設定してください。

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
