---
id: 61fd6343fb42013d99bcd7f3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`body` 要素の中に、`section` 要素の中に `main` 要素をネストしてください。

# --hints--

`body` 要素の中の最初の要素は `main` 要素であるべきです。

```js
assert.strictEqual(
  document.querySelector('body')?.children?.[0]?.localName,
  'main'
);
```

`section` 要素の中に `main` 要素があるはずです。

```js
assert.strictEqual(
  document.querySelector('main')?.children?.[0]?.localName,
  'section'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
