---
id: 61fd67a656743144844941cb
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

スクリーンリーダーはドキュメントの流れに基づいてHTML要素を読み上げます。最終的には貸借対照表に「Balance Sheet」という見出しと「AcmeWidgetCorp」というサブ見出しを付けたいです。しかし、この順序はスクリーンリーダーが読み上げる場合には適切ではありません。

既存の`span`に`class`属性を`flex`に設定し、その中に2つの`span`要素を追加してください。最初の要素にはテキスト`AcmeWidgetCorp`を、2番目の要素にはテキスト`Balance Sheet`を入れてください。ページ上のテキストの順序はCSSで逆にしますが、HTMLの順序はスクリーンリーダーにとってより意味のあるものになります。

# --hints--

既存の`span`要素には`class`属性を`flex`に設定してください。

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

既存の`span`要素の中に新しい`span`要素を2つ追加してください。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

新しい`span`要素には`class`属性を付けないでください。

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

最初の新しい`span`要素にはテキスト`AcmeWidgetCorp`を入れてください。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

2番目の新しい`span`要素にはテキスト`Balance Sheet`を入れてください。

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
