---
id: 61fd719788899952e67692b9
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

`thead` と `tbody` 要素は、テーブルのどの部分がヘッダーで、どの部分が主なデータやコンテンツを含むかを示すために使われます。

最初の `thead` の中で、`tbody` 要素の下に `table` と `caption` を追加してください。

# --hints--

最初の `table` 要素には `thead` 要素が必要です。

```js
const table = document.querySelectorAll('table')?.[0];
assert.isNotNull(table?.querySelector('thead'));
```

最初の `table` 要素には `tbody` 要素が必要です。

```js
const table = document.querySelectorAll('table')?.[0];
const tbody = table?.querySelector('tbody');
assert.isNotNull(tbody);
```

`thead`要素は`caption`要素のすぐ下に配置してください。

```js
assert.strictEqual(
  document.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

`tbody`要素は`thead`要素のすぐ下に配置してください。

```js
assert.strictEqual(
  document.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
