---
id: 61fd70336ebb3e4f62ee81ba
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

HTMLのテーブルは、テーブルの内容を説明するために`caption`要素を使います。`caption`要素は常に`table`の最初の子要素であるべきですが、`caption-side`のCSSプロパティで位置を調整できます。

最初の`caption`に`table`要素を追加し、テキストに`Assets`を設定してください。

# --hints--

最初の`table`要素には`caption`要素が必要です。

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

`caption`要素にはテキスト`Assets`を設定してください。

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
