---
id: 61fd6cc9475a784b7776233e
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

既存の `div` 要素の下に、`div` に設定された `class` を持つ新しい `table-wrap` 要素を追加してください。これがテーブルのコンテナになります。

その中に、3つの `table` 要素を追加してください。CSSを使ってこれらを1つのテーブルのようにスタイルしますが、別々のテーブルにすることでスクリーンリーダーが文書の流れを理解しやすくなります。

# --hints--

新しい `div` 要素を作成してください。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

新しい `div` 要素には、`class` を `table-wrap` に設定してください。

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

`.table-wrap` 要素は既存の `div` の後に配置してください。

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

`.table-wrap` 要素には3つの `table` 要素を含めてください。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
