---
id: 61fd990577d8227dd93fbeeb
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

4番目の`tr`要素に、テキスト`th`を持つ`Total Liabilities`要素を追加してください。テキスト`Liabilities`は、`span`属性が`class`に設定された`sr-only`要素で囲んでください。

その後、次のテキストを持つ3つの`td`要素を順に追加してください：`$750`、`$600`、`$475`。3番目の`td`要素には、`class`属性を`current`に設定してください。

# --hints--

4番目の`tr`には`th`要素が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.isNotNull(tableRow?.querySelector('th'));
```

`th`要素にはテキスト`Total Liabilities`が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Total Liabilities'
);
```

テキスト`Liabilities`は`span`要素で囲んでください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Liabilities'
);
```

あなたの `span` 要素は `class` 属性が `sr-only` に設定されているべきです。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('sr-only')
);
```

3つの `td` 要素を持つべきです。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

1番目の`td`要素にはテキスト`$750`が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$750');
```

2番目の`td`要素にはテキスト`$600`が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$600');
```

3番目の`td`要素にはテキスト`$475`が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$475');
```

3番目の `td` 要素は `class` が `current` に設定されているべきです。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[3];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
                <th>Expenses <span class="description">Annual anticipated expenses, such as payroll.</span></th>
                <td>$200</td>
                <td>$300</td>
                <td class="current">$400</td>
              </tr>
              <tr class="data">
                <th>Credit <span class="description">The outstanding balance on our credit card.</span></th>
                <td>$50</td>
                <td>$50</td>
                <td class="current">$75</td>
              </tr>
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
