---
id: 61fd933ba685de776a94997e
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

最初の `tr` の中に、テキスト `th` を持つ `Loans The outstanding balance on our startup loan.` 要素を追加してください。`Loans ` を除くそのテキストを、`span` が `class` に設定された `description` 要素で囲んでください。

その下に `td` 要素を3つ追加し、それぞれに順番に `$500`、`$250`、`$0` のテキストを設定してください。3つ目の `td` 要素には `class` が `current` に設定されている必要があります。

# --hints--

最初の `tr` には `th` 要素が必要です。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isNotNull(tableRow?.querySelector('th'));
```

`th` 要素にはテキスト `Loans The outstanding balance on our startup loan.` を設定してください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Loans The outstanding balance on our startup loan.'
);
```

テキスト `The outstanding balance on our startup loan.` は `span` 要素で囲む必要があります。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'The outstanding balance on our startup loan.'
);
```

`span` 要素には `class` 属性が `description` に設定されている必要があります。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

`td` 要素を3つ用意してください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

最初の `td` 要素にはテキスト `$500` を設定してください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$500');
```

2つ目の `td` 要素にはテキスト `$250` を設定してください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$250');
```

3つ目の `td` 要素にはテキスト `$0` を設定してください。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$0');
```

3つ目の `td` 要素には `class` が `current` に設定されている必要があります。

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
            --fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
