---
id: 61fd5a93fd62bb35968adeab
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin your project by giving your `head` element a `link` element for your stylesheet.

# --hints--

コードには `link` 要素が必要です。

```js
assert.match(code, /<link/);
```

Your `link` element should be within your `head` element.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

`link` 要素は値が `stylesheet` に設定されている `rel` 属性を持つ必要があります。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
