---
id: 61fd6b7c83dbf54a08cf0498
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`div` 要素の中に、3つの `span` 要素を追加してください。それぞれに `class` 属性を `year` に設定し、次のテキストを順番に追加してください：`2019`、`2020`、`2021`。

# --hints--

`div` 要素には3つの `span` 要素が必要です。

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

各 `span` 要素には `class` 属性を `year` に設定してください。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

最初の `span` にはテキスト `2019` を入れてください。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

2番目の `span` にはテキスト `2020` を入れてください。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

3番目の `span` にはテキスト `2021` を入れてください。

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
