---
id: 688141a8be6e6be03af945ea
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

사용자가 노트를 편집할 때, 편집된 버전을 추적하고 싶습니다.

`let`를 사용해 `currentContent`이라는 변수를 만들고 빈 문자열을 할당하세요.

*참고*: 실제 애플리케이션에서는 보통 노트를 데이터베이스에 저장합니다. 또는 로컬 저장소에 저장할 수도 있습니다. 하지만 데이터베이스와 로컬 저장소 작업은 이 워크숍의 범위를 벗어나며, 해당 개념은 나중에 배웁니다.

# --hints--

`let`를 사용해 `currentContent` 변수를 만들어야 합니다.

```js
assert.match(code, /let\s+currentContent/);
```

`currentContent` 변수는 문자열이어야 합니다.

```js
assert.isString(currentContent);
```

`currentContent` 변수는 빈 문자열로 초기화해야 합니다.

```js
assert.equal(currentContent, "");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

--fcc-editable-region--

--fcc-editable-region--
```
