---
id: 68814221fc4752e17b9696bd
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

사용자가 노트를 클릭하고 편집한 다음 노트 밖을 클릭하면, 새로 편집된 버전을 가져올 방법이 필요합니다. `blur` 이벤트를 사용하면 이 작업을 할 수 있습니다. 이 이벤트는 요소가 포커스를 잃을 때 발생합니다. 예시는 다음과 같습니다:

```js
element.addEventListener("blur", () => {
  // do something here
});
```

`addEventListener` 변수에 `noteEl` 메서드를 연결하세요. 첫 번째 인수는 `"blur"` 이벤트여야 하며, 두 번째 인수는 화살표 함수여야 합니다. 그 화살표 함수 본문 안에서 `const`를 사용해 `newContent`라는 새 변수를 만들고 `noteEl.innerHTML` 값을 할당하세요.

`innerHTML`는 노트의 간격과 형식을 유지하는 데 사용됩니다.

# --hints--

`addEventListener` 변수에 `noteEl` 메서드를 연결해야 합니다.

```js
assert.match(code, /noteEl\.addEventListener/);
```

`addEventListener` 메서드의 첫 번째 인수는 `"blur"` 이벤트여야 합니다.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

화살표 함수 본문 안에서 `const`를 사용해 `newContent`라는 새 변수를 만들고 `noteEl.innerHTML` 값을 할당해야 합니다.

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--
  
--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
