---
id: 688140eac18479dee90553e5
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

다른 `div` 요소를 사용하고 있으므로 화면 읽기 프로그램에 더 잘 접근할 수 있게 만들어야 합니다.

이전 수업에서 페이지에 라이브 영역을 만드는 데 사용되는 `aria-live` 속성에 대해 배웠습니다. 이 속성은 화면 읽기 프로그램 사용자가 라이브 영역의 내용이 변경될 때 수동으로 포커스를 맞추거나 상호작용하지 않아도 자동으로 알림을 받도록 합니다.

`#status` 요소에 `aria-live` 속성을 추가하고 그 값을 `"polite"`로 설정하세요.

`polite` 값은 화면 읽기 프로그램에 현재 알림이 끝나거나 사용자가 입력을 멈출 때까지 기다렸다가 업데이트를 알리도록 지시합니다.

# --hints--

`#status` 요소에 `aria-live` 속성이 있어야 합니다.

```js
assert.isTrue(document.getElementById("status").hasAttribute("aria-live"));
```

`#status` 요소에 `aria-live` 속성을 `"polite"`으로 설정해야 합니다.

```js
assert.equal(document.getElementById("status").getAttribute("aria-live"), "polite");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    --fcc-editable-region--
    <div id="status"></div>
    --fcc-editable-region--

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
