---
id: 688141e05e367de0e3ef7635
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

현재 `currentContent`는 빈 문자열 값으로 초기화되어 있습니다. 하지만 원하는 결과는 `currentContent`가 현재 노트 텍스트 값을 가지는 것입니다.

이전 수업에서 `DOMContentLoaded` 이벤트를 이렇게 다루는 방법을 배웠습니다:

```js
window.addEventListener("DOMContentLoaded", () => {
  // Do stuff here
});
```

HTML 문서의 모든 내용이 로드되고 파싱된 후에야 `noteEl`에 안전하게 접근할 수 있습니다.

먼저 `addEventListener` 객체에 `window` 메서드를 연결하세요. `addEventListener` 메서드의 첫 번째 인수는 `"DOMContentLoaded"` 이벤트여야 합니다. 두 번째 인수는 화살표 함수여야 합니다. 화살표 함수 본문 안에서 `noteEl.textContent` 값을 `currentContent`에 할당하세요.

# --hints--

`addEventListener` 객체에 `window` 메서드를 연결해야 합니다.

```js
assert.match(code, /window\.addEventListener/);
```

이벤트 리스너는 `"DOMContentLoaded"` 이벤트를 청취해야 합니다.

```js
assert.match(code, /window\.addEventListener\(['"]DOMContentLoaded['"]\,/);
```

화살표 함수 본문 안에서 `noteEl.textContent` 값을 `currentContent`에 할당하세요.

```js
const event = new Event("DOMContentLoaded");
window.dispatchEvent(event);

assert.equal(currentContent, noteEl.textContent);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--

--fcc-editable-region--
```
