---
id: 68813c1b28ec3fd41838daf6
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 `aria-live`와 `aria-label` 속성을 사용해 접근 가능한 노트 작성 앱을 만드는 연습을 합니다.

대부분의 HTML과 모든 CSS는 이미 제공되어 있습니다. 처음 몇 단계는 나머지 마크업을 추가하는 작업입니다.

미리보기를 보면 카드를 클릭해 편집하라는 문구가 있습니다. 하지만 실제로 클릭해도 작동하지 않습니다. `div` 요소를 편집 가능하게 만들려면 `contenteditable` 속성을 이렇게 사용하세요:

```html
<div contenteditable="true"></div>
```

열리는 `div` 태그에 `contenteditable` 속성을 추가하고 값을 `"true"`으로 설정하세요. 이제 요소를 클릭해 텍스트를 편집할 수 있습니다.

# --hints--

열리는 `div` 태그에 `contenteditable` 속성이 있어야 합니다.

```js
assert.isTrue(document.getElementById("note").hasAttribute("contenteditable"));
```

`div` 요소에 `contenteditable` 속성을 `"true"`으로 설정하세요.

```js
assert.equal(document.getElementById("note").getAttribute("contenteditable"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    --fcc-editable-region--
    <div id="note" class="note">
    --fcc-editable-region--
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```
