---
id: 695774002591bbc5f8cf3e53
title: 建置折扣計算器
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

在這個實驗中，你將撰寫一個函式，計算套用百分比折扣後的項目最終價格。

例如，如果一個項目的價格是 `50` 且套用 `20` 的折扣，折扣金額是 `10`，最終價格是 `40`。

**目標：** 完成以下使用者故事以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `apply_discount` 的函式（程式）。

2. `apply_discount` 函式（程式）應該精確地接受兩個參數：`price` 和 `discount`。

3. 如果 `price` 不是數字（`int` 或 `float`），該函式應該傳回字串 `The price should be a number`。

4. 如果 `discount` 不是數字（`int` 或 `float`），該函式應該傳回字串 `The discount should be a number`。

5. 如果 `price` 小於或相等於 `0`，函式應該傳回字串 `The price should be greater than 0`。

6. 如果 `discount` 小於 `0` 或大於 `100`，該函式（程式）應該傳回字串 `The discount should be between 0 and 100`。

7. 如果兩個輸入皆有效，該函式（程式）應該計算折扣為價格的百分比。

8. 該函式（程式）應在套用折扣後傳回最終價格。

# --hints--

你應該有一個名為 `apply_discount` 的函式（程式）。

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

你的 `apply_discount` 函式（程式）應該接受兩個參數：`price` 和 `discount`。

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

當 `apply_discount` 被呼叫且 `price`（第一個引數）不是數字（`int` 或 `float`）時，應該傳回 `The price should be a number`。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

當 `apply_discount` 被呼叫且 `discount`（第二個引數）不是數字（`int` 或 `float`）時，應該傳回 `The discount should be a number`。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

當 `apply_discount` 被呼叫且 `price` 小於或相等於 `0` 時，應該傳回 `The price should be greater than 0`。

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

當 `apply_discount` 被呼叫且 `discount` 小於 `0` 或大於 `100` 時，應該傳回 `The discount should be between 0 and 100`。

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` 應該傳回 `80`。

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` 應該傳回 `100`。

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` 應該傳回 `50`。

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

當以折扣 `100` 呼叫 `apply_discount` 時，應該傳回 `0`。

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` 應該傳回 `59.6`。

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
