---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Compilação de um contador em tempo real
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

Neste laboratório, você vai criar um contador de caracteres em tempo real usando JavaScript.O contador de caracteres exibirá o número de caracteres inseridos em um elemento `textarea`.O contador será atualizado em tempo real à medida que o usuário digita na `textarea`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `textarea` com o `id` de `text-input`.
2. Deve haver um elemento `p` com o `id` de `char-count`. Ele deve conter inicialmente o texto `Contagem de Caracteres: 0/50`. Este texto placeholder deve ser substituído pela contagem atual de caracteres no elemento `textarea`.
3. Quando o elemento `#text-input` contiver o texto `hello coder`, o elemento `#char-count` deverá conter o texto `"Character Count: 11/50"`
4. Quando o elemento `#text-input` contiver o texto `i am a programmer`, o elemento `#char-count` deverá conter o texto `"Character Count: 17/50"`
5. Quando o elemento `#text-input` contiver o texto `hello world`, o elemento `#char-count` deverá conter o texto `"Character Count: 11/50"`
6. Quando o elemento `#text-input` contiver o texto `I am learning a new language and it is called c++.` o elemento `#char-count` deve conter o texto `"Character Count: 50/50"`. **NOTA**: Embora o atributo `maxlength` alcance a funcionalidade correta, este laboratório requer uma solução apenas em JavaScript.  
7. O usuário não deve ser capaz de inserir mais de `50` caracteres no elemento `textarea`. Quando a contagem de caracteres atingir `50`, qualquer entrada extra deve ser automaticamente cortada e o texto `Contagem de Caracteres: 50/50` deve ser exibido em `vermelho`.

**Nota:** Certifique-se de vincular sua folha de estilos e o arquivo JavaScript no seu HTML.

# --hints--

Você deve ter um elemento `textarea` com o `id` de `text-input`. 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Você deve ter um elemento `p` com o `id` de `char-count`.

```js
const charCount = document.querySelector('p');
assert.equal(charCount.id, 'char-count');
```

O elemento `#char-count` deve conter inicialmente o texto `Contagem de Caracteres: 0/50`.

```js
const charCount = document.querySelector('p')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Quando o elemento `#text-input` contiver o texto `hello coder`, o elemento `#char-count` deverá conter o texto `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Quando a contagem de caracteres for `50`, o texto deve ser exibido na cor `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Se a contagem de caracteres for maior ou igual a `50`, a entrada extra deve ser cortada.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
