---
id: ac6993d51946422351508a41
title: تنفيذ خوارزمية تقصير سلسلة نصية
challengeType: 26
dashedName: truncate-a-string
---

# --description--

في هذا المختبر، ستتدرب على تقصير سلسلة نصية إلى طول معين.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `truncateString` تقبل معلمتين، الأولى سلسلة نصية، والثانية رقم.
1. إذا كان طول السلسلة النصية أكبر من الرقم المعطى، يجب تقصير السلسلة لتصبح بطول يساوي الرقم المعطى، ويُضاف `...` في نهاية السلسلة المختصرة.
1. إذا كان طول السلسلة النصية يساوي أو أقل من الرقم المعطى، يجب إرجاع السلسلة كما هي دون تغيير.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` يجب أن تُرجع السلسلة النصية `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` يجب أن تُرجع السلسلة النصية `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` يجب أن تُرجع السلسلة النصية `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` يجب أن تُرجع السلسلة النصية `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` يجب أن تُرجع السلسلة النصية `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` يجب أن تُرجع السلسلة النصية `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
