---
id: 66be24cb4144f955b6bcc550
title: Jenga ukurasa wa kadi za kuchezea
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kujenga ukurasa wa mtandao unaoonyesha angalau kadi tatu za kuchezea.
1. Unapaswa kuwa na kipengele cha `main` chenye kitambulisho cha `playing-cards`.
1. Ndani ya kipengele chako cha `#playing-cards`, unapaswa kuwa na angalau vipengele vitatu vya `div`, kila kikiwa na darasa la vitu la `card`.
1. Ndani ya kila kipengele cha `.card`, unapaswa kuwa na vipengele vitatu vya `div`, cha kwanza kikiwa na darasa la vitu la `left`, cha pili na darasa la vitu la `middle`, na cha tatu na darasa la vitu la `right`.
1. Kipengele chako cha `#playing-cards` kinapaswa kutumia flexbox kuweka watoto wake katikati kwa usawa, kuruhusu kuzunguka, na kuweka nafasi ya `20px` kati yao.
1. Kila kipengele chako cha `.card` kinapaswa kutumia flexbox kuoanisha watoto wake kwa kutumia `space-between`.
1. Kila kipengele chako cha `.left` kinapaswa kutumia sifa za kipengele cha flexbox kujipanga mwanzoni mwa kontena yake mzazi.
1. Kila kipengele chako cha `.middle` kinapaswa kutumia sifa za kipengele cha flexbox kujipanga katikati mwa kontena yake mzazi.
1. Kila kipengele chako cha `.right` kinapaswa kutumia sifa za kipengele cha flexbox kujipanga mwishoni mwa kontena yake mzazi.
1. Kila kipengele chako cha `.middle` kinapaswa kutumia flexbox kuonyesha watoto wake kwa safu.

Hapa kuna baadhi ya herufi unazoweza kunakili na kubandika kutumia kwenye kadi zako ikiwa unataka: `♠`, `♣`, `♥`, na `♦`.

# --hints--

Unapaswa kuwa na kipengele cha `main` chenye `id` ya `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

Unapaswa kuwa na vipengele vitatu vya angalau vya `div` vyenye darasa la vitu la `card` ndani ya kipengele chako cha `#playing-cards`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Kila kipengele chako cha `.card` kinapaswa kuwa na upana na urefu.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Kila kipengele chako cha `.card` kinapaswa kuwa na vipengele vitatu vya `div` kama watoto.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

Unapaswa kuwa na `div` yenye darasa la vitu la `left` ndani ya kila kipengele cha `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

Unapaswa kuwa na `div` yenye darasa la vitu la `middle` ndani ya kila kipengele cha `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

Unapaswa kuwa na `div` yenye darasa la vitu la `right` ndani ya kila kipengele cha `.card`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Unapaswa kuwa na kichaguzi cha `#playing-cards` kinachoweka `display` ya vipengele vyake kuwa `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Unapaswa kuwa na kichaguzi cha `#playing-cards` kinachoweka `justify-content` ya vipengele vyake kuwa `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Unapaswa kuwa na kichaguzi cha `#playing-cards` kinachoweka `flex-wrap` ya vipengele vyake kuwa `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Unapaswa kuwa na kichaguzi cha `#playing-cards` kinachoweka `gap` ya vipengele vyake kuwa `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Unapaswa kuwa na kichaguzi cha `.card` kinachoweka `display` ya vipengele vyake kuwa `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Unapaswa kuwa na kichaguzi cha `.card` kinachoweka `justify-content` ya vipengele vyake kuwa `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Unapaswa kuwa na kichaguzi cha `.left` kinachoweka `align-self` ya vipengele vyake kuwa `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Unapaswa kuwa na kichaguzi cha `.middle` kinachoweka `align-self` ya vipengele vyake kuwa `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Unapaswa kuwa na kichaguzi cha `.right` kinachoweka `align-self` ya vipengele vyake kuwa `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Unapaswa kuwa na kichaguzi cha `.middle` kinachoweka `display` ya vipengele vyake kuwa `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

Unapaswa kuwa na `flex-direction` iliyowekwa kuwa `column` katika darasa lako la `.middle`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
