---
id: 694bf6eb530e19a9c48c59f2
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora completa la oración concatenando la cadena ` years old` al final de `employee_info`. Recuerda incluir un espacio al principio de tu cadena.

Finalmente, imprime `employee_info`.

# --hints--

Debes concatenar `' years old'` al final de la variable `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Debes imprimir `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
