---
id: 694bf6eb530e19a9c48c59f5
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Cuando trabajas con cadenas, a menudo necesitarás extraer una porción específica de una cadena. Esto se llama **slicing**.

La sintaxis es `string[start:stop]`, donde:

* `start` es el índice donde comienza el slice (**inclusive**).
* `stop` es el índice donde termina la porción (**exclusivo**).

Por ejemplo, si `text = 'Python'`, entonces `text[0:2]` da `'Py'`.

Define `employee_code` como `'DEV-2026-JD-001'`. Después, crea una variable `department` y asígnale el segmento de `employee_code` desde el índice `0` hasta el `3`. Luego imprime `department` en la terminal.

# --hints--
Debes tener una variable llamada `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Debes asignar la cadena `DEV-2026-JD-001` a tu variable `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

Debes tener una variable llamada `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Debes cortar los primeros tres caracteres de `employee_code` y asignarlos a tu variable `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Debes imprimir `department` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
