---
id: 694bf6ea530e19a9c48c59f0
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Si concatenas dos cadenas como `'John' + 'Doe'`, el resultado es `'JohnDoe'` sin espacio. Para arreglar esto, necesitas concatenar una cadena que contenga un espacio (`' '`) entre ellas.


Actualiza tu variable `full_name` para que concatene `first_name`, un espacio y `last_name`.

# --hints--

Debes agregar un espacio (`' '`) entre `first_name` y `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
