---
id: 694bf6eb530e19a9c48c59f1
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Ahora, tu dirección parece incompleta. También quieres agregar el número del apartamento donde vive el empleado, así que debes modificar la variable.  

Cuando quieras agregar contenido al final de una variable de cadena existente, puedes usar el operador de **asignación aumentada** ,`+=`. Esto es más corto que escribir `var = var + 'new text'`. Por ejemplo:  

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Recuerda que las cadenas son inmutables, por lo tanto esta operación no cambia la cadena original. En cambio, crea una nueva cadena y la reasigna a la variable.

Usa el operador `+=` para agregar la cadena `, Apartment 4B` a tu variable `address`.

# --hints--

Debes usar el operador `+=` para agregar la cadena `, Apartment 4B` a la variable `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
