---
id: 694bf6eb530e19a9c48c59f7
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

También puedes usar números negativos para cortar desde el final de una cadena. Por ejemplo, `-1` se refiere al último carácter, `-2` se refiere al penúltimo carácter, y así sucesivamente.

Para obtener los últimos tres caracteres de una cadena, puedes usar `string[-3:]`. Observa cómo el parámetro stop se omite después de los dos puntos. Esto significa que el corte debe continuar hasta el límite de la cadena.

Crea una variable llamada `last_three`. Usa indexación negativa para extraer los últimos tres caracteres de `employee_code` (que representan el número de identificación). Finalmente, imprime `last_three` en la terminal.

Con eso, el taller de tarjetas de empleado está completo.

# --hints--

Debes tener una variable llamada `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Debes cortar los últimos tres caracteres de `employee_code` usando indexación negativa y asignarlos a tu variable `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Debes imprimir `last_three` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2024-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
