---
id: 694bf6eb530e19a9c48c59f4
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Concatenar muchas cadenas usando `+` y convertir números usando `str()` puede volverse desordenado y difícil de leer.

Python 3.6 introdujo las **f-strings** para resolver esto. Al agregar la letra `f` antes de la comilla de apertura, puedes poner variables y expresiones dentro de campos de reemplazo representados por llaves `{}`. Por ejemplo:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Crea una variable `employee_card` y asígnale un f-string que muestre `Employee:` seguido de un espacio y el valor de la variable `full_name`.

# --hints--

Debes tener una variable llamada `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Tu variable `employee_card` debe tener el valor `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
