---
id: 694bf3af05d10fd89dab05cc
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

En Python, puedes combinar cadenas usando el operador `+`. Esto se llama concatenación. Aquí tienes un ejemplo:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Crea una variable `full_name` concatenando `first_name` y `last_name`. Luego imprime `full_name`.

# --hints--

Debes tener una variable llamada `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Debes asignar una cadena formada por la concatenación de `first_name` y `last_name` a tu variable `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Debes imprimir la variable `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
