---
id: 694bf6eb530e19a9c48c59f6
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

Puedes cortar desde cualquier parte de la cadena, no solo desde el principio. Y es útil en muchos casos.

Crea una variable `year_code` y asígnale el segmento de `employee_code` desde el índice `4` hasta el `8`. Esto extraerá `2026`.

Luego crea una variable `initials` y asígnale el segmento de `employee_code` desde el índice `9` hasta el `11`. Esto extraerá `JD`.

Finalmente, imprime ambas variables en el terminal.

# --hints--

Debes tener una variable llamada `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Debes cortar `employee_code` desde el índice `4` hasta `8` y asignar el resultado a tu variable `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Debes tener una variable llamada `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Debes cortar `employee_code` desde el índice `9` hasta `11` y asignar el resultado a tu variable `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Debes imprimir `year_code` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Debes imprimir `initials` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
