---
id: 694bf6eb530e19a9c48c59f3
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora vas a usar la función `str()` una vez más. Al igual que con age, debes convertir cualquier variable numérica a una cadena antes de concatenarla con otro texto.

Crea una variable llamada `experience_years` y asígnale el entero `5`.

Luego, crea una variable `experience_info`. Asígnale una cadena formada por la concatenación de `'Experience: '`, la variable `experience_years` (convertida a cadena) y `' years'`. Imprime el resultado en el terminal.

# --hints--

Debes tener una variable llamada `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Debes asignar el entero `5` a tu variable `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Debes tener una variable llamada `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Debes asignar una cadena formada por la concatenación de `'Experience: '`, `str(experience_years)` y `' years'` a tu variable `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Debes imprimir `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
