---
id: 695c9f21a8b1d53cf1b97841
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora crea una variable llamada `employee_age` y asígnale el entero `28`. 


# --hints--

Debes tener una variable llamada `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Debes asignar el entero `28` a tu variable `employee_age`. No rodees el texto con comillas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
