---
id: 695c9c7c61d446f1451e0ae7
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

A continuación, crea una variable `address` para almacenar la dirección del empleado. Asígnale la cadena `123 Main Street` y finalmente imprime `address`.

# --hints--

Debes tener una variable llamada `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Debes asignar la cadena `123 Main Street` a tu variable `address`. Recuerda encerrar el texto entre comillas simples o dobles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Debes imprimir la variable `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

