---
id: 695c96f82af4fb78e59fa9ed
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

Las cadenas son secuencias de caracteres usadas para almacenar datos de texto. Como recordarás de lecciones anteriores, puedes crear una cadena encerrando texto dentro de comillas simples (`'`) o dobles (`"`). Por ejemplo:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Crea una variable `first_name` que almacene la cadena `John` y una variable `last_name` que almacene la cadena `Doe`. Luego imprime `first_name` y `last_name`.

# --hints--

Debes tener una variable llamada `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Debes asignar la cadena `John` a tu variable `first_name`. Recuerda encerrar el texto entre comillas simples o dobles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Debes tener una variable llamada `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Debes asignar la cadena `Doe` a tu variable `last_name`. Recuerda encerrar el texto entre comillas simples o dobles.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Debes imprimir tu variable `first_name` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Debes imprimir tu variable `last_name` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
