---
id: 695df6480de23758af5a29cd
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Ahora, quieres crear una cadena que muestre la edad del empleado.

Comienza creando una variable `employee_info` y asígnale el resultado de concatenar:

- la variable `full_name`.
- una cadena que consiste en los caracteres `is` precedidos y seguidos por un espacio.

# --hints--

Debes tener una variable llamada `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Debes asignar `full_name + ' is '` a tu variable `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
