---
id: 695ea1e55c0e619e597fc5c9
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Como puedes ver, Python generó un `TypeError: can only concatenate str (not "int") to str`. Esto sucede porque Python no permite concatenar texto (cadenas) y números (enteros) directamente.

Para solucionar esto, debes convertir el número a una cadena primero usando la función `str()`, que devuelve la versión en cadena de un objeto:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Actualiza tu asignación de `employee_info` para convertir `employee_age` a una cadena usando `str(employee_age)`.

# --hints--

Debes modificar la asignación de tu variable `employee_info` convirtiendo `employee_age` en una cadena con `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
