---
id: 695fd3dde03273875e5d6709
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Ahora intenta concatenar `employee_age` al final de tu cadena `employee_info`.

Una vez que lo hayas hecho, verás un `TypeError` en el terminal. En el siguiente paso, trabajarás en corregirlo.

# --hints--

Debes concatenar `employee_age` al final de `employee_info`.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
