---
id: 69727ea5c9606bedfe36d3d2
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Ahora es momento de agregar los detalles finales a la tarjeta.

Crea una variable llamada `position` con el valor de la cadena `Data Analyst` y una variable llamada `salary` con el valor del entero `75000`.

Luego, actualiza tu f-string `employee_card` para incluir el puesto y el salario. Debe seguir este formato exacto: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Reemplaza los marcadores con las variables correspondientes.

Finalmente, imprime `employee_card` para ver el resultado.

# --hints--

Debes tener una variable llamada `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Debes asignar la cadena `Data Analyst` a tu variable `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Debes tener una variable llamada `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Debes asignar el entero `75000` a tu variable `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Debes actualizar `employee_card` para usar la f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Debes imprimir `employee_card` en el terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
