---
id: 6980d3ecfe800d8a54130f30
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Actualmente, `employee_card` solo muestra el nombre del empleado. Ahora vas a agregar más información a esta estructura. 

Actualiza la asignación de `employee_card` para incluir la edad del empleado. La cadena final debe verse así: `Employee: [name] | Age: [age]` con `[name]` reemplazado por el nombre del empleado, y `[age]` reemplazado por la edad del empleado.

# --hints--

Debes actualizar `employee_card` para usar la f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
