---
id: 66629f407d679d3105e8317f
title: ジャンケンゲームを作ろう
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
あなたのゲームはコンピューターと対戦します。`"rock"`、`"paper"`、または`"scissors"`のいずれかをランダムに返す関数を書いてください。

ゲームのフロントエンド部分は気にする必要はありません。ゲームのロジックだけを書きます。`script.js`を開いて指示に従い始めてください。

**ユーザーストーリー：**

1. `getComputerChoice`という名前の関数を持っている必要があります。

1. `getComputerChoice`関数は`"rock"`、`"paper"`、または`"scissors"`のいずれかをランダムに返す必要があります。

**ヒント：** [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random)メソッドは0以上1未満のランダムな数を返します。これを使って複数の選択肢のうちの1つを条件付きで返す方法を考えてみてください。

ゲームは人間のプレイヤーが遊びます。ユーザーの選択を受け取り、それを返す関数を書いてください。

1. `getHumanChoice`という名前の関数を作成してください。

1. `getHumanChoice`がユーザーの入力に応じて有効な選択肢の1つを返すようにコードを書いてください。

**ヒント：** [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt)メソッドを使ってユーザーの入力を取得してください。

ゲームはプレイヤーのスコアを記録します。プレイヤーのスコアを記録する変数を書いてください。

1. グローバルスコープに`humanScore`と`computerScore`という2つの新しい変数を作成してください。

1. それらの変数を`0`の値で初期化してください。

ゲームはラウンドごとに進行します。人間とコンピューターの選択を引数に取り、1ラウンドをプレイし、ラウンドの勝者のスコアを増やし、勝者の発表をログに出す関数を書いてください。

1. `playRound`という新しい関数を作成してください。

1. `playRound`のために2つのパラメータを定義してください。1つ目のパラメータは`humanChoice`、2つ目は`computerChoice`です。これら2つのパラメータを使って人間とコンピューターの選択を引数として受け取ります。

1. 関数の`humanChoice`パラメータを大文字小文字を区別しないようにして、プレイヤーが`"rock"`、`"ROCK"`、`"RocK"`、またはその他のバリエーションを入力できるようにしてください。

1. `playRound`関数のコードを書いて、ラウンドの勝者を表す文字列の値を返してください。

- 引き分けの場合は`"It's a tie!"`を返すべきです。

- プレイヤーが勝った場合は`"You win! [player choice] beats [computer choice]"`を返すべきです。

- コンピューターが勝った場合は`"You lose! [computer choice] beats [player choice]"`を返すべきです。

1. ラウンドの勝者に応じて`humanScore`または`computerScore`変数を増やしてください。

ゲームは3ラウンド行います。`playGame`という関数を書いて、`playRound`を呼び出して3ラウンドをプレイし、スコアを記録し、最後に勝者を宣言してください。

1. `playGame`という新しい関数を作成してください。

1. 3ラウンドをプレイするループを作成し、毎回`playRound`関数を人間の選択とコンピューターの選択の関数を引数として呼び出してください。

1. ゲームの最後に、どちらが多くのラウンドに勝ったかに基づいてゲームの勝者を返してください。

- 人間のプレイヤーがコンピューターより多くのラウンドに勝った場合、`"You win the game!"`というメッセージを返してください。

- コンピューターが人間より多くのラウンドに勝った場合、`"You lose the game!"`というメッセージを返してください。

# --hints--

`getComputerChoice`という関数を持っている必要があります。

```js
assert.isFunction(getComputerChoice);
```

`getComputerChoice`関数は`"rock"`、`"paper"`、または`"scissors"`のいずれかをランダムに返す必要があります。

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

`getHumanChoice`という関数を持っている必要があります。

```js
assert.isFunction(getHumanChoice);
```

グローバルスコープに`humanScore`と`computerScore`という2つの変数を持っている必要があります。

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

`playRound`という関数を持っている必要があります。

```js
assert.isFunction(playRound);
```

`playRound`関数は人間とコンピューターの選択を`humanChoice`と`computerChoice`のパラメータで引数として受け取る必要があります。

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

`playRound`関数は大文字小文字を区別しないようにして、プレイヤーが`"rock"`、`"ROCK"`、`"RocK"`、またはその他のバリエーションを入力できるようにしてください。

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

`playRound`関数は人間とコンピューターの選択が同じ場合、文字列`"It's a tie!"`を返す必要があります。

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

`playRound`関数はプレイヤーが勝った場合、文字列`"You win! [player choice] beats [computer choice]"`を返す必要があります。

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

`playRound`関数はコンピューターが勝った場合、文字列`"You lose! [computer choice] beats [player choice]"`を返す必要があります。

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

`playRound`関数はラウンドの勝者に応じて`humanScore`または`computerScore`変数を増やす必要があります。

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

`playGame`という関数を持っている必要があります。

```js
assert.isFunction(playGame);
```

3ラウンドをプレイするためにループを使う必要があります。

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

最も多くのラウンドに勝った方を基にゲームの勝者を返す必要があります。

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```
