---
id: 646c48df8674cf2b91020ecc
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

基本的なボイラープレートが提供されています。まず、`link` 要素の中に `head` 要素を追加してください。その `link` 要素に対して、`rel` 属性を `stylesheet` に、`href` を `./styles.css` に設定してください。

# --hints--

`link` 要素があるはずです。

```js
assert.exists(document.querySelector("head>link"));
```

`link` 要素には `rel` と `href` 属性があり、それぞれ `stylesheet` と `./styles.css` に設定されている必要があります。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
