---
id: 646cfc2b8e6fe95c20a819d5
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

耳は、重なっている頭の部分の上に常に配置する必要があります。これは`z-index`プロパティで設定できます。

`z-index`は、重なり合うHTML要素の順序を定義するために使えるプロパティです。`z-index`が高い要素は、常に`z-index`が低い要素の上に配置されます。

`z-index`の動作を確認するには、左耳の`z-index`プロパティを`-1`に設定してください。

# --hints--

`.cat-left-ear`セレクターには`z-index`を`-1`に設定してください。セミコロンを忘れないでください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.zIndex, '-1')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
