---
id: 646ce9d790d2a44de5f99e04
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

`.cat-left-ear` 要素の中にクラス `cat-left-inner-ear` を持つ `div` 要素を1つ作成し、`.cat-right-ear` 要素の中にクラス `cat-right-inner-ear` を持つ `div` 要素をもう1つ作成してください。

# --hints--

既存のクラス `cat-left-ear` を持つ `div` 要素は変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

既存のクラス `cat-right-ear` を持つ `div` 要素は変更しないでください。

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

`.cat-left-ear` 要素の中に `div` 要素が1つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

`.cat-right-ear` 要素の中に `div` 要素が1つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

`.cat-left-ear` の中の新しい `div` 要素はクラス `cat-left-inner-ear` を持つべきです。

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

`.cat-right-ear` の中の新しい `div` 要素はクラス `cat-right-inner-ear` を持つべきです。

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
