---
id: 646c59e942f35541923104bf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

HTMLに戻り、`main` 要素を作成してください。その `main` 要素の中に、クラス `div` を持つ `cat-head` 要素を追加してください。

# --hints--

`main` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`div` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

`div` 要素にはクラス `cat-head` が付いているはずです。

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

`div` 要素は `main` タグの中にあるはずです。

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
