---
id: 646c5ace05e4be4211407935
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`.cat-head` 要素を見えるようにするために、`#5e5e5e` を `85%` に、`#45454f` を `100%` にした線形グラデーションの背景を付けてください。

この2つの色の違いは気づきにくいかもしれませんが、確かに存在しています。

# --hints--

`.cat-head` セレクターには `background` プロパティが必要です。

```js
assert.match(code, /background:/)
```

`background` プロパティには `linear-gradient` 関数を使ってください。

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

`linear-gradient` 関数では、最初の色を `#5e5e5e` に `85%` で設定してください。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

`linear-gradient` 関数では、2番目の色を `#45454f` に `100%` で設定してください。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
