---
id: 646c5d7057c45f432fcdd46c
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

クラスセレクターを使って、`.cat-head` 要素に幅 `205px` と高さ `180px` を指定してください。また、ボーダーを `1px solid #000`、`border-radius` を `46%` に設定してください。

# --hints--

`.cat-head` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

`.cat-head` セレクターには `width` を `205px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

`.cat-head` セレクターには `height` を `180px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

`.cat-head` セレクターには `border` を `1px solid #000` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

`.cat-head` セレクターには `border-radius` を `46%` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
