---
id: 646dde7dc20dc167489faa69
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

クラスセレクターを使って、`.cat-left-eye` 要素に `width` の `30px` と `height` の `40px` を付けてください。また、`background-color` の `#000` も付けてください。

# --hints--

`.cat-left-eye` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cat-left-eye'))
```

`.cat-left-eye` セレクターには `width` に設定された `30px` があるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-eye')?.width, '30px')
```

`.cat-left-eye` セレクターには `height` に設定された `40px` があるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-eye')?.height, '40px')
```

`.cat-left-eye` セレクターには `background-color` に設定された `#000` があるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-eye')?.backgroundColor, 'rgb(0, 0, 0)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>

        <div class="cat-eyes">
          <div class="cat-left-eye">
            <div class="cat-left-inner-eye"></div>
          </div>
          <div class="cat-right-eye">
            <div class="cat-right-inner-eye"></div>
          </div>
        </div>        
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-left-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

.cat-right-inner-ear {
  position: absolute;
  top: 22px;
  left: -20px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-bottom-right-radius: 40%;
  border-bottom-left-radius: 40%;
  border-left: 20px solid transparent;
  border-right: 20px solid transparent;
  border-bottom: 40px solid #3b3b4f;
}

--fcc-editable-region--

--fcc-editable-region--
```
