---
id: 6476fd4213318f6ee211028a
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

今、`position` の `.cat-head` プロパティの値を `fixed` に置き換えてください。`top` と `left` はそのままにしてください。

その後、上下にスクロールして `fixed` の値がどのように動作するか確認してください。

# --hints--

`.cat-head` セレクターには `position` プロパティを `fixed` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'fixed');
```

`.cat-head` セレクターには `top` プロパティを `100px` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

`.cat-head` セレクターには `left` プロパティを `100px` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
