---
id: 6476f5c17f99146071ee884c
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

何も起こらなかったことがわかります。`.cat-head` 要素は、それぞれ `top` の `left` と `100px` を設定しても動きませんでした。しかし、`relative` の位置指定ではそうではありません。

`relative` の値を使うと、要素はまだ文書の通常のフローに従って配置されますが、`top`、`left`、`bottom`、および `right` の値が有効になります。

`static` の代わりに、`.cat-head` に `relative` の位置を指定し、`top` と `left` の両方のプロパティはそのままにしてください。

# --hints--

`.cat-head` セレクターには、`position` プロパティを `relative` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

`.cat-head` セレクターには `top` プロパティを `100px` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

`.cat-head` セレクターには `left` プロパティを `100px` に設定してください。セミコロンを忘れずに追加してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
