---
id: 6476fc5cf14b276e6d04e82a
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

クラスセレクターを使って、`.box` 要素に幅 `200px`、高さ `600px`、背景色 `#000` を設定してください。また、`position` を `absolute`、`top` を `800px`、`left` を `650px` に設定してください。

# --hints--

`.box` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.box'));
```

`.box` セレクターには `width` プロパティを `200px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.width, '200px');
```

`.box` セレクターには `height` プロパティを `600px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.height, '600px');
```

`.box` セレクターには `background-color` プロパティを `#000` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor, 'rgb(0, 0, 0)');
```

`.box` セレクターには `position` プロパティを `absolute` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.position, 'absolute');
```

`.box` セレクターには `left` プロパティを `650px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.left, '650px');
```

`.box` セレクターには `top` プロパティを `800px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.top, '800px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

