---
id: 6476f7a4827bcc61682f2347
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`absolute` プロパティに `position` の値を使うと、その要素は文書の通常の流れから外れ、位置は `top`、`right`、`bottom`、および `left` プロパティによって決まります。

`.cat-head` 要素の `position` プロパティを `absolute` に設定してください。

# --hints--

`.cat-head` セレクターには `position` プロパティを `absolute` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

`.cat-head` セレクターには `top` の任意の数値に設定された `px` プロパティが必要です。セミコロンを忘れずに追加してください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

`.cat-head` セレクターには `left` の任意の数値に設定された `px` プロパティが必要です。セミコロンを忘れずに追加してください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>fCC Cat Painting</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class="cat-head"></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
--fcc-editable-region--
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
