---
id: 6476fb7cbaafa36d65e9cf35
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` は、ユーザーがページをどこまでスクロールしても要素をページに固定できる `position` プロパティ値です。

`fixed` の位置指定がどのように機能するかを見るために、もう少しマークアップを行う必要があります。HTML内で、クラス `div` を持つ `box` 要素を作成してください。

# --hints--

`div` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

`div` 要素にはクラス `box` を付けてください。

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
