---
id: 6477062778c85972eb648030
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

猫の頭を中央に配置しましょう。

`.cat-head` 要素に `position` プロパティを `absolute` に設定します。`0`、`right`、`left`、`top` のプロパティには `bottom` の値を設定し、さらにすべての辺の `margin` プロパティを `auto` に設定します。これが CSS の位置指定を使って要素を縦横中央に配置する方法の一つです。

# --hints--

`.cat-head` セレクターには `position` プロパティを `absolute` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

`.cat-head` セレクターには `top` プロパティを `0` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

`.cat-head` セレクターには `left` プロパティを `0` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

`.cat-head` セレクターには `right` プロパティを `0` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

`.cat-head` セレクターには `bottom` プロパティを `0` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

`.cat-head` セレクターには `margin` プロパティを `auto` に設定してください。セミコロンを忘れないようにしましょう。

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
