---
id: 69a92f0b9cc04bb0d5327bb5
title: implementare uno stack
challengeType: 26
dashedName: implement-a-stack
---

# --description--

In questo laboratorio, implementerai una struttura dati stack usando funzioni. Uno stack è una struttura dati Last-In-First-Out (LIFO) in cui gli elementi vengono aggiunti e rimossi dalla cima.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `push` che aggiunge un elemento in cima allo stack.
1. Dovresti avere una funzione `pop` che rimuove e restituisce l'elemento in cima allo stack, oppure `undefined` se non ce n'è uno.
1. Dovresti avere una funzione `peek` che restituisce l'elemento in cima allo stack senza rimuoverlo, oppure `undefined` se non ce n'è uno.
1. Dovresti avere una funzione `isEmpty` che restituisce `true` se lo stack non contiene elementi, e `false` altrimenti.
1. Dovresti avere una funzione `clear` che rimuove tutti gli elementi dallo stack.

**Nota:** La maggior parte dei test dipende dalla funzione `push`. Implementa prima `push`, poiché i test per `pop`, `peek`, `isEmpty` e `clear` richiedono di aggiungere elementi allo stack.

# --hints--

Dovresti avere una funzione `push`.

```js
assert.isFunction(push);
```

La funzione `push` dovrebbe aggiungere un elemento in cima allo stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Dovresti avere una funzione `pop`.

```js
assert.isFunction(pop);
```

La funzione `pop` dovrebbe rimuovere e restituire l'elemento in cima allo stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

La funzione `pop` dovrebbe restituire `undefined` se lo stack è vuoto.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

La funzione `pop` dovrebbe restituire correttamente valori falsy.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Dovresti avere una funzione `peek`.

```js
assert.isFunction(peek);
```

La funzione `peek` dovrebbe restituire l'elemento in cima allo stack senza rimuoverlo.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

La funzione `peek` dovrebbe restituire `undefined` se lo stack è vuoto.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

La funzione `peek` dovrebbe restituire correttamente valori falsy.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Dovresti avere una funzione `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La funzione `isEmpty` dovrebbe restituire `true` per uno stack vuoto.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

La funzione `isEmpty` dovrebbe restituire `false` per uno stack non vuoto.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

La funzione `isEmpty` dovrebbe restituire `false` quando l'elemento in cima è falsy.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Dovresti avere una funzione `clear`.

```js
assert.isFunction(clear);
```

La funzione `clear` dovrebbe rimuovere tutti gli elementi dallo stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
