---
id: 587d824e367417b2b2512c58
title: 使用 Chai-HTTP 在 API 端上运行功能测试
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

Mocha 允许你使用名为 `chai-http` 的插件测试异步操作，例如调用 API 端点。

以下是使用 `chai-http` 对名为 `'GET /hello?name=[name] => "hello [name]"'` 的测试套件进行测试的示例：

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

该测试向服务器发送一个 `GET` 请求，并将名称作为 URL 查询字符串（`?name=John`）。 在`end` 方法的回调函数中，接收到响应对象（`res`）并包含 `status` 属性。

第一个 `assert.equal` 检查状态是否为 `200`。 第二个 `assert.equal` 检查响应字符串（`res.text`）是否为 `"hello John"`。

同时，请注意测试的回调函数中的 `done` 参数。 在测试结束时，调用它且不带参数，是发出异步操作完成所必需的信号。

最后，请注意 `request` 方法后面的 `keepOpen` 方法。 通常，你会从命令行中运行你的测试，或者作为自动集成过程的一部分，你可以让 `chai-http` 自动启动和停止你的服务器。

然而，当你提交项目链接时运行的测试需要你的服务器是正常的，所以你需要使用 `keepOpen` 方法来防止 `chai-http` 停止你的服务器。

# --instructions--

在 `tests/2_functional-tests.js` 中，修改 `'Test GET /hello with no name'` 测试（`// #1`），对响应的 `status` 和 `text` 使用断言来通过测试。 不要改变传递给断言的参数。

不应该有任何 URL 查询。 如果没有名称 URL 查询，端点将使用 `hello Guest` 进行响应。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

你应该测试 `res.status` == 200。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

你应该测试 `res.text` == `'hello Guest'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```

