---
id: 587d824d367417b2b2512c51
title: 测试数组是否包含项目
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

在 `Arrays` 套件中标记为 `#12` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.include` 或 `assert.notInclude` 以使测试通过（应评估为 `true`）。 不要修改传给断言的参数。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`include` 或 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

应该为第二个断言选择正确的方法：`include` 或 `notInclude`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a backend language !!'
);
```

