---
id: 587d824d367417b2b2512c52
title: 测试某个值是否为字符串
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

`isString` 或 `isNotString` 断言一个值是否为字符串。

# --instructions--

在 `Strings` 套件中标记为 `#13` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.isString` 或 `assert.isNotString` 以使测试通过（应评估为 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`isString` 或 `isNotString`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

应该为第二个断言选择正确的方法：`isString` 或 `isNotString`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

应该为第三个断言选择正确的方法：`isString` 或 `isNotString`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```

