---
id: 587d824b367417b2b2512c48
title: 使用 Assert.isOk() 和 Assert.isNotOK()
challengeType: 2
forumTopicId: 301607
dashedName: use-assert-isok-and-assert-isnotok
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

`isOk()` 用来测试值是否为真值，`isNotOk()` 用来测试值是否为假值。

要了解有关真值和假值的更多信息，请尝试我们的 <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-algorithm-scripting/falsy-bouncer" target="_blank" rel="noopener noreferrer nofollow">Falsy Bouncer</a> 挑战。

# --instructions--

在 `tests/1_unit-tests.js` 中，`Basic Assertions` 套件中标注为 `#3` 的测试下，修改每个 `assert` 为 `assert.isOk()` 或 `assert.isNotOk()`，通过测试（结果应为 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

应该第一个断言选择正确的方法：`isOk` 或 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNotOk', 'Null is falsy');
```

应该第二个断言选择正确的方法：`isOk` 或 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isOk', 'A string is truthy');
```

应该第三个断言选择正确的方法：`isOk` 或 `isNotOk`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=2');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[2].method, 'isOk', 'true is truthy');
```

