---
id: 587d824e367417b2b2512c57
title: 测试对象是否是构造函数的实例
challengeType: 2
forumTopicId: 301605
dashedName: test-if-an-object-is-an-instance-of-a-constructor
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

`#instanceOf` 断言一个对象是一个构造器的实例。

# --instructions--

在 `Objects` 套件中标记为 `#18` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.instanceOf` 或 `assert.notInstanceOf` 以使测试通过（应评估为 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`instanceOf` 或 `notInstanceOf`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInstanceOf',
  'myCar is not an instance of Plane'
);
```

应该为第二个断言选择正确的方法：`instanceOf` 或 `notInstanceOf`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'instanceOf',
  'airlinePlane is an instance of Plane'
);
```

应该为第三个断言选择正确的方法：`instanceOf` 或 `notInstanceOf`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[2].method,
  'instanceOf',
  'everything is an Object in JavaScript...'
);
```

应该为第四个断言选择正确的方法：`instanceOf` 或 `notInstanceOf`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=17');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'notInstanceOf',
  'myCar.wheels is not an instance of String'
);
```

