---
id: 587d824b367417b2b2512c47
title: 测试变量或函数是否已定义
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

# --instructions--

在 `Basic Assertions` 套件中标记为 `#2` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.isDefined()` 或 `assert.isUndefined()` 以使测试通过（应评估为 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`isDefined` 或 `isUndefined`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

应该为第二个断言选择正确的方法：`isDefined` 或 `isUndefined`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

应该为第三个断言选择正确的方法：`isDefined` 或 `isUndefined`。

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```

