---
id: 587d824c367417b2b2512c4e
title: 测试一个值是否小于或等于另一个值
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

# --instructions--

在 `tests/1_unit-tests.js` 中，在 `Comparisons` 套件里标有 `#9` 的测试中，将每个 `assert` 改成 `assert.isBelow` 或 `assert.isAtLeast` 方法，让测试通过（结果应该返回 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`isBelow` 或 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

应该为第二个断言选择正确的方法：`isBelow` 或 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

应该为第三个断言选择正确的方法：`isBelow` 或 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

应该为第四个断言选择正确的方法：`isBelow` 或 `isAtLeast`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```

