---
id: 587d824e367417b2b2512c55
title: 测试对象是否具有某个属性
challengeType: 2
forumTopicId: 301604
dashedName: test-if-an-object-has-a-property
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`property` 断言一个对象含有给定属性。

# --instructions--

在 `tests/1_unit-tests.js` 中，在测试标签 `#16`下，将每个`Objects` 下的`assert` 修改成`assert.property` 或`assert.notProperty`以确保通过测试 (结果应返回“ `true`”). 不要修改传给断言的参数。

# --hints--

不应有未通过的测试

```js
  $.get(code + '/_api/get-tests?type=unit&n=15').then(
    (data) => {
      assert.equal(data.state, 'passed');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

应该为第一个断言选择正确的方法：`property` 或 `notProperty`。

```js
  $.get(code + '/_api/get-tests?type=unit&n=15').then(
    (data) => {
      assert.equal(
        data.assertions[0].method,
        'notProperty',
        'A car has not wings'
      );
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

应该为第二个断言选择正确的方法：`property` 或 `notProperty`。

```js
  $.get(code + '/_api/get-tests?type=unit&n=15').then(
    (data) => {
      assert.equal(
        data.assertions[1].method,
        'property',
        'planes have engines'
      );
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

应该为第三个断言选择正确的方法：`property` 或 `notProperty`。

```js
  $.get(code + '/_api/get-tests?type=unit&n=15').then(
    (data) => {
      assert.equal(data.assertions[2].method, 'property', 'Cars have wheels');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

