---
id: 587d8250367417b2b2512c5d
title: 使用无头浏览器运行功能测试
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

在页面上有一个输入表单。 它将数据作为 AJAX 请求发送到 `PUT /travellers` 端点。

当请求成功完成时，客户代码将包含响应中信息的 `<div>` 追加到 DOM。

下面是一个如何使用 Zombie.js 与表单互动的例子。

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

首先，`browser` 对象的 `fill` 方法用值 `'Polo'` 填充表单的 `surname` 字段。`fill` 返回一个 promise，因此在它后面链式调用了 `then`。

在 `then` 回调中，`browser` 对象的 `pressButton` 方法用于调用表单的 `submit` 的事件侦听器。 `pressButton` 方法是异步的。

然后，一旦收到来自 AJAX 请求的响应，就会做出一些断言来确认：

1.  响应状态是 `200`
2.  `<span id='name'></span>` 元素内的文本匹配 `'Marco'`
3.  `<span id='surname'></span>` 元素内的文本匹配 `'Polo'`
4.  有 `1` 个 `<span id='dates'></span>` 元素。

最后，执行 `done`，这是异步测试所必需的。

# --instructions--

在 `tests/2_functional-tests.js` 中，在 `'Submit the surname "Colombo" in the HTML form'` 测试（`// #5`），自动执行以下操作：

1.  在表格中填写姓氏 `Colombo`。
2.  点击提交按钮

在 `pressButton` 回调中：

1.  断言状态是正常的 `200`。
2.  断言元素 `span#name` 中的文本是 `'Cristoforo'`。
3.  断言元素 `span#surname` 中的文本是 `'Colombo'`。
4.  断言有 `span#dates` 元素，它们的计数是 `1`。

不要忘记删除 `assert.fail()` 调用。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应断言无头浏览器成功执行请求。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

应断言元素 `span#name` 中的文本是 `'Cristoforo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

应断言元素 `span#surname` 中的文本是 `'Colombo'`。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

应该断言元素 `span#dates` 存在，且它的值为 1。

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```

