---
id: 587d824c367417b2b2512c4f
title: 测试某个值是否在特定范围内
challengeType: 2
forumTopicId: 301598
dashedName: test-if-a-value-falls-within-a-specific-range
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

```javascript
.approximately(actual, expected, delta, [message])
```

断言 `actual` 等于 `expected`，在 +/- `delta` 的范围内。

# --instructions--

在 `Comparisons` 套件中标记为 `#10` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.approximately` 以使测试通过（应评估为 `true`）。

选择最小范围（第三个参数）来通过所有测试。 它应该小于 1。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的范围——`approximately(actual, expected, range)`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'approximately');
assert.equal(
  data.assertions[0].args[2],
  0.5,
  "weirdNumbers(0.5) is in the range (0.5, 1.5]. It's within 1 +/- 0.5"
);
```

应该为第二个断言选择正确的范围——`approximately(actual, expected, range)`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=9');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'approximately');
assert.equal(
  data.assertions[1].args[2],
  0.8,
  "weirdNumbers(0.2) is in the range (0.2, 1.2]. It's within 1 +/- 0.8"
);
```

