---
id: 587d824d367417b2b2512c54
title: 使用正则表达式测试字符串
challengeType: 2
forumTopicId: 301608
dashedName: use-regular-expressions-to-test-a-string
---

# --description--

提醒你，这个项目是基于从 <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a> 克隆的以下启动项目构建的。

`match()` 断言一个值匹配一个正则表达式（第二个参数）。

# --instructions--

在 `Strings` 套件中标记为 `#15` 的测试下的 `tests/1_unit-tests.js` 中，将每个 `assert` 更改为 `assert.match` 或 `assert.notMatch` 以使测试通过（应评估为 `true`）。 不要改变传递给断言的参数。

# --hints--

应通过所有测试。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

应该为第一个断言选择正确的方法：`match` 或 `notMatch`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'match',
  "'# name:John Doe, age:35' matches the regex"
);
```

应该为第二个断言选择正确的方法：`match` 或 `notMatch`。

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notMatch',
  "'# name:Paul Smith III, age:twenty-four' does not match the regex (the age must be numeric)"
);
```

