---
id: 61537d86bdc3dd343688fceb
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Geben Sie als Nächstes jedem `img` ein `src`-Attribut entsprechend seiner Reihenfolge im Dokument. Das erste `img` sollte ein `src` mit dem Wert `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` haben. Die übrigen sollten gleich sein, außer dass Sie das `1` durch die Zahl ersetzen, an der sich das `img` im Dokument befindet.

# --hints--

Alle neun Ihrer `img`-Elemente sollten ein `src`-Attribut haben.

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

Ihr erstes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ihr zweites `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Ihr drittes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Ihr viertes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Ihr fünftes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Ihr sechstes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Ihr siebtes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Ihr achtes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Ihr neuntes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` als Wert des `src`-Attributs haben.

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
--fcc-editable-region--
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
