---
id: 61537c9eecea6a335db6da79
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Unter deinem `.header` Element erstelle ein neues `div` Element und weise ihm eine `class` von `gallery` zu. Dieses `div`-Element dient als Container für die Galeriebilder.

Erstelle innerhalb deines `.gallery`-Elements neun `img`-Elemente.

# --hints--

Du solltest ein `div`-Element in deinem `body`-Element erstellen.

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

Dein neues `div`-Element sollte eine `class` mit `gallery` als Wert haben.

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

Dein neues `div`-Element sollte nach deinem `.header`-Element stehen.

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Dein `.gallery`-Element sollte neun `img`-Elemente haben.

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
