---
id: 61537bb9b1a29430ac15ad38
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Starten Sie dieses Projekt, indem Sie Ihre `styles.css`-Datei mit der Seite verknüpfen.

# --hints--

Sie sollten ein `link`-Element innerhalb Ihres `head`-Elements haben.

```js
assert.exists(document.querySelector("head>link"));
```

Ihr `link`-Element sollte `rel="stylesheet"` und `href="styles.css"` enthalten.

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
