---
id: 64dd076a1b33c86d84b4232f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Um besser zu visualisieren, wie Ihre Elemente dimensioniert sind, kann es hilfreich sein, einen Rahmen hinzuzufügen.

Geben Sie Ihrem `.gallery`-Element einen `width` von `50%` und eine `border`, die auf `5px solid red` gesetzt ist.

Geben Sie dann Ihren `img`-Elementen einen `width` von `100%`, eine `padding`, die auf `5px` gesetzt ist, und eine `border`, die auf `5px solid blue` gesetzt ist.

# --hints--

Sie sollten einen `.gallery`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Ihr `.gallery`-Selektor sollte eine `width`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Ihr `.gallery`-Selektor sollte eine `width`-Eigenschaft haben, die auf `50%` gesetzt ist.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Ihr `.gallery`-Selektor sollte eine `border`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Ihr `.gallery`-Selektor sollte eine `border`-Eigenschaft haben, die auf `5px solid red` gesetzt ist.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Sie sollten einen `img`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Ihr `img`-Selektor sollte eine `width`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Ihr `img`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `100%` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Ihr `img`-Selektor sollte eine `padding`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Ihr `img`-Selektor sollte eine `padding`-Eigenschaft haben, die auf `5px` gesetzt ist.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Ihr `img`-Selektor sollte eine `border`-Eigenschaft haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Ihr `img`-Selektor sollte eine `border`-Eigenschaft haben, die auf `5px solid blue` gesetzt ist.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
