---
id: 63ec3287b182ec0efe8a3135
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

使用 `products` 將此表中剩餘的產品（依序）添加到你的 `new Dessert(...)` 陣列。每加入一個產品，`id` 都要遞增計數。

| `name`                         | `price` | `category`  |
| ------------------------------ | ------- | ----------- |
| `French Macaron`               | `3.99`  | `Macaron`   |
| `Pumpkin Cupcake`              | `3.99`  | `Cupcake`   |
| `Chocolate Cupcake`            | `5.99`  | `Cupcake`   |
| `Chocolate Pretzels (4 Pack)`  | `10.99` | `Pretzel`   |
| `Strawberry Ice Cream`         | `2.99`  | `Ice Cream` |
| `Chocolate Macarons (4 Pack)`  | `9.99`  | `Macaron`   |
| `Strawberry Pretzel`           | `4.99`  | `Pretzel`   |
| `Butter Pecan Ice Cream`       | `2.99`  | `Ice Cream` |
| `Rocky Road Ice Cream`         | `2.99`  | `Ice Cream` |
| `Vanilla Macarons (5 Pack)`    | `11.99` | `Macaron`   |
| `Lemon Cupcakes (4 Pack)`      | `12.99` | `Cupcake`   |

# --hints--

你的第二個產品應該有一個 `id` 屬性，設定為數字 `2`。

```js
assert.strictEqual(products[1].id, 2);
```

你的第二個產品應該有一個 `name` 屬性，設定為 `French Macaron`。

```js
assert.strictEqual(products[1].name, 'French Macaron');
```

你的第二個產品應該有一個 `price` 屬性，設定為數字 `3.99`。

```js
assert.strictEqual(products[1].price, 3.99);
```

你的第二個產品應該有一個 `category` 屬性，設定為 `Macaron`。

```js
assert.strictEqual(products[1].category, 'Macaron');
```

你的第三個產品應該有一個 `id` 屬性，設定為數字 `3`。

```js
assert.strictEqual(products[2].id, 3);
```

你的第三個產品應該有一個 `name` 屬性，設定為 `Pumpkin Cupcake`。

```js
assert.strictEqual(products[2].name, 'Pumpkin Cupcake');
```

你的第三個產品應該有一個 `price` 屬性，設定為數字 `3.99`。

```js
assert.strictEqual(products[2].price, 3.99);
```

你的第三個產品應該有一個 `category` 屬性，設定為 `Cupcake`。

```js
assert.strictEqual(products[2].category, 'Cupcake');
```

你的第四個產品應該有一個 `id` 屬性，設定為數字 `4`。

```js
assert.strictEqual(products[3].id, 4);
```

你的第四個產品應該有一個 `name` 屬性，設定為 `Chocolate Cupcake`。

```js
assert.strictEqual(products[3].name, 'Chocolate Cupcake');
```

你的第四個產品應該有一個 `price` 屬性，設定為數字 `5.99`。

```js
assert.strictEqual(products[3].price, 5.99);
```

你的第四個產品應該有一個 `category` 屬性，設定為 `Cupcake`。

```js
assert.strictEqual(products[3].category, 'Cupcake');
```

你的第五個產品應該有一個 `id` 屬性，設定為數字 `5`。

```js
assert.strictEqual(products[4].id, 5);
```

你的第五個產品應該有一個 `name` 屬性，設定為 `Chocolate Pretzels (4 Pack)`。

```js
assert.strictEqual(products[4].name, 'Chocolate Pretzels (4 Pack)');
```

你的第五個產品應該有一個 `price` 屬性，設定為數字 `10.99`。

```js
assert.strictEqual(products[4].price, 10.99);
```

你的第五個產品應該有一個 `category` 屬性，設定為 `Pretzel`。

```js
assert.strictEqual(products[4].category, 'Pretzel');
```

你的第六個產品應該有一個 `id` 屬性，設定為數字 `6`。

```js
assert.strictEqual(products[5].id, 6);
```

你的第六個產品應該有一個 `name` 屬性，設定為 `Strawberry Ice Cream`。

```js
assert.strictEqual(products[5].name, 'Strawberry Ice Cream');
```

你的第六個產品應該有一個 `price` 屬性，設定為數字 `2.99`。

```js
assert.strictEqual(products[5].price, 2.99);
```

你的第六個產品應該有一個 `category` 屬性，設定為 `Ice Cream`。

```js
assert.strictEqual(products[5].category, 'Ice Cream');
```

你的第七個產品應該有一個 `id` 屬性，設定為數字 `7`。

```js
assert.strictEqual(products[6].id, 7);
```

你的第七個產品應該有一個 `name` 屬性，設定為 `Chocolate Macarons (4 Pack)`。

```js
assert.strictEqual(products[6].name, 'Chocolate Macarons (4 Pack)');
```

你的第七個產品應該有一個 `price` 屬性，設定為數字 `9.99`。

```js
assert.strictEqual(products[6].price, 9.99);
```

你的第七個產品應該有一個 `category` 屬性，設定為 `Macaron`。

```js
assert.strictEqual(products[6].category, 'Macaron');
```

你的第八個產品應該有一個 `id` 屬性，設定為數字 `8`。

```js
assert.strictEqual(products[7].id, 8);
```

你的第八個產品應該有一個 `name` 屬性，設定為 `Strawberry Pretzel`。

```js
assert.strictEqual(products[7].name, 'Strawberry Pretzel');
```

你的第八個產品應該有一個 `price` 屬性，設定為數字 `4.99`。

```js
assert.strictEqual(products[7].price, 4.99);
```

你的第八個產品應該有一個 `category` 屬性，設定為 `Pretzel`。

```js
assert.strictEqual(products[7].category, 'Pretzel');
```

你的第九個產品應該有一個 `id` 屬性，設定為數字 `9`。

```js
assert.strictEqual(products[8].id, 9);
```

你的第九個產品應該有一個 `name` 屬性，設定為 `Butter Pecan Ice Cream`。

```js
assert.strictEqual(products[8].name, 'Butter Pecan Ice Cream');
```

你的第九個產品應該有一個 `price` 屬性，設定為數字 `2.99`。

```js
assert.strictEqual(products[8].price, 2.99);
```

你的第九個產品應該有一個 `category` 屬性，設定為 `Ice Cream`。

```js
assert.strictEqual(products[8].category, 'Ice Cream');
```

你的第十個產品應該有一個 `id` 屬性，設定為數字 `10`。

```js
assert.strictEqual(products[9].id, 10);
```

你的第十個產品應該有一個 `name` 屬性，設定為 `Rocky Road Ice Cream`。

```js
assert.strictEqual(products[9].name, 'Rocky Road Ice Cream');
```

你的第十個產品應該有一個 `price` 屬性，設定為數字 `2.99`。

```js
assert.strictEqual(products[9].price, 2.99);
```

你的第十個產品應該有一個 `category` 屬性，設定為 `Ice Cream`。

```js
assert.strictEqual(products[9].category, 'Ice Cream');
```

你的第十一個產品應該有一個 `id` 屬性，設定為數字 `11`。

```js
assert.strictEqual(products[10].id, 11);
```

你的第十一個產品應該有一個 `name` 屬性，設定為 `Vanilla Macarons (5 Pack)`。

```js
assert.strictEqual(products[10].name, 'Vanilla Macarons (5 Pack)');
```

你的第十一個產品應該有一個 `price` 屬性，設定為數字 `11.99`。

```js
assert.strictEqual(products[10].price, 11.99);
```

你的第十一個產品應該有一個 `category` 屬性，設定為 `Macaron`。

```js
assert.strictEqual(products[10].category, 'Macaron');
```

你的第十二個產品應該有一個 `id` 屬性，設定為數字 `12`。

```js
assert.strictEqual(products[11].id, 12);
```

你的第十二個產品應該有一個 `name` 屬性，設定為 `Lemon Cupcakes (4 Pack)`。

```js
assert.strictEqual(products[11].name, 'Lemon Cupcakes (4 Pack)');
```

你的第十二個產品應該有一個 `price` 屬性，設定為數字 `12.99`。

```js
assert.strictEqual(products[11].price, 12.99);
```

你的第十二個產品應該有一個 `category` 屬性，設定為 `Cupcake`。

```js
assert.strictEqual(products[11].category, 'Cupcake');
```

你的 `products` 陣列應該有 12 個值，每個都是 `Dessert` 的實體。

```js
assert.strictEqual(products.length, 12);
products.forEach(p => assert.instanceOf(p, Dessert));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp Shopping Cart</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1 class="title">Desserts Page</h1>
    </header>
    <main>
      <button id="cart-btn" type="button" class="btn">
        <span id="show-hide-cart">Show</span> Cart
      </button>
      <div id="cart-container">
        <button class="btn" id="clear-cart-btn">Clear Cart</button>
        <div id="products-container"></div>
        <p>Total number of items: <span id="total-items">0</span></p>
        <p>Subtotal: <span id="subtotal">$0</span></p>
        <p>Taxes: <span id="taxes">$0</span></p>
        <p>Total: <span id="total">$0</span></p>
      </div>
      <div id="dessert-card-container"></div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.title {
  color: var(--light-grey);
  text-align: center;
  margin: 25px 0;
}

#dessert-card-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
}

.dessert-card {
  background-color: var(--light-grey);
  padding: 15px;
  text-align: center;
  border-radius: 15px;
  margin: 20px 10px;
}

.dessert-price {
  font-size: 1.2rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 100px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#cart-btn {
  position: absolute;
  top: 0;
  right: 0;
}

.add-to-cart-btn {
  margin: 30px auto 10px;
}

#cart-container {
  display: none;
  position: absolute;
  top: 60px;
  right: 0;
  background-color: var(--light-grey);
  width: 200px;
  height: 400px;
  border: 8px double var(--black);
  border-radius: 15px;
  text-align: center;
  font-size: 1.2rem;
  overflow-y: scroll;
}

.product {
  margin: 25px 0;
}

.product-count {
  display: inline-block;
  margin-right: 10px;
}

.product-category {
  margin: 10px 0;
}

@media (min-width: 768px) {
  #dessert-card-container {
    flex-direction: row;
  }

  .dessert-card {
    flex: 1 0 21%;
  }

  #cart-container {
    width: 300px;
  }
}
```

```js
const cartContainer = document.getElementById("cart-container");
const productsContainer = document.getElementById("products-container");
const dessertCards = document.getElementById("dessert-card-container");
const cartBtn = document.getElementById("cart-btn");
const clearCartBtn = document.getElementById("clear-cart-btn");
const totalNumberOfItems = document.getElementById("total-items");
const cartSubTotal = document.getElementById("subtotal");
const cartTaxes = document.getElementById("taxes");
const cartTotal = document.getElementById("total");
const showHideCartSpan = document.getElementById("show-hide-cart");
let isCartShowing = false;

class Dessert {
  constructor(id, name, price, category) {
    this.id = id;
    this.name = name;
    this.price = price;
    this.category = category;
  }
}

const products = [
  new Dessert(1, "Vanilla Cupcakes (6 Pack)", 12.99, "Cupcake"),
--fcc-editable-region--
  
--fcc-editable-region--
];
```
