---
id: 61969c487ced6f12db8fef94
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

將 `.left-mountain` 元素設為標的，並將其 `width` 和 `height` 設為 `300px`。然後，將 `background` 設為從 `rgb(203, 241, 228)` 開始到 `rgb(80, 183, 255)` 結束的線性漸層。

# --hints--

你應該使用 `.left-mountain` 選擇器。

```js
assert.match(code, /\.left-mountain\s*\{/);
```

你應該給 `.left-mountain` 一個 `width` 為 `300px`。預期 `--fcc-actual--` 為 `--fcc-expected--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.width, '300px');
```

你應該給 `.left-mountain` 一個 `height` 為 `300px`。預期 `--fcc-actual--` 為 `--fcc-expected--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.height, '300px');
```

你應該給 `.left-mountain` 一個 `background` 為 `linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255))`。

```js
assert.include(['linear-gradient(rgb(203,241,228),rgb(80,183,255))', 'rgba(0,0,0,0)linear-gradient(rgb(203,241,228),rgb(80,183,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
