---
id: 61967e74a8e3690ab6292daa
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

將 `body` 元素設定為 `background`，使用順時針 45 度角的線性漸層，起始於 `rgb(118, 201, 255)`，結束於 `rgb(247, 255, 222)`。

# --hints--

你應該使用 `body` 元素選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你應該在 `body` 選擇器中使用 `background` 屬性。

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

你應該將 `background` 設定為 `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`。

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
