---
id: 61968df2acd5550bf1616c34
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

透過移除 `body` 元素的 `margin` 和 `padding` 來標準化你的頁面尺寸。

# --hints--

你應該給 `body` 一個 `margin` 為 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

你應該給 `body` 一個 `padding` 為 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.padding, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  
}
--fcc-editable-region--
```
