---
id: 61969aa6acef5b12200f672e
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

為你的背景添加兩座山以創建一些風景。

在 `.penguin` 元素上方，添加一個具有 `class` 為 `left-mountain` 的 `div`。

# --hints--

你應該在 `body` 中添加一個新的 `div`。預期會看到 `--fcc-expected--` 個 `div` 元素，但實際發現 `--fcc-actual--` 個。

```js
assert.equal(document.querySelectorAll('body > div')?.length, 3);
```

你應該給 `div` 一個 `class` 為 `left-mountain`。

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

你應該將 `.left-mountain` 放在 `.penguin` 之前。

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
