---
id: 61968f8877c6720d6d61aaf5
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中，為你的 `div` 添加一個 `class` 為 `ground`。

# --hints--

你應該為你的 `div` 添加一個新的。

```js
assert.exists(document.querySelector('div'));
```

你應該給 `div` 一個 `class` 為 `ground`。

```js
assert.include(document.querySelector('div')?.className, 'ground');
```

你應該將 `div` 放置在 `body` 內。

```js
assert.exists(document.querySelector('body > div.ground'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}
```
