---
id: 619665c9abd72906f3ad30f9
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你將會建造一隻快樂的 Flappy Penguin，並在此過程中進一步探索 CSS 變形和動畫。

從你的基本 HTML 範本開始。包含 `DOCTYPE` 宣告、語言設為英文的 `html` 元素、適當的 `meta` 頁籤、`head`、`body` 和 `title` 元素。並且，將你的樣式表聯結到頁面。

# --hints--

你的程式碼應該有一個 `<!DOCTYPE html>` 宣告式。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的程式碼應該包含一個 `html` 元素。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

你的 `html` 元素應該有一個開啟標籤，並且帶有 `lang` 屬性且值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的程式碼應該在 `html` 元素內有一個 `head` 元素。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

你的程式碼應該在 `html` 元素內有一個 `body` 元素。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

你的 `head` 元素應該位於你的 `body` 元素之前。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

一個 `meta` 元素應該設定 `name` 為 `viewport`，並將 `content` 設定為 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設為 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 應該包含一些文字。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

你的程式碼應該有一個 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
