---
id: 61a8fe15a6a31306e60d1e89
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

透過將 `width` 設為 `100%`，並將 `height` 設為 `100vh`，來標準化你的頁面。

# --hints--

你應該給 `body` 一個 `--fcc-expected--` 的 `width`，但發現是 `--fcc-actual--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你應該給 `body` 一個 `height` 為 `--fcc-expected--`，但發現是 `--fcc-actual--`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;

}
--fcc-editable-region--
```
