---
id: 699b61ca8c9cb9bc1968364e
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Ukifungua konsoli, unapaswa kuona ujumbe wa hitilafu ufuatao:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

Katika JavaScript ya kawaida, ulikuwa na ruhusa ya kuongeza sifa na thamani yoyote unayotaka kwenye kitu cha `profile`. Lakini katika TypeScript, mkusanyaji wa msimbo hukagua kwamba kitu hicho kina sifa tu ulizozitaja (`username`, `age`, na `isLoggedIn` katika kesi hii). Kwa kuwa `mood` haikuingizwa katika aina, TypeScript huichukulia kama sifa isiyo halali na hupiga hitilafu.

Ondoa `mood: null` kutoka kwenye msimbo wako.

# --hints--

Haupaswi tena kuwa na `mood: null` katika msimbo wako.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
