---
id: 699b70b10ae344c802d106b0
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Sasa ni wakati wa kuongeza sifa chache zaidi kwenye kitu cha `userRoles`.

Ongeza sifa ya `moderator` yenye thamani ya `"medium-access"`.

Kisha chini yake, ongeza sifa ya `guest` yenye thamani ya `3`.

# --hints--

Kitu chako cha `userRoles` kinapaswa kuwa na sifa ya `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Sifa yako ya `moderator` inapaswa kuwa na thamani ya `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Kitu chako cha `userRoles` kinapaswa kuwa na sifa ya `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

Sifa yako ya `guest` inapaswa kuwa na thamani ya nambari `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
