---
id: 699b65caf17a8e8d340af27d
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Sasa ni wakati wa kuongeza sifa ya mwisho kwenye kitu cha `profile`.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na aina za safu ya data. Ulijifunza kuhusu aina tofauti za safu ya data kama `string[]`, au `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

Katika mfano huu, safu ya `scores` inakubali nambari tu. Ikiwa utajaribu kuongeza aina nyingine kama `string`, au `null` basi TypeScript itatupa kosa.

Sahihisha kitu chako cha `profile` ili kuongeza sifa ya `programmingLanguages`, yenye aina ya `string[]`. Ndani ya kitu hicho, weka sifa hii kuwa `["JavaScript", "Python", "C++"]`.

# --hints--

Unapaswa kuwa na sifa ya `programmingLanguages` ndani ya kitu chako cha `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Sifa yako ya `programmingLanguages` inapaswa kuwa na thamani ya `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Sifa yako ya `programmingLanguages` inapaswa kuwa na aina ya `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
--fcc-editable-region--
  
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
  
--fcc-editable-region--
};

console.log(profile);
```
