---
id: 699b6dd1a2961bd38491ef2e
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Kipengee cha `profile` kilikuwa na funguo na aina kali. Ikiwa ulijaribu kuongeza sifa ambayo haikuainishwa, au kuweka thamani ya aina isiyolingana, TypeScript ingetupa kosa.

Itakuwa vizuri kama kipengee cha `userRoles` kingekuwa na funguo zisizo kali na zenye kubadilika zaidi.

Kama unavyokumbuka katika mafundisho yaliyopita, unaweza kufanikisha matokeo haya kwa kutumia aina ya `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

Katika mfano huu wa `studentGrades`, TypeScript haitaji kujua masomo yote mapema. Hii inakuwezesha kuongeza masomo zaidi kama `Art` au `Music` baadaye. TypeScript inajali tu kwamba thamani za kila funguo ni nambari.

Sahihisha kipengee chako cha `userRoles` kwa kuongeza aina ya `Record` ndani ya mstari. Iweke kuwa `<string, string>` ili kila funguo iwe mfuatano wa herufi wowote, na kila thamani iwe mfuatano wa herufi.

# --hints--

Kipengee chako cha `userRoles` kinapaswa kuwa na aina ya `Record` ndani ya mstari iliyowekwa kuwa `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
