---
id: 699b5370a43e299b4673bcaa
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Itakuwa na msaada kuona mali za kitu cha `profile` kwenye konsoli unapoendelea kuijenga.

Ongeza `console.log` inayorekodi `profile` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kurekodi `profile` kwenye konsoli.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
