---
id: 699b64096301fd75e3efeb1b
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ukifungua tena konsoli, utaona ujumbe tofauti wa kosa:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Kama unavyokumbuka katika mafundisho ya awali, unaweza kuweka sifa kuwa hiari kwa kutumia alama ya `?` karibu na jina la sifa kama hii:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Katika mfano huu, `numberOfAlbums` imewekwa kuwa hiari hivyo TypeScript haitatupa kosa katika kesi hii.

Sahihisha sifa yako ya `bio` ili iwe imewekwa kuwa hiari.

# --hints--

Sifa yako ya `bio` inapaswa kuwekwa kuwa hiari.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
