---
id: 699b68f94a0d154c2a7a5ce0
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa baada ya kitu cha `profile` kukamilika, ni wakati wa kuunda kitu kipya kinachowakilisha majukumu ya watumizi.

Unda kigezo kipya kinachoitwa `userRoles` na ukae thamani ya kitu. Kitu chako kinapaswa kuwa na sifa na thamani zifuatazo:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `userRoles`.

```js
assert.isDefined(userRoles);
```

Kigezo chako cha `userRoles` kinapaswa kuwa kitu.

```js
assert.isObject(userRoles);
```

Kitu chako cha `userRoles` kinapaswa kuwa na sifa ya `admin`.

```js
assert.property(userRoles, "admin");
```

Sifa yako ya `admin` inapaswa kuwa na thamani ya `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Kitu chako cha `userRoles` kinapaswa kuwa na sifa ya `editor`.

```js
assert.property(userRoles, "editor");
```

Sifa yako ya `editor` inapaswa kuwa na thamani ya `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Kitu chako cha `userRoles` kinapaswa kuwa na sifa ya `viewer`.

```js
assert.property(userRoles, "viewer");
```

Sifa yako ya `viewer` inapaswa kuwa na thamani ya `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
