---
id: 699b4e284291adcfcb90df47
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utazoea kufanya kazi na aina za msingi pamoja na aina za vitu na safu za data kwa kujenga wasifu wa mtumizi salama wa aina.

Anza kwa kuunda kigezo kinachoitwa `profile` na kisha weka thamani ya kitu. Kitu chako kinapaswa kuwa na sifa hizi tatu na thamani zao:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `profile`.

```js
assert.exists(profile);
```

Kigezo chako cha `profile` kinapaswa kuwa kitu.

```js
assert.isObject(profile);
```

Unapaswa kuwa na sifa ya `username` katika kitu chako cha `profile`.

```js
assert.property(profile, "username");
```

Sifa yako ya `username` inapaswa kuwa na thamani ya `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Kitu chako cha `profile` kinapaswa kuwa na sifa ya `age`.

```js
assert.property(profile, "age");
```

Sifa yako ya `age` inapaswa kuwa na thamani ya `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Kitu chako cha `profile` kinapaswa kuwa na sifa ya `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Sifa yako ya `isLoggedIn` inapaswa kuwa na thamani ya `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
