---
id: 699b71d98c820ac7ab9eea2a
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Ukifungua tena konsoli, utaona ujumbe mpya wa kosa:

```md
Type 'number' is not assignable to type 'string'.
```

Hii inatokea kwa sababu unajaribu kuongeza sifa ya `guest` yenye thamani ya nambari.

Ili kuondoa kosa la TypeScript, badilisha thamani ya sifa ya `guest` kuwa `"read-only"`.

Baada ya mabadiliko hayo ya mwisho, warsha yako ya wasifu wa mtumizi itakuwa imekamilika!

# --hints--

Sifa yako ya `guest` kwa kitu cha `userRoles` inapaswa kuwa na thamani ya `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
