---
id: 699b62c99ec8c63470313365
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa ni wakati wa kuongeza sifa nyingine kwenye kitu chako cha `profile`.

Anza kwa kusahihisha aina ya kitu ndani ya mstari ili kujumuisha sifa ya `bio` yenye aina iliyowekwa kuwa `string`.

# --hints--

Unapaswa kusahihisha aina ya kitu ndani ya mstari ili kujumuisha sifa ya `bio` iliyowekwa kuwa `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
