---
id: 699b5b7cb68b6d35afca61b7
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Ingawa inawezekana kuendelea kuongeza sifa kama hizi kwenye kitu cha `profile`, itakuwa vizuri kuweka kikomo kwa aina za sifa zinazoruhusiwa. Itakuwa bora kufafanua muundo wa kitu cha `profile` kabla.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na aina za vitu kama hivi:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Katika mfano huu, `composer` ni kitu chenye aina ya kitu ndani ya mstari. Kitu hiki kina sifa tatu muhimu za `name`, `instrument` na `isActive`. Sifa mbili za kwanza lazima ziwe za aina ya `string` na `isActive` lazima iwe ya aina ya `boolean`.

Sahihisha kitu chako cha `profile` kilichopo ili kijumuishe aina ya kitu ndani ya mstari. `username` inapaswa kuwa na aina ya `string`, `age` inapaswa kuwa na aina ya `number` na `isLoggedIn` inapaswa kuwa na aina ya `boolean`. 

# --hints--

Kitu chako cha `profile` kinapaswa kuwa na aina ya kitu ndani ya mstari yenye sifa ya `username` iliyowekwa kuwa `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Kitu chako cha `profile` kinapaswa kuwa na aina ya kitu ndani ya mstari yenye sifa ya `age` iliyowekwa kuwa `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Kitu chako cha `profile` kinapaswa kuwa na aina ya kitu ndani ya mstari yenye sifa ya `isLoggedIn` iliyowekwa kuwa `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
--fcc-editable-region--
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}

console.log(profile);
```
