---
id: 61fd5a93fd62bb35968adeab
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comece seu projeto dando ao seu elemento `head` um elemento `link` para sua folha de estilos.

# --hints--

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
