---
id: 61fd7a160ed17960e971f28b
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

No quarto elemento `tr`, adicione um elemento `th` com o texto `Total Assets`. Envolva o texto `Assets` em um elemento `span` com o atributo `class` definido como `sr-only`.

Depois disso, adicione três elementos `td` com o seguinte texto (em ordem): `$579`, `$736`, `$809`. Dê ao terceiro elemento `td` a `class` com o valor de `current`.

# --hints--

O quarto `tr` deve ter um elemento `th`.

```js
assert.exists(
  document
    .querySelector('tbody')
    ?.querySelectorAll('tr')?.[3]
    ?.querySelector('th')
);
```

O elemento `th` deve conter o texto `Total Assets`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableHeader = tableRow.querySelector('th');
assert.strictEqual(tableHeader.innerText.trim(), 'Total Assets');
```

Você deve envolver o texto `Assets` dentro de um elemento `span`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.strictEqual(span.textContent.trim(), 'Assets');
```

O elemento `span` deve ter o atributo `class` com o valor `sr-only`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.isTrue(span?.classList?.contains('sr-only'));
```

Você deve ter três elementos `td`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableDatas = tableRow.querySelectorAll('td');
assert.lengthOf(tableDatas, 3);
```

O primeiro elemento `td` deve ter o texto `$579`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '$579');
```

O segundo elemento `td` deve ter o texto `$736`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[1];
assert.strictEqual(tableData?.textContent.trim(), '$736');
```

O terceiro elemento `td` deve ter o texto `$809`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.strictEqual(tableData.textContent.trim(), '$809');
```

O terceiro elemento `td` deve ter o atributo `class` definido como `current`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.isTrue(tableData?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
