---
id: 61fd67a656743144844941cb
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Screen readers announce HTML elements based on the document flow. We will eventually want the balance sheet to have a heading of "Balance Sheet" and a subheading of "AcmeWidgetCorp". However, this order does not make sense if announced by a screen reader.

Dê ao `span` existente o atributo `class` definido como `flex` e adicione dois elementos `span` dentro dele. Dê ao primeiro o texto `AcmeWidgetCorp`. Dê ao segundo o texto `Balance Sheet`. Você usará o CSS para inverter a ordem do texto na página, mas a ordem do HTML fará mais sentido para um leitor de tela.

# --hints--

O elemento `span` existente deve ter o atributo `class` definido como `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

O elemento `span` existente deve ter dois novos elementos `span` dentro dele.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Os novos elementos `span` não devem ter um atributo `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

O primeiro novo elemento `span` deve ter o texto `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

O segundo novo elemento `span` deve ter o texto `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
