---
id: 69757cc0faae0152c1418aad
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 숫자와 수학 연산을 사용해 계산기 분할기를 만드는 연습을 합니다. 이 도구는 식사 비용과 팁을 더한 후 각 사람이 얼마를 내야 하는지 계산합니다.

먼저, 비용이 추가될 때마다 총액을 추적할 방법이 필요합니다. Python에서는 시간이 지나면서 변하는 정수(온전한 수)를 저장하기 위해 변수를 사용할 수 있습니다.

예를 들어, 다음과 같이 작성할 수 있습니다:

```py
my_number = 2
```

`running_total`라는 변수를 만들고 값으로 `0`를 할당하세요.

# --hints--

`running_total`라는 변수가 있어야 합니다.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('running_total')
    `))
})
```

정수 `0`를 `running_total` 변수에 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('running_total').is_equivalent('running_total = 0')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
