---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

각 강좌에는 비용이 있습니다. 나중에 사용하려면 이 금액들을 변수에 저장해야 합니다. 이 금액에는 센트 단위가 포함되어 있으므로 소수를 나타내는 데 사용하는 float 타입을 사용할 것입니다. 다음은 float 값을 가진 변수의 예입니다:

```py
change = 2.35
```

네 개의 변수를 만드세요: `appetizers`는 `37.89`로, `main_courses`는 `57.34`로, `desserts`는 `39.39`로, `drinks`는 `64.21`로 각각 할당하세요.

# --hints--

`appetizers`라는 이름의 변수를 정의하고 `37.89` 값을 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('appetizers').is_equivalent('appetizers = 37.89')
    `))
})
```

`main_courses`라는 이름의 변수를 정의하고 `57.34` 값을 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('main_courses').is_equivalent('main_courses = 57.34')
    `))
})
```

`desserts`라는 이름의 변수를 정의하고 `39.39` 값을 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('desserts').is_equivalent('desserts = 39.39')
    `))
})
```

`drinks`라는 이름의 변수를 정의하고 `64.21` 값을 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('drinks').is_equivalent('drinks = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
