---
id: 6976cc0b0c135686a4620b94
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

다음으로, 계산서를 함께 나누는 사람 수를 고려해야 합니다. 이전 단계에서 했던 것처럼 이 값을 변수에 저장하세요.

`num_of_friends`라는 변수를 만들고 `4` 값을 할당하세요. 이 변수는 워크숍에서 최종 분할 금액을 계산할 때 사용됩니다.

# --hints--

`num_of_friends`라는 변수가 있어야 합니다.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_variable('num_of_friends')
    `))
})
```

정수 `4`를 `num_of_friends` 변수에 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('num_of_friends').is_equivalent('num_of_friends = 4')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

--fcc-editable-region--

--fcc-editable-region--
```
