---
id: 6977a8bbcbdd557886bd0718
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

팁이 포함되었으니, 이제 전체 그룹의 최종 금액이 있습니다. 총 청구 금액을 친구 수로 나누어 각자가 얼마를 내야 하는지 계산하세요.

Python에서는 나누기를 수행할 때 슬래시 `/`를 사용합니다. 예를 들어:

```py
half = 10 / 2
```

`final_bill`를 `running_total`으로 나눈 결과를 `num_of_friends`이라는 변수에 할당하세요.

마지막으로 `print()` 함수를 사용해 문자열 `Bill per person:`와 공백, 그리고 `final_bill` 값을 출력하세요.

# --hints--

`final_bill`라는 변수를 정의해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

`/` 연산자를 사용해 `running_total`를 `num_of_friends`으로 나누고 결과를 `final_bill` 변수에 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

문자열 `Bill per person:`와 공백, 그리고 `final_bill` 변수를 출력해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
