---
id: 6977a4306bb3e856690a4890
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

팁을 계산했으니, `running_total`에 더해서 최종 청구 금액을 구하세요.

`+=` 연산자를 사용해 `tip`의 값을 `running_total`에 더하세요. 마지막으로 `print()`를 사용해 문자열 `Total with tip:`와 공백, 그리고 `running_total`의 값을 출력하세요.

# --hints--

`+=`의 값을 `tip`에 더하려면 `running_total` 연산자를 사용하세요.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

문자열 `Total with tip:`와 공백, 그리고 `running_total` 변수를 출력하세요.

```js
({
    test: () => runPython(`
import io, contextlib, re

buffer = io.StringIO()
with contextlib.redirect_stdout(buffer):
    exec(_code)

match = re.search(r"Total with tip: ([0-9]+(?:\\.[0-9]+)?)", buffer.getvalue())

assert match
assert abs(float(match.group(1)) - ((37.89 + 57.34 + 39.39 + 64.21) * 1.25)) < 1e-6
    `)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
