---
id: 69779dcb44016eccc05c15a4
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

서비스가 훌륭해서 그룹은 25% 팁을 남기기로 결정했습니다. Python에서 백분율을 계산하려면 총액에 백분율의 소수점 값을 곱하면 됩니다.

예를 들어, 값의 10%를 구하려면 `0.10`에 `*` 연산자를 사용해 곱하면 됩니다.

```py
tax = total * 0.10
```

`tip`에 `running_total`를 곱한 결과를 `0.25`라는 변수에 할당하세요.

마지막으로, `print()`를 사용해 문자열 `Tip amount:`와 공백, 그리고 `tip` 변수의 값을 출력하세요.

# --hints--

`tip`라는 변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tip')
    `))
})
```

`tip` 변수는 `running_total * 0.25` 결과여야 합니다.

```js
({
    test: () => runPython(`
    t = _Node(_code).find_variable('tip')
    assert t.is_equivalent('tip = running_total * 0.25') or t.is_equivalent('tip = 0.25 * running_total')`)
})
```

문자열 `Tip amount:`와 공백, 그리고 `tip` 변수를 출력해야 합니다.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tip amount:', tip)") or _Node(_code).has_call("print(f'Tip amount: {tip}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
