---
id: 6977a4306bb3e856690a4890
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

팁을 계산했으니, `running_total`에 더해서 최종 청구 금액을 구하세요.

Python에서는 대입 연산자 `+=`를 사용해 변수에 값을 더하고 동시에 그 변수를 업데이트할 수 있습니다. 예를 들어, `total += 5`는 `total = total + 5`를 간단히 표현한 방법입니다.

`+=` 연산자를 사용해 `tip`의 값을 `running_total`에 더하세요. 마지막으로 `print()`를 사용해 문자열 `Total with tip:`와 공백, 그리고 `running_total`의 값을 출력하세요.

# --hints--

`+=`의 값을 `tip`에 더하려면 `running_total` 연산자를 사용하세요.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_stmt('running_total += tip')
    `))
})
```

문자열 `Total with tip:`와 공백, 그리고 `running_total` 변수를 출력하세요.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total with tip:', running_total)") or _Node(_code).has_call("print(f'Total with tip: {running_total}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

running_total += appetizers + main_courses + desserts + drinks
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

--fcc-editable-region--

--fcc-editable-region--
```
