---
id: 6982684f3a25f379e195a5fc
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

개별 비용을 저장했으니 이제 총합을 계산할 수 있습니다. Python에서는 덧셈 연산자 `+`를 사용해 값을 더합니다.

`+=` 연산자는 기존 변수에 값을 더하고 동시에 업데이트합니다. 예를 들어:

```py
total = 10
total += 2 + 2 + 1
print(total)  # total is now 15
```

`+=` 연산자를 한 번 사용해 `appetizers`, `main_courses`, `desserts`, `drinks`를 `running_total`에 더하세요.

마지막으로 `print()`를 사용해 문자열 `Total bill so far:`와 공백, 그리고 `running_total`의 값을 출력하세요.

**참고:** 출력에 예상보다 더 많은 10진수 숫자가 표시될 수 있습니다. 이전 수업에서 배운 것처럼, 숫자는 이진수로 저장되며 많은 10진수 값이 이 형식으로 정확히 나타낼 수 없기 때문에 반올림 오류가 발생합니다.

# --hints--

`running_total`를 한 번만 `+=` 연산자를 사용해 네 개의 강좌 비용을 모두 더해 업데이트하세요.

```js
({
       test: () => runPython(`
import itertools  
perms = itertools.permutations(['+ appetizers', '+ main_courses', '+ desserts', '+ drinks'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'running_total += {v}' for v in values)  
assert any(_Node(_code).has_stmt(s) for s in solutions)`)
})
```

문자열 `Total bill so far:`와 공백, 그리고 `running_total`의 값을 출력해야 합니다.

```js
({
    test: () => runPython(`
    sol1 = "print('Total bill so far:', running_total)"
    sol2 = "print(f'Total bill so far: {running_total}')"
    assert _Node(_code).has_call(sol1) or _Node(_code).has_call(sol2)`)
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

appetizers = 37.89
main_courses = 57.34
desserts = 39.39
drinks = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
