---
id: a6e40f1041b06c996f7b2406
title: Erstellen Sie einen Finder für das erste Element
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die ein Array durchsucht und das erste Element zurückgibt, das eine Testfunktion besteht (auch bekannt als „Wahrheitstest“).

Die Funktion soll das Array durchlaufen und jedes Element mit der bereitgestellten Testfunktion prüfen. Am Ende gibt sie das erste Element zurück, bei dem die Testfunktion `true` ergibt.

Beispiel:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `findElement` haben, die ein Array und eine Funktion als Argumente akzeptiert.
2. Die Funktion sollte das erste Element im Array zurückgeben, das einen Wahrheitstest besteht. Das bedeutet, dass beim Aufrufen der übergebenen Funktion `func` mit einem Element `x` der Wahrheitstest bestanden ist, wenn `func(x)` `true` ergibt.
3. Wenn kein Element den Test besteht, sollte die Funktion `undefined` zurückgeben.

# --hints--

Sie sollten eine `findElement`-Funktion haben.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` sollte `8` zurückgeben.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` sollte `undefined` zurückgeben.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` sollte `3` zurückgeben.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` sollte `"javascript"` zurückgeben.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` sollte `undefined` zurückgeben.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` sollte `2` zurückgeben.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` sollte `undefined` zurückgeben.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
