---
id: 587d7dbd367417b2b2512bb5
title: SassでCSSをネストする
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

SassではCSSルールをネストすることができ、スタイルシートを整理するのに便利な方法です。

通常、各要素は別々の行でターゲットにしてスタイルを指定します。例えば次のようにします。

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

大きなプロジェクトでは、CSSファイルに多くの行とルールが含まれます。ここでネストを使うと、子のスタイルルールをそれぞれの親要素の中に配置してコードを整理できます。

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

上記のネスト手法を使って、`.blog-post`要素の両方の子のCSSルールを再整理してください。テストのために、`h1`は`p`の前に来るようにしてください。

# --hints--

コードは`h1`と`p`が`.blog-post`の親要素の中にネストされるように再整理してください。

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
