---
id: 587d7dbe367417b2b2512bb8
title: スタイルにロジックを追加するために@ifと@elseを使う
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Sassの`@if`ディレクティブは特定のケースをテストするのに便利で、JavaScriptの`if`文と同じように動作します。

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

そしてJavaScriptと同様に、`@else if`と`@else`ディレクティブはさらに多くの条件をテストします：

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

`border-stroke`という名前のミックスインを作成し、パラメータ`$val`を受け取ってください。ミックスインは`@if`、`@else if`、および`@else`ディレクティブを使って以下の条件をチェックする必要があります：

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

もし`$val`パラメータの値が`light`、`medium`、または`heavy`でなければ、`border`プロパティを`none`に設定してください。

# --hints--

`border-stroke`という名前のミックスインを宣言し、`$val`というパラメータを持つコードを書く必要があります。

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

ミックスインには`@if`文があり、`$val`が`light`かどうかをチェックし、`border`を`1px solid black`に設定してください。

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

ミックスインには`@else if`文があり、`$val`が`medium`かどうかをチェックし、`border`を`3px solid black`に設定してください。

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

ミックスインには`@else if`文があり、`$val`が`heavy`かどうかをチェックし、`border`を`6px solid black`に設定してください。

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

ミックスインには`@else`文があり、`border`を`none`に設定してください。

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
