---
id: 587d7fa5367417b2b2512bbd
title: あるCSSスタイルセットを別の要素に拡張する
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sassには`extend`という機能があり、ある要素のCSSルールを簡単に借りて、別の要素でそれを基にスタイルを作成できます。

例えば、以下のCSSルールのブロックは`.panel`クラスにスタイルを適用しています。`background-color`、`height`、`border`があります。

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

今、`.big-panel`という別のパネルが必要です。`.panel`と同じ基本プロパティを持ちますが、`width`と`font-size`も必要です。`.panel`から最初のCSSルールをコピー＆ペーストすることもできますが、パネルの種類が増えるとコードが繰り返しになってしまいます。`extend`ディレクティブは、ある要素に書かれたルールを再利用し、別の要素に追加のスタイルを加える簡単な方法です。

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel`は`.panel`と同じプロパティに加えて、新しいスタイルも持ちます。

# --instructions--

`.info-important`クラスを作成し、`.info`を拡張し、さらに`background-color`をマゼンタに設定してください。

# --hints--

`info-important`クラスは`background-color`を`magenta`に設定する必要があります。

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

`info-important`クラスは`@extend`を使って`info`クラスのスタイルを継承してください。

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
