---
id: 587d7dbf367417b2b2512bba
title: リスト内のアイテムに対して @each を使ってマッピングする
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

前回のチャレンジでは、`@for` ディレクティブが開始値と終了値を使って指定回数ループする方法を示しました。Sass にはリストやマップの各アイテムに対してループする `@each` ディレクティブもあります。各イテレーションで、変数にはリストやマップの現在の値が代入されます。

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

マップは少し異なる構文を持ちます。例を示します：

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

マップのキーを参照するには `$key` 変数が必要なことに注意してください。そうしないと、コンパイルされた CSS に `color1`、`color2`... といったものが入ってしまいます。上記のコード例はどちらも次の CSS に変換されます：

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

`@each` ディレクティブを使ってリスト `blue, black, red` をループし、各変数を `.color-bg` クラスに代入してください。`color` の部分はそれぞれの色に応じて変わります。各クラスは `background-color` をそれぞれの色に設定してください。

# --hints--

コードには `@each` ディレクティブを使う必要があります。

```js
assert(code.match(/@each /g));
```

`.blue-bg` クラスは `background-color` が blue である必要があります。

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

`.black-bg` クラスは `background-color` が black である必要があります。

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

`.red-bg` クラスは `background-color` が red である必要があります。

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
