---
id: 587d7dbe367417b2b2512bb9
title: "@for を使って Sass のループを作成する"
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

`@for` ディレクティブは、JavaScript の `for` ループに非常によく似たループでスタイルを追加します。

`@for` は「start through end」と「start to end」の2通りの使い方があります。主な違いは、「start **to** end」がカウントに終了番号を含まないのに対し、「start **through** end」は終了番号をカウントに含むことです。

こちらは start **through** end の例です：

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

`#{$i}` の部分は、変数（`i`）とテキストを組み合わせて文字列を作る構文です。Sass ファイルが CSS に変換されると、次のようになります：

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

これはグリッドレイアウトを作成する強力な方法です。これで CSS クラスとして12種類のカラム幅のオプションが使えるようになりました。

# --instructions--

`@for` ディレクティブを使って、変数 `$j` が 1 **to** 6 まで変化するループを書いてください。

`.text-1` から `.text-5` までの5つのクラスを作成し、それぞれの `font-size` はインデックスに15pxを掛けた値にしてください。

# --hints--

コードには `@for` ディレクティブを使う必要があります。

```js
assert(code.match(/@for /g));
```

`.text-1` クラスは `font-size` が 15px であるべきです。

```js
assert($('.text-1').css('font-size') == '15px');
```

`.text-2` クラスは `font-size` が 30px であるべきです。

```js
assert($('.text-2').css('font-size') == '30px');
```

`.text-3` クラスは `font-size` が 45px であるべきです。

```js
assert($('.text-3').css('font-size') == '45px');
```

`.text-4` クラスは `font-size` が 60px であるべきです。

```js
assert($('.text-4').css('font-size') == '60px');
```

`.text-5` クラスは `font-size` が 75px であるべきです。

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
