---
id: 587d7dbd367417b2b2512bb6
title: ミックスインで再利用可能なCSSを作成する
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

Sassでは、<dfn>ミックスイン</dfn>とは、スタイルシート全体で再利用できるCSS宣言のグループです。定義は`@mixin`アットルールで始まり、その後にカスタム名が続きます。ミックスインは`@include`アットルールを使って適用します。

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

コンパイルすると以下のようになります：

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

ミックスインは引数を取ることもでき、その動作をカスタマイズできます。ミックスインを使う際には引数が必要です。

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

パラメータにデフォルト値を設定すると、引数を省略可能にできます。

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

`shape`という名前のミックスインを作成し、3つのパラメータ：`$w`、`$h`、`$bg-color`を指定してください。

`shape`ミックスインを使って、`#square`要素に幅と高さを`50px`、背景色を`red`に設定してください。`#rect-a`要素には幅を`100px`、高さを`50px`、背景色を`blue`に設定してください。最後に、`#rect-b`要素には幅を`50px`、高さを`100px`、背景色を`orange`に設定してください。

# --hints--

`shape`という名前のミックスインを3つのパラメータ：`$w`、`$h`、`$bg-color`で宣言してください。

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

ミックスインには`width`プロパティを含め、`$w`パラメータを使ってください。

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

ミックスインには`height`プロパティを含め、`$h`パラメータを使ってください。

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

ミックスインには`background-color`プロパティを含め、`$bg-color`パラメータを使ってください。

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

`#square`セレクター内のスタイルを、`shape`ミックスインの`@include`キーワードを使った呼び出しに置き換えてください。幅と高さは`50px`、背景色は`red`に設定します。

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

`#rect-a`セレクター内のスタイルを、`shape`ミックスインの`@include`キーワードを使った呼び出しに置き換えてください。幅は`100px`、高さは`50px`、背景色は`blue`に設定します。

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

`#rect-b`セレクター内のスタイルを、`shape`ミックスインの`@include`キーワードを使った呼び出しに置き換えてください。幅は`50px`、高さは`100px`、背景色は`orange`に設定します。

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
