---
id: 587d7dbf367417b2b2512bbc
title: "パーシャルでスタイルを小さなチャンクに分割する"
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>パーシャル</dfn>とは、SassでCSSコードの一部を保持する別のファイルのことです。これらは他のSassファイルにインポートして使います。似たコードをモジュールとしてまとめて整理するのに最適な方法です。

パーシャルの名前はアンダースコア（`_`）で始まり、Sassにこれは小さなCSSの一部であり、CSSファイルに変換しないことを伝えます。また、Sassファイルは`.scss`の拡張子で終わります。パーシャルのコードを別のSassファイルに取り込むには、`@import`ディレクティブを使います。

例えば、すべてのミックスインが「\_mixins.scss」というパーシャルに保存されていて、「main.scss」ファイルで使いたい場合、メインファイルでの使い方は次の通りです。

```scss
@import 'mixins'
```

アンダースコアと拡張子は`import`文には不要です。Sassはそれがパーシャルだと理解します。パーシャルがファイルにインポートされると、すべての変数、ミックスイン、その他のコードが使えるようになります。

# --instructions--

main.scssファイルに「`@import`」という名前のパーシャルをインポートする`_variables.scss`文を書いてください。

# --hints--

コードは`@import`ディレクティブを使い、ファイル名にアンダースコアは含めないでください。

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
