---
id: 587d7dbd367417b2b2512bb4
title: Sass変数でデータをストアする
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

SassのCSSと異なる機能の1つは、変数を使うことです。変数は宣言されてデータをストアするために設定され、JavaScriptに似ています。

JavaScriptでは、変数は`let`と`const`キーワードを使って定義されます。Sassでは、変数は`$`に続けて変数名を書きます。

いくつか例を示します。

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

そして変数を使うには：

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

変数が役立つ例の1つは、多くの要素を同じ色にしたい場合です。その色を変えたいときは、コードの編集は変数の値だけで済みます。

# --instructions--

`$text-color`という変数を作成し、`red`に設定してください。それから`color`と`.blog-post`の`h2`プロパティの値を`$text-color`変数に変更してください。

# --hints--

`$text-color`のSass変数が`red`の値で宣言されている必要があります。

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

`$text-color`変数を使って、`color`と`.blog-post`のアイテムの`h2`を変更してください。

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

`.blog-post`要素の`color`はredである必要があります。

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

`h2`要素の`color`はredである必要があります。

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
