---
id: 587d7dbf367417b2b2512bbb
title: 条件が満たされるまでスタイルを適用する @while
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

`@while` ディレクティブは、JavaScript の `while` ループと似た機能を持つオプションです。条件が満たされるまで CSS ルールを作成します。

`@for` チャレンジでは、シンプルなグリッドシステムを作成する例がありました。これは `@while` でも同様に機能します。

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

まず、変数 `$x` を定義し、1 に設定します。次に、`@while` ディレクティブを使って、`$x` が 13 未満の間グリッドシステムを作成します。`width` の CSS ルールを設定した後、無限ループを避けるために `$x` を 1 増やします。

# --instructions--

`@while` を使って、異なる `font-sizes` を持つ一連のクラスを作成してください。

`text-1` から `text-5` までの 5 つの異なるクラスが必要です。次に、`font-size` を現在のインデックス番号に `15px` を掛けた値に設定してください。無限ループにならないように注意してください！

# --hints--

コードには `@while` ディレクティブを使う必要があります。

```js
assert(code.match(/@while /g));
```

コードには、インデックス変数を使い、1 から始めてください。

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

カウンター変数をインクリメントしてください。

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

`.text-1` クラスは `font-size` が `15px` である必要があります。

```js
assert($('.text-1').css('font-size') == '15px');
```

`.text-2` クラスは `font-size` が `30px` である必要があります。

```js
assert($('.text-2').css('font-size') == '30px');
```

`.text-3` クラスは `font-size` が `45px` である必要があります。

```js
assert($('.text-3').css('font-size') == '45px');
```

`.text-4` クラスは `font-size` が `60px` である必要があります。

```js
assert($('.text-4').css('font-size') == '60px');
```

`.text-5` クラスは `font-size` が `75px` である必要があります。

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
