---
id: 68e707257ea91c2e12d489fb
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви потренуєтеся додавати кастомні стилі до радіокнопок, створюючи форму для конференції батьків і вчителів. Шаблонний код HTML уже надано.

Почніть із додавання `main` елемента з класом `container`.

# --hints--

У вас має бути `main` елемент.

```js
assert.exists(document.querySelector("main"));
```

Ваш `main` елемент повинен мати клас `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
