---
id: 68e7e729f6d14f69a23786ef
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Далі, всередині вашого елемента `main`, додайте елемент `h1` з текстом `Parent Teacher Conference Form` та класами `title` і `center`. 

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Ваш елемент `h1` має містити текст `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Ваш елемент `h1` має мати класи `title` і `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Ваш елемент `h1` має бути всередині елемента `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
