---
id: 68ea9f01c51b66b937793611
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Далі додайте ще одну мітку з класом `"contact-method"` та атрибутом `for`, встановленим у `"phone"`. Текст мітки має бути `Phone: `.

Під вашим елементом `label` додайте радіокнопку з атрибутами `id` і `value`, встановленими у `"phone"`. Атрибут `name` має бути встановлений у `"contact-method"`, а клас — у `"contact-method-radio-btn"`.

# --hints--

У вас має бути два елементи `label` у останньому `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

Елемент `label` має мати клас `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Атрибут `label` елемента `for` має бути встановлений у `"phone"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

Елемент `label` має містити текст `Phone: `.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

У вас має бути два елементи `input` всередині останнього `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

Ваш елемент `input` має мати атрибут `type`, встановлений у `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Елемент `input` має мати атрибут `id`, встановлений у `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

Елемент `input` має мати атрибут `value`, встановлений у `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

Елемент `input` має мати атрибут `name`, встановлений у `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Елемент `input` має мати клас `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
