---
id: 68ea8c330ed07b60879f3d38
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Тепер настав час додати елементи форми для збору інформації про оцінку студента.

Почніть із додавання ще одного `label` елемента з текстом `Student Grade: ` і атрибутом `for`, встановленим у `"grade"`.

Потім, під вашим `label` елементом, додайте `input` елемент з атрибутом `type`, встановленим у `"number"`. Атрибути `name` і `id` повинні бути встановлені у `"grade"`. Атрибут `placeholder` має бути встановлений у `"E.g., 4"`. Нарешті, ваш числовий `input` повинен бути обов’язковим.

# --hints--

У вас має бути загалом два `label` елементи на сторінці.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Ваш другий `label` елемент має містити текст `Student Grade: `. Не забудьте пробіл після двокрапки.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Ваш другий `label` елемент має мати атрибут `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Атрибут `label` другого `for` елемента має мати значення `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

У вас має бути два `input` елементи на сторінці.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Ваш другий `input` елемент має бути розташований під другим `label` елементом.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ваш другий `input` елемент має мати атрибут `type`, встановлений у `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Ваш другий `input` елемент має мати атрибут `name`, встановлений у `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Ваш другий `input` елемент має мати атрибут `id`, встановлений у `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Ваш другий `input` елемент має мати атрибут `placeholder`, встановлений у `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Ваш другий `input` елемент має мати атрибут `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
