---
id: 68ea8a5db73b155708992687
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Наступний крок — додати пов’язаний елемент `input` для інформації про студента.

Почніть із додавання елемента `input` з атрибутом `type`, встановленим у `"text"`. Потім додайте атрибути `name` і `id`, обидва встановлені в `student-name`.

Далі додайте атрибут `placeholder`, встановлений у `E.g., Jane Doe`. І нарешті, додайте атрибут `required`.

# --hints--

У вас має бути елемент `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Ваш елемент `input` має мати атрибут `type`, встановлений у `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ваш елемент `input` має мати атрибут `name`, встановлений у `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Ваш елемент `input` має мати атрибут `id`, встановлений у `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Ваш елемент `input` має мати атрибут `placeholder`, встановлений у `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

Ваш елемент `input` має мати атрибут `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
