---
id: 68ea7c41f2555e16174e7cfa
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Далі додайте елемент абзацу з текстом `Please fill out the form below to help schedule your parent-teacher conference.`. Ваш елемент абзацу також має містити класи `description` і `center`.

# --hints--

У вас має бути елемент `p`.

```js
assert.exists(document.querySelector("p"));
```

Ваш елемент `p` має містити текст `Please fill out the form below to help schedule your parent-teacher conference.`.

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

Ваш елемент `p` має містити класи `description` і `center`.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
