---
id: 68ea8fdf2c9b167307e7f671
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Тепер час додати елементи форми для збору інформації про батьківський елемент.

Почніть із додавання `label` елемента з текстом `Parent/Guardian Name: ` та атрибутом `for`, встановленим у `"parent-name"`.

Потім, під вашим `label` елементом, додайте `input` елемент з атрибутом `type`, встановленим у `"text"`. Атрибути `name` і `id` мають бути встановлені у `"parent-name"`. Атрибут `placeholder` має бути встановлений у `"E.g., Nancy Doe"`. Нарешті, ваш `input` має бути обов’язковим.

# --hints--

У вас має бути загалом три `label` елементи на сторінці.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Ваш третій `label` елемент має містити текст `Parent/Guardian Name: `. Не забудьте про пробіл після двокрапки.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Ваш третій `label` елемент має мати атрибут `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Атрибут `label` вашого третього `for` елемента має мати значення `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

У вас має бути три `input` елементи на сторінці.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Ваш третій `input` елемент має бути розташований під вашим третім `label` елементом.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Ваш третій `input` елемент має мати атрибут `type`, встановлений у `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Ваш третій `input` елемент має мати атрибут `name`, встановлений у `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Ваш третій `input` елемент має мати атрибут `id`, встановлений у `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Ваш третій `input` елемент має мати атрибут `placeholder`, встановлений у `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Ваш третій `input` елемент має мати атрибут `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
