---
id: 68ea8666b8b473455f5b33ef
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Перший розділ форми зосередиться на інформації про студента, такій як його ім’я та клас.

Всередині елемента `form` додайте елемент `fieldset`. Усередині цього елемента `fieldset` додайте елемент `legend` з текстом `Student Information`.

Під елементом `legend` додайте елемент `label` з текстом `Full Name: ` та атрибутом `for` зі значенням `student-name`.


# --hints--

У вас має бути елемент `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

Ваш елемент `fieldset` має бути всередині елемента `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

У вас має бути елемент `legend`.

```js
assert.exists(document.querySelector("legend"));
```

Ваш елемент `legend` має бути всередині елемента `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Ваш елемент `legend` має містити текст `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

У вас має бути елемент `label`.

```js
assert.exists(document.querySelector("label"));
```

Ваш елемент `label` має містити текст `Full Name: `. Не забудьте пробіл після двокрапки.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Ваш елемент `label` має мати атрибут `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

Атрибут `label` вашого елемента `for` має мати значення `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Ваш елемент `label` має бути розташований нижче елемента `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
