---
id: 68ea975c7cbbb894a386e402
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Наступний крок — додати елементи `label` і `input` для методу контакту через е-пошту.

Почніть із додавання мітки з класом `"contact-method"` та атрибутом `for`, встановленим у `"email"`. Текст мітки має бути `Email: `.

Під вашим елементом `label` додайте радіокнопку з атрибутами `id` і `value`, встановленими у `"email"`. Атрибут `name` має бути встановлений у `"contact-method"`, а клас — у `"contact-method-radio-btn"`.

Нарешті, переконайтеся, що ця радіокнопка вибрана за замовчуванням.

# --hints--


Ваш останній набір полів має містити елемент `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

Цей новий елемент `label` (для е-пошти) має мати клас `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

Атрибут `for` нового елемента `label` має бути встановлений у `"email"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

Елемент `label` має містити текст `Email: `. Переконайтеся, що після двокрапки є пробіл.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

Усередині останнього `fieldset` має бути елемент `input`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

Ваш елемент `input` має мати атрибут `type`, встановлений у `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

Елемент `input` має мати атрибут `id`, встановлений у `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

Елемент `input` має мати атрибут `value`, встановлений у `"email"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

Елемент `input` має мати атрибут `name`, встановлений у `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

Елемент `input` має мати клас `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

Ваша кнопка `radio` має бути `checked` за замовчуванням.

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
