---
id: 68eb3a06c6adde47584e2721
title: –Ъ—А–Њ–Ї 35
challengeType: 0
dashedName: step-35
---

# --description--

–І–∞—Б —Б—В–Є–ї—Ц–Ј—Г–≤–∞—В–Є –Ї–љ–Њ–њ–Ї—Г –≤—Ц–і–њ—А–∞–≤–Ї–Є.

–Т–Є–±–µ—А—Ц—В—М –µ–ї–µ–Љ–µ–љ—В —Ц–Ј –Ї–ї–∞—Б–Њ–Љ `submit-btn`. –Т—Б—В–∞–љ–Њ–≤—Ц—В—М `cursor` —Г `pointer` (—Й–Њ–± –њ–Њ–Ї–∞–Ј–∞—В–Є, —Й–Њ –≤—Ц–љ –∞–Ї—В–Є–≤–љ–Є–є), `background-color` —Г `royalblue`, `color` —Г `whitesmoke`, `border` —Г `none`, `border-radius` —Г `6px` —Ц `padding` —Г `12px 20px`.

# --hints--

–Т–Є –њ–Њ–≤–Є–љ–љ—Ц –љ–∞—Ж—Ц–ї–Є—В–Є—Б—П –љ–∞ –µ–ї–µ–Љ–µ–љ—В —Ц–Ј –Ї–ї–∞—Б–Њ–Љ `submit-btn`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.submit-btn'));
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `cursor` —Г `pointer`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.strictEqual(style?.cursor, 'pointer');
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `background-color` —Г `royalblue`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.strictEqual(style?.backgroundColor, 'royalblue');
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `color` —Г `whitesmoke`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.strictEqual(style?.color, 'whitesmoke');
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `border` —Г `none`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.oneOf(style?.getPropVal('border'), ['none', 'medium', 'medium none']);
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `border-radius` —Г `6px`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.strictEqual(style?.borderRadius, '6px');
```

–Т—Б—В–∞–љ–Њ–≤—Ц—В—М –≤–ї–∞—Б—В–Є–≤—Ц—Б—В—М `padding` —Г `12px 20px`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.submit-btn');
assert.strictEqual(style?.padding, '12px 20px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}

.contact-method-radio-btn::before {
  display: block;
  content: " ";
  width: 10px;
  height: 10px;
  border-radius: 50%;
  transform: translate(3px, 3px) scale(0);
  transition: all 0.3s ease-in;
}

.contact-method-radio-btn:checked::before {
  transform: translate(3px, 3px) scale(1);
  background-color: lightgreen;
}

--fcc-editable-region--

--fcc-editable-region--
```
