---
id: 68eb058081504963447bfa60
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Далі слід стилізувати головний контейнер. Виберіть елемент із класом `container` і встановіть його `background-color` у шістнадцятковий код `#ffffff1a`.

Шістнадцятковий код — це шестизначне поєднання цифр і літер, яке використовується в HTML і CSS для представлення кольорів. Він починається з `#`, за яким ідуть три пари:

- Перші дві цифри представляють червоний колір
- Наступні дві — зелений
- Останні дві — синій

Наприклад, `#ffffff` — це чистий білий, бо має максимальне значення для червоного, зеленого та синього.

Також можна додати дві додаткові цифри в кінці, щоб контролювати прозорість (це називається альфа-канал). У `#ffffff1a` `1a` робить білий колір напівпрозорим. Чим нижче значення альфа, тим прозорішим здається колір. Ви дізнаєтеся більше про шістнадцяткові коди пізніше!

# --hints--

Слід націлитися на елемент із класом `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Встановіть властивість `background-color` у `#ffffff1a`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
