---
id: 684aaf9ec670c68d20efd0d0
title: Costruisci un gestore di configurazione utente
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

In questo laboratorio costruirai un gestore di configurazione utente che permette agli utenti di gestire le loro impostazioni come tema, lingua e notifiche. Implementerai funzioni per aggiungere, aggiornare, eliminare e visualizzare le impostazioni utente.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti definire una funzione chiamata `add_setting` con due parametri che rappresentano un dizionario di impostazioni e una tupla contenente una coppia chiave-valore

1. La funzione `add_setting` dovrebbe:
    
   - Convertire la chiave e il valore in minuscolo.
   - Se l'impostazione della chiave esiste, restituisci `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - Se l'impostazione della chiave non esiste, aggiungi la coppia chiave-valore al dizionario di impostazioni fornito e restituisci `Setting '[key]' added with value '[value]' successfully!`.
   - I messaggi restituiti dovrebbero avere la chiave e il valore in minuscolo.

1. Dovresti definire una funzione chiamata `update_setting` con due parametri che rappresentano un dizionario di impostazioni e una tupla contenente una coppia chiave-valore.

1. La funzione `update_setting` dovrebbe:
       
   - Convertire la chiave e il valore in minuscolo.
   - Se l'impostazione della chiave esiste, aggiorna il suo valore nel dizionario di impostazioni fornito e restituisci: `Setting '[key]' updated to '[value]' successfully!`
   - Se l'impostazione della chiave non esiste, restituisci `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - I messaggi restituiti dovrebbero avere la chiave e il valore in minuscolo.


1. Dovresti definire una funzione chiamata `delete_setting` con due parametri che rappresentano un dizionario di impostazioni e una chiave.

1. La funzione `delete_setting` dovrebbe:
    
   - Convertire la chiave passata in minuscolo.
   - Se l'impostazione della chiave esiste, rimuovi la coppia chiave-valore dal dizionario di impostazioni fornito e restituisci `Setting '[key]' deleted successfully!`
   - Se l'impostazione della chiave non esiste, restituisci `Setting not found!`
   - I messaggi restituiti dovrebbero avere la chiave in minuscolo.

1. Dovresti definire una funzione chiamata `view_settings` con un parametro che rappresenta un dizionario di impostazioni.

1. La funzione `view_settings` dovrebbe:
   
   - Restituire `No settings available.` se il dizionario di impostazioni fornito è vuoto. 
   - Se il dizionario contiene delle impostazioni, restituisci una stringa che mostra le impostazioni. La stringa dovrebbe iniziare con `Current User Settings:` seguita dalle coppie chiave-valore, ciascuna su una nuova riga e con la chiave maiuscola. Per esempio, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` dovrebbe restituire:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Per testare il codice, dovresti creare un dizionario chiamato `test_settings` per archiviare alcune preferenze di configurazione utente.

# --hints--

Dovresti creare un dizionario chiamato `test_settings` e aggiungere alcuni valori.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Dovresti definire una funzione chiamata `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

La funzione `add_setting` dovrebbe avere due parametri.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` dovrebbe convertire la chiave in minuscolo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` dovrebbe convertire il valore in minuscolo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` dovrebbe restituire il messaggio di errore `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` dovrebbe aggiungere una nuova coppia chiave-valore e restituire il messaggio di successo `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` dovrebbe aggiungere correttamente la coppia chiave-valore fornita al dizionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Dovresti definire una funzione chiamata `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

La funzione `update_setting` dovrebbe avere due parametri.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

La funzione `update_setting` dovrebbe convertire la chiave in minuscolo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

La funzione `update_setting` dovrebbe convertire il valore in minuscolo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` dovrebbe aggiornare una chiave esistente e restituire il messaggio di successo `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` dovrebbe restituire il messaggio di errore `Setting 'volume' does not exist! Cannot update a non-existing setting.` quando la chiave non esiste.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` dovrebbe aggiornare correttamente la coppia chiave-valore fornita nel dizionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Dovresti definire una funzione chiamata `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

La funzione `delete_setting` dovrebbe avere due parametri.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` dovrebbe convertire la chiave in minuscolo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` dovrebbe rimuovere la chiave esistente e restituire il messaggio di successo `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` dovrebbe restituire il messaggio di errore `Setting not found!` quando la chiave non esiste.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` dovrebbe rimuovere correttamente la chiave fornita dal dizionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Dovresti definire una funzione chiamata `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

La funzione `view_settings` dovrebbe avere un parametro.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` dovrebbe restituire il messaggio `No settings available.` se il dizionario fornito è vuoto.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` dovrebbe restituire le impostazioni formattate per un dizionario non vuoto.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` dovrebbe mettere in maiuscolo la prima lettera di ogni nome di impostazione.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` dovrebbe mostrare i risultati corretti e terminare con un carattere di nuova riga.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
