---
id: 6718d2d59337c822ecb697ff
title: 은행 계좌 관리 프로그램 만들기
challengeType: 26
dashedName: build-a-bank-account-management-program
saveSubmissionToDB: true
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 생성자가 있는 `BankAccount`라는 클래스를 정의해야 합니다. 생성자는 기본 `balance`을 `0`로 설정하고 거래 기록을 객체로 저장할 빈 배열 `transactions`을 초기화해야 합니다.

2. `transactions` 배열에 저장된 각 거래는 `type`와 `amount` 두 개의 속성을 가진 객체여야 합니다. `type` 속성은 `deposit` 또는 `withdraw` 중 하나여야 하며, `amount` 속성은 입금 또는 출금된 금액이어야 합니다.

3. 입금 금액을 매개변수로 받는 `deposit`라는 메서드를 정의해야 합니다. 입금 금액이 `0`보다 클 때는 다음을 수행해야 합니다:

   - 거래 배열에 `deposit` 유형과 입금 금액을 가진 새 객체를 추가합니다.
   - 잔액을 업데이트합니다.
   - `"Successfully deposited $[amount]. New balance: $[balance]"`를 반환합니다.

4. 금액이 `0` 이하일 경우 `deposit` 메서드는 `"Deposit amount must be greater than zero."`를 반환해야 합니다.

5. 금액을 매개변수로 받는 `withdraw`라는 메서드를 정의해야 합니다. 이 메서드는 출금에 따라 현재 잔액을 업데이트해야 합니다. 출금 금액이 `0`보다 크고 현재 잔액 이하일 때는 다음을 수행해야 합니다:

   - 거래 배열에 `withdraw` 유형과 출금 금액을 가진 새 객체를 추가합니다.
   - 잔액을 업데이트합니다.
   - `"Successfully withdrew $[amount]. New balance: $[balance]"`를 반환합니다.

6. 출금 금액이 `0` 이하이거나 현재 잔액보다 클 경우 `withdraw` 메서드는 `"Insufficient balance or invalid amount."`를 반환해야 합니다.

7. 현재 잔액을 `checkBalance` 형식으로 반환하는 `"Current balance: $[balance]"`라는 메서드를 정의해야 합니다.

8. 거래 배열을 반복하며 모든 입금을 `listAllDeposits` 형식으로 반환하는 `"Deposits: amount,amount,..."`라는 메서드를 정의해야 합니다.

9. 거래 배열을 반복하며 모든 출금을 `listAllWithdrawals` 형식으로 반환하는 `"Withdrawals: amount,amount,..."`라는 메서드를 정의해야 합니다.

10. `BankAccount`의 새 인스턴스 `myAccount`를 생성해야 합니다.

11. `myAccount` 은행 계좌에는 최소 다섯 개의 거래가 있어야 합니다.

12. `myAccount` 은행 계좌에는 최소 두 개의 입금이 있어야 합니다.

13. `myAccount` 은행 계좌에는 최소 두 개의 출금이 있어야 합니다.

14. `myAccount` 은행 계좌의 잔액은 `$100`보다 커야 합니다.

# --hints--

`BankAccount`라는 클래스를 정의해야 합니다.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

`BankAccount` 객체는 처음에 `balance`이 `0`이고 거래 기록을 저장할 빈 배열 `transactions`을 가져야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

입금 금액을 매개변수로 받는 `deposit` 메서드를 만들어야 합니다.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

출금 금액을 매개변수로 받는 `withdraw` 메서드를 만들어야 합니다.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

현재 잔액을 확인하는 `checkBalance` 메서드를 만들어야 합니다.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

모든 입금을 나열하는 `listAllDeposits` 메서드를 만들어야 합니다.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

모든 출금을 나열하는 `listAllWithdrawals` 메서드를 만들어야 합니다.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

`BankAccount.deposit(100)`는 `"Successfully deposited $100. New balance: $100"`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)`는 `"Deposit amount must be greater than zero."`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)`는 `"Deposit amount must be greater than zero."`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

계좌 잔액이 `100`일 때 `BankAccount.withdraw(150)`는 `"Insufficient balance or invalid amount."`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)`는 `"Insufficient balance or invalid amount."`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)`는 `"Insufficient balance or invalid amount."`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

계좌 잔액이 `200`일 때 `BankAccount.withdraw(150)`는 `"Successfully withdrew $150. New balance: $50"`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

계좌 잔액이 `200`일 때 `BankAccount.checkBalance()`는 `"Current balance: $200"`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

`10`, `35`, `90`를 입금할 때 `listAllDeposits` 메서드는 `"Deposits: 10,35,90"`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

`20`, `50`, `100`를 출금할 때 `listAllWithdrawals` 메서드는 `"Withdrawals: 20,50,100"`를 반환해야 합니다.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

`BankAccount`의 인스턴스 `myAccount`를 만들어야 합니다.

```js
assert.exists(myAccount);
assert.instanceOf(myAccount, BankAccount);
```

`myAccount` 은행 계좌에는 최소 다섯 개의 거래가 있어야 합니다.

```js
assert.isAtLeast(myAccount.transactions.length, 5);
```

`myAccount` 은행 계좌에는 최소 두 개의 입금이 있어야 합니다.

```js
const depositCount = myAccount.transactions.filter(transaction => transaction.type === "deposit").length;
assert.isAtLeast(depositCount, 2);
```

`myAccount` 은행 계좌에는 최소 두 개의 출금이 있어야 합니다.

```js
const withdrawalCount = myAccount.transactions.filter(transaction => transaction.type === "withdraw").length;
assert.isAtLeast(withdrawalCount, 2);
```

`myAccount` 은행 계좌의 잔액은 `$100`보다 커야 합니다.

```js
assert.isAbove(myAccount.balance, 100);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor() {
        this.balance = 0;
        this.transactions = [];
    }

    deposit(amount) {
        if (amount > 0) {
            this.transactions.push({ type: 'deposit', amount });
            this.balance += amount;
            return `Successfully deposited $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Deposit amount must be greater than zero.';
        }
    }

    withdraw(amount) {
        if (amount > 0 && amount <= this.balance) {
            this.transactions.push({ type: 'withdraw', amount });
            this.balance -= amount;
            return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
        } else {
            return 'Insufficient balance or invalid amount.';
        }
    }

    checkBalance() {
        return `Current balance: $${this.balance}`;
    }

    listAllDeposits() {
        const deposits = this.transactions
            .filter(transaction => transaction.type === 'deposit')
            .map(transaction => transaction.amount);
        return `Deposits: ${deposits}`;
    }

    listAllWithdrawals() {
        const withdrawals = this.transactions
            .filter(transaction => transaction.type === 'withdraw')
            .map(transaction => transaction.amount);
        return `Withdrawals: ${withdrawals}`;
    }
}

const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(20)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(55));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
