---
id: 68e6aff5325410118e76e7be
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafanya mazoezi ya kutumia ARIA kwa kujenga kidhibiti sauti kinachopatikana kwa urahisi.

Kuanza, ongeza kipengele cha `h1` chenye maandishi `Audio Controls`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
