---
id: 68eac57fadcc662ce917e75e
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ifuatayo, ongeza kipengele cha `button` chenye sifa ya kipengele `type` iliyowekwa kwa thamani ya `button`.

Maandishi ya kipengele cha `button` yanapaswa kuwa `Play`.

# --hints--

Unapaswa kuwa na kipengele cha `button`.

```js
assert.exists(document.querySelector('button'));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele `type`.

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele `type` iliyowekwa kwa `button`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Kipengele chako cha `button` kinapaswa kuwa na maandishi ya `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
