---
id: 68eb7f2397da251a40e04aa2
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ifuatayo, ongeza kipengele cha `div` ili kuhusisha vipengele vyote vya udhibiti wa sauti.

Kipengele cha `span` ni kontena la kawaida ndani ya mstari kwa ajili ya maudhui. Kinatumika sana kupamba kipande cha maandishi tofauti na maandishi yanayozunguka, au kufunika maandishi na kupewa `id` ambayo vipengele vingine (kama sifa za ARIA) vinaweza kuitumia kama rejea.

Ndani ya kipengele cha `div`, panga ndani kipengele cha `span` chenye sifa ya `id` iliyowekwa kuwa `volume-label`.
Pia, maandishi ya kipengele cha `span` yanapaswa kuwa `Volume`.

`id` baadaye itarejea na sifa za ARIA ili teknolojia za msaada ziweze kutambua maandishi haya kama lebo ya udhibiti wa sauti.

# --hints--

Unapaswa kuwa na kipengele cha `div` chini ya kipengele chako cha `button`.

```js
assert.exists(document.querySelector('button + div'));
```

Unapaswa kuwa na kipengele cha `span` ndani ya kipengele chako cha `div`.

```js
assert.exists(document.querySelector('button + div > span'));
```

Kipengele chako cha `span` kinapaswa kuwa na sifa ya `id`.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Kipengele chako cha `span` kinapaswa kuwa na sifa ya `id` yenye thamani iliyowekwa kuwa `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Kipengele chako cha `span` kinapaswa kuwa na maandishi ya `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
