---
id: 68ec0a48c9d6db3594d768c3
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sasa ni wakati wa kuongeza kipengele kingine cha `span`.

Kipengele hiki cha `span` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `volume-description`.
Pia, maandishi ya kipengele hiki cha `span` yanapaswa kuwa `Adjust the sound level`.

Kama ilivyo kwa kipengele cha `span` kilichopita, rejea ya `id` ya `span` itatazamwa baadaye na sifa za ARIA ili teknolojia za msaada ziweze kutambua maandishi haya kama maelezo ya udhibiti wa sauti.

# --hints--

Unapaswa kuwa na vipengele viwili vya `span` ndani ya kipengele cha `div`.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na sifa ya kipengele ya `id`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na sifa ya kipengele ya `id` yenye thamani ya `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Kipengele chako cha pili cha `span` kinapaswa kuwa na maandishi ya `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
