---
id: 68ed3a96e0ebcc36565b39e5
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ifuatayo, unahitaji kuongeza kipengele cha `input` chenye sifa ya kipengele ya `type` iliyowekwa kuwa `range` ili kuunda kipimaji sauti. Weka sifa ya kipengele ya `min` kuwa `0`, sifa ya kipengele ya `max` kuwa `100`, na sifa ya kipengele ya `value` kuwa `50` ili kufafanua kiwango cha sauti cha msingi.

# --hints--

Unapaswa kuwa na kipengele cha `input` chini ya kipengele chako cha pili cha `span`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `type`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani ya `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `min`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `min` yenye thamani ya `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `max`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `max` yenye thamani ya `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `value`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele ya `value` yenye thamani ya `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>

  </body>
</html>
```
