---
id: 68ed488821845990c911a223
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Kudhibiti sauti yako yenye upatikanaji karibu imekamilika lakini kuna jambo moja la mwisho la kuongeza.

Ongeza kipengele cha `button` chenye sifa ya kipengele ya `type` iliyowekwa kuwa `button` na na maandishi ya `Mute` chini ya kipengele cha `div`.

Na kwa mabadiliko hayo, umehitimisha warsha hii!

# --hints--

Unapaswa kuwa na kipengele cha `button` chini ya kipengele cha `div`.

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele ya `type`.

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani ya `button`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

Kipengele chako cha `button` kinapaswa kuwa na maandishi ya `Mute`.

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
        aria-labelledby="volume-label volume-description">
    </div>

  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
        aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
