---
id: 65f569725359e10d345bc52a
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

你將要建置一個程式，計算並繪製給定丟擲角度、速度和高度的拋射物軌跡。

首先匯入 `math`，你會在這個專案中大量使用它，因為它有像 `math.radians`、`math.cos`、`math.sin` 等有用的行為。

也創建這些變數，使其具有重力加速度的值和一些稍後會用到的特殊符號（請使用複製和貼上）。

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

使用 `import math` 來匯入 `math` 模組。

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

你應該有一個值為 `9.81` 的 `GRAVITATIONAL_ACCELERATION` 變數。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

你應該有一個值為 `"∙"` 的 `PROJECTILE` 變數。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

你應該有一個 `x_axis_tick` 變數，其值為 `"T"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

你應該有一個值為 `"⊣"` 的 `y_axis_tick` 變數。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
