---
id: 65f569725359e10d345bc52a
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

你將編寫一個程序，根據投擲的角度、速度和高度計算並繪製拋射物的軌跡。

首先導入 `math`，你將在該項目中經常使用它，因爲它具有 `math.radians`、`math.cos`、`math.sin` 等有用的方法。

還需要創建這些變量以表示重力加速度的數值以及一些以後會用到的特殊符號（請使用複製和粘貼來完成這些操作）。

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

使用 `import math` 導入 `math` 模塊。

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

你應該有一個 `GRAVITATIONAL_ACCELERATION` 變量，其值爲 `9.81`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

你應該有一個 `PROJECTILE` 變量，其值爲 `"∙"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

你應該有一個 `x_axis_tick` 變量，其值爲 `"T"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

你應該有一個 `y_axis_tick` 變量，其值爲 `"⊣"`。

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
