---
id: 65fd4bd84561a14d3e720061
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

類別變數 `__slots__` 在 Python 類別中有特殊的用法。宣告 `__slots__` 並將其指定為一組字串序列，會限制屬性的產生僅限於該序列中包含的屬性。此外，它也會防止產生特殊屬性 `__dict__`，並且允許更高效的屬性存取。

你應該在類別內使用 `__slots__` 變數來定義類別擁有的屬性：將一個包含 3 個字串的元組指定給 `__slots__`，每個字串都相等於 `__init__` 中定義的其中一個屬性名稱。

# --hints--

你應該將包含 `'__height'`、`'__speed'` 和 `'__angle'` 的元組指定給 `__slots__`。

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
