---
id: 65fd4b3bedee044b4b957d33
title: 步驟2
challengeType: 20
dashedName: step-2
---

# --description--

創建一個 `Projectile` 類和一個 `__init__` 方法，該方法使用三個參數初始化該類：起始速度、起始高度和投擲物起始角度（按此順序）。

在 `__init__` 方法內部，將這些參數分配給三個私有屬性：`__speed`、`__height` 和 `__angle`。 提供的角度以度爲單位；然而，它應該以弧度爲單位進行內部存儲。 爲此，在將角度分配給 `__angle` 時，使用 `math.radians()` 函數將角度從度轉換爲弧度。

屬性名稱前使用兩個下劃線會觸發 Python 中的名稱修改。 這意味着不能使用給定的名稱從類外部直接訪問屬性，並且如果外部需要則必須使用像 `ball._Projectile__height` 這樣的混亂名稱進行訪問，表明這些屬性僅供內部使用。

# --hints--

你應該使用 `class Projectile:` 來創建該類。

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

你應該在 `Projectile` 類內有一個 `__init__` 方法。

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

新的 `Projectile` 類實例應該正確初始化這三個屬性的值。

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
