---
id: 65fd4bd84561a14d3e720061
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

類變量 `__slots__` 在 Python 類中具有特殊用途。 聲明 `__slots__` 併爲其分配一個字符串序列會將屬性的創建限制爲該序列中包含的屬性。 此外，它可以防止創建 `__dict__` 特殊屬性，並允許更高效的屬性訪問。

你應該使用類內的 `__slots__` 變量來定義該類具有哪些屬性：爲 `__slots__` 分配一個包含 3 個字符串的元組，每個字符串等於 `__init__` 中定義的屬性名稱之一。

# --hints--

你應該爲 `__slots__` 分配一個包含 `'__height'`、`'__speed'` 和 `'__angle'` 的元組。

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
