---
id: 65fd4b3bedee044b4b957d33
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

創建一個 `Projectile` 類別，並使用一個 `__init__` 方法，該方法依序使用三個引數初始化類別：起始速度、起始高度和投擲物的起始角度。 

在 `__init__` 方法中，將這些引數指定到三個私有屬性：`__speed`、`__height` 和 `__angle`。提供的角度將以度為單位；然而，應該以徑度形式在內部儲存。為了達成此目的，請在指定給 `__angle` 時使用 `math.radians()` 函式將角度從度轉換為徑度。 

在屬性名稱前使用兩個底線會觸發 Python 的名稱改寫。這表示這些屬性無法直接使用其給定名稱從類別外部存取，若需要從外部存取，必須使用改寫後的名稱，例如 `ball._Projectile__height`，表示這些屬性僅供內部使用。

# --hints--

你應該使用 `class Projectile:` 來創建該類別。

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

你應該在 `Projectile` 類別中有一個 `__init__` 方法。

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

新的 `Projectile` 類別實體應該正確初始化這三個屬性的值。

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
