---
id: 660fcf3e1b9bb056b2edb567
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

此時你已準備好創建字串表述。  
首先創建 `Projectile` 類別的實體。在類別定義式下方，創建一個 `ball` 變數，並將使用 `10, 3, 45` 作為引數的 `Projectile` 呼叫指定給它。

在類別中定義一個 `__str__` 方法，為你的實體提供適當的字串表述，並使它傳回具有以下格式的字串：

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

應該以換行字元開始和結束，角度必須以整數度數表示，位移應以一位小數列印。

你會發現有用的 `math.degrees` 用來將角度從徑度轉換為度數。

當你準備好時，你可以 `print(ball)` 來測試你的函式（程式）。

# --hints--

你應該有一個 `ball` 變數。

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

`ball` 變數應該有一個值為 `Projectile(10, 3, 45)`。

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

`Projectile(45, 45, 45)` 的字串表述應該是正確的。

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

字串表述也應該對其他實體正確。

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

你應該有一個 `print(ball)` 呼叫。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
