---
id: 6616ec0bf9ee8548ce6b0f08
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

現在是時候計算軌跡的座標了！為 `Projectile` 類別創建一個名為 `__calculate_y_coordinate` 的方法，該方法除了 `self` 外，還應該有一個 `x` 參數。

\\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \\]

以上是計算任意水平位置 `x` 的垂直位置 $y$ 的公式，給定起始角度 $\theta$、速度 $v_0$ 和高度 $y_0$。

你需要使用 `math.tan()` 和 `math.cos()`，並且記住 `x ** y` 是寫出 $x^y$ 的方式，而 $g$ 的值在變數 `GRAVITATIONAL_ACCELERATION` 中。

實作該方法，使其傳回 $y$ 座標。

如果你想測試該方法，從類別外部，在宣告 `ball` 後，你可以使用一個數字作為引數呼叫 `ball._Projectile__calculate_y_coordinate()`。

# --hints--

你的 `Projectile` 類別應該有一個名為 `__calculate_y_coordinate` 的方法。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

`__calculate_y_coordinate` 方法應該有 `self, x` 作為引數。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

該方法應該傳回正確的值。對於 `Projectile(12, 12, 12)`，給定 4 個 $x$ 座標 0、1、2 和 3，y 座標應該約為 12、12.18、12.28、12.32。

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

該方法應該傳回正確的值。對於 `Projectile(45, 12, 22)`，給定 $x$ 座標為 6、9、10 和 12，y 座標應約為 14.32、15.41、15.76、16.44。

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

`__calculate_y_coordinate` 方法不應該對結果進行四捨五入。

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
