---
id: 662fa453ec8033322b3ab2e3
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

現在是時候談談封裝和取值函式了。你已經使用前置雙底線將三個實體屬性設為私有。請注意，這些屬性是依慣例稱為私有：雖然它們仍然可以從外部存取，但大家約定不這麼做。

取值函式是用來從外部取得值的方式。要定義取值函式，你需要定義一個傳回所需屬性值的函式，並給它一個 `@property` 裝飾器：

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 裝飾器將方法變成屬性，這表示該方法不會像一般方法那樣被呼叫，而是像屬性一樣使用：

```py
n = Nest()
print(n.number_of_eggs)
```

在上述範例中，私有屬性 `__number of eggs` 是透過 `n` 的 `number_of_eggs` 屬性來存取的。

創建一個名為 `speed` 的取值函式，並使其傳回私有屬性 `__speed` 的值。

# --hints--

你應該定義一個 `speed` 方法。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

你的 `speed` 方法應該只有一個參數，`self`。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

你的 `speed` 方法應該有一個 `@property` 裝飾器。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

你的 `speed` 方法應該傳回私有屬性 `__speed` 的值。

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
