---
id: 6634fa2417d7194b0c9425f9
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

現在為 `__height` 和 `__angle` 創建兩個取值函式。請記得角度的值是以度數提供給類別，但內部是以徑度存檔，因此取值函式應傳回度數值。你可以使用 `math.degrees` 和 `round` 函式來達成這個目的。

# --hints--
你應該定義一個 `height` 方法。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('height')
`))})
```

你的 `height` 方法應該只有一個參數，`self`。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_args('self')
`))})
```

你的 `height` 方法應該有一個 `@property` 裝飾器。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_decorators('property')
`))})
```

你的 `height` 方法應該傳回私有屬性 `__height` 的值。

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.height == 23
`))})
```

你應該定義一個 `angle` 方法。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('angle')
`))})
```

你的 `angle` 方法應該有一個 `@property` 裝飾器。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_decorators('property')
`))})
```

你的 `angle` 方法應該只有一個參數，`self`。

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_args('self')
`))})
```

你的 `angle` 方法應該傳回私有屬性 `__angle` 轉換成度數的值。

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.angle == 24
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    @property
    def speed(self):
        return self.__speed

    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
