---
id: 691f7bbf76229bb97c60827d
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用 `for...of` 迴圈，透過建置一個函式來計算字串在字串陣列中出現的次數。

在撰寫該函式之前，你會先創建一個較簡單的函式，將字串的每個字元記錄到主控台。

首先定義一個名為 `printCharacters` 的空函式，並帶有參數 `str`。

# --hints--

你應該創建一個 `printCharacters` 函式。

```js
assert.isFunction(printCharacters);
```

你的 `printCharacters` 函式應該有一個名為 `str` 的參數。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.printCharacters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), 'str');
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
