---
id: 6920d189e36b6b60a6330fee
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

現在，在迴圈內將 `char` 變數記錄到主控台。

# --hints--

你應該在迴圈內有一個 `console.log` 述語，並且 `char` 是它的引數。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

當 `printCharacters` 函式被呼叫並帶入引數 `'test'` 時，函式內的 `console.log()` 述語應該依序以以下引數被呼叫四次：`'t'`、`'e'`、`'s'` 和 `'t'`。但實際上它被呼叫時帶入了以下引數：`--fcc-actual--`。

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
