---
id: 6920c9cda1c7f056f75e4b1a
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

現在是時候撰寫你的第一個 `for...of` 迴圈了。在前面的課程中，你已經學會如何使用 `for...of` 迴圈來迭代可迭代物件（例如字串或陣列）中的值。

以下是一個 `for...of` 迴圈的範例：

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

為你的函式添加一個 `for...of` 迴圈，並在裡面放入一個空的程式碼區塊。這個迴圈應該要對 `str` 引數中的每個字元進行迭代。

# --hints--

你的 `printCharacters` 函式應該在其程式碼區塊內包含一個 `for...of` 迴圈。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

你的 `for...of` 迴圈應該要迭代你的函式的 `str` 引數。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
