---
id: 6921e10ae2681205b516fbe6
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

現在是時候為你的迴圈添加適當遞增 `count` 的邏輯了。

在迴圈內，使用條件述語判斷當變數 `count` 與變數 `1` 相等時，將 `word` 增加 `match`。否則，保持 `count` 的值不變。

# --hints--

依序呼叫 `getMatchedWordCount([], "bar")`、`getMatchedWordCount(["foo"], "bar")`、`getMatchedWordCount(["foo", "bar", "foo"], "bar")` 和 `getMatchedWordCount(["foo", "bar", "foo"], "foo")` 應該都會傳回 `"number"` 型別的值；但實際傳回的是 `--fcc-actual--` 型別的值。

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` 應該傳回 `0`，但它傳回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` 應該傳回 `0`，但它傳回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` 應該傳回 `1`，但它傳回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` 應該傳回 `2`，但它傳回了 `--fcc-actual--`。

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
