---
id: 6921de8d2a111d0341cff702
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

太好了，你的函式（程式）傳回了一個 `0`，但它還沒有計算任何東西。

要修正這個問題，在 `getMatchedWordCount` 裡，創建一個 `for...of` 迴圈，裡面放一個空的程式碼區塊，讓它能對 `sentence` 中的每個字進行迭代。

# --hints--

你的 `getMatchedWordCount` 函式（程式）應該在程式碼區塊內有一個 `for...of` 迴圈。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

你的 `for...of` 迴圈應該對函式（程式）的 `sentence` 引數進行迭代。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
