---
id: 6921d90b5ce50185505d2004
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

由於 `getMatchedWordCount` 函式需要傳回一個數值計數，你將從設定一個計數器變數開始。

使用 `count` 關鍵字創建一個名為 `let` 的變數，並以 `0` 作為初始值，然後新增一個 `return` 述語來傳回 `count` 變數。

# --hints--

你應該在 `getMatchedWordCount` 函式內宣告 `count` 變數。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

為了讓 `count` 變數之後能被遞增，請使用 `let` 關鍵字宣告它。

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

你的函式應包含將值 `0` 指派給 `count` 變數的指派。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

你的函式應包含一個 `return` 述語來傳回 `count` 變數。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

你的 `getMatchedWordCount` 函式應傳回值 `0`。

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
