---
id: 6921d688b8d58581d34459b9
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

在接下來的幾個步驟中，你將建置一個函式，用來計算字串在字串陣列中出現的次數。

首先，定義一個空的函式，名稱為 `getMatchedWordCount`，並依序帶入參數 `sentence` 和 `match`。

# --hints--

你應該創建一個 `getMatchedWordCount` 函式。

```js
assert.isFunction(getMatchedWordCount);
```

你的 `getMatchedWordCount` 函式應該有兩個參數：`sentence` 和 `match`，依序排列。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
