---
id: 6921dae10700e300f26994e3
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

要查看你目前 `getMatchedWordCount` 函式的輸出結果，請在它下面添加一個 `console.log` 述語。在 `console.log()` 中，使用引數 `["I", "really", "really", "really", "like", "to", "code"]` 呼叫該函式，作為 `sentence` 參數，並使用 `"really"` 作為 `match` 參數。

# --hints--

你應該在 `getMatchedWordCount` 述語中呼叫函式 `console.log`。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

你應該使用引數 `getMatchedWordCount` 呼叫函式 `["I", "really", "really", "really", "like", "to", "code"]`，作為 `sentence` 參數，並使用 `"really"` 作為 `match` 參數。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
