---
id: 6921d688b8d58581d34459b9
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

在接下來的幾個步驟中，你將建置一個函式，用來計算字串在字串陣列中出現的次數。

首先，定義一個空的函式，名稱為 `getMatchedWordCount`，並依序帶入參數 `sentence` 和 `match`。

# --hints--

你應該創建一個 `getMatchedWordCount` 函式。

```js
assert.isFunction(getMatchedWordCount);
```

你的 `getMatchedWordCount` 函式應該有兩個參數：`sentence` 和 `match`，依序排列。

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
