---
id: 61537485c4f2a624f18d7794
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comience con su plantilla HTML estándar. Añade una declaración `DOCTYPE`, un elemento `html`, un elemento `head` y un elemento `body`.

Agrega el atributo `lang` a la etiqueta de apertura `<html>` con `en` como valor.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir `html` como tipo de documento.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu `html` debe tener una etiqueta de apertura.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu etiqueta de apertura `<html>` debe tener el atributo `lang` con `en` como valor.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener solo una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Las etiquetas `head` y `body` deben ser hermanas.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

La etiqueta `head` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

La etiqueta `body` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
