---
id: a97fd23d9b809dac9921074f
title: بناء دالة جمع معلمات اختيارية
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

في هذا المختبر ستبني دالة تقبل حتى معلمتين، وتجمعهما، ولكن إذا كان هناك معلمة واحدة فقط تُرجع دالة تنتظر الرقم الثاني للجمع.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. أنشئ دالة `addTogether`.
2. إذا استقبلت الدالة معلمتين، يجب أن تُرجع الدالة مجموع المعلمتين.
3. إذا تم توفير معلمة واحدة فقط، يجب أن تُرجع `addTogether` دالة.
   - عندما تُستدعى الدالة المُرجعة بمعلمة واحدة، يجب أن تُرجع المجموع.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. إذا لم تكن أي من المعلمتين رقمًا صالحًا، أرجع `undefined`.

# --hints--

يجب أن تمتلك دالة `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` يجب أن تُرجع 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` يجب أن تُرجع 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` يجب أن تُرجع دالة.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` يجب أن تُرجع 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` يجب أن تُرجع `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
