---
id: 699b6dd1a2961bd38491ef2e
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

L'oggetto `profile` aveva chiavi e tipi rigidi. Se provavi ad aggiungere una proprietà non definita o ad assegnare un tipo non compatibile, TypeScript lanciava un errore.

Sarebbe utile se l'oggetto `userRoles` non fosse così rigido e avesse chiavi più flessibili.

Come ricorderai dalle lezioni precedenti, puoi ottenere questo risultato usando il tipo `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

In questo esempio `studentGrades`, TypeScript non ha bisogno di conoscere tutte le materie in anticipo. Questo ti permette di aggiungere altre materie come `Art` o `Music` in seguito. A TypeScript interessa solo che i valori per ogni chiave siano numeri.

Aggiorna il tuo oggetto `userRoles` aggiungendo un tipo `Record` in linea. Impostalo su `<string, string>` in modo che ogni chiave possa essere una stringa qualsiasi e ogni valore deve essere una stringa.

# --hints--

Il tuo oggetto `userRoles` dovrebbe avere un tipo `Record` in linea impostato su `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
