---
id: 699b70b10ae344c802d106b0
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Ora è il momento di aggiungere un paio di altre proprietà all'oggetto `userRoles`.

Aggiungi una proprietà `moderator` con un valore di `"medium-access"`.

Poi, subito sotto, aggiungi una proprietà `guest` con un valore di `3`.

# --hints--

Il tuo oggetto `userRoles` dovrebbe avere una proprietà `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

La tua proprietà `moderator` dovrebbe avere un valore di `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Il tuo oggetto `userRoles` dovrebbe avere una proprietà `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

La tua proprietà `guest` dovrebbe avere un valore del numero `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
