---
id: 699b71d98c820ac7ab9eea2a
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Se apri di nuovo la console, vedrai un nuovo messaggio di errore:

```md
Type 'number' is not assignable to type 'string'.
```

Questo succede perché stai tentando di aggiungere una proprietà `guest` con un valore numerico.

Per rimuovere l'errore di TypeScript, cambia il valore della proprietà `guest` in `"read-only"`.

Con quest'ultima modifica, il tuo workshop sul profilo utente è completo!

# --hints--

La tua proprietà `guest` per l'oggetto `userRoles` dovrebbe avere il valore `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
