---
id: 699b64096301fd75e3efeb1b
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Se apri di nuovo la console, vedrai un messaggio di errore diverso:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Come ricorderai dalle lezioni precedenti, puoi indicare che una proprietà è opzionale usando il simbolo `?` accanto al nome della proprietà in questo modo:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

In questo esempio, `numberOfAlbums` è indicato come opzionale quindi TypeScript non lancerà un errore in questo caso.

Aggiorna la tua proprietà `bio` per farla diventare opzionale.

# --hints--

La tua proprietà `bio` dovrebbe essere indicata come opzionale.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
