---
id: 699b65caf17a8e8d340af27d
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Ora Ă¨ il momento di aggiungere l'ultima proprietĂ  all'oggetto `profile`.

Nelle lezioni precedenti, hai imparato a lavorare con i tipi array. Hai conosciuto diversi tipi di array come `string[]` o `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

In questo esempio, l'array `scores` accetta solo numeri. Se provi ad aggiungere un tipo diverso come un `string` o un `null`, TypeScript lancerĂ  un errore.

Aggiorna il tuo oggetto `profile` per aggiungere una proprietĂ  `programmingLanguages`, tipizzata come un `string[]`. All'interno del letterale oggetto, imposta questa proprietĂ  su `["JavaScript", "Python", "C++"]`.

# --hints--

Dovresti avere una proprietĂ  `programmingLanguages` all'interno del tuo oggetto `profile`.

```js
assert.property(profile, "programmingLanguages");
```

La tua proprietĂ  `programmingLanguages` dovrebbe avere il valore `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

La tua proprietĂ  `programmingLanguages` dovrebbe essere tipizzata come un `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
