---
id: 699b61ca8c9cb9bc1968364e
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Se apri la console, dovresti vedere il seguente messaggio di errore:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

Nel JavaScript puro, potevi aggiungere qualsiasi proprietà e valore volessi all'oggetto `profile`. Ma in TypeScript, il compilatore verifica che l'oggetto abbia solo le proprietà che hai definito (`username`, `age` e `isLoggedIn` in questo caso). Poiché `mood` non era inclusa nel tipo, TypeScript la considera una proprietà non valida e lancia un errore.

Rimuovi `mood: null` dal tuo codice.

# --hints--

Non dovresti più avere `mood: null` nel tuo codice.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
