---
id: 699b5370a43e299b4673bcaa
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Sarebbe utile vedere le proprietà dell'oggetto `profile` nella console mentre lo costruisci.

Aggiungi un `console.log` che registri `profile` nella console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti registrare `profile` nella console.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
