---
id: 699b4e284291adcfcb90df47
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso di tipi primitivi oltre a tipi oggetto e array costruendo un profilo utente con tipi sicuri.

Inizia creando una variabile chiamata `profile` e assegnale un oggetto. Il tuo oggetto dovrebbe avere queste tre proprietà e valori:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Dovresti avere una variabile chiamata `profile`.

```js
assert.exists(profile);
```

La tua variabile `profile` dovrebbe essere un oggetto.

```js
assert.isObject(profile);
```

Dovresti avere una proprietà `username` nel tuo oggetto `profile`.

```js
assert.property(profile, "username");
```

La tua proprietà `username` dovrebbe avere un valore `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Il tuo oggetto `profile` dovrebbe avere una proprietà `age`.

```js
assert.property(profile, "age");
```

La tua proprietà `age` dovrebbe avere un valore `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Il tuo oggetto `profile` dovrebbe avere una proprietà `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

La tua proprietà `isLoggedIn` dovrebbe avere un valore `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
