---
id: 699b5b7cb68b6d35afca61b7
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Anche se Ă¨ possibile continuare ad aggiungere proprietĂ  in questo modo all'oggetto `profile`, sarebbe utile limitare i tipi di proprietĂ  che dovrebbero essere consentite. Sarebbe meglio definire in anticipo la forma dell'oggetto `profile`.

Nelle lezioni precedenti, hai imparato a lavorare con tipi di oggetto come questo:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

In questo esempio, `composer` Ă¨ un oggetto con un tipo di oggetto in linea. Questo oggetto ha tre proprietĂ  obbligatorie `name`, `instrument` e `isActive`. Le prime due proprietĂ  devono essere di tipo `string` e `isActive` deve essere di tipo `boolean`.

Aggiorna il tuo oggetto `profile` esistente per includere un tipo di oggetto in linea. `username` dovrebbe essere tipizzato come `string`, `age` dovrebbe essere tipizzato come `number` e `isLoggedIn` dovrebbe essere tipizzato come `boolean`.

# --hints--

Il tuo oggetto `profile` dovrebbe avere un tipo di oggetto in linea con una proprietĂ  `username` impostata su `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Il tuo oggetto `profile` dovrebbe avere un tipo di oggetto in linea con una proprietĂ  `age` impostata su `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Il tuo oggetto `profile` dovrebbe avere un tipo di oggetto in linea con una proprietĂ  `isLoggedIn` impostata su `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
