---
id: 699b62c99ec8c63470313365
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Ora è il momento di aggiungere un'altra proprietà al tuo oggetto `profile`.

Inizia aggiornando il tipo di oggetto in linea per includere una proprietà `bio` con un tipo impostato su `string`.

# --hints--

Dovresti aggiornare il tipo di oggetto in linea per includere una proprietà `bio` impostata su `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
