---
id: 699b68f94a0d154c2a7a5ce0
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora che l'oggetto `profile` è completo, è il momento di buildare un nuovo oggetto che rappresenta i ruoli utente.

Crea una nuova variabile chiamata `userRoles` e assegnale un oggetto. Il tuo oggetto dovrebbe avere le seguenti proprietà e valori:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Dovresti avere una variabile chiamata `userRoles`.

```js
assert.isDefined(userRoles);
```

La tua variabile `userRoles` dovrebbe essere un oggetto.

```js
assert.isObject(userRoles);
```

Il tuo oggetto `userRoles` dovrebbe avere una proprietà `admin`.

```js
assert.property(userRoles, "admin");
```

La tua proprietà `admin` dovrebbe avere un valore `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Il tuo oggetto `userRoles` dovrebbe avere una proprietà `editor`.

```js
assert.property(userRoles, "editor");
```

La tua proprietà `editor` dovrebbe avere un valore `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Il tuo oggetto `userRoles` dovrebbe avere una proprietà `viewer`.

```js
assert.property(userRoles, "viewer");
```

La tua proprietà `viewer` dovrebbe avere un valore `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
