---
id: 61437d575fb98f57fa8f7f36
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

從你的標準 HTML 範本開始。為你的文件添加 `DOCTYPE` 宣告、指定此頁面為英文的 `html` 元素、`head` 元素和 `body` 元素。

在 `head` 元素中為你的 `<meta>` 頁籤添加適當的 `charset` 以及用於行動響應的 `<meta>` 頁籤。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該定義文件型別為 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你應該有一個開啟的 `<html>` 頁籤，且它的 `lang` 應該是 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開啟的 `<head>` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `</head>` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `<body>` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `</body>` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同層元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一個 `meta` 元素應該設定 `name` 為 `viewport`，並將 `content` 設定為 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設為 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
