---
id: 61438b5b66d76a6264430f2a
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

在你的每個新的 `a` 元素中，添加一個 `i` 元素並給它們以下類別：

- 你的第一個 `i` 元素應該有 `fab fa-facebook-f` 類別
- 你的第二個 `i` 元素應該有類別 `fab fa-twitter`
- 你的第三個 `i` 元素應該有類別 `fab fa-instagram`
- 你的第四個 `i` 元素應該有類別 `fab fa-linkedin-in`
- 你的第五個 `i` 元素應該有 `fab fa-youtube` 類別

# --hints--

你應該有五個 `i` 元素。

```js
assert(document.querySelectorAll('i')?.length === 5);
```

每個 `a` 元素應該只包含一個 `i` 元素。

```js
const aaaaa = [...document.querySelectorAll('.social-icons a')];
assert(aaaaa?.every(a => a?.children?.length === 1 && a?.children?.[0]?.localName === 'i'));
```

每個 `i` 元素都應該有一個包含 `fab` 的 `class` 屬性。

```js
const iiiii = [...document.querySelectorAll('i')];
assert(iiiii?.every(i => i?.classList?.contains('fab')));
```

第一個 `i` 元素應該有一個包含 `fa-facebook-f` 的 `class` 屬性。

```js
assert(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

第二個 `i` 元素應該有一個包含 `fa-twitter` 的 `class` 屬性。

```js
assert(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

第三個 `i` 元素應該有一個包含 `fa-instagram` 的 `class` 屬性。

```js
assert(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

第四個 `i` 元素應該有一個包含 `fa-linkedin-in` 的 `class` 屬性。

```js
assert(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

第五個 `i` 元素應該有一個包含 `fa-youtube` 的 `class` 屬性。

```js
assert(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
