---
id: 614385513d91ae5c251c2052
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

為你的文件添加一個 `title` 元素，文字為 `Magazine`，一個 `link` 元素用於 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 字型樣式表，一個 `link` 用於 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome 樣式表，以及一個 `link` 用於你的 `./styles.css` 樣式表。

你的字型樣式表將載入三個獨立的字型：`Anton`、`Baskervville` 和 `Raleway`。

# --hints--

你的程式碼應該有三個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

你的 `link` 元素應該在你的 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

你的三個 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

你的其中一個聯結元素的 `href` 應設為 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

你的其中一個聯結元素的 `href` 應設為 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

你的其中一個 `link` 元素的 `href` 應該設定為 `styles.css`。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

你的 `title` 元素應該在你的 `head` 元素內。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

你的專案應該有一個標題為 `Magazine`。請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
