---
id: 6165d3b702a5d92ad970b30c
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在你的 `img` 元素之後，為你的 `h1` 元素添加 `class` 設為 `hero-title` 並將文字設為 `OUR NEW CURRICULUM`，接著添加一個 `p` 元素，`class` 設為 `hero-subtitle` 並將文字設為 `Our efforts to restructure our curriculum with a more project-based focus`。

# --hints--

你應該創建一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該位於你的 `img` 元素之後。

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

你的 `h1` 元素應該將 `class` 設定為 `hero-title`。

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

你的 `h1` 元素應該將文字設定為 `OUR NEW CURRICULUM`。

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

你應該創建一個新的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素應該位於你的 `h1` 元素之後。

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

你的 `p` 元素應該將 `class` 設定為 `hero-subtitle`。

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

你的 `p` 元素應該將文字設為 `Our efforts to restructure our curriculum with a more project-based focus`。

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
