---
id: 66a83fec026a7a4631e084d2
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Beim Arbeiten mit `fieldset`-Elementen ist es üblich, eine Bildunterschrift zu verwenden, um die Gruppe von Eingaben zu beschreiben. Dafür können Sie das `legend`-Element verwenden.

Hier ist ein Beispiel für die Verwendung eines `legend`-Elements:

```html
<form action="/example-url">
  <fieldset>
    <legend>Personal Information</legend>
    <!-- inputs go inside here-->
  </fieldset>
</form>
```

Fügen Sie innerhalb Ihres `fieldset`-Elements ein `legend`-Element mit dem Text `Personal Information` hinzu.

# --hints--

Sie sollten ein öffnendes `legend`-Tag haben.

```js
assert.match(code, /<legend>/);
```

Sie sollten ein schließendes `legend`-Tag haben.

```js
assert.match(code, /<\/legend>/);
```

Ihr `legend`-Element sollte sich innerhalb des `fieldset`-Elements befinden.

```js
assert.isNotNull(document.querySelector('fieldset legend'));
```

Ihr `legend`-Element sollte den Text `Personal Information` enthalten.

```js
const legend = document.querySelector('fieldset legend');
assert.strictEqual(legend?.innerText.trim(), 'Personal Information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
