---
id: 66a83bdcf425e7446900b7c4
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie mit dem `form` Element so arbeiten:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

Das `action` Attribut wird verwendet, um anzugeben, wohin die Formulardaten gesendet werden sollen, wenn das Formular abgesendet wird.

Das `method` Attribut wird verwendet, um die HTTP-Methode anzugeben, die beim Senden der Formulardaten verwendet werden soll. Die gebräuchlichsten Methoden sind `GET` und `POST`.

**HINWEIS**: Wie HTTP-Methoden funktionieren, lernen Sie in späteren Modulen.

Fügen Sie innerhalb Ihres `main` Elements ein `form` Element mit einem `action` Attribut hinzu, das auf `"https://hotel-feedback.freecodecamp.org"` gesetzt ist, und einem `method` Attribut, das auf `"POST"` gesetzt ist.

# --hints--

Sie sollten einen öffnenden `form` Tag haben.

```js
assert.match(code, /<form.*>/i);
```

Sie sollten einen schließenden `form` Tag haben.

```js
assert.match(code, /<\/form>/i);
```

Ihr `form` Element sollte ein `action` Attribut mit dem Wert `"https://hotel-feedback.freecodecamp.org"` haben.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Ihr `form` Element sollte ein `method` Attribut mit dem Wert `"POST"` haben.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement?.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
