---
id: 66a8380d911e3f4270d5cadc
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, das `main`-Element hinzuzufügen, das den Hauptinhalt der Seite darstellt.

# --hints--

Sie sollten ein öffnendes `main`-Tag haben.

```js
assert.match(code, /<main>/i);
```

Sie sollten ein schließendes `main`-Tag haben.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
