---
id: 66a83601cd819e37f0dccd14
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit HTML-Formularen, indem Sie ein Hotel-Feedback-Formular erstellen.

Für den Einführungstext möchten Sie den Haupttitel gefolgt von einer kurzen Notiz zum Hinterlassen von Feedback anzeigen.

Fügen Sie innerhalb Ihres `body`-Elements ein `header`-Element hinzu.

Fügen Sie innerhalb des `header`-Elements ein `h1`-Element hinzu. Der Text des `h1`-Elements sollte `Hotel Feedback Form` sein.

Fügen Sie unter Ihrem `h1`-Element ein `p`-Element hinzu. Der Text des `p`-Elements sollte `Thank you for staying with us. Please provide feedback on your recent stay.` sein.

# --hints--

Sie sollten ein öffnendes `header`-Tag haben.

```js
assert.match(code, /<header>/i);
```

Sie sollten ein schließendes `header`-Tag haben.

```js
assert.match(code, /<\/header>/i);
```

Sie sollten ein öffnendes `h1`-Tag haben.

```js
assert.match(code, /<h1>/i);
```

Sie sollten ein schließendes `h1`-Tag haben.

```js
assert.match(code, /<\/h1>/i);
```

Ihr `h1`-Element sollte den Text `Hotel Feedback Form` haben.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Ihr `h1`-Element sollte sich innerhalb des `header`-Elements befinden.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Sie sollten ein öffnendes `p`-Tag haben.

```js
assert.match(code, /<p>/i);
```

Sie sollten ein schließendes `p`-Tag haben.

```js
assert.match(code, /<\/p>/i);
```

Ihr Absatz-Element sollte den Text `Thank you for staying with us. Please provide feedback on your recent stay.` haben.

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Ihr Absatz-Element sollte sich innerhalb Ihres `header` befinden.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Ihr Absatz-Element sollte nach Ihrem `h1` kommen.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
