---
id: 66a84111965a0c46df6bbd0a
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man ein `label`-Element mit einem `input` verknüpft, und zwar so:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

Das `for`-Attribut am `label`-Element sollte mit dem `id`-Attribut am `input`-Element übereinstimmen. Dies wird als explizite Verknüpfung bezeichnet.

Fügen Sie unter Ihrem `legend`-Element ein `label`-Element mit dem Text von `Name (required):` hinzu. Setzen Sie dessen `for`-Attribut auf den Wert von `"full-name"`.

Fügen Sie dann unter Ihrem `label`-Element ein `input`-Element ohne Attribute hinzu. In den nächsten Schritten werden Sie die notwendigen Attribute ergänzen.

# --hints--

Sie sollten ein öffnendes `label`-Tag haben.

```js
assert.match(code, /<label.*>/);
```

Sie sollten ein schließendes `label`-Tag haben.

```js
assert.match(code, /<\/label>/);
```

Ihr `label`-Element sollte unter Ihrem `legend` innerhalb des `fieldset`-Elements platziert sein.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

Ihr `label`-Element sollte ein `for`-Attribut mit dem Wert von `"full-name"` haben.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Ihr `label`-Element sollte den Text von `Name (required):` enthalten.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText.trim(), 'Name (required):');
```

Sie sollten ein `input`-Element unter Ihrem `label`-Element haben.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
