---
id: 66a83e5e491625454b6f62c3
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Formulare bestehen aus `inputs`, in denen Benutzer ihre Daten eingeben können. Sie können verwandte Eingaben mit dem `fieldset`-Element gruppieren.

Hier ist ein Beispiel für die Verwendung eines `fieldset`-Elements:

```html
<form action="/example-url">
  <fieldset>
  <!-- inputs go inside here-->
  </fieldset>
</form>
```

Fügen Sie innerhalb Ihres `form`-Elements ein `fieldset`-Element hinzu.

# --hints--

Sie sollten ein öffnendes `fieldset`-Tag haben.

```js
assert.match(code, /<fieldset>/i);
```

Sie sollten ein schließendes `fieldset`-Tag haben.

```js
assert.match(code, /<\/fieldset>/i);
```

Ihr `fieldset`-Element sollte sich innerhalb des `form`-Elements befinden.

```js
assert.isNotNull(document.querySelector('form fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">        
--fcc-editable-region--
        
--fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
