---
id: 66a975c259525b7bc2d5c776
title: Schritt 28
challengeType: 0
dashedName: step-28
---

# --description--

Ihr Feedbackformular für das Hotel sollte den Nutzenden auch die Möglichkeit geben, das Essen zu bewerten.

Beginnen Sie damit, ein `label`-Element mit dem Text `How was the food?` hinzuzufügen. Dieses `label`-Element sollte ein `for`-Attribut mit dem Wert `"food"` haben.

Fügen Sie unter Ihrem `label`-Element ein `select`-Element mit einem `id` und `name` hinzu, die auf `"food"` gesetzt sind.

# --hints--

Sie sollten ein `label`-Element mit dem `for`-Attribut haben, das auf `"food"` gesetzt ist.

```js
assert.exists(document.querySelector('label[for="food"]'));
```

Ihr `label` sollte den Text `How was the food?` enthalten.

```js
assert.strictEqual(document.querySelector('label[for="food"]')?.textContent.trim(), 'How was the food?');
```

Sie sollten ein `select`-Element mit dem `id` haben, das unter Ihrem `"food"` auf `label` gesetzt ist.

```js
assert.exists(document.querySelector('label + select#food'));
```

Ihr `select`-Element sollte das `name`-Attribut mit dem Wert `"food"` haben.

```js
assert.exists(document.querySelector('label + select[name="food"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>

          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
