---
id: 66a962954f4e0d76223b37ed
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Fügen Sie als Nächstes ein weiteres Kontrollkästchen `input` hinzu, bei dem die `id`- und `value`-Attribute auf `"location"` gesetzt sind und das `name`-Attribut auf `"choice"`.

Für das `label`-Element sollten der Text von `Location` und das `for`-Attribut auf `"location"` gesetzt werden.

Unterhalb dieses `label`-Elements fügen Sie ein weiteres Kontrollkästchen `input` hinzu, bei dem die `id`- und `value`-Attribute auf `"reputation"` gesetzt sind und das `name`-Attribut auf `"choice"`.

Für das `label`-Element sollten der Text von `Reputation` und das `for`-Attribut auf `"reputation"` gesetzt werden.

# --hints--

Sie sollten ein drittes Kontrollkästchen `input` haben.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)"));
```

Ihr drittes Kontrollkästchen `input` sollte ein `id`-Attribut haben, das auf `"location"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('id'), 'location');
```

Ihr drittes Kontrollkästchen `input` sollte ein `name`-Attribut haben, das auf `"choice"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('name'), 'choice');
```

Ihr drittes Kontrollkästchen `input` sollte ein `value`-Attribut haben, das auf `"location"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox']:nth-of-type(3)")?.getAttribute('value'), 'location');
```

Sie sollten ein `label`-Element nach Ihrem Kontrollkästchen `input` mit dem Text `Location` haben.

```js 
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[value='location'] + label")?.textContent?.trim(), 'Location');
```

Ihr `label` sollte das `for`-Attribut haben, das auf `"location"` gesetzt ist.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'location');
```

Sie sollten ein viertes Kontrollkästchen `input` unter Ihrem `label` haben.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']"));
```

Sie sollten ein viertes Kontrollkästchen `input` mit einem `id`-Attribut haben, das auf `"reputation"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('id'), 'reputation');
```

Ihr viertes Kontrollkästchen `input` sollte ein `name`-Attribut haben, das auf `"choice"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('name'), 'choice');
```

Ihr viertes Kontrollkästchen `input` sollte ein `value`-Attribut haben, das auf `"reputation"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label:nth-of-type(3) + input[type='checkbox']")?.getAttribute('value'), 'reputation');
```

Sie sollten ein `label`-Element mit dem Text `Reputation` unter dem Kontrollkästchen für Reputation haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.textContent?.trim(), 'Reputation');
```

Ihr `label` sollte das `for`-Attribut haben, das auf `"reputation"` gesetzt ist.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="reputation"] + label')?.getAttribute('for'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
