---
id: 66a9419e2d18476c645ce693
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Die Zahl `input` wird verwendet, um ein numerisches Eingabefeld zu erstellen.

Hier ist ein Beispiel für ein Zahlen-Eingabefeld:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Die `min`- und `max`-Attribute werden verwendet, um die minimalen und maximalen Werte festzulegen, die im Eingabefeld eingegeben werden können.

Fügen Sie unter Ihrem `label`-Element ein Eingabefeld mit dem `type`-Attribut, das auf `"number"` gesetzt ist, und einem `id` von `"age"` hinzu.

Das `name`-Attribut sollte auf `"age"` gesetzt sein, das `min`-Attribut auf `"3"` und das `max`-Attribut auf `"100"`.

# --hints--

Sie sollten ein Zahlen-`input` unter Ihrem `label`-Element haben.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Ihr Zahlen-`input` sollte ein `id` von `age` haben.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Ihr Zahlen-`input` sollte ein `name`-Attribut mit dem Wert `age` haben.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Ihr Zahlen-`input` sollte ein `min`-Attribut mit dem Wert `"3"` haben.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Ihr Zahlen-`input` sollte ein `max`-Attribut mit dem Wert `"100"` haben.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
