---
id: 66a95d0eff8168747805f1f3
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Wenn Sie möchten, dass ein Benutzer mehrere Optionen aus einer Liste auswählt, können Sie Kontrollkästchen verwenden.

Hier ist ein Beispiel, wie man mit Kontrollkästchen für Essensoptionen arbeitet:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

Das `value`-Attribut wird verwendet, um den Wert anzugeben, der an den Server gesendet wird, wenn das Formular abgesendet wird.

Fügen Sie unter Ihrem `legend`-Element ein Kontrollkästchen-`input` mit den `id`- und `value`-Attributen, die auf `"ads"` gesetzt sind, und dem `name`-Attribut, das auf `"choice"` gesetzt ist, hinzu.

Fügen Sie unter Ihrem Kontrollkästchen `input` ein `label`-Element mit dem Text `Social Media Ads` hinzu. Das `for`-Attribut sollte auf `"ads"` gesetzt sein.

# --hints--

Sie sollten ein Kontrollkästchen `input` unter Ihrem `legend`-Element haben.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

Ihr Kontrollkästchen `input` sollte ein `id`-Attribut haben, das auf `"ads"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

Ihr Kontrollkästchen `input` sollte ein `name`-Attribut haben, das auf `"choice"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

Ihr Kontrollkästchen `input` sollte ein `value`-Attribut haben, das auf `"ads"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Sie sollten ein `label`-Element unter Ihrem Kontrollkästchen `input` haben.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Ihr `label`-Element sollte den Text `Social Media Ads` enthalten.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Ihr `label`-Element sollte ein `for`-Attribut haben, das auf `"ads"` gesetzt ist.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
