---
id: 66a9521bc70162712caf118d
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Wenn Sie möchten, dass Benutzer eine Option aus einer Liste von Optionen auswählen, können Sie eine Gruppe von Optionsschaltflächen verwenden.

Hier ist ein Beispiel für zwei Optionsschaltflächen:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

In diesem Beispiel sind die Optionsschaltflächen durch die Verwendung desselben `name`-Attributwerts gruppiert. Das bedeutet, dass jeweils nur eine Optionsschaltfläche ausgewählt werden kann.

Fügen Sie unter Ihrem `legend`-Element eine `radio`-Schaltfläche mit dem `id`, das auf `"yes-option"` gesetzt ist, dem `name`-Attribut, das auf `"hotel-stay"` gesetzt ist, und einem `value`-Attribut, das auf `"yes"` gesetzt ist, hinzu.

Fügen Sie unter Ihrer `radio`-Schaltfläche ein `label`-Element mit dem Text `Yes` und einem `for`-Attribut, das auf `"yes-option"` gesetzt ist, hinzu.

# --hints--

Sie sollten eine `radio`-Schaltfläche unter Ihrem `legend`-Element haben.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]'));
```

Ihre `radio`-Schaltfläche sollte ein `id`-Attribut haben, das auf `"yes-option"` gesetzt ist.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.id, 'yes-option');
```

Ihre `radio`-Schaltfläche sollte ein `name`-Attribut haben, das auf `"hotel-stay"` gesetzt ist.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.name, 'hotel-stay');
```

Ihre Optionsschaltfläche sollte ein `value`-Attribut haben, das auf `"yes"` gesetzt ist.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) legend + input[type="radio"]')?.value,
  'yes'
);
```

Sie sollten ein `label`-Element unter Ihrer `radio`-Schaltfläche haben.

```js
assert.exists(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Ihr `label`-Element sollte ein `for`-Attribut haben, das auf `"yes-option"` gesetzt ist.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, 'yes-option');
```

Ihr `label`-Element sollte den Text `Yes` haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.innerText?.trim(), 'Yes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
