---
id: 66a937e74920ba68ebe5e86d
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit den Attributen `placeholder` und `required` so arbeitet:

```html
<input type="text" placeholder="e.g., John Doe" required>
```

Fügen Sie für Ihr bestehendes `input`-Element ein `placeholder`-Attribut mit dem Wert `"e.g., John Doe"` hinzu.

Fügen Sie außerdem das `required`-Attribut zum `input`-Element hinzu.

# --hints--

Sie sollten das bestehende `input`-Element nicht entfernen.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Ihr `input` sollte ein `placeholder`-Attribut mit dem Wert `"e.g., John Doe"` haben.

```js
assert.strictEqual(document.querySelector('fieldset input')?.getAttribute('placeholder'), 'e.g., John Doe');
```

Ihr `input` sollte ein `required`-Attribut haben.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
