---
id: 66a969951120be7818d8ee49
title: Schritt 24
challengeType: 0
dashedName: step-24
---

# --description--

Der nächste Abschnitt des Formulars gibt den Nutzenden die Möglichkeit, eine Bewertung für das Hotel abzugeben.

Beginnen Sie damit, ein neues `fieldset`-Element hinzuzufügen, in das ein `legend`-Element verschachtelt ist. Das `legend`-Element sollte den Text `Ratings` enthalten.

Fügen Sie unter dem `legend`-Element ein `label`-Element mit dem Text `How was the service?` hinzu. Das `for`-Attribut sollte auf `"service"` gesetzt werden.

# --hints--

Sie sollten ein `fieldset`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 4);
```

Ihr `fieldset` sollte ein `legend`-Element enthalten.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend'), 1);
```

Ihre Legende sollte den Text `Ratings` haben.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent.trim(), 'Ratings');
```

Ihr `fieldset` sollte ein `label`-Element unter dem `legend` haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend + label'), 1);
```

Ihr `label` sollte den Text `How was the service?` enthalten.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent.trim(), 'How was the service?');
```

Ihr `label` sollte das `for`-Attribut auf `"service"` gesetzt haben.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
