---
id: 66a93c95bc58e26a8fe95818
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Fügen Sie als Nächstes ein `input` mit dem Typ `"email"` unter Ihrem E-Mail-`label` hinzu. Dieses `input` sollte ein `id`-Attribut mit dem Wert `"email"` haben. Setzen Sie außerdem das `name`-Attribut auf den Wert `"email"`.

Dieses `input` ist ebenfalls erforderlich, also stellen Sie sicher, dass Sie das `required`-Attribut hinzufügen.

Fügen Sie schließlich ein `placeholder`-Attribut mit dem Wert `"example@email.com"` hinzu.

# --hints--

Sie sollten ein `input`-Element unter Ihrem E-Mail-`label` haben.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Ihr `input`-Element sollte ein `type` von `"email"` haben.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Ihr `input`-Element sollte ein `id`-Attribut mit dem Wert `"email"` haben.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Ihr `input`-Element sollte ein `name`-Attribut mit dem Wert `"email"` haben.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Ihr `input`-Element sollte ein `required`-Attribut haben.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Ihr `input`-Element sollte ein `placeholder`-Attribut mit dem Wert `"example@email.com"` haben.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
