---
id: 66a94ea5df66236ebed260e8
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Der nächste Abschnitt im Formular wird dafür zuständig sein, die Nutzer zu fragen, ob sie schon einmal in dem Hotel übernachtet haben.

Beginnen Sie damit, ein `fieldset`-Element hinzuzufügen.

Fügen Sie innerhalb des `fieldset`-Elements ein `legend`-Element mit dem Text von `Was this your first time at our hotel?` hinzu.

# --hints--

Sie sollten ein `fieldset`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

Sie sollten ein `legend`-Element innerhalb Ihres `fieldset`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

Ihr `legend`-Element sollte den Text von `Was this your first time at our hotel?` enthalten.

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
