---
id: 66a9577022877d72d8f43b4f
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Der nächste Abschnitt des Formulars wird die Nutzer fragen, warum sie genau dieses Hotel gewählt haben. Die Nutzer haben die Möglichkeit, mehrere Optionen auszuwählen.

Beginnen Sie damit, ein weiteres `fieldset`-Element hinzuzufügen.

Fügen Sie innerhalb des `fieldset`-Elements ein `legend`-Element mit dem Text `Why did you choose to stay at our hotel? (Check all that apply)` hinzu.

# --hints--

Sie sollten ein `fieldset`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

Sie sollten ein `legend`-Element innerhalb des `fieldset`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

Ihr `legend`-Element sollte den Text `Why did you choose to stay at our hotel? (Check all that apply)` enthalten. Überprüfen Sie Rechtschreibung und Abstände sorgfältig.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
