---
id: 66a9824ac5d9f77ec304969f
title: Schritt 31
challengeType: 0
dashedName: step-31
---

# --description--

Wenn Sie den Nutzenden mehr Platz zum Schreiben ihrer Kommentare geben möchten, können Sie ein `textarea`-Element verwenden.

Das `textarea`-Element ist ein mehrzeiliges Texteingabefeld, das es den Nutzenden erlaubt, Text einzugeben, der länger als eine einzelne Zeile ist. Es kann verwendet werden, um ein Kommentarfeld, eine Nachrichteneingabe oder andere Texteingaben zu erstellen, die mehrere Zeilen benötigen.

Hier ist ein Beispiel für ein `textarea`-Element:

```html
<textarea id="comments" name="comments" rows="4" cols="50"></textarea>
```

Das `rows`-Attribut wird verwendet, um die sichtbare Höhe des `textarea` festzulegen, und das `cols`-Attribut wird verwendet, um die sichtbare Breite des `textarea` zu bestimmen.

Fügen Sie unter Ihrem `label`-Element ein `textarea`-Element hinzu. Im nächsten Schritt fügen Sie die notwendigen Attribute hinzu.

# --hints--

Sie sollten ein `textarea`-Element unter Ihrer Bezeichnung haben.

```js
assert.exists(document.querySelector('label + textarea'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="choice" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="choice"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="choice" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="choice"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="choice" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>

          <label for="food">How was the food?</label>

          <select name="food" id="food">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>
        </fieldset>

        <label for="comments">Other Comments?</label>
        
        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
