---
id: 66a930b20f589b6664c51cb0
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Wenn ein Benutzer seinen vollständigen Namen angibt, akzeptiert das `input` Klartext.

In vorherigen Lektionen haben Sie gelernt, wie Sie mit dem `type`-Attribut so arbeiten:

```html
<input type="text">
```

Fügen Sie für Ihr bestehendes `input`-Element ein `type`-Attribut hinzu, das auf `"text"` gesetzt ist.

Geben Sie außerdem dem `input`-Element ein `id`-Attribut mit dem Wert `"full-name"`.

# --hints--

Sie sollten das bestehende `input`-Element nicht entfernen.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Sie sollten kein neues `input`-Element erstellen.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ihr `input`-Element sollte ein `type`-Attribut mit dem Wert `"text"` haben.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Ihr `input`-Element sollte ein `id`-Attribut mit dem Wert `"full-name"` haben.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
