---
id: 66a93730719e1f68410cce54
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Das Attribut `name` wird verwendet, um Formulardaten zu identifizieren, nachdem sie an den Server übermittelt wurden.

Hier ist ein Beispiel, wie man das Attribut `name` verwendet:

```html
<input type="email" name="email">
```

Fügen Sie dem Element `name` ein Attribut `input` mit dem Wert `"name"` hinzu.

# --hints--

Sie sollten das vorhandene Element `input` nicht entfernen.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Sie sollten kein neues Element `input` erstellen.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Ihr Element `input` sollte ein Attribut `name` mit dem Wert `"name"` haben.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
