---
id: 66a93bbe65a26169dbf3bc39
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Ihr Hotel-Feedback-Formular sollte auch eine E-Mail-Adresse vom Benutzer erfassen.

Beginnen Sie damit, dem Formular ein neues `label`-Element mit dem Text `Email address (required):` hinzuzufügen. Ihr `label`-Element sollte ein `for`-Attribut mit dem Wert `"email"` haben.

# --hints--

Sie sollten ein `label`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Ihr `label`-Element sollte den Text `Email address (required):` enthalten.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Ihr `label`-Element sollte ein `for`-Attribut mit dem Wert `"email"` haben.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
