---
id: 66a954b2bcddba72076c1857
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Fügen Sie unter Ihrem `label`-Element eine `radio`-Schaltfläche mit dem `id`-Attribut hinzu, das auf `"no-option"` gesetzt ist, dem `name`-Attribut mit dem Wert `"hotel-stay"` und einem `value`-Attribut mit dem Wert `"no"`.

Fügen Sie unter Ihrer neuen `radio`-Schaltfläche ein weiteres `label`-Element mit dem `for`-Attribut hinzu, das auf `"no-option"` gesetzt ist. Der Text für das `label` sollte `No` sein.

Wenn Sie fertig sind, können Sie die Optionsschaltflächen ausprobieren, indem Sie jeweils eine Option auswählen.

# --hints--

Sie sollten ein `input`-Element unter Ihrem ersten `label`-Element haben.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Ihr `input`-Element sollte ein `type`-Attribut mit dem Wert `radio` haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Ihr `input`-Element sollte ein `id`-Attribut mit dem Wert `"no-option"` haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Ihr `input`-Element sollte ein `name`-Attribut mit dem Wert `"hotel-stay"` haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Ihr Eingabe-Element sollte ein `value`-Attribut mit dem Wert `"no"` haben.

```js
assert.strictEqual(
  document.querySelector('fieldset:nth-of-type(2) label + input')?.value,
  'no'
);
```

Sie sollten ein zweites `label`-Element unter Ihrem zweiten `input`-Element haben.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Ihr `label`-Element sollte ein `for`-Attribut mit dem Wert `"no-option"` haben.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Ihr `label`-Element sollte den Text `No` enthalten.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          --fcc-editable-region--
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
