---
id: 67a51d3a8a6fe123b77b0c6e
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

`input`-Elemente können ein `size`-Attribut haben. Dieses Attribut definiert die Anzahl der Zeichen, die sichtbar sein sollen, während der Benutzer in das `input` tippt. Der Wert von `size` sollte eine positive ganze Zahl sein. Wenn `size` nicht angegeben ist oder mit einem ungültigen Wert angegeben wird, erhält das `input` die vom Browser festgelegte Standardbreite.

```html
<label for="last-name">Last Name:</label>
<input id="last-name" name="last-name" type="text" size="10" />
```

Geben Sie den Eingaben für Name und E-Mail ein `size`-Attribut mit dem Wert `"20"`.

# --hints--

Sie sollten dem `input` für den Namen ein `size`-Attribut mit dem Wert `"20"` geben.

```js
const nameInput = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(nameInput?.getAttribute('size'), '20');
```

Sie sollten dem `input` für die E-Mail ein `size`-Attribut mit dem Wert `"20"` geben.

```js
const emailInput = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(emailInput?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
