---
id: 594810f028c0303b75339acc
title: ABC 问题
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

你将获得一组 ABC 块（例如，儿童字母块）。 有 20 个块，每个块上有两个字母。 在所有块的所有边之中，保证有完整的字母表。 块的样本集合：

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

实现一个函数，该函数接受一个字符串（单词），并确定该单词是否可以用给定的块集合拼写。

要记住的一些规则：

<ul>
  <li>Once a letter on a block is used, that block cannot be used again.</li>
  <li>The function should be case-insensitive.</li>
</ul>

# --hints--

`canMakeWord` 应该是一个函数。

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` 应该返回一个布尔值。

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` 应该返回 true。

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` 应该返回 false。

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` 应该返回 true。

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` 应返回 false。

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` 应该返回 true。

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` 应该返回 true。

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --after-user-code--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
