---
id: 594810f028c0303b75339acf
title: 阿克曼函数
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

阿克曼函数是递归函数的一个经典例子，值得注意的是它不是原始递归函数。 它的价值增长非常快，其调用树的大小也是如此。

阿克曼函数通常定义如下：

$A(m, n) = \\begin{cases} n+1 & \\mbox{如果 } m = 0 \\\\ A(m-1, 1) & \\mbox{如果 } m > 0 \\mbox{ 且 } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{如果 } m > 0 \\mbox{ 且 } n > 0. \\end{cases}$

它的参数从不是负的，而且总是有终点。

# --instructions--

编写一个函数，返回 $A(m, n)$ 的值。 任意精度是首选（因为函数增长如此之快），但不是必需的。

# --hints--

`ack` 应该是一个函数。

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` 应该返回 1。

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` 应该返回 3。

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` 应该返回 13。

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` 应该返回 61。

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
