---
id: 594810f028c0303b75339ace
title: 累加器工厂
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Paul Graham提出的一个问题是创建一个函数，它接受一个 (数字) 参数，并返回另一个累加器函数。 返回的累加器函数又接受一个数字参数，并返回到目前为止传递给该累加器的所有数值的总和（包括创建累加器时传递的初始值）。

# --instructions--

创建一个函数，它接受一个数字 $n$ 并生成累加器函数，该函数返回传递给它们的每个数字的总和。

**规则：**

请勿使用全局变量。

**提示：**

关闭可保存外部状态。

# --hints--

`accumulator` 应该是一个函数。

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` 应该返回一个函数。

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` 应该返回一个数字。

```js
assert(typeof accumulator(0)(2) === 'number');
```

传入值 3、-4、1.5 和 5 应返回 5.5。

```js
assert(testFn(5) === 5.5);
```

# --seed--

## --after-user-code--

```js
const testFn = typeof accumulator(3) === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
