---
id: 594da033de4190850b893874
title: 平均值/均方根
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

计算包含 1 到 10 的数字的均方根（RMS）。

<abbr title="Root mean square">RMS</abbr> 是通过对数字的平方的平均值取平方根计算得出的，公式为：

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms` 应该是一个函数。

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` 应该等于 `6.2048368229954285`。

```js
assert.equal(rms(arr1), answer1);
```

# --seed--

## --after-user-code--

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
```

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
