---
id: 59622f89e4e137560018a40e
title: Hofstadter 图-图序列
challengeType: 1
forumTopicId: 302286
dashedName: hofstadter-figure-figure-sequences
---

# --description--

Hofstadter Figure-Figure 序列 $R_n$ 和 $S_n$ 定义为

$R_1 = 1\\ ;\\ S_1 = 2 \\\\R_n = R_{n-1} + S_{n-1}, \\quad n>1.$

具体来说，序列 $R_n$ 包含值

<pre>1, 3, 7, 12, 18, ...</pre>

且序列 $S_n$ 包含这些值

<pre>2, 4, 5, 6, 8, ...</pre>

序列 $R_n$ 由递推关系 $R_n = R_{n-1} + S_{n-1}$ 定义，而 $S_n$ 定义为不包含在序列 $R_n$ 中的正整数序列。

# --instructions--

创建两个名为 `ffr` 和 `ffs` 的函数，分别返回任意索引 `n` 的 `R(n)` 或 `S(n)`。请注意，霍夫斯塔特图形-图形序列是从 1 开始索引的，且 $R_1 = 1$ 和 $S_1 = 2$。

不应假定 `n` 的最大值。

**引用**

 <p>Rosetta: <a href='https://rosettacode.org/wiki/Hofstadter_Figure-Figure_sequences' target='_blank'>Hofstadter Figure-Figure sequences</a></p>.


# --hints--

`ffr` 应该是一个函数。

```js
assert(typeof ffr === 'function');
```

`ffs` 应该是一个函数。

```js
assert(typeof ffs === 'function');
```

`ffr` 应该返回整数。

```js
assert(Number.isInteger(ffr(1)));
```

`ffs` 应该返回整数。

```js
assert(Number.isInteger(ffs(1)));
```

`ffr(10)` 应该返回 `69`

```js
assert.equal(ffr(ffrParamRes[0][0]), ffrParamRes[0][1]);
```

`ffr(50)` 应该返回 `1509`

```js
assert.equal(ffr(ffrParamRes[1][0]), ffrParamRes[1][1]);
```

`ffr(100)` 应该返回 `5764`

```js
assert.equal(ffr(ffrParamRes[2][0]), ffrParamRes[2][1]);
```

`ffr(1000)` 应该返回 `526334`

```js
assert.equal(ffr(ffrParamRes[3][0]), ffrParamRes[3][1]);
```

`ffs(10)` 应该返回 `14`

```js
assert.equal(ffs(ffsParamRes[0][0]), ffsParamRes[0][1]);
```

`ffs(50)` 应该返回 `59`

```js
assert.equal(ffs(ffsParamRes[1][0]), ffsParamRes[1][1]);
```

`ffs(100)` 应该返回 `112`

```js
assert.equal(ffs(ffsParamRes[2][0]), ffsParamRes[2][1]);
```

`ffs(1000)` 应该返回 `1041`

```js
assert.equal(ffs(ffsParamRes[3][0]), ffsParamRes[3][1]);
```

# --seed--

## --after-user-code--

```js
const ffrParamRes = [[10, 69], [50, 1509], [100, 5764], [1000, 526334]];
const ffsParamRes = [[10, 14], [50, 59], [100, 112], [1000, 1041]];
```

## --seed-contents--

```js
function ffr(n) {
  return n;
}

function ffs(n) {
  return n;
}
```

# --solutions--

```js
const R = [null, 1];
const S = [null, 2];

function extendSequences (n) {
  let current = Math.max(R[R.length - 1], S[S.length - 1]);
  let i;
  while (R.length <= n || S.length <= n) {
    i = Math.min(R.length, S.length) - 1;
    current += 1;
    if (current === R[i] + S[i]) {
      R.push(current);
    } else {
      S.push(current);
    }
  }
}

function ffr (n) {
  extendSequences(n);
  return R[n];
}

function ffs (n) {
  extendSequences(n);
  return S[n];
}
```
