---
id: 59694356a6e7011f7f1c5f4e
title: 空当接龙的交易卡
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* 是保罗·阿尔菲尔（Paul Alfille）于 1978 年引入 PLATO 系统的单人纸牌游戏。微软的吉姆·霍恩（Jim Horne）将其命名为 FreeCell，并重新实现了适用于 DOS 和 Windows 的版本。这个版本引入了编号为 32000 的发牌。

随着游戏的流行，Jim Horne 公开了算法，FreeCell 的其他实现也开始复制微软的交易。 这些交易编号从 1 到 32000。 微软的较新版本有 100 万笔交易，编号从 1 到 1000000；某些实现允许超出该范围的数字。

该算法使用来自 Microsoft C 的此线性同余生成器：

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$ 的范围是 0 到 32767。</li>
</ul>

算法如下：

<ol>
  <li>用发牌编号作为随机数生成器的种子。
  </li><li>创建一个包含 52 张牌的数组：梅花 A、方块 A、红心 A、黑桃 A，梅花 2、方块 2，依此类推，牌面顺序为 A、2、3、4、5、6、7、8、9、10、J、Q、K。数组索引从 0 到 51，梅花 A 在 0，黑桃 K 在 51。</li>
  <li>直到数组为空：</li>
    <ul>
      <li>选择一个随机牌，其<i>索引</i> ≡ <i>下一个随机数</i>（模<i>数组长度</i>）。</li>
      <li>将这张随机牌与数组最后一张牌交换。</li>
      <li>从数组中移除这张随机牌。（数组长度减 1。）</li>
      <li>发出这张随机牌。</li>
    </ul>
  <li>将所有 52 张牌正面朝上，分布在 8 列中。前 8 张牌分别放入 8 列，接下来的 8 张牌叠放在前 8 张牌上，依此类推。</li>
</ol>

**示例：**

**发牌顺序**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**游戏 #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**游戏 #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

编写一个函数，以与该算法相同的顺序获取发牌号码和发牌。 该函数必须返回一个代表空当接龙板的二维数组。

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` 应该是一个函数。

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` 应该返回一个对象。

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` 应该返回一个长度为 7 的数组。

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` 应该返回一个与示例“Game #1”相同的数组

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` 应该返回一个与示例“Game #617”相同的数组

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
