---
id: 598ee8b91b410510ae82efef
title: 可扩展素数生成器
challengeType: 1
forumTopicId: 302262
dashedName: extensible-prime-generator
---

# --description--

按顺序编写一个素数生成器，它会自动调整以适应任何合理的高素数的生成。

生成器应该能够：

<ul>
  <li>显示前 <code>n</code> 个素数</li>
  <li>显示一个范围内的素数</li>
  <li>显示一个范围内素数的数量</li>
  <li>显示第 <code>n<sup>th</sup></code> 个素数</li>
</ul>

该函数应有两个参数。第一个参数接收 `n` 或作为数组的范围。第二个参数接收一个布尔值，用于指定函数是返回素数数组还是单个数字（范围内素数的数量或第 <code>n<sup>th</sup></code> 个素数）。函数应根据参数返回一个数组。

# --hints--

`primeGenerator` 应该是一个函数。

```js
assert(typeof primeGenerator === 'function');
```

`primeGenerator(20, true)` 应该返回 `[2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71]`。

```js
assert.deepEqual(primeGenerator(20, true), [
  2,
  3,
  5,
  7,
  11,
  13,
  17,
  19,
  23,
  29,
  31,
  37,
  41,
  43,
  47,
  53,
  59,
  61,
  67,
  71
]);
```

`primeGenerator([100, 150], true)` 应该返回 `[101, 103, 107, 109, 113, 127, 131, 137, 139, 149]`。

```js
assert.deepEqual(primeGenerator([100, 150], true), [
  101,
  103,
  107,
  109,
  113,
  127,
  131,
  137,
  139,
  149
]);
```

`primeGenerator([7700, 8000], false)` 应该返回 `30`。

```js
assert.equal(primeGenerator([7700, 8000], false), 30);
```

`primeGenerator(10000, false)` 应该返回 `104729`。

```js
assert.equal(primeGenerator(10000, false), 104729);
```

# --seed--

## --seed-contents--

```js
function primeGenerator(num, showPrimes) {

}
```

# --solutions--

```js
function primeGenerator(num, showPrimes) {
  let i,
    arr = [];

  function isPrime(num) {
    // try primes <= 16
    if (num <= 16) { return (
      num == 2 || num == 3 || num == 5 || num == 7 || num == 11 || num == 13
    ); }
    // cull multiples of 2, 3, 5 or 7
    if (num % 2 == 0 || num % 3 == 0 || num % 5 == 0 || num % 7 == 0)
      { return false; }
    // cull square numbers ending in 1, 3, 7 or 9
    for (let i = 10; i * i <= num; i += 10) {
      if (num % (i + 1) == 0) return false;
      if (num % (i + 3) == 0) return false;
      if (num % (i + 7) == 0) return false;
      if (num % (i + 9) == 0) return false;
    }
    return true;
  }

  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (isPrime(i)) arr.push(i);
    // first x primes
    if (showPrimes) return arr;
    // xth prime
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (isPrime(i)) arr.push(i);
    // primes between x .. y
    if (showPrimes) return arr;
    // number of primes between x .. y
    return arr.length;
  }
}
```
