---
id: 598eef80ba501f1268170e1e
title: 斐波那契 n 步数列
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

这些数字串是对常规斐波那契数列的扩展，其中：

<ol>
  <li>当 $n = 2$ 时，我们得到斐波那契数列；初始值为 $[1, 1]$，且 $F_k^2 = F_{k-1}^2 + F_{k-2}^2$</li>
  <li>当 $n = 3$ 时，我们得到三项斐波那契数列；初始值为 $[1, 1, 2]$，且 $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$</li>
  <li>当 $n = 4$ 时，我们得到四项斐波那契数列；初始值为 $[1, 1, 2, 4]$，且 $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$…</li>
  <li>对于一般的 $n>2$，我们有斐波那契 n 步数列 - $F_k^n$；初始值为前一个 $(n-1)$ 阶斐波那契 n 步数列 $F_k^{n-1}$ 的前 $n$ 个值；该第 $n$ 阶数列的第 $k$ 个值为 $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$</li>
</ol>

对于较小的 $n$，有时会使用希腊数字前缀来单独命名每个数列。

斐波那契 n 步数列：

| $n$ | 数列名称   | 数值                                                    |
| --- | ---------- | ------------------------------------------------------- |
| 2   | fibonacci  | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...          |
| 3   | tribonacci | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...     |
| 4   | tetranacci | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...   |
| 5   | pentanacci | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...   |
| 6   | hexanacci  | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...   |
| 7   | heptanacci | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ...  |
| 8   | octonacci  | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ...  |
| 9   | nonanacci  | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ...  |
| 10  | decanacci  | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ...  |

可以生成相关数列，通过更改初始值：Lucas 数列像 $n=2$ 的斐波那契数列一样求和前述的两个值，但使用 $\[2, 1]$ 作为初始值。

# --instructions--

编写一个函数来生成斐波那契 n 步数列和 Lucas 数列。第一个参数是 $n$。第二个参数是要返回的元素数量。第三个参数指定输出斐波那契数列还是 Lucas 数列。如果参数是 `"f"`，则返回斐波那契数列；如果是 `"l"`，则返回 Lucas 数列。数列必须以数组形式返回。

# --hints--

`fib_luc` 应该是一个函数。

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` 应该返回 `[1,1,2,3,5,8,13,21,34,55]`。

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` 应该返回 `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`。

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` 应该返回 `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`。

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` 应该返回 `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`。

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` 应该返回 `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`。

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` 应该返回 `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`。

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` 应该返回 `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`。

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
