---
id: 599d1566a02b571412643b84
title: 埃塞俄比亚乘法
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

埃塞俄比亚乘法是一种仅使用加法、加倍和减半将整数相乘的方法。

**方法：**

<ol>
  <li>取两个要相乘的数字，写在两列的顶部</li>
  <li>在左侧列中反复将最后一个数字除以二取整，丢弃任何余数，并将结果写在同一列的上一个数字下方，直到写出值为 <code>1</code></li>
  <li>在右侧列中反复将最后一个数字加倍，并将结果写在下方。当你添加的结果与左侧列显示 <code>1</code> 的行在同一行时停止</li>
  <li>检查生成的表格，丢弃左侧列中值为偶数的任何行</li>
  <li>将右侧列中剩余的值相加，得到原始两个数字相乘的结果</li>
</ol>

**例如：** `17 × 34`

<pre>17   34
</pre>

将第一列减半：

<pre>17   34
8
4
2
1
</pre>

将第二列加倍：

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

第一个单元格为偶数的删除行：

<pre>17   34
8    <s>68</s>
4   <s>136</s>
2   <s>272</s>
1   544
</pre>

对右侧栏中的剩余数字求和：

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

所以`17`乘以`34`，按照埃塞俄比亚的方法是`578`。

# --instructions--

任务是定义三个命名的函数/方法/过程/子例程：

<ol>
  <li>一个用于将整数除以二取整，</li>
  <li>一个用于将整数加倍，和</li>
  <li>一个用于判断整数是否为偶数</li>
</ol>

使用这些函数创建一个执行埃塞俄比亚乘法的函数。

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` 应该是一个函数。

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` 应该返回 `578`。

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` 应该返回 `1058`。

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` 应该返回 `324`。

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` 应该返回 `5488`。

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` 应该返回 `4662`。

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
