---
id: 59e0a8df964e4540d5abe599
title: 执行 Brain****
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

编写一个函数来实现 Brain**** 解释器。该函数将接受一个字串作为参数，并应返回一个字串作为输出。详细信息如下：

RCBF 是一组<a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">Brainf***</a> 编译器和解释器，使用多种语言为 Rosetta Code 编写。

下面是 RCBF 各版本的链接。

实现只需正确实现以下命令：

| 命令              | 描述                                                                              |
| ----------------- | --------------------------------------------------------------------------------- |
| <code>></code>    | 将指针向右移动                                                                   |
| <code>&lt;</code> | 将指针向左移动                                                                   |
| <code>+</code>    | 增加指针下内存单元的值                                                           |
| <code>-</code>    | 减少指针下内存单元的值                                                           |
| <code>.</code>    | 输出指针下内存单元表示的字符                                                     |
| <code>,</code>    | 输入一个字符并保存到指针下的内存单元                                             |
| <code>\[</code>   | 如果指针下内存单元为 0，则跳过匹配的 <code>]</code>                              |
| <code>]</code>    | 如果指针下内存单元非零，则跳回匹配的 <code>\[</code>                            |

允许任意内存单元大小，EOF（文件结束）支持为可选，内存是否有边界也为可选。

# --hints--

`brain(globalThis.hello)` 应返回一个字串

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` 应返回 "A"

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` 应返回 `Hello World!\n`

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` 应返回 `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89`

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
