---
id: 5a23c84252665b21eecc8041
title: 系列的总和
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

计算一个<em>串行</em>的第 **n**<sup>项</sup>，即对应<em>序列</em>的前 **n** 项之和。非正式地，这个值或当 **n** 趋近于无穷大时的极限，也称为*串行的和*，这也是本任务标题的由来。对于本任务，使用公式：$S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$。

# --instructions--

编写一个函数，以 $a$ 和 $b$ 作为参数，并将 $a^{th}$ 的总和返回给序列的 $b^{th}$ 成员。

# --hints--

`sum` 应该是一个函数。

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` 应该返回一个数字。

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` 应该返回 `1.6349839001848923`。

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` 应该返回 `0.009262256361481223`。

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` 应该返回 `0.044086990748706555`。

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` 应该返回 `0.08619778593108679`。

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` 应该返回 `1.5497677311665408`。

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
