---
id: 5a23c84252665b21eecc7edf
title: 最小公倍数
challengeType: 1
forumTopicId: 302301
dashedName: least-common-multiple
---

# --description--

12 和 18 的最小公倍数是 36，因为 12 是因数（12 × 3 = 36），18 是因数（18 × 2 = 36），且没有比 36 小的正整数同时具有这两个因数。作为特殊分支，如果 $m$ 或 $n$ 为零，则最小公倍数为零。计算最小公倍数的一种方法是迭代所有 $m$ 的倍数，直到找到一个也是 $n$ 的倍数。如果你已经有了 <a href="https://rosettacode.org/wiki/Greatest_common_divisor" target="_blank" rel="noopener noreferrer nofollow">最大公因数</a> 的 $gcd$，那么这个公式可以计算 $lcm$。 

$$
\\operatorname{lcm}(m, n) = \\frac{|m \\times n|}{\\operatorname{gcd}(m, n)}
$$

# --instructions--

计算整数数组的最小公倍数。 给定 *m* 和 *n*，最小公倍数是同时具有 <em x-id ="3">m</em> 和 *n* 作为因子。

# --hints--

`LCM` 应该是一个函数。

```js
assert(typeof LCM == 'function');
```

`LCM([2, 4, 8])` 应该返回一个数字。

```js
assert(typeof LCM([2, 4, 8]) == 'number');
```

`LCM([2, 4, 8])` 应该返回 `8`。

```js
assert.equal(LCM([2, 4, 8]), 8);
```

`LCM([4, 8, 12])` 应该返回 `24`。

```js
assert.equal(LCM([4, 8, 12]), 24);
```

`LCM([3, 4, 5, 12, 40])` 应该返回 `120`。

```js
assert.equal(LCM([3, 4, 5, 12, 40]), 120);
```

`LCM([11, 33, 90])` 应该返回 `990`。

```js
assert.equal(LCM([11, 33, 90]), 990);
```

`LCM([-50, 25, -45, -18, 90, 447])` 应该返回 `67050`。

```js
assert.equal(LCM([-50, 25, -45, -18, 90, 447]), 67050);
```

# --seed--

## --seed-contents--

```js
function LCM(A) {

}
```

# --solutions--

```js
function LCM(A) {
  var n = A.length,
    a = Math.abs(A[0]);
  for (var i = 1; i < n; i++) {
    var b = Math.abs(A[i]),
      c = a;
    while (a && b) {
      a > b ? (a %= b) : (b %= a);
    }
    a = Math.abs(c * A[i]) / (a + b);
  }
  return a;
}
```
