---
id: 5a23c84252665b21eecc8036
title: 从字符串中剥离控制代码和扩展字符
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

任务是从字符串中去除控制代码和扩展字符。 该解决方案应演示如何实现以下每个结果： 去除了控制代码和扩展字符的字符串。 在 ASCII 中，控制代码具有从 0 到 31 和 127 的十进制代码。 在基于 ASCII 的系统上，如果控制代码被剥离，则结果字符串的所有字符都在 ASCII 表上的十进制 32 到 126 范围内。 在基于非 ASCII 的系统上，出于此任务的目的，我们将在 ASCII 表（在 ASCII 范围为 32 到 126 十进制）中没有相应字形的字符视为扩展字符。

# --hints--

`strip` 应该是一个函数。

```js
assert(typeof strip == 'function');
```

`strip("abc")` 应该返回一个字符串。

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` 应该返回 `"abcd"`。

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` 应该返回 `" abcdef"`。

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` 应该返回 `"abdf"`。

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` 应该返回 `"Franais."`。

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` 应该返回 `"123abcDEF+-*/"`。

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
