---
id: 5a23c84252665b21eecc7ec5
title: 约瑟夫斯问题
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

约瑟夫问题是一个带有严峻描述的数学难题：$n$ 个囚犯站成一个圆圈，循序编号从 $0$ 到 $n-1$。

`josephus`应该是一个功能。

随着过程的继续，圆圈变得越来越小，直到只剩下一个囚犯，然后被释放。

例如，如果$ n = 5 $囚犯和$ k = 2 $，囚犯被杀的命令（我们称之为“杀戮序列”）将是1,3,0和4，幸存者将是＃2。

给定任意 $n, k > 0$，查找最终存活的囚犯。

在一次这样的事件中，有 41 名囚犯，每第 3<sup>rd</sup> 名囚犯被处死（$k=3$）。

</big>其中有一个聪明的名字约瑟夫斯，他解决了这个问题，站在幸存的位置，并继续讲述这个故事。

他是哪个号码？

# --instructions--

编写一个函数，该函数以初始囚犯人数和 `k` 作为参数，并返回存活囚犯的编号。

# --hints--

`josephus` 应该是一个函数。

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)` 应该返回一个数字。

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)` 应该返回 `28`。

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)` 应该返回 `2`。

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)` 应该返回 `8`。

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)` 应该返回 `1`。

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)` 应该返回 `1`。

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
