---
id: 5a23c84252665b21eecc7e06
title: 剪一个长方形
challengeType: 1
forumTopicId: 302242
dashedName: cut-a-rectangle
---

# --description--

给定的矩形由 *m* × *n* 个正方形组成。 如果 *m* 和 *n* 不是奇数，那么可以沿着矩形切割一条路径方形边缘，使得矩形分成两个连接的具有相同形状的部分（将其中一个部分旋转 180° 后）。 2 × 2 和 4 × 3 矩形的所有此类路径如下所示。

<img src="https://cdn.freecodecamp.org/curriculum/coding-interview-prep/rosetta-cut-a-rectangle.png" alt="Diagram showing the possible paths for 2 by 2 and 4 by 3 rectangles" />

# --instructions--

编写一个函数来计算切割 *m* × *n* 矩形的不同方法的数量。

# --hints--

`cutRectangle` 应该是一个函数。

```js
assert(typeof cutRectangle == 'function');
```

`cutRectangle(2, 2)` 应该返回一个数字。

```js
assert(typeof cutRectangle(2, 2) == 'number');
```

`cutRectangle(2, 2)` 应该返回 `2`。

```js
assert.equal(cutRectangle(2, 2), 2);
```

`cutRectangle(4, 3)` 应该返回 `9`。

```js
assert.equal(cutRectangle(4, 3), 9);
```

`cutRectangle(4, 4)` 应该返回 `22`。

```js
assert.equal(cutRectangle(4, 4), 22);
```

`cutRectangle(8, 3)` 应该返回 `53`。

```js
assert.equal(cutRectangle(8, 3), 53);
```

`cutRectangle(7, 4)` 应该返回 `151`。

```js
assert.equal(cutRectangle(7, 4), 151);
```

# --seed--

## --seed-contents--

```js
function cutRectangle(w, h) {

}
```

# --solutions--

```js
function cutRectangle(w, h) {
  if (w % 2 == 1 && h % 2 == 1) return;

  var dirs = [[0, -1], [-1, 0], [0, 1], [1, 0]];

  var grid = new Array(h);
  for (var i = 0; i < grid.length; i++) grid[i] = new Array(w);
  var stack = [];

  var half = Math.floor((w * h) / 2);
  var bits = Math.pow(2, half) - 1;
  var result = 0;
  for (; bits > 0; bits -= 2) {
    for (var i = 0; i < half; i++) {
      var r = Math.floor(i / w);
      var c = i % w;
      grid[r][c] = (bits & (1 << i)) != 0 ? 1 : 0;
      grid[h - r - 1][w - c - 1] = 1 - grid[r][c];
    }

    stack.push(0);
    grid[0][0] = 2;
    var count = 1;
    while (stack.length != 0) {
      var pos = stack.pop();
      var r = Math.floor(pos / w);
      var c = pos % w;

      for (var dir of dirs) {
        var nextR = r + dir[0];
        var nextC = c + dir[1];

        if (nextR >= 0 && nextR < h && nextC >= 0 && nextC < w) {
          if (grid[nextR][nextC] == 1) {
            stack.push(nextR * w + nextC);
            grid[nextR][nextC] = 2;
            count++;
          }
        }
      }
    }

    if (count == half) {
      result++;
    }
  }

  return result;
}
```
