---
id: 5e4ce2f5ac708cc68c1df261
title: 线性同余生成器
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

线性同余生成器（LCG）是一种<em>算法</em>，它通过一个不连续分段线性方程计算出伪随机数序列。所有线性同余生成器都使用此公式：

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

其中：

<ul>
<li>$ r_0 $ is a seed.</li>
<li>$r_1$, $r_2$, $r_3$, ..., are the random numbers.</li>
<li>$a$, $c$, $m$ are constants.</li>
</ul>

如果仔细选择 $a$、$c$ 和 $m$ 的值，那么生成器将产生从 $0$ 到 $m - 1$ 的均匀整数发行（版本）。

<abbr title="linear congruential generator">LCG</abbr> 数字质量较差。$r_n$ 和 $r\_{n + 1}$ 并不像真正的随机数那样独立。任何知道 $r_n$ 的人都可以预测 $r\_{n + 1}$，因此 <abbr title="linear congruential generator">LCG</abbr> 不是密码学安全的。<abbr title="linear congruential generator">LCG</abbr> 仍然足够用于像 Miller-Rabin 素数测试或 FreeCell 发牌这样简单的任务。<abbr title="linear congruential generator">LCG</abbr> 的优点之一是可以轻松地从相同的 $r_0$ 重现一系列数字。由于公式非常简单，也可以用不同的编程语言重现这样的序列。

# --instructions--

编写一个函数，接受 $r_0,a,c,m,n$ 作为参数并返回 $r_n$。

# --hints--

`linearCongGenerator` 应该是一个函数。

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` 应该返回一个数字。

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` 应该返回 `855`。

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` 应该返回 `1110`。

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` 应该返回 `62217`。

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` 应该返回 `12345`。

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` 应该返回 `1406932606`。

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
