---
id: 5e6dd1278e6ca105cde40ea9
title: 最长公共子序列
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**最长公共子序列**（或 **LCS**）是组 A 和组 B 中最长的元素组，这些元素在两个组中都存在且顺序相同。例如，序列 `1234` 和 `1224533324` 的 LCS 是 `1234`：

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

对于一个字串示例，考虑序列 `thisisatest` 和 `testing123testing`。一个 LCS 将是 `tsitest`：

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing。

您的代码只需要处理字符串。

# --instructions--

编写一个区分大小写的函数，返回两个字符串的 LCS。 你不需要显示多个 LCS。

# --hints--

`lcs` 应该是一个函数。

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")` 应该返回一个字符串。

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")` 应该返回 `"tsitest"`。

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")` 应该返回 `"ADH"`。

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")` 应该返回 `"GTAB"`。

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")` 应该返回 `"BDCB"`。

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")` 应该返回 `"ABAD"`。

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
