---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

编写一个程序，生成一个从 1 到 100（含）的整数数组。 但：

<ul>
    <li>for multiples of 3, add <code>"Fizz"</code> to the array instead of the number</li>
    <li>for multiples of 5, add <code>"Buzz"</code> to the array instead of the number</li>
    <li>for multiples of 3 and 5, add <code>"FizzBuzz"</code> to the array instead of the number</li>
</ul>

# --instructions--

您的程序应该返回一个包含基于上述规则的结果的数组。

# --hints--

`fizzBuzz` 应该是一个函数。

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` 应该返回一个数组。

```js
assert(Array.isArray(fizzBuzz()) == true);
```

只能被 3 整除的数字应该返回 `"Fizz"`。

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

只能被 5 整除的数字应该返回 `"Buzz"`。

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

可被 3 和 5 整除的数字应返回 `"FizzBuzz"`。

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

不能被 3 或 5 整除的数字应返回数字本身。

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
