---
id: 69a5f35669099ed52f8563b1
title: Construir un reabastecedor inteligente de despensa
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

En este laboratorio, construirás un pequeño programa de gestión de despensa usando conceptos básicos de JavaScript como arreglos, objetos, ciclos y condicionales.

Simularás la recepción de un envío de artículos para la despensa, decidirás qué hacer con cada artículo y organizarás los resultados para su almacenamiento.

**Objetivo:** Cumple las historias de usuario a continuación y haz que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes implementar una función `parseShipment(rawData)` que reciba un arreglo de cadenas y devuelva un arreglo de objetos con propiedades `{ sku, name, qty, expires, zone }`.

2. Debes implementar una función `planRestock(pantry, shipment)` que compare la despensa actual con el envío entrante y devuelva un arreglo de acciones en la forma `{ type: "restock" | "discard" | "donate", item }`.

3. Debes implementar una función `groupByZone(actions)` que agrupe las acciones en zonas de almacenamiento según la propiedad `zone` de cada artículo.

4. Debes implementar una función `clonePantry(pantry)` que devuelva una copia profunda de la despensa para que los cambios en la planificación no afecten la lista original. Una copia profunda significa crear un nuevo arreglo con nuevos objetos, de modo que modificar la copia no cambie la despensa original.

5. Debes usar todas las funciones juntas para procesar un envío y registrar el objeto final agrupado de resultados en la consola.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes definir una función llamada `parseShipment` que acepte un arreglo de cadenas llamado `rawData` como parámetro.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Tu función `parseShipment` debe convertir las cadenas del envío en el arreglo en objetos con las propiedades: `sku`, `name`, `qty`, `expires` y `zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

Se deben ignorar los SKU duplicados en el envío.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Cuando `zone` no esté presente en una cadena del envío, debes asignarle `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

El valor `qty` debe convertirse en un número.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Debes definir una función llamada `planRestock` que acepte dos parámetros: `pantry` y `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Tu función `planRestock` debe devolver un arreglo de acciones con las propiedades `type` y `item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Los artículos con una cantidad de `0` o menos en el envío deben crear una acción con tipo `discard`, sin importar si el artículo está en la despensa.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Cuando un artículo del envío ya existe en la despensa (mismo `sku`), la acción debe ser `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Si un artículo del envío no existe en la despensa, la acción debe ser `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Debes definir una función llamada `groupByZone` que acepte un parámetro llamado `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

Tu función `groupByZone` debe devolver las acciones agrupadas por la propiedad `zone` de cada objeto.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

Debes definir una función llamada `clonePantry` que acepte un parámetro llamado `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` debe devolver un arreglo nuevo en lugar del arreglo original de la despensa.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Los objetos dentro de la despensa clonada también deben ser objetos nuevos (copia profunda).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Las funciones deben trabajar juntas para procesar un envío y agrupar las acciones resultantes.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Debes registrar las acciones resultantes agrupadas por zonas.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
