---
id: 66e2edb23013ca596f8d9568
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

現在是測試你的 `getPunctuationCount` 函式（程式）的時候了。

創建一個 `punctuationCount` 變數並將呼叫 `getPunctuationCount` 函式（程式）且引數為 `"WHAT?!?!?!?!?"` 的結果指定給它。

接著，將以下內容記錄到主控台：`"Punctuation Count: [Punctuation count goes here]"`。將 `[Punctuation count goes here]` 替換為實際的變數名稱。你可以選擇使用模板字串或使用 `+` 運算子進行字串串接。

# --hints--

你應該創建一個 `punctuationCount` 變數。

```js
assert.isNotNull(punctuationCount)
```

你的 `punctuationCount` 變數應該設定為 `getPunctuationCount("WHAT?!?!?!?!?")` 的結果。

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

你應該將以下內容記錄到主控台：`"Punctuation Count: [Punctuation count goes here]"`。將 `[Punctuation count goes here]` 替換為實際的變數名稱。請確保在此使用正確的字串串接語法。

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

