---
id: 66e2d680e129e1423116a541
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將建置一個句子分析器，該分析器會接收一個句子並取得字、母音、子音和標點符號的數量。

首先，創建一個名為 `getVowelCount` 的函式（程式），並帶有一個名為 `sentence` 的參數。你的函式（程式）應該傳回句子中母音的總數。

# --hints--

你應該創建一個 `getVowelCount` 函式（程式）。

```js
assert.isFunction(getVowelCount);
```

你的 `getVowelCount` 函式（程式）應該有一個 `sentence` 參數。

```js
assert.match(getVowelCount.toString(), /sentence/);
```

你的 `getVowelCount` 函式（程式）應該傳回一個數字。

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

當句子是 `"Apples are tasty fruits"` 時，`getVowelCount` 函式（程式）應該傳回 `7`。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

當句子是 `"Hello, World!"` 時，該函式應該傳回 `3`。

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

你的母音計數應該不區分大小寫。

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

你的 `getVowelCount` 函式應該傳回任何句子的正確母音數量。

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
