---
id: 66e2f06b191e305c00574e4d
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

最後，透過建立一個帶有 `sentence` 參數的 `getWordCount` 函式（程式）來計算字數。當呼叫該函式時，應傳回傳入句子中的字數總和。

# --hints--

你應該創建一個 `getWordCount` 函式（程式）。

```js
assert.isFunction(getWordCount);
```

你的 `getWordCount` 函式應該有一個 `sentence` 參數。

```js
assert.match(getWordCount.toString(), /sentence/);
```

你的 `getWordCount` 函式應該傳回一個數字。

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

當句子是 `"When are you gonna start learning to code?"` 時，`getWordCount` 函式（程式）應該傳回 `8`。

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

當句子是 `"What's going on?"` 時，`getWordCount` 函式（程式）應該傳回 `3`。

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

你的字數計算應該不區分大小寫。

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

你的 `getWordCount` 函式（程式）應該傳回任何句子的正確字數。

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both frontend and backend development"), 7);
```

你的 `getWordCount` 函式應該傳回空字串或僅包含空格的字串的正確字數。

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
