---
id: 66e2df8c5676fe51cead8ec8
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在是測試你的 `getVowelCount` 函式（程式）的時候了。

創建一個 `vowelCount` 變數並將呼叫 `getVowelCount` 函式（程式）且引數為 `"Apples are tasty fruits"` 的結果指定給它。

接著，將以下內容記錄到主控台：`"Vowel Count: [vowel count goes here]"`。將 `[vowel count goes here]` 取代為實際的變數名稱。你可以選擇使用模板字串或使用 `+` 運算子進行字串串接。

# --hints--

你應該創建一個 `vowelCount` 變數。

```js
assert.isNotNull(vowelCount)
```

你的 `vowelCount` 變數應該設定為 `getVowelCount("Apples are tasty fruits")` 的結果。

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

你應該使用參數 `"Vowel Count: ${vowelCount}"` 將 `vowelCount` 變數記錄到主控台。

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
