---
id: 66e2eab8a5638f57b637b7cc
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

你現在應該計算標點符號的數量。

創建一個帶有 `sentence` 參數的 `getPunctuationCount` 函式（程式）。

在函式（程式）內，創建一個迴圈來計算在呼叫該函式（程式）時會傳入的 `sentence` 中標點字元的數量。就我們的用途而言，標點字元是指任何不是空格（`" "`）或字母的字元。

你的 `getPunctuationCount` 函式必須傳回一個數字。

# --hints--

你應該創建一個 `getPunctuationCount` 函式（程式）。

```js
assert.isFunction(getPunctuationCount);
```

你的 `getPunctuationCount` 函式（程式）應該有一個 `sentence` 參數。

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

你的 `getPunctuationCount` 函式應該傳回一個數字。

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

當句子是 `"What's going on here?"` 時，`getPunctuationCount` 函式（程式）應該傳回 `2`。

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

當句子是 `"What????!"` 時，`getPunctuationCount` 函式（程式）應該傳回 `5`。

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

你的 `getPunctuationCount` 函式應該傳回任何句子的正確標點符號數量。

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
