---
id: 66e2e5d890ac8a54c9d703bd
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

是時候計算子音了。為你的 `sentence` 參數創建一個 `getConsonantCount` 函式（程式）。

在函式（程式）內，使用迴圈來計算在呼叫時會傳入函式（程式）的 `sentence` 中的子音數量。子音是指不屬於以下字元的任何字母：`"aeiou"`。 

你的 `getConsonantCount` 函式必須傳回一個數字。

# --hints--

你應該創建一個 `getConsonantCount` 函式（程式）。

```js
assert.isFunction(getConsonantCount);
```

你的 `getConsonantCount` 函式（程式）應該有一個 `sentence` 參數。

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

你的 `getConsonantCount` 函式應該傳回一個數字。

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

當句子是 `"Coding is fun"` 時，`getConsonantCount` 函式（程式）應該傳回 `7`。

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

當句子是 `"Hello, World!"` 時，`getConsonantCount` 函式（程式）應該傳回 `7`。

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

你的子音計數應該不區分大小寫。

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

你的 `getConsonantCount` 函式應該傳回任何句子的正確子音數量。

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
