---
id: 66e2e99caa0cb3570f4822b4
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在是測試你的 `getConsonantCount` 函式（程式）的時候了。

創建一個 `consonantCount` 變數並將呼叫 `getConsonantCount` 函式（程式）且引數為 `"Coding is fun"` 的結果指定給它。

接著，將以下內容記錄到主控台：`"Consonant Count: [Consonant count goes here]"`。將 `[Consonant count goes here]` 替換為實際的變數名稱。你可以選擇使用模板字串或使用 `+` 運算子進行字串串接。

# --hints--

你應該創建一個 `consonantCount` 變數。

```js
assert.isNotNull(consonantCount)
```

你的 `consonantCount` 變數應該設定為 `getConsonantCount("Coding is fun")` 的結果。

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

你應該使用參數 `"Consonant Count: ${consonantCount}"` 將 `consonantCount` 變數記錄到主控台。

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
