---
id: a5deed1811a43193f9f1c841
title: Tekeleza kitendakazi cha kuruka kipengee
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

Katika maabara hii utaunda kitendakazi kinachoruka vipengee katika safu ya data hadi kipate kipengee kinachokubalika kulingana na kitendakazi cha mtihani maalum.

Kwa mfano, kwa safu ya data kama `[1, 1, 1, 2, 1, 1, 1]` na kitendakazi cha mtihani `function(n) {return n === 2}`, kipengee cha kwanza kinachokubalika ni kile kilicho kwenye kielezo 3, hivyo vipengee vyote kabla yake vinapaswa kuachwa, na matokeo yanapaswa kuwa vipengee vilivyobaki `[2, 1, 1, 1]`.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi `dropElements` kinachokubali safu ya data (`arr`) na kitendakazi (`func`) kama hoja.
1. Kitendakazi `dropElements` kinapaswa kurudia kupitia safu ya data na kuondoa vipengee kuanzia kipengee cha kwanza hadi `func` irudishe `true` kwa kipengee fulani.
1. Kitendakazi `dropElements` kinapaswa kurudisha vipengee vilivyobaki katika safu ya data ikiwa sharti limekamilika.
1. Ikiwa sharti halitakamilika kamwe, kinapaswa kurudisha safu ya data tupu.

# --hints--

Unapaswa kuwa na kitendakazi `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` inapaswa kurudisha `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` inapaswa kurudisha `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` inapaswa kurudisha `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` inapaswa kurudisha `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` inapaswa kurudisha `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` inapaswa kurudisha `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
