---
id: 6814d8e1516e86b171929de4
title: "挑战 1：元音平衡"
challengeType: 28
dashedName: challenge-1
---

# --description--

给定一个字串，判断字串前半部分的元音数量是否与后半部分的元音数量相等。

- 该字串可以包含任何字符。
- 字母 `a`、`e`、`i`、`o` 和 `u`，无论是大写还是小写，都被视为元音。
- 如果字串中有奇数个字符，忽略中间的字符。

# --hints--

`isBalanced("racecar")` 应该返回 `true`。

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` 应该返回 `true`。

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` 应返回 `false`。

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` 应返回 `false`。

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` 应该返回 `true`。

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` 应该返回 `false`。

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` 应该返回 `true`。

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
