---
id: 681cb1afdab50c87ddb2e515
title: "挑战 4：S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

给定一个字串，移除字串中的所有空格，在每个字符之间插入两个空格，将所有字母转换为大写，并返回结果。

- 非字母字符应保持不变（空格除外）。

# --hints--

`spaceJam("freeCodeCamp")` 应该返回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` 应该返回 `"F  R  E  E  C  O  D  E  C  A  M  P"`。

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` 应该返回 `"H  E  L  L  O  W  O  R  L  D  ?  !"`。

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` 应该返回 `"C  @  T  $  &  D  0  G  $"`。

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` 应该返回 `"A  L  L  Y  O  U  R  B  A  S  E"`。

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
