---
id: 681cb1b0dab50c87ddb2e51b
title: "挑战 10：3 次击打"
challengeType: 28
dashedName: challenge-10
---

# --description--

给定一个介于 1 和 10,000 之间的整数，返回从 1 到该整数中其平方包含至少一个数字 3 的数字个数。

# --hints--

`squaresWithThree(1)` 应返回 `0`。

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` 应该返回 `1`。

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` 应返回 `19`。

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` 应该返回 `326`。

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` 应该返回 `4531`。

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
