---
id: 681cb1b0dab50c87ddb2e519
title: "挑战 8：阶乘器"
challengeType: 28
dashedName: challenge-8
---

# --description--

给定一个从零到 20 的整数，返回该数字的阶乘。一个数字的阶乘是从 1 到该数字之间所有数字的乘积。

- 零的阶乘是 1。

# --hints--

`factorial(0)` 应该返回 `1`。

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` 应该返回 `120`。

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` 应该返回 `2432902008176640000`。

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
