---
id: 681cb1afdab50c87ddb2e517
title: "挑战 6：变位词检查器"
challengeType: 28
dashedName: challenge-6
---

# --description--

给定两个字串，判断它们是否是彼此的变位词（包含相同的字符，顺序可以不同）。

- 忽略大小写和空白。

# --hints--

`areAnagrams("listen", "silent")` 应该返回 `true`。

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` 应该返回 `true`。

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` 应该返回 `true`。

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` 应该返回 `false`。

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` 应该返回 `false`。

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` 应该返回 `false`。

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
