---
id: 681cb05adab50c87ddb2e513
title: "挑战 2：基础 `Check`"
challengeType: 28
dashedName: challenge-2
---

# --description--

给定一个表现数字的字串，以及一个从 2 到 36 的整数进制，确定该数字在该进制中是否有效。

- 该字串可能包含整数，以及大写或小写字符。
- 该检查应不区分分支。
- 基数可以是 2-36 之间的任何数字。
- 如果每个字符都是给定进制中的有效数字，则数字是有效的。
- 有效进制的数字示例：
  - 基数 2：0-1
  - 八进制：0-7
  - 十进制：0-9
  - 基数 16：0-9 和 A-F
  - 基数 36：0-9 和 A-Z

# --hints--

`isValidNumber("10101", 2)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` 应该返回 `false`。

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` 应该返回 `false`。

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` 应该返回 `false`。

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` 应该返回 `false`。

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` 应该返回 `false`。

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` 应该返回 `true`。

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
