---
id: 681cb1b0dab50c87ddb2e51a
title: "挑战 9：平方和"
challengeType: 28
dashedName: challenge-9
---

# --description--

给定一个最大为 1,000 的正整数，返回从 1 到该数字所有整数的平方和。

# --hints--

`sumOfSquares(5)` 应该返回 `55`。

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` 应该返回 `385`。

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` 应返回 `5525`。

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` 应该返回 `41791750`。

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` 应该返回 `333833500`。

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
