---
id: 681cb1b0dab50c87ddb2e518
title: "挑战 7：目标和"
challengeType: 28
dashedName: challenge-7
---

# --description--

给定一个数字数组和一个整数目标，查找数组中两个相加等于目标值的唯一数字。返回包含这两个数字索引的数组，如果没有两个数字的和等于目标，则返回 `"Target not found"`。

- 返回的数组应按索引升序排列。

# --hints--

`findTarget([2, 7, 11, 15], 9)` 应该返回 `[0, 1]`。

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` 应该返回 `[1, 2]`。

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` 应该返回 `[4, 5]`。

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` 应该返回 `"Target not found"`。

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
