---
id: 681cb1afdab50c87ddb2e516
title: "挑战 5：Jbelmud 文本"
challengeType: 28
dashedName: challenge-5
---

# --description--

给定一个字串，返回该字串的一个混乱版本，其中每个单词都使用以下约束进行转换：

- 单词的首尾字母保持不变
- 第一个和最后一个字母之间的所有字母按字母顺序排序。
- 输入字串将不包含标点符号，并且全部为小写。

# --hints--

`jbelmu("hello world")` 应该返回 `"hello wlord"`。

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` 应该返回 `"i love jbelmud text"`。

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` 应该返回 `"faccdeeemorp is my faiortve pacle to laern to cdoe"`。

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` 应该返回 `"the qciuk borwn fox jmpus oevr the lazy dog"`。

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
