---
id: 6821ebc9237de8297eaee78f
title: "挑战 13：不自然的素数"
challengeType: 28
dashedName: challenge-13
---

# --description--

给定一个整数，判断该数字是素数还是负素数。

- 素数是大于 1 且仅能被 1 和其自身整除的正整数。
- 负素数是正素数的负数版本。
- `1` 和 `0` 不被视为素数。

# --hints--

`isUnnaturalPrime(1)` 应该返回 `false`。

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` 应该返回 `false`。

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` 应该返回 `true`。

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` 应该返回 `true`。

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` 应该返回 `false`。

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` 应该返回 `true`。

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` 应该返回 `true`。

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` 应该返回 `false`。

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` 应该返回 `false`。

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
