---
id: 6821ebf8237de8297eaee798
title: "挑战 22：Tribonacci 序列"
challengeType: 28
dashedName: challenge-22
---

# --description--

Tribonacci 序列是一串数字，其中每个数字都是前述的三个数字之和。当以 `0`、`0` 和 `1` 开始时，序列中的前 10 个数字是 `0`、`0`、`1`、`1`、`2`、`4`、`7`、`13`、`24`、`44`。

给定一个包含 Tribonacci 序列前三个数字的数组，以及一个表现序列长度的整数，返回一个包含给定长度序列的数组。

- 你的函数应处理长度大于或相等于 0 的任意序列。
- 如果长度为 0，返回一个空数组。
- 请注意，起始数字是序列的一部分。

# --hints--

`tribonacciSequence([0, 0, 1], 20)` 应该返回 `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`。

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` 应该返回 `[21]`。


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` 应返回 `[]`。

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` 应返回 `[10, 20]`。

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` 应该返回 `[10, 20, 30]`。

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` 应该返回 `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`。

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
