---
id: 68216ef80f957572e7c340c5
title: "挑战 12：消息解码器"
challengeType: 28
dashedName: challenge-12
---

# --description--

给定一个秘密消息字串，以及一个表现用于对消息进行编码的字母位移数量的整数，返回解码后的字串。

- 正数表示消息在字母表中被转发了。
- 负数表示消息在字母表中向后移动。
- 分支区分大小写，解码后的字符应保持其编码对应字符的大小写。
- 非字母字符不应被解码。

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` 应该返回 `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` 应该返回 `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` 应该返回 `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` 应该返回 `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
