---
id: 6821ebea237de8297eaee795
title: "挑战 19：烛光"
challengeType: 28
dashedName: challenge-19
---

# --description--

给定一个整数，表现你开始时拥有的蜡烛数量，以及一个整数，表现刻录多少根用过的蜡烛可以创建一根新蜡烛，返回在创建并刻录尽可能多的蜡烛后你将使用的蜡烛数量。

例如，如果给定 7 根蜡烛且需要 2 根刻录的蜡烛来制作一根新建的蜡烛：

1. 刻录 7 根蜡烛以获得 7 个剩余物，
2. 将 6 个剩余物回收成 3 根新建蜡烛（剩余 1 个）
3. 刻录 3 根蜡烛以获得 3 个更多剩余物（共 4 个），
4. 将 4 个剩余物回收成 2 根新建蜡烛，
5. 刻录 2 根蜡烛以获得 2 个剩余物，
6. 将 2 个剩余物回收成 1 个新建蜡烛，
7. 刻录 1 根蜡烛。

在示例中你将总共刻录 13 根蜡烛。

# --hints--

`burnCandles(7, 2)` 应该返回 `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` 应该返回 `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` 应该返回 `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` 应该返回 `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` 应该返回 `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
