---
id: 6821ebfd237de8297eaee799
title: "挑战 23：RGB 转 Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

给定一个 CSS `rgb(r, g, b)` 颜色字串，返回其十六进制等价物。

以下是给定 `input` 的一些示例输出：

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- 将所有字母转换为小写。
- 返回一个 `#`，后跟六个字符。不要使用任何简写值。

# --hints--

`rgbToHex("rgb(255, 255, 255)")` 应该返回 `"#ffffff"`。

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` 应该返回 `"#010b6f"`。

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` 应该返回 `"#add8e6"`。

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` 应该返回 `"#4f7bc9"`。

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
