---
id: 6821ebd4237de8297eaee791
title: "挑战 15：驼峰式大小写"
challengeType: 28
dashedName: challenge-15
---

# --description--

给定一个字串，使用以下规则返回其驼峰式写法版本：

- 字串参数中的单词由以下集合中的一个或多个字符分隔：空格（` `）、破折号（`-`）或下划线（`_`）。将这些字符的任意序列视为单词分隔符。
- 第一个单词应全部小写。
- 每个后续单词应以大写字母开头，其余部分为小写。
- 所有空格和分隔符都应删除。

# --hints--

`toCamelCase("hello world")` 应该返回 `"helloWorld"`。

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` 应该返回 `"helloWorld"`。

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` 应该返回 `"secretAgentX"`。

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` 应该返回 `"freeCodeCamp"`。

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` 应该返回 `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`。

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
