---
id: 6821ebee237de8297eaee796
title: "挑战 20：数组副本"
challengeType: 28
dashedName: challenge-20
---

# --description--

给定一个整数数组，返回一个在初始数组中出现超过一次的整数数组，按升序排序。如果没有值出现超过一次，返回一个空数组。

- 返回的数组中只包含每个值的一个实例。

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` 应该返回 `[]`。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` 应该返回 `[1, 2]`。

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` 应该返回 `[-6, 0, 2, 4, 5, 23]`。

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
