---
id: 68216eb60f957572e7c340c4
title: "挑战 11：英里配速"
challengeType: 28
dashedName: challenge-11
---

# --description--

给定跑过的英里数和跑完这些英里所用的时间（格式为 `"MM:SS"`（分钟:秒）），返回一个表示每英里平均用时的字串，格式为 `"MM:SS"`。

- 在需要时添加前导零。

# --hints--

`milePace(3, "24:00")` 应该返回 `"08:00"`。

```js
assert.equal(milePace(3, "24:00"), "08:00");
```

`milePace(1, "06:45")` 应该返回 `"06:45"`。

```js
assert.equal(milePace(1, "06:45"), "06:45");
```

`milePace(2, "07:00")` 应该返回 `"03:30"`。

```js
assert.equal(milePace(2, "07:00"), "03:30");
```

`milePace(26.2, "120:35")` 应该返回 `"04:36"`。

```js
assert.equal(milePace(26.2, "120:35"), "04:36");
```

# --seed--

## --seed-contents--

```js
function milePace(miles, duration) {

  return miles;
}
```

# --solutions--

```js
function milePace(miles, duration) {
  const [minutes, seconds] = duration.split(':').map(Number);
  const totalSeconds = minutes * 60 + seconds;
  const avgSecondsPerMile = totalSeconds / miles;

  const avgMinutes = Math.floor(avgSecondsPerMile / 60);
  const avgSeconds = Math.round(avgSecondsPerMile % 60);

  const paddedMinutes = avgMinutes.toString().padStart(2, '0');
  const paddedSeconds = avgSeconds.toString().padStart(2, '0');

  return `${paddedMinutes}:${paddedSeconds}`;
}
```
