---
id: 6821ebdf237de8297eaee793
title: "挑战 17：操作顺序"
challengeType: 28
dashedName: challenge-17
---

# --description--

给定一个整数数组和一个字串操作符数组，将操作从左到右循序地应用到数字上。根据需要重复操作，直到所有数字都被使用。返回最终结果。

例如，给定 `[1, 2, 3, 4, 5]` 和 `['+', '*']`，返回从左到右计算 `1 + 2 * 3 + 4 * 5` 的结果，忽略标准的操作顺序。

- 有效的操作符是 `+`、`-`、`*`、`/` 和 `%`。

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` 应该返回 `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` 应该返回 `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` 应该返回 `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` 应该返回 `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` 应该返回 `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
