---
id: 68adce01c0e1144d0a90295a
title: "挑战 27：矩阵旋转"
challengeType: 28
dashedName: challenge-27
---

# --description--

给定一个矩阵（一个数组的数组），将矩阵顺时针旋转 90 度并 `return` 它。例如，给定 `[[1, 2], [3, 4]]`，其看起来像这样：

| 1 | 2 |
|---|---|
| 3 | 4 |

你应该返回 `[[3, 1], [4, 2]]`，它看起来像这样：

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` 应该返回 `[[1]]`。

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` 应该返回 `[[3, 1], [4, 2]]`。

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` 应该返回 `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`。

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` 应该返回 `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`。

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
