---
id: 68adce01c0e1144d0a902958
title: "挑战 26：IPv4 验证器"
challengeType: 28
dashedName: challenge-26
---

# --description--

给定一个字串，判断它是否是有效的 IPv4 `Address`。有效的 IPv4 `Address` 由四个由点号（`.`）分隔的整数构成。每个数字必须满足以下条件：

- 它的取值范围是 0 到 255（含）。
- 它没有前导零（例如允许 0，不允许 01）。
- 只允许数字字符。

# --hints--

`isValidIPv4("192.168.1.1")` 应该返回 `true`。

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` 应该返回 `true`。

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` 应该返回 `false`。

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` 应返回 `false`。

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` 应该返回 `false`。

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` 应返回 `false`。

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` 应该返回 `false`。

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
