---
id: 68adce01c0e1144d0a90295c
title: "挑战 28：罗马数字解析器"
challengeType: 28
dashedName: challenge-28
---

# --description--

给定一个表现罗马数字的字串，返回其整数值。

罗马数字由以下符号和数值组成：

| Symbol | Value |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- 数字从左到右读取。如果较小的数字出现在较大的数字之前，则该值被减去。否则，值被相加。

# --hints--

`parseRomanNumeral("III")` 应该返回 `3`。

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` 应该返回 `4`。

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` 应该返回 `26`。

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` 应该返回 `99`。

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` 应该返回 `460`。

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` 应该返回 `504`。

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` 应该返回 `2025`。

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
