---
id: 68adce01c0e1144d0a90295e
title: "挑战 29：首字母缩略词生成器"
challengeType: 28
dashedName: challenge-29
---

# --description--

给定一个包含一个或多个单词的字串，使用以下约束返回这些单词的首字母缩写：

- 缩写应由每个单词的首字母大写组成，除非另有说明。
- 缩写应忽略这些字的第一个字母，除非它们是给定字串的第一个字：`a`、`for`、`an`、`and`、`by` 和 `of`。
- 缩写字母应按给定顺序返回。
- 缩写不应包含任何空格。

# --hints--

`buildAcronym("Search Engine Optimization")` 应该返回 `"SEO"`。

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` 应该返回 `"FAQ"`。

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` 应该返回 `"NASA"`。

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` 应该返回 `"FBI"`。

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` 应该返回 `"FYI"`。

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` 应该返回 `"BTW"`。

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` 应该返回 `"AUHWPOTIMSH"`。

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
