---
id: 68adce01c0e1144d0a902956
title: "挑战 25：元音重复器"
challengeType: 28
dashedName: challenge-25
---

# --description--

给定一个字串，返回该字串的新建版本，其中每个元音的重复次数比你遇到的前一个元音多一次。例如，句子中的第一个元音保持不变。第二个元音应连续出现两次。第三个元音应连续出现三次，依此类推。

- 字母 `a`、`e`、`i`、`o` 和 `u`，无论是大写还是小写，都被视为元音。
- 原始元音应保持其分支。
- 重复的元音应为小写。
- 所有非元音字符应保持其原始大小写。

# --hints--

`repeatVowels("hello world")` 应该返回 `"helloo wooorld"`。

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` 应该返回 `"freeeCooodeeeeCaaaaamp"`。

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` 应该返回 `"AEeIiiOoooUuuuu"`。

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` 应该返回 `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`。

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
