---
id: 68af0687ef34c76c28ffa547
title: "挑战 30：唯一字符"
challengeType: 28
dashedName: challenge-30
---

# --description--

给定一个字串，判断字串中的所有字符是否唯一。

- 大写字母和小写字母应被视为不同的字符。

# --hints--

`allUnique("abc")` 应该返回 `true`。

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` 应该返回 `true`。

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` 应该返回 `true`。

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` 应该返回 `true`。

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` 应返回 `false`。

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` 应该返回 `false`。

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` 应该返回 `false`。

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
