---
id: 68af0687ef34c76c28ffa54b
title: "挑战 32：反转句子"
challengeType: 28
dashedName: challenge-32
---

# --description--

给定一个字串，返回一个新建字串，其中文字的顺序颠倒。例如，第一个单词应位于返回字串的末尾，最后一个单词应位于返回字串的开头。

- 在给定的字串中，单词可以由一个或多个空格分隔。
- 返回的字串中单词之间应只有一个空格。

# --hints--

`reverseSentence("world hello")` 应该返回 `"hello world"`。

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` 应该返回 `"git commit push"`。

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` 应该返回 `"sudo install npm"`。

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` 应该返回 `"export function default import"`。

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
