---
id: 68b06e589bf2273243814775
title: "挑战 38：Slug 生成器"
challengeType: 28
dashedName: challenge-38
---

# --description--

给定一个字串，使用以下约束返回该字串的 URL 友好版本：

- 所有字母应为小写。
- 所有不是字母、数字或空格的字符都应被移除。
- 所有空格都应替换为 URL 编码的空格代码 `%20`。
- 连续空格应替换为单个 `%20`。
- 返回的字串不应有前导或尾随的 `%20`。

# --hints--

`generateSlug("helloWorld")` 应该返回 `"helloworld"`。

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` 应该返回 `"hello%20world"`。

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` 应该返回 `"helloworld"`。

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` 应该返回 `"hello%20world"`。

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` 应该返回 `"h3110%20w0r1d"`。

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
