---
id: 68b06e589bf2273243814777
title: "挑战 39：加满油箱"
challengeType: 28
dashedName: challenge-39
---

# --description--

给定油箱的容量、当前油量和每加仑价格，返回将油箱加满的费用。

- `tankSize` 是油箱的总容量，单位为加仑。
- `fuelLevel` 是油箱中当前的燃料量，单位为加仑。
- `pricePerGallon` 是一加仑燃料的成本。
- 返回值应四舍五入到两位小数，格式为：`"$d.dd"`。

# --hints--

`costToFill(20, 0, 4.00)` 应该返回 `"$80.00"`。

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` 应该返回 `"$17.50"`。

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` 应该返回 `"$29.25"`。

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` 应该返回 `"$0.00"`。

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` 应该返回 `"$21.89"`。

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
