---
id: 68b06e589bf227324381476f
title: "挑战 35：词频"
challengeType: 28
dashedName: challenge-35
---

# --description--

给定一个段落，返回出现频率最高的三个单词的数组。

- 段落中的单词将被空格分开。
- 忽略给定段落中的分支。例如，将 `Hello` 和 `hello` 视为相同的单词。
- 忽略给定段落中的标点符号。标点符号包括逗号（`,`）、句号（`.`）和感叹号（`!`）。
- 返回的数组应包含所有小写单词。
- 返回的数组应按降序排列，最常出现的单词排在最前面。

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` 应该返回 `["coding", "python", "in"]`。

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` 应该返回 `["i", "like", "coding"]`。

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` 应该返回 `["debug", "test", "deploy"]`。

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
