---
id: 68b06e589bf2273243814771
title: "挑战 36：恒温器调节器"
challengeType: 28
dashedName: challenge-36
---

# --description--

给定房间的当前温度和目标温度，返回一个字串，指示如何根据这些约束调整房间温度：

- 如果当前温度低于目标，则返回 `"heat"`。
- 如果当前温度高于目标，则返回 `"cool"`。
- 如果当前温度与目标相等，则返回 `"hold"`。

# --hints--

`adjustThermostat(68, 72)` 应该返回 `"heat"`。

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` 应该返回 `"cool"`。

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` 应该返回 `"hold"`。

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` 应该返回 `"heat"`。

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` 应该返回 `"cool"`。

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` 应返回 `"hold"`。

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
