---
id: 68b1f72371a5ac895ac70a0a
title: "挑战 44：字串 镜射"
challengeType: 28
dashedName: challenge-44
---

# --description--

给定两个字串，判断第二个字串是否是第一个字串的镜射。

- 如果一个字串包含相同的字母且顺序相反，则该字串被视为镜射。
- 将大写字母和小写字母视为不同。
- 忽略所有非字母字符。

# --hints--

`isMirror("helloworld", "helloworld")` 应该返回 `false`。

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` 应该返回 `true`。

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` 应该返回 `true`。

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` 应该返回 `false`。

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` 应该返回 `false`。

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` 应该返回 `true`。

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` 应该返回 `true`。

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
