---
id: 68b1f72371a5ac895ac70a08
title: "挑战 43：数字 对 字母"
challengeType: 28
dashedName: challenge-43
---

# --description--

给定一个字串，如果字串中数字多于字母，则返回 `"digits"`，如果字母多于数字，则返回 `"letters"`，如果数字和字母数量相同，则返回 `"tie"`。

- 数字由 `0-9` 组成。
- 字母由大写或小写的 `a-z` 组成。
- 忽略任何其他字符。

# --hints--

`digitsOrLetters("abc123")` 应该返回 `"tie"`。

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` 应该返回 `"letters"`。

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` 应该返回 `"digits"`。

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` 应该返回 `"letters"`。

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` 应该返回 `"digits"`。

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` 应该返回 `"tie"`。

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
