---
id: 68b1f72371a5ac895ac70a04
title: "挑战 41：文件存储"
challengeType: 28
dashedName: challenge-41
---

# --description--

给定一个文件大小、文件大小的单位以及以千兆字节（GB）为单位的硬盘容量，返回硬盘可以保存的文件数量，使用以下约束：

- 文件大小的单位可以是字节（`"B"`）、千字节（`"KB"`）或兆字节（`"MB"`）。
- 返回驱动器可以容纳的完整文件数。
- 使用以下转换：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

例如，给定 `500`、`"KB"` 和 `1` 作为参数，确定在 1 GB 硬盘上可以容纳多少个 500 KB 文件。

# --hints--

`numberOfFiles(500, "KB", 1)` 应该返回 `2000`。

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` 应该返回 `20000`。

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` 应该返回 `200`。

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` 应该返回 `366210`。

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` 应该返回 `453514`。

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` 应该返回 `166666`。

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
