---
id: 68b1f72371a5ac895ac70a06
title: "挑战 42：视频存储"
challengeType: 28
dashedName: challenge-42
---

# --description--

给定一个视频大小、视频大小的单位、硬盘容量以及硬盘的单位，使用以下约束返回硬盘可以保存的视频数量：

- 视频大小的单位可以是字节（`"B"`）、千字节（`"KB"`）、兆字节（`"MB"`）或千兆字节（`"GB"`）。
- 如果未给出上述视频单位之一，返回 `"Invalid video unit"`。
- 硬盘容量的单位可以是千兆字节（`"GB"`）或太字节（`"TB"`）。
- 如果未给出上述硬盘单位之一，返回 `"Invalid drive unit"`。
- 返回驱动器可以容纳的完整 `video` 数量。
- 使用以下转换：

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

例如，给定 `500`、`"MB"`、`100` 和 `"GB"` 作为参数，确定 多少 个 500 MB 视频 可以存储在 100 GB 硬盘上。

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` 应该返回 `200`。

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` 应该返回 `"Invalid video unit"`。

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` 应该返回 `"Invalid drive unit"`。

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` 应该返回 `4000`。

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` 应该返回 `1466`。

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
