---
id: 68b1f72371a5ac895ac70a02
title: "挑战 40：照片存储"
challengeType: 28
dashedName: challenge-40
---

# --description--

给定以兆字节（MB）为单位的照片大小和以千兆字节（GB）为单位的硬盘容量，使用以下约束返回硬盘可以保存的照片数量：

- 1 吉字节等于 1000 兆字节。
- 返回驱动器可以保存的完整照片数量。

# --hints--

`numberOfPhotos(1, 1)` 应该返回 `1000`。

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` 应该返回 `500`。

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` 应该返回 `64000`。

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` 应该返回 `214285`。

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` 应该返回 `1571`。

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
