---
id: 68b7687dded630607aceccab
title: "挑战 45：完全平方"
challengeType: 28
dashedName: challenge-45
---

# --description--

给定一个整数，判断它是否是一个完全平方数。

- 如果你可以将一个整数乘以它自身得到该数字，那么该数字就是一个完全平方数。例如，9 是一个完全平方数，因为你可以将 3 乘以它自身得到它。

# --hints--

`isPerfectSquare(9)` 应该返回 `true`。

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` 应该返回 `true`。

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` 应该返回 `true`。

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` 应返回 `false`。

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` 应返回 `false`。

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` 应返回 `false`。

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` 应返回 `true`。

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` 应返回 `true`。

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
