---
id: 68b7687dded630607aceccb3
title: "挑战 49：CSV 头部解析器"
challengeType: 28
dashedName: challenge-49
---

# --description--

给定逗号分隔值（CSV）文件的第一行，返回包含标题的数组。

- CSV `file` 的第一行包含由逗号分隔的标题。
- 删除每个标题的任何前导或尾随空白。

# --hints--

`getHeadings("name,age,city")` 应该返回 `["name", "age", "city"]`。

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` 应该返回 `["first name", "last name", "phone"]`。

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` 应该返回 `["username", "email", "signup date"]`。

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
