---
id: 68b7cadffed0e75a517da677
title: "挑战 54：P@ssw0rd Str3ngth！"
challengeType: 28
dashedName: challenge-54
---

# --description--

给定一个密码字串，根据密码强度返回 `"weak"`、`"medium"` 或 `"strong"`。

密码根据以下规则进行评估：

- 它至少有 8 个字符长。
- 它包含大写和小写字母。
- 它至少包含 1 个数字。
- 它至少包含来自此集合的一个特殊字符：`!`、`@`、`#`、`$`、`%`、`^`、`&` 或 `*`。

如果密码满足的规则少于 2 条，则返回 `"weak"`。  
如果密码满足 2 条或 3 条规则，则返回 `"medium"`。  
如果密码满足所有 4 条规则，则返回 `"strong"`。

# --hints--

`checkStrength("123456")` 应该返回 `"weak"`。

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` 应该返回 `"weak"`。

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` 应该返回 `"weak"`。

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` 应该返回 `"weak"`。

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` 应该返回 `"medium"`。

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` 应该返回 `"medium"`。

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` 应该返回 `"medium"`。

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` 应该返回 `"strong"`。

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` 应该返回 `"strong"`。

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
