---
id: 68b7cadffed0e75a517da673
title: "挑战 52：二进制转十进制"
challengeType: 28
dashedName: challenge-52
---

# --description--

给定一个表现二进制数的字串，返回其对应的十进制数字。

二进制数字只使用数字 `0` 和 `1` 来表现任何数字。要将二进制转换为十进制，将每个数字乘以 `2` 的幂并将它们相加。首先将最右边的数字乘以 `2^0`，左边的下一个数字乘以 `2^1`，依此类推。一旦所有数字都乘以了 `2` 的幂，将结果相加。

例如，二进制数 `101` 在十进制中等于 `5`，因为：

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` 应该返回 `5`。

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` 应该返回 `10`。

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` 应该返回 `18`。

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` 应该返回 `85`。

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
