---
id: 68b7cadffed0e75a517da675
title: "挑战 53：十进制转二进制"
challengeType: 28
dashedName: challenge-53
---

# --description--

给定一个非负整数，返回其二进制表现形式的字串。

二进制数字只使用数字 `0` 和 `1` 来表现任何数字。要将十进制数字转换为二进制，反复将数字除以 `2` 并记录余数。重复直到数字为零。按最后记录的余数到最先记录的余数顺序读取。例如，要将 `12` 转换为二进制：

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` 的二进制是 `1100`。

# --hints--

`toBinary(5)` 应该返回 `"101"`。

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` 应该返回 `"1100"`。

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` 应该返回 `"110010"`。

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` 应返回 `"1100011"`。

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
