---
id: 68b7687dded630607aceccb1
title: "挑战 48：垃圾邮件检测器"
challengeType: 28
dashedName: challenge-48
---

# --description--

给定格式为 `"+A (BBB) CCC-DDDD"` 的电话号码，其中每个字母代表一个数字，具体如下：

- `A` 表现国家代码，可以是任意位数的数字。
- `BBB` 表现区号，且始终为三位数字。
- `CCC` 和 `DDDD` 表现局部的数字，且长度分别始终为三位和四位。

根据以下标准确定它是否为垃圾邮件号码：

- 国家代码超过 2 位数字长度或不以零 (`0`) 开头。
- 区号大于 900 或小于 200。
- 局部的号码的前三位数字之和出现在局部的号码的最后四位数字中。
- 该数字中有相同的数字连续出现四次或更多次（忽略格式字符）。

# --hints--

`isSpam("+0 (200) 234-0182")` 应该返回 `false`。

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` 应该返回 `true`。

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` 应该返回 `true`。

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` 应返回 `true`。

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` 应返回 `true`。

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` 应该返回 `true`。

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` 应该返回 `true`。

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` 应返回 `false`。

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
