---
id: 68c1a929005bf54d342aa8d3
title: "挑战 56：太空周第 2 天：系外行星查找"
challengeType: 28
dashedName: challenge-56
---

# --description--

在太空周的第二天，你会得到一个字串，其中每个字符表现一颗恒星的亮度读数。判断这些读数是否通过过境方法检测到了系外行星。过境方法是指当行星经过恒星前方时，减少了观测到的亮度。

- 亮度读数仅由字符 `0-9` 和 `A-Z` 组成，每个读数对应以下数值：
- 字符 `0-9` 对应于亮度级别 `0-9`。
- 字符 `A-Z` 对应亮度级别 `10-35`。

如果任何单次 `read` 小于或相等于所有 `read` 平均值的 80％，则该恒星被认为拥有一个系外行星。例如，如果恒星的平均亮度为 10，则如果任何单次 `read` 小于或相等于 8，则认为该恒星拥有一个系外行星。

# --hints--

`hasExoplanet("665544554")` 应返回 `false`。

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` 应该返回 `true`。

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` 应该返回 `false`。

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` 应该返回 `true`。

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` 应该返回 `false`。

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` 应该返回 `true`。

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
