---
id: 68c1a929005bf54d342aa8d6
title: "挑战 59：太空周 第 5 天：适居带"
challengeType: 28
dashedName: challenge-59
---

# --description--

在太空周的第 5 天，你将计算恒星的“Goldilocks zone”——恒星周围条件“恰到好处”以存在液态水的区域。

给定恒星的质量，返回一个包含其金发姑娘区起始和结束距离（以天文单位为单位）的数组。

计算 Goldilocks Zone：

1. 通过将恒星的质量提升到 3.5 次方来查找其光度。
2. 该区域的起点是其亮度的 0.95 倍平方根。
3. 该区域的末端是其亮度的 1.37 倍平方根。

- 返回四舍五入到两位小数的距离。

例如，给定 `1` 作为质量，返回 `[0.95, 1.37]`。

# --hints--

`goldilocksZone(1)` 应返回 `[0.95, 1.37]`。

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` 应该返回 `[0.28, 0.41]`。

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` 应该返回 `[21.85, 31.51]`。

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` 应该返回 `[9.38, 13.52]`。

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` 应返回 `[179.69, 259.13]`。

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
