---
id: 68c1a929005bf54d342aa8d2
title: "挑战 55：太空周 第 1 天：恒星分类"
challengeType: 28
dashedName: challenge-55
---

# --description--

10 月 4 日标志着世界航天周的开始。接下来的 7 天将为你带来天文学主题的编码挑战。

对于今天的挑战，你将获得一颗恒星的表面温度，单位为开尔文（K），并需要根据以下范围确定其恒星分类：

- `"O"`：30,000 K 或更高
- `"B"`：10,000 K - 29,999 K
- `"A"`：7,500 K - 9,999 K
- `"F"`：6,000 K - 7,499 K
- `"G"`：5,200 K - 5,999 K
- `"K"`：3,700 K - 5,199 K
- `"M"`：0 K - 3,699 K

- 返回给定恒星的分类。

# --hints--

`classification(5778)` 应该返回 `"G"`。

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` 应该返回 `"M"`。

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` 应该返回 `"A"`。

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` 应该返回 `"K"`。

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` 应该返回 `"M"`。

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` 应该返回 `"O"`。

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` 应该返回 `"F"`。

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` 应该返回 `"B"`。

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
