---
id: 68c1a929005bf54d342aa8d4
title: "挑战 57：太空周 第 3 天：回家打电话"
challengeType: 28
dashedName: challenge-57
---

# --description--

在太空周的第三天，你将获得一个表示你自己、卫星和你的母星之间通信路线中距离（以公里为单位）的数字数组。使用以下约束确定通过该路线发送的消息到达目标星球所需的时间：

- 数组中的第一个值是你的位置到第一个卫星的距离。
- 除最后一个外，每个后续值都是到下一个卫星的距离。
- 数组中的最后一个值是从前一个卫星到你的母星的距离。
- 该消息以 300,000 公里/秒 的速度传播。
- 每经过一个卫星，消息 **传递** 会增加 0.5 秒的传输延迟。
- 返回一个四位小数的数字，去除末尾的零。

# --hints--

`sendMessage([300000, 300000])` 应该返回 `2.5`。

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` 应该返回 `3.0627`。

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` 应该返回 `364.5`。

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` 应该返回 `1674.3333`。

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` 应该返回 `2.4086`。

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` 应该返回 `21.1597`。

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
