---
id: 68c497f3aaefc9fd9f1b0e24
title: "挑战 60：太空周第 6 天：月相"
challengeType: 28
dashedName: challenge-60
---

# --description--

在太空周的第 6 天，你将获得一个格式为 `"YYYY-MM-DD"` 的日期，并需要使用以下规则确定该日期的月相：

使用简化的 28 天月周期，分为四个相等的阶段：

- `"New"`：第 1 天至第 7 天
- `"Waxing"`：第 8 天至第 14 天
- `"Full"`：第 15 天至第 21 天
- `"Waning"`：第 22 天至第 28 天

在第 28 天之后，周期从第 1 天开始重复，新月。

- 使用 `"2000-01-06"` 作为参考新月（周期的第 1 天）来确定给定日期的月相。
- 你不会被给予早于 `reference date` 的任何日期。
- 返回正确的阶段作为字串。

**注意：** 第 1 天表现新月当天，意味着自上一个新月以来已过 0 天。

# --hints--

`moonPhase("2000-01-12")` 应该返回 `"New"`。

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` 应该返回 `"Waxing"`。

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` 应该返回 `"Full"`。

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` 应该返回 `"Waning"`。

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` 应该返回 `"New"`。

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
