---
id: 68cae5b538ff798bbd4da002
title: "挑战 64：24 到 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

给定一个表现一天时间的 24 小时格式的字串 `"HHMM"`，返回其等效的 12 小时格式 `"H:MM AM"` 或 `"H:MM PM"`。

- 给定的输入将始终是一个四位数的字串，采用 24 小时制格式，范围从 `"0000"` 到 `"2359"`。

# --hints--

`to12("1124")` 应该返回 `"11:24 AM"`。

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` 应该返回 `"9:00 AM"`。

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` 应该返回 `"2:55 PM"`。

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` 应该返回 `"11:46 PM"`。

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` 应该返回 `"12:30 AM"`。

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
