---
id: 68cae5b538ff798bbd4da009
title: "挑战 71：小费计算器"
challengeType: 28
dashedName: challenge-71
---

# --description--

给定你的餐费价格和一个自定义小费百分比，返回一个包含三个小费值的数组；15％、20％和自定义金额。

- 价格将以 `"$N.NN"` 形式给出。
- 自定义小费百分比将以这种形式给出：`"25%"`。
- 以相同的 `"$N.NN"` 形式返回金额，四舍五入到两位小数。

例如，给定一个 `"$10.00"` 的餐价和一个 `"25%"` 的自定义小费值，返回 `["$1.50", "$2.00", "$2.50"]`。

# --hints--

`calculateTips("$10.00", "25%")` 应该返回 `["$1.50", "$2.00", "$2.50"]`。

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` 应该返回 `["$13.45", "$17.93", "$23.31"]`。

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` 应该返回 `["$2.98", "$3.97", "$1.79"]`。

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
