---
id: 68cae5b538ff798bbd4da005
title: "挑战 67：电子邮件验证器"
challengeType: 28
dashedName: challenge-67
---

# --description--

给定一个字串，使用以下约束确定它是否是有效的电子邮件地址：

- 它必须包含恰好一个 `@` 符号。
- `@` 之前的局部的部分：
  - 只能包含字母（`a-z`、`A-Z`）、数字（`0-9`）、点号（`.`）、下划线（`_`）或连字符（`-`）。
  - 不能以点号开头或结尾。
- `@` 后面的域部分：
  - 必须包含至少一个点。
  - 必须以一个点号后跟至少两个字母结尾。
- 局部的部分和域部分都不能有两个连续的点。

# --hints--

`validate("a@b.cd")` 应该返回 `true`。

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` 应该返回 `true`。

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` 应该返回 `false`。

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` 应该返回 `false`。

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` 应该返回 `false`。

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` 应该返回 `true`。

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` 应返回 `false`。

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` 应该返回 `false`。

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` 应返回 `false`。

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` 应该返回 `false`。

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
