---
id: 68cae5b538ff798bbd4da007
title: "挑战 69：丢失的袜子"
challengeType: 28
dashedName: challenge-69
---

# --description--

给定一个整数，表现你开始时的袜子点对数量，以及另一个整数，表现你经历了多少洗涤周期，使用以下约束返回你当前拥有的完整袜子点对数量：

- 每 2 个洗涤周期，你会丢失一只袜子。
- 每 3 个洗涤周期，你会查找一个丢失的袜子。
- 每 5 个洗涤周期，会有一只袜子磨损，必须被抛弃。
- 每 10 个洗涤周期，你购买一双袜子。
- 你的袜子总数永远不会少于零。
- 规则可以重叠。例如，在洗涤 `cycle` 10 上，你会丢失一只袜子，抛出一只袜子，并购买一双新袜子。
- 返回完整点对袜子的数量。

# --hints--

`sockPairs(2, 5)` 应该返回 `1`。

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` 应该返回 `0`。

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` 应该返回 `4`。

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` 应该返回 `3`。

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` 应该返回 `0`。

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
