---
id: 68cae5b538ff798bbd4da008
title: "挑战 70：`HTML` 属性提取器"
challengeType: 28
dashedName: challenge-70
---

# --description--

给定一个有效 `HTML` 元素的字串，使用以下标准返回该元素的属性：

- 你将只获得一个元素。
- 属性将采用以下形式：`attribute="value"`。
- 返回一个字串数组，每个属性属性和值用逗号分隔，格式为：`["attribute1, value1", "attribute2, value2"]`。
- 按给定顺序返回属性。
- 如果未找到任何属性，返回一个空数组。

# --hints--

`extractAttributes('<span class="red"></span>')` 应该返回 `["class, red"]`。

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` 应该返回 `["charset, UTF-8"]`。

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` 应该返回 `[]`。

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` 应该返回 `["name, email", "type, email", "required, true"]`。

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` 应该返回 `["id, submit", "class, btn btn-primary"]`。

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
