---
id: 68cae5b538ff798bbd4da004
title: "挑战 66：超文本标记语言 标签剥离器"
challengeType: 28
dashedName: challenge-66
---

# --description--

给定一段 `HTML` 代码字串，移除标签并返回纯文本内容。

- 输入字串将仅包含有效的 `HTML`。
- `HTML` 标签可以嵌套。
- 删除标签和任何属性。

例如，`'<a href="#">Click here</a>'` 应该返回 `"Click here"`。

# --hints--

`stripTags('<a href="#">Click here</a>')` 应该返回 `"Click here"`。

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` 应该返回 `"Hello World!"`。

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` 应返回空字串（`""`）。

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` 应该返回 `sectionsection`。

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
