---
id: 68cae5b538ff798bbd4da003
title: "挑战 65：字串计数"
challengeType: 28
dashedName: challenge-65
---

# --description--

给定两个字串，确定第二个字串在第一个字串中出现了多少次。

- 模式字串可以在第一个字串中重叠。例如，`"aaa"` 包含两次 `"aa"`。第一次是前两个 `a`，第二次是后两个。

# --hints--

`count('abcdefg', 'def')` 应该返回 `1`。

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` 应该返回 `0`。

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` 应该返回 `2`。

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` 应该返回 `3`。

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` 应该返回 `10`。

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
