---
id: 68d2ba1468508398389487d0
title: "挑战 75：隐藏的宝藏"
challengeType: 28
dashedName: challenge-75
---

# --description--

给定一个表现海底地图的二维数组，其中包含一个隐藏的宝藏，以及一个包含你下一次宝藏查找潜水坐标（`[row, column]`）的数组，使用以下规则返回 `"Empty"`、`"Found"` 或 `"Recovered"`。

- 给定的二维数组将包含且仅包含一个未被发现的宝藏，该宝藏将占据多个单元格。
- 二维数组中的每个单元格将包含以下值之一：
  - `"-"`：无宝藏。
  - `"O"`：尚未找到的宝藏部分。
  - `"X"`：已经找到的宝藏的一部分。

- 如果潜水地点没有宝藏，返回 `"Empty"`。
- 如果潜水地点查找到了宝藏，但至少还有一部分宝藏未被发现，返回 `"Found"`。
- 如果潜水地点查找到了宝藏的最后一部分未发现部分，返回 `"Recovered"`。

例如，给定：

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

并且 `[2, 1]` 作为潜水地点的坐标，返回 `"Recovered"`，因为潜水找到了最后一部分未发现的宝藏。

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` 应该返回 `"Recovered"`。

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` 应该返回 `"Empty"`。

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` 应该返回 `"Found"`。

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` 应该返回 `"Found"`。

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` 应该返回 `"Recovered"`。

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` 应返回 `"Empty"`。

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
