---
id: 68d30fc57588d97fd3027b30
title: "挑战 79：导航器"
challengeType: 28
dashedName: challenge-79
---

# --description--

1994 年 10 月 28 日，`Netscape Navigator` 发布，帮助数百万人探索早期网络。

给定你在 `Netscape Navigator` 上执行的一组浏览器命令数组，按照以下规则执行所有命令后，返回你当前所在的页面。

- 你总是从 `"Home"` 页面开始，该页面不会包含在命令数组中。
- 有效的命令有：
  - `"Visit Page"`：其中 `"Page"` 是你正在访问的页面名称。例如，`"Visit About"` 会带你到 `"About"` 页面。当你访问一个新建页面时，确保丢弃你所有的转发历史。
  - `"Back"`：带你返回历史中的上一页，如果没有上一页则停留在当前页面。
  - `"Forward"`：将你转发到历史记录中你之前访问的页面，如果没有则停留在当前页面。

例如，给定 `["Visit About Us", "Back", "Forward"]`，返回 `"About Us"`。

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` 应该返回 `"About Us"`。

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` 应该返回 `"Home"`。

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` 应该返回 `"Home"`。

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` 应该返回 `"Gallery"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` 应该返回 `"Home"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` 应该返回 `"Contact"`。

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` 应该返回 `"Visit Us"`。

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
