---
id: 68d30845cc08266018fc46bc
title: "挑战 76：互补 DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

给定一个表现 DNA 序列的字串，使用以下规则返回其补集链：

- DNA 由字母 `"A"`、`"C"`、`"G"` 和 `"T"` 组成。
- 字母 `"A"` 和 `"T"` 互为补集。
- 字母 `"C"` 和 `"G"` 互为补集。

例如，给定 `"ACGT"`，返回 `"TGCA"`。

# --hints--

`complementaryDNA("ACGT")` 应该返回 `"TGCA"`。

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` 应该返回 `"TACGCATGCAATCG"`。

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` 应该返回 `"CCGAATGCTAGCTTC"`。

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` 应该返回 `"CTAGATCGATCCGATCGATC"`。

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
