---
id: 68d30845cc08266018fc46bd
title: "挑战 77：持续时间格式化器"
challengeType: 28
dashedName: challenge-77
---

# --description--

给定一个整数秒数，返回一个表现相同时长的字串，格式为 `"H:MM:SS"`，其中 `"H"` 是小时数，`"MM"` 是分钟数，`"SS"` 是秒数。使用以下规则返回时间：

- 秒：应始终为两位数字。
- 分钟：当不需要时应省略前导零。如果持续时间少于 1 分钟，则使用 `"0"`。
- 小时：仅当其大于零时才应包含。

# --hints--

`format(500)` 应该返回 `"8:20"`。

```js
assert.equal(format(500), "8:20");
```

`format(4000)` 应该返回 `"1:06:40"`。

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` 应该返回 `"0:01"`。

```js
assert.equal(format(1), "0:01");
```

`format(5555)` 应该返回 `"1:32:35"`。

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` 应该返回 `"27:46:39"`。

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
