---
id: 68e39ed6106dac2f0a98fd63
title: "挑战 81：第 N 个素数"
challengeType: 28
dashedName: challenge-81
---

# --description--

素数是大于 1 的正整数，只能被 1 和它本身整除。前五个素数是 `2`、`3`、`5`、`7` 和 `11`。

给定一个正整数 `n`，返回第 `n` 个素数。例如，给定 `5` 返回第 5 个素数：`11`。

# --hints--

`nthPrime(5)` 应该返回 `11`。

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` 应该返回 `29`。

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` 应该返回 `53`。

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` 应该返回 `523`。

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` 应该返回 `7919`。

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
