---
id: 68e39ed6106dac2f0a98fd66
title: "挑战 84：感染"
challengeType: 28
dashedName: challenge-84
---

# --description--

1988 年 11 月 2 日，第一个主要的互联网蠕虫被释放，仅仅一天后就感染了约 10% 连接到互联网的计算机。

在此挑战中，你将获得自互联网蠕虫发布以来经过的天数，你需要根据以下规则确定有多少台计算机被感染：

- 在第 0 天，第一台计算机被感染。
- 每 增加 一 天，感染的计算机数量 翻倍。
- 每隔 3 天，在病毒传播后应用一个补丁，将感染的计算机数量减少 20%。将补丁应用的计算机数量向上取整到最接近的整数。

例如，在：

- 第 0 天：共有 1 台计算机被感染。
- 第 1 天：共有 2 台计算机被感染。
- 第 2 天：共有 4 台计算机被感染。
- 第 3 天：共有 8 台计算机被感染。然后，应用补丁：8 台被感染 * 20% = 1.6 台已补丁。将 1.6 向上取整为 2。8 台计算机被感染 - 2 台已补丁 = 第 3 天后共有 6 台计算机被感染。

在经过给定天数后，返回感染的计算机总数。

# --hints--

`infected(1)` 应该返回 `2`。

```js
assert.equal(infected(1), 2);
```

`infected(3)` 应返回 `6`。

```js
assert.equal(infected(3), 6);
```

`infected(8)` 应该返回 `152`。

```js
assert.equal(infected(8), 152);
```

`infected(17)` 应返回 `39808`。

```js
assert.equal(infected(17), 39808);
```

`infected(25)` 应该返回 `5217638`。

```js
assert.equal(infected(25), 5217638);
```

# --seed--

## --seed-contents--

```js
function infected(days) {

  return days;
}
```

# --solutions--

```js
function infected(days) {
  let infected = 1;

  for (let day = 1; day <= days; day++) {
    infected *= 2;

    if (day % 3 === 0) {
      let patched = Math.ceil(infected * 0.2);
      infected -= patched;
    }
  }

  return infected;
}
```
