---
id: 68e39ed6106dac2f0a98fd64
title: "挑战 82：SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

给定一个表现变量名的字串，使用以下约束将其转换为“spooky case”：

- 将所有下划线（`_`）和连字符（`-`）替换为波浪号（`~`）。
- 将字串的第一个字母和之后的每隔一个字母大写。计数时忽略波浪号字符。将所有其他字母转换为小写。

例如，给定 `hello_world`，返回 `HeLlO~wOrLd`。

# --hints--

`spookify("hello_world")` 应该返回 `"HeLlO~wOrLd"`。

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` 应该返回 `"SpOoKy~CaSe"`。

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` 应该返回 `"TrIcK~oR~tReAt"`。

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` 应该返回 `"C~a~N~d~Y~~b~O~w~L"`。

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` 应该返回 `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`。

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
