---
id: 68e39ed6106dac2f0a98fd62
title: "挑战 80：电子邮件分类器"
challengeType: 28
dashedName: challenge-80
---

# --description--

1971 年 10 月 29 日，发送了第一封电子邮件，介绍了我们仍在使用的 `username@domain` 形式。现在，有数十亿个电子邮件地址。

在此挑战中，你将获得一个电子邮件地址列表，需要先按域名（`@` 后面的部分）字母顺序排序，再按用户名（`@` 前面的部分）排序。

- 排序应忽略分支大小写。
- 如果多个电子邮件具有相同的 `Domain`，则按它们的用户名排序。
- 返回排序后的地址数组。
- 返回的地址应保留其原始分支。

例如，给定 `["jill@mail.com", "john@example.com", "jane@example.com"]`，返回 `["jane@example.com", "john@example.com", "jill@mail.com"]`。

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` 应该返回 `["jane@example.com", "john@example.com", "jill@mail.com"]`。

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` 应该返回 `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`。

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` 应该返回 `["user@x.com", "user@y.com", "user@z.com"]`。

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` 应该返回 `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`。

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` 应该返回 `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`。

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
