---
id: 68e39ed6106dac2f0a98fd65
title: "挑战 83：签名验证"
challengeType: 28
dashedName: challenge-83
---

# --description--

给定一个消息字串、一个密钥字串和一个签名数字，使用此编码方法确定签名是否有效：

- 消息和密钥中的字母具有以下值：
  - `a` 到 `z` 分别具有值 `1` 到 `26`。
  - `A` 到 `Z` 分别具有值 `27` 到 `52`。
- 所有其他字符没有价值。
- 通过将消息的和与密钥的和相加来计算签名。

例如，给定消息 `"foo"` 和密钥 `"bar"`，签名将是 `57`：

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

查看计算的签名是否与提供的签名匹配。

# --hints--

`verify("foo", "bar", 57)` 应该返回 `true`。

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` 应该返回 `false`。

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` 应该返回 `true`。

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` 应该返回 `false`。

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` 应该返回 `true`。

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` 应该返回 `true`。

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
