---
id: 68ee9e3066cfd4eb2328e8a5
title: "挑战 86：图像查找"
challengeType: 28
dashedName: challenge-86
---

# --description--

2001 年 11 月 4 日，Google 推出了其 `image search`，允许人们使用搜索词查找图像。在本次挑战中，你将模仿该 `image search`。

给定一个图像名称的数组和一个查找词，返回包含该查找词的图像名称数组。

- 匹配搜索词时忽略分支。
- 以它们在输入数组中出现的相同顺序返回图像。

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` 应该返回 `["dog.png"]`。

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` 应该返回 `["Sunset.jpg", "sunflower.jpeg"]`。

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` 应该返回 `["Moon.png", "stars.png"]`。

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` 应该返回 `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`。

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
