---
id: 68ee9e3066cfd4eb2328e8a4
title: "挑战 85：单词计数器"
challengeType: 28
dashedName: challenge-85
---

# --description--

给定一个句子字串，返回句子中单词的数量。

- 单词是任何非空格字符的序列，并由单个空格分隔。

# --hints--

`countWords("Hello world")` 应该返回 `2`。

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` 应该返回 `9`。

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` 应该返回 `4`。

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` 应该返回 `7`。

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` 应该返回 `7`。

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
