---
id: 68ee9e3066cfd4eb2328e8a7
title: "挑战 88：工作日查找器"
challengeType: 28
dashedName: challenge-88
---

# --description--

给定格式为 `YYYY-MM-DD` 的字串日期，返回星期几。

有效的返回天数有：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

务必忽略时区。

# --hints--

`getWeekday("2025-11-06")` 应该返回 `Thursday`。

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` 应该返回 `Friday`。

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` 应该返回 `Saturday`。

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` 应该返回 `Wednesday`。

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` 应该返回 `Monday`。

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
