---
id: 68ee9e3066cfd4eb2328e8a6
title: "挑战 87：矩阵构建器"
challengeType: 28
dashedName: challenge-87
---

# --description--

给定两个整数（行数和列数），返回一个指定大小的由零（`0`）填充的矩阵（数组的数组）。

例如，给定 `2` 和 `3`，返回：

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` 应该返回 `[[0, 0, 0], [0, 0, 0]]`。

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` 应该返回 `[[0, 0], [0, 0], [0, 0]]`。

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` 应该返回 `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`。

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` 应该返回 `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`。

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
