---
id: 68ee9e3066cfd4eb2328e8a8
title: "挑战 89：计数牌"
challengeType: 28
dashedName: challenge-89
---

# --description--

一副标准扑克牌的每个花色有 13 张独特的牌。给定一个表示从牌堆中抽取牌数的整数，返回你可以抽取的牌的独特组合数。

- 顺序无关紧要。先选卡片 A 再选卡片 B 与先选卡片 B 再选卡片 A 是一样的。

例如，给定 `52`，返回 `1`。从 52 张牌的牌堆中选择 52 张牌只有一种组合。并且给定 `2`，返回 `1326`，从牌堆中选择 2 张牌时，你可以得到 1326 种牌的组合。

# --hints--

`combinations(52)` 应返回 `1`。

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` 应该返回 `52`。

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` 应该返回 `1326`。

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` 应该返回 `2598960`。

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` 应返回 `15820024220`。

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` 应该返回 `1326`。

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
