---
id: 68f6587287ad1f4ad39b0c81
title: "挑战 95：数组移位"
challengeType: 28
dashedName: challenge-95
---

# --description--

给定一个数组和一个表示要移动数组多少位置的整数，返回移动后的数组。

- 正整数将数组向左移动。
- 负整数将数组向右移动。
- 移位会在数组中循环。

例如，给定 `[1, 2, 3]` 和 `1`，将数组向左移动 1 位，返回 `[2, 3, 1]`。

# --hints--

`shiftArray([1, 2, 3], 1)` 应该返回 `[2, 3, 1]`。

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` 应该返回 `[3, 1, 2]`。

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` 应该返回 `["charlie", "alpha", "bravo"]`。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` 应该返回 `["bravo", "charlie", "alpha"]`。

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` 应该返回 `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`。

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
