---
id: 68f6587287ad1f4ad39b0c7c
title: "挑战 90：字符限制"
challengeType: 28
dashedName: challenge-90
---

# --description--

在此挑战中，你将获得一个字串，需要判断它是否适合发布在社交媒体上。根据给定规则返回以下字串：

- `"short post"` 如果它符合 40 字符限制。
- `"long post"` 如果它超过 40 个字符且适合在 80 个字符限制内。
- `"invalid post"` 如果它太长而无法符合任一限制。

# --hints--

`canPost("Hello world")` 应该返回 `"short post"`。

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` 应该返回 `"long post"`。

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` 应该返回 `"invalid post"`。

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```
