---
id: 68f6587287ad1f4ad39b0c7f
title: "挑战 93：元音和构造"
challengeType: 28
dashedName: challenge-93
---

# --description--

给定一个字串，返回一个数组，其中包含字串中元音和辅音的数量。

- 元音由任意大小写的 `a`、`e`、`i`、`o`、`u` 组成。
- 辅音由任何分支中的所有其他字母组成。
- 忽略任何非字母字符。

例如，给定 `"Hello World"`，返回 `[3, 7]`。

# --hints--

`count("Hello World")` 应该返回 `[3, 7]`。

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` 应该返回 `[3, 7]`。

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` 应该返回 `[1, 5]`。

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` 应该返回 `[5, 7]`。

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` 应该返回 `[3, 7]`。

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` 应该返回 `[11, 24]`。

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
