---
id: 68f6587287ad1f4ad39b0c82
title: "挑战 96：今天是周末吗？"
challengeType: 28
dashedName: challenge-96
---

# --description--

给定格式为 `"YYYY-MM-DD"` 的日期，返回距离周末剩余的天数。

- 周末从星期六开始。
- 如果给定的日期是星期六或星期日，返回 `"It's the weekend!"`。
- 否则，返回 `"X days until the weekend."`，其中 `X` 是距离星期六的天数。
- 如果 `X` 是 `1`，使用 `"day"`（单数）而不是 `"days"`（复数）。
- 确保计算忽略你的局部时区。

# --hints--

`daysUntilWeekend("2025-11-14")` 应该返回 `"1 day until the weekend."`。

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` 应该返回 `"3 days until the weekend."`。

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` 应该返回 `"It's the weekend!"`。

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` 应该返回 `"4 days until the weekend."`。

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` 应该返回 `"5 days until the weekend."`。

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` 应该返回 `"It's the weekend!"`。

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
