---
id: 68f6587287ad1f4ad39b0c84
title: "挑战 98：矩形计数"
challengeType: 28
dashedName: challenge-98
---

# --description--

给定两个表现矩形宽度和高度的正整数，确定在给定矩形中可以容纳多少个矩形。

- 只计算宽度和高度为整数的矩形。

例如，给定 `1` 和 `3`，返回 `6`。三个 1x1 矩形，两个 1x2 矩形，和一个 1x3 矩形。

# --hints--

`countRectangles(1, 3)` 应该返回 `6`。

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` 应该返回 `18`。

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` 应该返回 `3`。

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` 应该返回 `150`。

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` 应该返回 `12540`。

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
