---
id: 68f6587287ad1f4ad39b0c7e
title: "挑战 92：扩展名提取器"
challengeType: 28
dashedName: challenge-92
---

# --description--

给定一个表现文件名的字串，返回该文件的扩展名。

- 扩展名是文件名中最后一个句点（`.`）之后的部分。
- 如果文件名不包含句点或以句点结尾，返回 `"none"`。
- 扩展名应保持原样返回，保留分支。

# --hints--

`getExtension("document.txt")` 应该返回 `"txt"`。

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` 应该返回 `"none"`。

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` 应该返回 `"PNG"`。

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` 应该返回 `"gitignore"`。

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` 应该返回 `"gz"`。

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` 应该返回 `"none"`。

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
