---
id: 68f6587287ad1f4ad39b0c85
title: "挑战 99：指纹测试"
challengeType: 28
dashedName: challenge-99
---

# --description--

给定两个表现指纹的字串，使用以下规则确定它们是否匹配：

- 每个指纹将仅由小写字母组成（`a-z`）。
- 如果两个指纹满足以下条件，则视为匹配：
  - 它们长度相同。
  - 不同的字符数不超过指纹长度的 10%。

# --hints--

`isMatch("helloworld", "helloworld")` 应该返回 `true`。

```js
assert.isTrue(isMatch("helloworld", "helloworld"));
```

`isMatch("helloworld", "helloworlds")` 应该返回 `false`。

```js
assert.isFalse(isMatch("helloworld", "helloworlds"));
```

`isMatch("helloworld", "jelloworld")` 应该返回 `true`。

```js
assert.isTrue(isMatch("helloworld", "jelloworld"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog")` 应该返回 `true`。

```js
assert.isTrue(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthelazydog"));
```

`isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog")` 应该返回 `true`。

```js
assert.isTrue(isMatch("theslickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazydog"));
```

`isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat")` 应该返回 `false`。

```js
assert.isFalse(isMatch("thequickbrownfoxjumpsoverthelazydog", "thequickbrownfoxjumpsoverthehazycat"));
```

# --seed--

## --seed-contents--

```js
function isMatch(fingerprintA, fingerprintB) {

  return fingerprintA;
}
```

# --solutions--

```js
function isMatch(fingerprintA, fingerprintB) {
  if (fingerprintA.length !== fingerprintB.length) return false;

  const length = fingerprintA.length;
  let mismatches = 0;

  for (let i = 0; i < length; i++) {
    if (fingerprintA[i] !== fingerprintB[i]) {
      mismatches++;
      if (mismatches > length * 0.1) return false;
    }
  }

  return true;
}
```
