---
id: 68ffb91507a5b645769328c9
title: "挑战 106：消息验证器"
challengeType: 28
dashedName: challenge-106
---

# --description--

给定一个消息字串和一个验证字串，确定该消息是否有效。

- 如果消息中的每个单词都按顺序以验证字串中对应的字母开头，则该消息是有效的。
- 字母不区分大小写。
- 消息中的单词由单个空格分隔。

# --hints--

`isValidMessage("hello world", "hw")` 应该返回 `true`。

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` 应该返回 `true`。

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` 应该返回 `false`。

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` 应该返回 `true`。

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` 应该返回 `false`。

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
