---
id: 68ffb91507a5b645769328c7
title: "挑战 104：配方缩放器"
challengeType: 28
dashedName: challenge-104
---

# --description--

给定一个配方原料的数组和一个用于缩放配方的数字，返回一个相应缩放数量的数组。

- 给定数组中的每一项都是 `"quantity unit ingredient"` 形式的字串。例如 `"2 C Flour"`。
- 将数量按给定数字缩放。不要包含任何尾随零，也不要转换任何单位。
- 以给定的相同顺序返回缩放后的项。

# --hints--

`scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2)` 应该返回 `["4 C Flour", "3 T Sugar"]`。

```js
assert.deepEqual(scaleRecipe(["2 C Flour", "1.5 T Sugar"], 2), ["4 C Flour", "3 T Sugar"]);
```

`scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5)` 应该返回 `["6 T Flour", "1.5 C Milk", "3 T Oil"]`。

```js
assert.deepEqual(scaleRecipe(["4 T Flour", "1 C Milk", "2 T Oil"], 1.5), ["6 T Flour", "1.5 C Milk", "3 T Oil"]);
```

`scaleRecipe(["3 C Milk", "2 C Oats"], 0.5)` 应该返回 `["1.5 C Milk", "1 C Oats"]`。

```js
assert.deepEqual(scaleRecipe(["3 C Milk", "2 C Oats"], 0.5), ["1.5 C Milk", "1 C Oats"]);
```

`scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5)` 应该返回 `["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]`。

```js
assert.deepEqual(scaleRecipe(["2 C All-purpose Flour", "1 t Baking Soda", "1 t Salt", "1 C Butter", "0.5 C Sugar", "0.5 C Brown Sugar", "1 t Vanilla Extract", "2 C Chocolate Chips"], 2.5), ["5 C All-purpose Flour", "2.5 t Baking Soda", "2.5 t Salt", "2.5 C Butter", "1.25 C Sugar", "1.25 C Brown Sugar", "2.5 t Vanilla Extract", "5 C Chocolate Chips"]);
```

# --seed--

## --seed-contents--

```js
function scaleRecipe(ingredients, scale) {

  return ingredients;
}
```

# --solutions--

```js
function scaleRecipe(ingredients, scale) {
  return ingredients.map(item => {
    const [quantityStr, unit, ...rest] = item.split(" ");
    const ingredientName = rest.join(" ");
    
    const quantity = parseFloat(quantityStr) * scale;
    
    return `${quantity} ${unit} ${ingredientName}`;
  });
}
```
