---
id: 68ffb91507a5b645769328c5
title: "挑战 102：最长单词"
challengeType: 28
dashedName: challenge-102
---

# --description--

给定一个句子字串，返回句子中最长的单词。

- 单词 由 单个 空格 分隔。
- 只有字母（`a-z`，不区分大小写）计入单词长度。
- 如果有多个长度相同的单词，返回第一个出现的单词。
- 返回字串中出现的单词，去除标点符号。

# --hints--

`longestWord("The quick red fox")` 应该返回 `"quick"`。

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` 应该返回 `"challenge"`。

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` 应该返回 `"This"`。

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` 应该返回 `"exclamation"`。

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` 应该返回 `"tie"`。

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` 应该返回 `"Wouldnt"`。

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
