---
id: 68ffb91507a5b645769328c4
title: "挑战 101：Markdown 标题转换器"
challengeType: 28
dashedName: challenge-101
---

# --description--

给定一个表现 Markdown 标题的字串，返回等效的超文本标记语言标题。

有效的 Markdown 标题必须：

- 以零个或多个空格开头，后跟
- 1 到 6 个连续的哈希字符（`#`），然后
- 至少一个空格。最后，
- 标题文本。

哈希符号的数量决定标题级别。例如，一个哈希符号对应一个 `h1` 标签，六个哈希符号对应一个 `h6` 标签。

如果给定的字串没有上述的精确形式，返回 `"Invalid format"`。

例如，给定 `"# My level 1 heading"`，返回 `"<h1>My level 1 heading</h1>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`convert("# My level 1 heading")` 应该返回 `"<h1>My level 1 heading</h1>"`。

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` 应该返回 `"Invalid format"`。

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` 应该返回 `"<h5>My level 5 heading</h5>"`。

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` 应返回 `"Invalid format"`。

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` 应该返回 `"<h3>My level 3 heading</h3>"`。

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` 应返回 `"Invalid format"`。

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` 应该返回 `"<h2>My #2 heading</h2>"`。

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
