---
id: 68ffb91507a5b645769328cc
title: "挑战 109：我到底几岁了？"
challengeType: 28
dashedName: challenge-109
---

# --description--

给定某人生日的日期，格式为 `YYYY-MM-DD`，返回该人在 2025 年 11 月 27 日的年龄。

- 假设所有生日均为 2025 年 11 月 27 日之前的有效日期。
- 将年龄作为整数返回。
- 请确保考虑该人在 2025 年是否已经过生日。

# --hints--

`calculateAge("2000-11-20")` 应该返回 `25`。

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` 应该返回 `24`。

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` 应该返回 `11`。

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` 应该返回 `31`。

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` 应该返回 `30`。

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
