---
id: 68ffb91507a5b645769328c8
title: "挑战 105：字符计数"
challengeType: 28
dashedName: challenge-105
---

# --description--

给定一个句子字串，返回一个按字母顺序排列的每个字符计数的数组。

- 计算时将大写字母和小写字母视为相同的字母。
- 忽略数字、空格、标点符号等。
- 以 `"letter count"` 形式返回计数和字母。例如，`"a 3"`。
- 所有返回的字母都应为小写。
- 不要返回不在给定 `string` 中的字母计数。

# --hints--

`countCharacters("hello world")` 应该返回 `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`。

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` 应该返回 `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`。

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` 应该返回 `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`。

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
