---
id: 68ffb91507a5b645769328c6
title: "挑战 103：最小公倍数"
challengeType: 28
dashedName: challenge-103
---

# --description--

给定两个整数，返回这两个数字的最小公倍数（LCM）。

两个数的最小公倍数是同时为这两个数的倍数的最小正整数。例如，给定 `4` 和 `6`，返回 `12`，因为：

- `4` 的倍数是 `4`、`8`、`12` 等等。
- `6` 的倍数是 `6`、`12`、`18` 等等。
- `12` 是同时为两者倍数的最小数字。

# --hints--

`lcm(4, 6)` 应该返回 `12`。

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` 应返回 `18`。

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` 应该返回 `100`。

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` 应该返回 `221`。

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` 应该返回 `630`。

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
