---
id: 68ffb91507a5b645769328c3
title: "挑战 100：100 个 `character`"
challengeType: 28
dashedName: challenge-100
---

# --description--

欢迎来到第 100 个每日编码挑战！

给定一个字串，重复其字符直到结果正好为 100 个字符长。如果你的重复超过 100 个字符，裁剪多余部分使其正好为 100 个字符。

# --hints--

`oneHundred("One hundred ")` 应该返回 `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`。

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` 应该返回 `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`。

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` 应该返回 `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`。

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` 应该返回 `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`。

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
