---
id: 69162d64f96574d9bb629efd
title: "挑战 112：AI 检测器"
challengeType: 28
dashedName: challenge-112
---

# --description--

今天的挑战灵感来自于 2022 年 11 月 30 日 ChatGPT 的版本。

给定一个包含一个或多个句子的字串，使用以下规则判断它是否可能由 AI 生成：

- 它包含两个或更多的破折号（`-`）。
- 它包含两个或更多组括号（`()`）。文本可以位于括号内。
- 它包含三个或更多单词，每个单词有 7 个或更多字母。

- 单词由单个空格分隔且仅由字母组成（`A-Z`）。不要将标点符号或其他非字母字符作为单词的一部分。

如果给定的句子符合上述任何规则，返回 `"AI"`，否则，返回 `"Human"`。

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` 应该返回 `"Human"`。

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` 应该返回 `"Human"`。

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` 应该返回 `"AI"`。

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` 应该返回 `"AI"`。

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` 应该返回 `"AI"`。

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
