---
id: 69162d64f96574d9bb629f01
title: "挑战 116：排列计数"
challengeType: 28
dashedName: challenge-116
---

# --description--

给定一个字串，返回可以由其字符组成的不同排列数。

- 排列是字串中字符的任意重新排序。
- 不要计算副本排列。
- 如果字串包含重复的字符，重复的排列只应计数一次。
- 该字串将仅包含字母（`A-Z`、`a-z`）。

例如，给定 `"abb"`，返回 `3`，因为有三种不同的方式排列这些字母：`"abb"`、`"bab"` 和 `"bba"`。

# --hints--

`countPermutations("abb")` 应该返回 `3`。

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` 应该返回 `6`。

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` 应该返回 `630`。

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` 应该返回 `39916800`。

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
