---
id: 69162d64f96574d9bb629efe
title: "挑战 113：英里转公里"
challengeType: 28
dashedName: challenge-113
---

# --description--

给定以英里为单位的距离数字，返回等效的公里数。

- 输入将始终是非负数。
- 1 英里相等 1.60934 公里。
- 将结果四舍五入到两位小数。
- 从四舍五入的结果中移除不必要的尾随零。

# --hints--

`convertToKm(1)` 应该返回 `1.61`。

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` 应该返回 `33.8`。

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` 应该返回 `5.63`。

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` 应该返回 `0`。

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` 应返回 `1`。

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
