---
id: 69162d64f96574d9bb629f00
title: "挑战 115：Markdown 有序列表项转换器"
challengeType: 28
dashedName: challenge-115
---

# --description--

给定一个表示 Markdown 中有序列表项的字串，返回等效的超文本标记语言字串。

Markdown 中有效的有序列表项必须：

- 以零个或多个空格开头，后跟
- 一个数字（1 或更大）和一个句点（`.`），后跟
- 至少一个空格，然后
- 列表项文本。

如果字串没有上述的精确形式，返回 `"Invalid format"`。否则，将列表项文本包裹在 `li` 标签中并返回该字串。

例如，给定 `"1. My item"`，返回 `"<li>My item</li>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`convertListItem("1. My item")` 应该返回 `"<li>My item</li>"`。

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` 应该返回 `"<li>Another item</li>"`。

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` 应该返回 `"Invalid format"`。

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` 应该返回 `"<li>list item text</li>"`。

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` 应该返回 `"Invalid format"`。

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` 应该返回 `"Invalid format"`。

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
