---
id: 69162d64f96574d9bb629f03
title: "挑战 118：日期格式化器"
challengeType: 28
dashedName: challenge-118
---

# --description--

给定格式为 `"Month day, year"` 的日期，返回格式为 `"YYYY-MM-DD"` 的日期。

- 给定的月份将是完整的英文月份名称。例如：`"January"`、`"February"` 等。
- 在返回值中，如果需要，使用前导零填充月份和日期以确保两位数。

例如，给定 `"December 6, 2025"`，返回 `"2025-12-06"`。

# --hints--

`formatDate("December 6, 2025")` 应该返回 `"2025-12-06"`。

```js
assert.equal(formatDate("December 6, 2025"), "2025-12-06");
```

`formatDate("January 1, 2000")` 应该返回 `"2000-01-01"`。

```js
assert.equal(formatDate("January 1, 2000"), "2000-01-01");
```

`formatDate("November 11, 1111")` 应该返回 `"1111-11-11"`。

```js
assert.equal(formatDate("November 11, 1111"), "1111-11-11");
```

`formatDate("September 7, 512")` 应该返回 `"512-09-07"`。

```js
assert.equal(formatDate("September 7, 512"), "512-09-07");
```

`formatDate("May 4, 1950")` 应该返回 `"1950-05-04"`。

```js
assert.equal(formatDate("May 4, 1950"), "1950-05-04");
```

`formatDate("February 29, 1992")` 应该返回 `"1992-02-29"`。

```js
assert.equal(formatDate("February 29, 1992"), "1992-02-29");
```

# --seed--

## --seed-contents--

```js
function formatDate(dateString) {

  return dateString;
}
```

# --solutions--

```js
function formatDate(dateString) {
  const months = {
    January: "01",
    February: "02",
    March: "03",
    April: "04",
    May: "05",
    June: "06",
    July: "07",
    August: "08",
    September: "09",
    October: "10",
    November: "11",
    December: "12"
  };

  const [month, day, year] = dateString.replace(",", "").split(" ");

  return `${year}-${months[month]}-${day.padStart(2, "0")}`;
}
```
