---
id: 69162d64f96574d9bb629eff
title: "挑战 114：驼峰命名转蛇形命名"
challengeType: 28
dashedName: challenge-114
---

# --description--

给定一个驼峰式字串，使用以下规则返回该字串的蛇形分支版本：

- 输入字串将仅包含字母（`A-Z` 和 `a-z`），且始终以小写字母开头。
- 驼峰字串中的每个大写字母都表示一个新建单词的开始。
- 将所有字母转换为小写。
- 用下划线（`_`）分隔单词。

# --hints--

`toSnake("helloWorld")` 应该返回 `"hello_world"`。

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` 应该返回 `"my_variable_name"`。

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` 应该返回 `"freecodecamp_daily_challenges"`。

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
