---
id: 69162d64f96574d9bb629efb
title: "挑战 110：单词猜测器"
challengeType: 28
dashedName: challenge-110
---

# --description--

给定两个长度相同的字串，一个是秘密字串，一个是猜测字串，使用以下规则将猜测字串与秘密字串进行比较：

- 秘密词和猜测将仅由大写字母（`"A"` 到 `"Z"`）组成；
- 对于猜测中的每个字母，根据它与秘密单词的匹配方式将其替换为一个数字：
  - 如果字母在秘密单词中且位置正确，则返回 `"2"`。
  - 如果字母在秘密单词中但位置错误，则为 `"1"`。
  - 如果字母不在秘密单词中，则为 `"0"`。
- 秘密词中的每个字母最多只能使用一次。
- 完全匹配（`"2"`）首先被赋值，然后部分匹配（`"1"`）从左到右为剩余字母赋值。
- 如果一个字母在猜测中出现多次，它只能匹配秘密词中出现的相同次数。

例如，给定一个秘密词为 `"APPLE"` 且猜测为 `"POPPA"`，返回 `"10201"`：

第一个 `"P"` 不在正确位置（`"1"`），`"O"` 不在秘密单词中（`"0"`），第二个 `"P"` 在正确位置（`"2"`），第三个 `"P"` 是零（`"0"`），因为秘密单词中的两个 `"P"` 已被使用，`"A"` 不在正确位置（`"1"`）。

# --hints--

`compare("APPLE", "POPPA")` 应该返回 `"10201"`。

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` 应该返回 `"11221"`。

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` 应该返回 `"000000"`。

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` 应该返回 `"0000222222"`。

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` 应该返回 `"110200"`。

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` 应该返回 `"10021000"`。

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
