---
id: 69162d64f96574d9bb629f04
title: "挑战 119：字串压缩"
challengeType: 28
dashedName: challenge-119
---

# --description--

给定一个字串 sentence，返回该句子的压缩版本，其中连续的副本词被替换为该词，后跟其重复次数的括号。

- 只有连续的副本会被压缩。
- 单词由单个空格分隔。

例如，给定 `"yes yes yes please"`，返回 `"yes(3) please"`。

# --hints--

`compressString("yes yes yes please")` 应该返回 `"yes(3) please"`。

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` 应该返回 `"I have(3) apples"`。

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` 应该返回 `"one(2) three and to the(4)"`。

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` 应该返回 `"route(6) tee(6)"`。

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
