---
id: 69162d64f96574d9bb629f02
title: "挑战 117：对称差集"
challengeType: 28
dashedName: challenge-117
---

# --description--

给定两个数组，返回一个包含它们对称差的新建数组。

- 两个集合的对称差是出现在任一集合中但不同时出现在两个集合中的值的集合。
- 返回值按它们在输入数组中首次出现的顺序。

# --hints--

`difference([1, 2, 3], [3, 4, 5])` 应该返回 `[1, 2, 4, 5]`。

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` 应该返回 `["a", "c"]`。

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` 应该返回 `[1, "b"]`。

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` 应该返回 `[2, 4, 6, 8]`。

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
