---
id: 691b559495c5cb5a37b9b485
title: "挑战 125：生命游戏"
challengeType: 28
dashedName: challenge-125
---

# --description--

给定一个表现康威生命游戏当前状态的矩阵（数组的数组），使用以下规则返回矩阵的下一状态：

- 每个单元格要么是 `1`（存活），要么是 `0`（死亡）。
- 一个单元格的邻居是最多八个周围的单元格（垂直、水平和对角线方向）。
- 边缘的单元格邻居少于八个。

更新每个单元格的规则：

- 任何活细胞如果活邻居少于 2 个则死亡（人口过少）。
- 任何有两个或三个活邻居的活细胞将继续存活。
- 任何有超过 3 个活邻居的活细胞都会死亡（过度繁殖）。
- 任何恰好有三个活邻居的死细胞会变为活细胞（繁殖）。

例如，给定：

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

返回：

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

每个单元格根据存活邻居的数量进行更新。例如，`[0][0]` 保持死亡（2 个存活邻居），`[0][1]` 保持存活（2 个存活邻居），`[0][2]` 死亡（3 个存活邻居），依此类推。

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` 应该返回 `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`。

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` 应该返回 `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`。

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` 应该返回 `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`。

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` 应该返回 `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`。

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
