---
id: 691b559495c5cb5a37b9b487
title: "挑战 127：速度查看"
challengeType: 28
dashedName: challenge-127
---

# --description--

根据你以每小时英里数（MPH）行驶的速度和以每小时公里数（KPH）表示的限速，判断你是否超速以及你会收到警告还是罚单。

- 1 英里相等 1.60934 公里。
- 如果你的速度小于或相等于限速，返回 `"Not Speeding"`。
- 如果你超速不超过 5 KPH，返回 `"Warning"`。
- 如果你超速超过限速 5 KPH，返回 `"Ticket"`。

# --hints--

`speedCheck(30, 70)` 应该返回 `"Not Speeding"`。

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` 应该返回 `"Warning"`。

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` 应该返回 `"Not Speeding"`。

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` 应该返回 `"Ticket"`。

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` 应该返回 `"Warning"`。

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` 应该返回 `"Ticket"`。

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
