---
id: 691b559495c5cb5a37b9b486
title: "挑战 126：首字母大写"
challengeType: 28
dashedName: challenge-126
---

# --description--

给定一个字串 `title`，返回一个按照以下规则格式化为标题形式的新建字串：

- 将每个单词的首字母大写。
- 将每个单词中的其他所有字母转换为小写。
- 单词总是由单个空格分隔。

# --hints--

`titleCase("hello world")` 应该返回 `"Hello World"`。

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` 应该返回 `"The Quick Brown Fox"`。

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` 应该返回 `"Javascript And Python"`。

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` 应该返回 `"Avocado Toast For Breakfast"`。

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
