---
id: 691b559495c5cb5a37b9b481
title: "挑战 121：最频繁"
challengeType: 28
dashedName: challenge-121
---

# --description--

给定一个元素数组，返回出现频率最高的元素。

- 总是会有一个出现频率最高的元素。

# --hints--

`mostFrequent(["a", "b", "a", "c"])` 应该返回 `"a"`。

```js
assert.equal(mostFrequent(["a", "b", "a", "c"]), "a");
```

`mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])` 应该返回 `2`。

```js
assert.equal(mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2);
```

`mostFrequent([true, false, "false", "true", false])` 应该返回 `false`。

```js
assert.isFalse(mostFrequent([true, false, "false", "true", false]));
```

`mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])` 应该返回 `40`。

```js
assert.equal(mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40);
```

# --seed--

## --seed-contents--

```js
function mostFrequent(arr) {
  return arr;
}
```

# --solutions--

```js
function mostFrequent(arr) {
  const freq = {};
  let maxCount = 0;
  let mostElem = null;

  for (const el of arr) {
    freq[el] = (freq[el] || 0) + 1;
    if (freq[el] > maxCount) {
      maxCount = freq[el];
      mostElem = el;
    }
  }

  return mostElem;
}
```
