---
id: 691b559495c5cb5a37b9b488
title: "挑战 128：辅音计数"
challengeType: 28
dashedName: challenge-128
---

# --description--

给定一个字串和一个目标数字，确定该字串是否恰好包含目标数量的构造。

- 辅音是除任何大小写的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"` 之外的所有字母字符。
- 忽略数字、标点符号、空格和其他非字母字符时计数。

# --hints--

`hasConsonantCount("helloworld", 7)` 应该返回 `true`。

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` 应该返回 `false`。

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` 应该返回 `true`。

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` 应返回 `false`。

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` 应该返回 `true`。

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
