---
id: 691b559495c5cb5a37b9b480
title: "挑战 120：磅到千克"
challengeType: 28
dashedName: challenge-120
---

# --description--

给定一个以磅为单位的权值数字，返回字串 `"(lbs) pounds equals (kgs) kilograms."`。

- 将 `"(lbs)"` 替换为输入数字。
- 将 `"(kgs)"` 替换为转换为千克的输入，四舍五入到两位小数且始终包含两位小数。
- 1 磅等于 0.453592 千克。
- 如果 `input` 是 `1`，则使用 `"pound"` 而不是 `"pounds"`。
- 如果转换后的值是 `1`，使用 `"kilogram"` 而不是 `"kilograms"`。

# --hints--

`convertToKgs(1)` 应该返回 `"1 pound equals 0.45 kilograms."`。

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` 应该返回 `"0 pounds equals 0.00 kilograms."`。

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` 应该返回 `"100 pounds equals 45.36 kilograms."`。

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` 应该返回 `"2.5 pounds equals 1.13 kilograms."`。

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` 应该返回 `"2.20462 pounds equals 1.00 kilogram."`。

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
