---
id: 691b559495c5cb5a37b9b482
title: "挑战 122：Markdown 粗体解析器"
challengeType: 28
dashedName: challenge-122
---

# --description--

给定一个可能包含一些 Markdown 粗体文本的字串，返回等效的超文本标记语言字串。

- Markdown 中的粗体文本是任何以两个星号（`**`）或两个下划线（`__`）开始和结束的文本。
- 文本和星号或下划线之间不能有任何空格，但文本本身可以有空格。
- 将所有加粗内容转换为超文本标记语言 `b` 标签并返回该字串。

例如，给定 `"**This is bold**"`，返回 `"<b>This is bold</b>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`parseBold("**This is bold**")` 应该返回 `"<b>This is bold</b>"`。

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` 应该返回 `"<b>This is also bold</b>"`。

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` 应该返回 `"**This is not bold **"`。

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` 应该返回 `"__ This is also not bold__"`。

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` 应该返回 `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`。

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
