---
id: 691b5597f0f3e85a588a5c94
title: "挑战 130：Checkerboard"
challengeType: 28
dashedName: challenge-130
---

# --description--

给定一个包含两个数字的数组，第一个是行数，第二个是列数，返回一个由指定大小的 `"X"` 和 `"O"` 字符填充的矩阵（数组的数组）。

- 字符应像棋盘格一样交替排列。
- 左上角单元格必须始终是 `"X"`。

例如，给定 `[3, 3]`，返回：

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])` 应该返回 `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`。

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])` 应该返回 `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`。

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])` 应该返回 `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`。

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])` 应该返回 `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`。

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
