---
id: 6925e2068081f40f549ced1b
title: "挑战 137：雪花生成器"
challengeType: 28
dashedName: challenge-137
---

# --description--

给定一个使用换行字符（`\n`）表现换行的多行字串，返回一个新字串，其中每一行都水平镜射并附加到原始行的末尾。

- 通过反转所有字符（包括空格）来镜射一行。

例如，给定 `"* \n *\n* "`，它会记录到控制台为：

```sh
* 
 *
* 
```

返回 `"*  *\n ** \n*  *"`，它会将日志输出到控制台，如下所示：

```sh
*  *
 ** 
*  *
```

请仔细注意给定和返回的字串中的空格。确保不要删除任何空格。

# --hints--

`generateSnowflake("* \n *\n* ")` 应该返回 `"*  *\n ** \n*  *"`。

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` 应该返回 `"X=~~=X"`。

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` 应该返回 `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`。

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` 应该返回 `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`。

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` 应该返回 `"*  --  *\n * -- * \n*  --  *"`。

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
