---
id: 6925e2068081f40f549ced1d
title: "挑战 139：石头、剪刀、布"
challengeType: 28
dashedName: challenge-139
---

# --description--

给定两个字串，第一个表现玩家 1，第二个表现玩家 2，确定石头剪刀布比赛的获胜者。

- 输入字串将始终是 `"Rock"`、`"Paper"` 或 `"Scissors"`。
- `"Rock"` 胜过 `"Scissors"`。
- `"Paper"` 胜过 `"Rock"`。
- `"Scissors"` 胜过 `"Paper"`。

返回：

- 如果 Player 1 获胜，则显示 `"Player 1 wins"`。
- 如果 Player 2 获胜，则返回 `"Player 2 wins"`。
- `"Tie"` 如果双方玩家选择相同的选项。

# --hints--

`rockPaperScissors("Rock", "Rock")` 应该返回 `"Tie"`。

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` 应该返回 `"Player 2 wins"`。

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` 应该返回 `"Player 1 wins"`。

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` 应该返回 `"Player 1 wins"`。

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` 应该返回 `"Tie"`。

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` 应该返回 `"Player 2 wins"`。

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
