---
id: 6925e2068081f40f549ced1c
title: "挑战 138：除数之和"
challengeType: 28
dashedName: challenge-138
---

# --description--

给定一个正整数，返回它所有除数的和。

- 除数是任何能将该数字整除（余数为 `0`）的整数。
- 每个除数只计算一次。

例如，给定 `6`，返回 `12`，因为 `6` 的除数是 `1`、`2`、`3` 和 `6`，它们的和是 `12`。

# --hints--

`sumDivisors(6)` 应该返回 `12`。

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` 应该返回 `14`。

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` 应该返回 `56`。

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` 应该返回 `224`。

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` 应该返回 `806`。

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` 应该返回 `23520`。

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
