---
id: 69272dcf1c24b44fd79137c6
title: "挑战 143：Markdown 斜体解析器"
challengeType: 28
dashedName: challenge-143
---

# --description--

给定一个可能包含一些 Markdown 斜体文本的字串，返回等效的超文本标记语言字串。

- Markdown 中的斜体文本是以单个星号（`*`）或单个下划线（`_`）开始和结束的任何文本。
- 文本和星号或下划线之间不能有任何空格，但文本本身可以有空格。
- 将所有斜体内容转换为超文本标记语言 `i` 标签并返回该字串。

例如，给定 `"*This is italic*"`，返回 `"<i>This is italic</i>"`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`parseItalics("*This is italic*")` 应该返回 `"<i>This is italic</i>"`。

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` 应该返回 `"<i>This is also italic</i>"`。

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` 应该返回 `"*This is not italic *"`。

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` 应该返回 `"_ This is also not italic_"`。

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` 应该返回 `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`。

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
