---
id: 69272dcf1c24b44fd79137c5
title: "挑战 142：求字串之和"
challengeType: 28
dashedName: challenge-142
---

# --description--

给定一个包含数字和其他字符的字串，返回字串中所有数字的总和。

- 将连续数字视为一个数字。例如，`"13"` 计为 13，而不是 1 + 3。
- 忽略所有非数字字符。

# --hints--

`stringSum("3apples2bananas")` 应该返回 `5`。

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` 应该返回 `17`。

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` 应返回 `125344`。

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` 应该返回 `321`。

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` 应该返回 `1653`。

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
