---
id: 69272dcf1c24b44fd79137c3
title: "挑战 140：SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

给定一个表现变量名的字串，返回转换为 SCREAMING_SNAKE_CASE 的变量名。

给定的变量名将以以下一种形式书写：

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

在上述形式中，单词由下划线（`_`）、连字符（`-`）分隔，或新单词以大写字母开头。

转换为 SCREAMING_SNAKE_CASE：

- 将所有字母转换为大写
- 用下划线（`_`）分隔单词

# --hints--

`toScreamingSnakeCase("userEmail")` 应该返回 `"USER_EMAIL"`。

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` 应该返回 `"USER_PASSWORD"`。

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` 应该返回 `"USER_ID"`。

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` 应该返回 `"USER_ADDRESS"`。

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` 应该返回 `"USERNAME"`。

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();;
}
```
