---
id: 69272dcf1c24b44fd79137c4
title: "挑战 141：起飞燃油"
challengeType: 28
dashedName: challenge-141
---

# --description--

给定你飞机当前的燃料加仑数和到达目的地所需的燃料升数，确定你应该添加多少额外的燃料加仑。

- 1 加仑相等于 3.78541 升。
- 如果飞机已经有足够的燃料，返回 `0`。
- 你只能添加整加仑。
- 返回数字中不要包含小数。

# --hints--

`fuelToAdd(0, 1)` 应该返回 `1`。

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` 应该返回 `6`。

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` 应返回 `0`。

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` 应该返回 `4520`。

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` 应该返回 `12209`。

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
