---
id: 69306364df283fcaff2e1ad8
title: "挑战 147：闰年计算器"
challengeType: 28
dashedName: challenge-147
---

# --description--

给定一个整数年份，判断它是否为闰年。

如果一年满足以下规则，则该年是闰年：

- 该年份能被 4 整除，且
- 该年份不能被 100 整除，除非
- 该年份能被 400 整除。

# --hints--

`isLeapYear(2024)` 应该返回 `true`。

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` 应该返回 `false`。

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` 应该返回 `false`。

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` 应该返回 `true`。

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` 应该返回 `false`。

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` 应该返回 `true`。

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` 应该返回 `false`。

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
