---
id: 69306364df283fcaff2e1ad6
title: "挑战 145：第 N 个斐波那契数"
challengeType: 28
dashedName: challenge-145
---

# --description--

给定一个整数 `n`，返回斐波那契序列中的第 `n` 个数字。

斐波那契序列是一串数字，其中每个数字都是前述的两个数字之和。序列中的前 10 个数字是 `0`、`1`、`1`、`2`、`3`、`5`、`8`、`13`、`21`、`34`。

# --hints--

`nthFibonacci(4)` 应该返回 `2`。

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` 应该返回 `34`。

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` 应该返回 `377`。

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` 应该返回 `63245986`。

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` 应该返回 `1304969544928657`。

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
