---
id: 69306364df283fcaff2e1ad9
title: "挑战 148：轮胎气压"
challengeType: 28
dashedName: challenge-148
---

# --description--

给定一个包含四个数字的数组，表现你车辆四个轮胎的胎压（单位为 psi），以及另一个包含两个数字的数组，表现你轮胎的最小和最大压力（单位为 bar），返回一个包含四个字串的数组，描述每个轮胎的状态。

- 1 bar 相等 14.5038 psi。

返回一个包含每个轮胎以下值的数组：

- `"Low"` 如果轮胎压力低于允许的最低值。
- `"Good"` 如果它在允许的最小值和最大值之间。
- `"High"` 如果它超过最大允许值。

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` 应该返回 `["Good", "Low", "Good", "Low"]`。

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` 应该返回 `["Good", "Low", "High", "Good"]`。

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` 应该返回 `["Low", "Low", "Good", "Low"]`。

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` 应该返回 `["High", "High", "High", "Good"]`。

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` 应该返回 `["Good", "Good", "Good", "Good"]`。

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
