---
id: 69306364df283fcaff2e1ada
title: "挑战 149：vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

给定一个字串，返回一个新建字串，其中所有元音字母转换为大写，所有其他字母字符转换为小写。

- 元音是任何分支中的 `"a"`、`"e"`、`"i"`、`"o"` 和 `"u"`。
- 非字母字符应保持不变。

# --hints--

`vowelCase("vowelcase")` 应该返回 `"vOwElcAsE"`。

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` 应该返回 `"cOdIng Is fUn"`。

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` 应该返回 `"hEllO, wOrld!"`。

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` 应该返回 `"gIt chErry-pIck"`。

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` 应该返回 `"hEAd~1"`。

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
