---
id: 69373793f5a867f769cde136
title: "挑战 151：排序数组？"
challengeType: 28
dashedName: challenge-151
---

# --description--

给定一个数字数组，判断这些数字是按升序排列、降序排列还是都不是。

如果给定的数组是：

- 按升序（从低到高）返回 `"Ascending"`。
- 按降序（从高到低）返回 `"Descending"`。
- 未按升序或降序排序，返回 `"Not sorted"`。

# --hints--

`isSorted([1, 2, 3, 4, 5])` 应该返回 `"Ascending"`。

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` 应该返回 `"Descending"`。

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` 应该返回 `"Not sorted"`。

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` 应该返回 `"Descending"`。

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` 应该返回 `"Ascending"`。

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` 应该返回 `"Not sorted"`。

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
