---
id: 69373793f5a867f769cde13a
title: "挑战 155：种植作物"
challengeType: 28
dashedName: challenge-155
---

# --description--

给定一个表现你农场字段大小的整数，以及表现农场字段大小单位的 `"acres"` 或 `"hectares"`，还有一种作物类型，确定你可以在字段中种植多少该类型的植物。

- 1 英亩相等于 4046.86 平方米。
- 1 公顷等于 10,000 平方米。

以下是作为 `input` 提供的作物 `list` 以及单株植物占用的空间：

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

返回适合该字段的植物数量，向下取整到最接近的整数植物数。

# --hints--

`getNumberOfPlants(1, "acres", "corn")` 应该返回 `4046`。

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` 应该返回 `100000`。

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` 应该返回 `161874`。

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` 应该返回 `150000`。

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` 应该返回 `271139`。

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
