---
id: 69373793f5a867f769cde139
title: "挑战 154：洞口标准杆数"
challengeType: 28
dashedName: challenge-154
---

# --description--

给定两个整数，一个是高尔夫球洞的标准杆数，另一个是高尔夫球手在该洞的击球次数，使用高尔夫术语返回球手的得分。

返回：

- 如果只用了一杆，则为 `"Hole in one!"`。
- `"Eagle"` 如果比标准杆少打了 2 杆。
- `"Birdie"` 如果比标准杆少打一杆。
- `"Par"` 如果击球次数与标准杆相同。
- `"Bogey"` 如果比标准杆多打一杆。
- 如果比标准杆多用了两杆，则为 `"Double bogey"`。

# --hints--

`golfScore(3, 3)` 应该返回 `"Par"`。

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` 应该返回 `"Birdie"`。

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` 应该返回 `"Hole in one!"`。

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` 应该返回 `"Double bogey"`。

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` 应该返回 `"Bogey"`。

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` 应该返回 `"Eagle"`。

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
