---
id: 69373793f5a867f769cde137
title: "挑战 152：循环素数"
challengeType: 28
dashedName: challenge-152
---

# --description--

给定一个整数，判断它是否是循环素数。

循环素数是一个整数，其所有数字的旋转都是素数。

例如，`197` 是一个循环素数，因为其所有数字的旋转：`197`、`971` 和 `719`，都是素数。

# --hints--

`isCircularPrime(197)` 应该返回 `true`。

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` 应该返回 `false`。

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` 应该返回 `true`。

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` 应该返回 `false`。

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` 应该返回 `true`。

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
