---
id: 69373793f5a867f769cde138
title: "挑战 153：井字棋"
challengeType: 28
dashedName: challenge-153
---

# --description--

给定一个 3×3 矩阵（一个数组的数组），表现一个已完成的井字棋游戏，确定赢家。

- 给定矩阵中的每个元素要么是 `"X"`，要么是 `"O"`。

如果玩家有三个他们的字符在一行——水平、垂直或对角线方向——则玩家获胜。

返回：

- 如果玩家 X 连成三子，则显示 `"X wins"`。
- `"O wins"` 如果玩家 O 连成三子。
- 如果没有玩家连成三子，则为 `"Draw"`。

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` 应该返回 `"X wins"`。

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` 应该返回 `"X wins"`。

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` 应该返回 `"Draw"`。

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` 应该返回 `"O wins"`。

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` 应该返回 `"X wins"`。

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` 应该返回 `"Draw"`。

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
