---
id: 6939b873185d8e00d453563f
title: "挑战 160：骑士移动"
challengeType: 28
dashedName: challenge-160
---

# --description--

给定骑士在棋盘上的位置，返回骑士可以移动到的有效格子数。

标准棋盘是 8x8，列从左到右标记为 `A` 到 `H`，行从下到上标记为 `1` 到 `8`。它看起来像这样：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

骑士以“L”形移动：在一个方向（水平方向或垂直方向）移动两平方，在垂直方向移动一平方。

这意味着骑士最多可以移动到八个可能的位置，但靠近棋盘边缘时会更少。例如，如果骑士在 `A1`，它只能移动到 `B3` 或 `C2`。

# --hints--

`knightMoves("A1")` 应该返回 `2`。

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` 应该返回 `8`。

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` 应该返回 `6`。

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` 应该返回 `3`。

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` 应该返回 `4`。

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
