---
id: 6939b873185d8e00d453563c
title: "挑战 157：Markdown 链接解析器"
challengeType: 28
dashedName: challenge-157
---

# --description--

给定 Markdown 中的链接字串，返回等效的超文本标记语言字串。

Markdown 图像具有以下形式：`"[link_text](link_url)"`。返回超文本标记语言 `a` 标签的字串，其 `href` 设置为 `link_url`，并且 `link_text` 作为标签内容。

例如，给定 `"[freeCodeCamp](https://freecodecamp.org/)"` 返回 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`；

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签——查看浏览器控制台以查看包含标签的日志。

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` 应该 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`。

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` 应该返回 `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`。

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` 应该 `return` `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`。

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
