---
id: 6939b873185d8e00d453563d
title: "挑战 158：数组交换"
challengeType: 28
dashedName: challenge-158
---

# --description--

给定一个包含两个值的数组，返回一个值已交换的数组。

例如，给定 `["A", "B"]` 返回 `["B", "A"]`。

# --hints--

`arraySwap(["A", "B"])` 应该返回 `["B", "A"]`。

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` 应该返回 `[20, 25]`。

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` 应该返回 `[false, true]`。

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` 应该返回 `[1, "1"]`。

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
