---
id: 6939b873185d8e00d453563b
title: "挑战 156：奇数还是偶数？"
challengeType: 28
dashedName: challenge-156
---

# --description--

给定一个正整数，如果它是奇数，则返回 `"Odd"`，如果是偶数，则返回 `"Even"`。

# --hints--

`oddOrEven(1)` 应该返回 `"Odd"`。

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` 应该返回 `"Even"`。

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` 应该返回 `"Odd"`。

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` 应返回 `"Even"`。

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` 应该返回 `"Odd"`。

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
