---
id: 6939b873185d8e00d453563e
title: "挑战 159：整数斜边"
challengeType: 28
dashedName: challenge-159
---

# --description--

给定两个表现直角三角形两条直角边（两条较短边）长度的正整数，确定斜边是否为整数。

斜边的长度通过将两条直角边长度的平方相加，然后对该总和取平方根来计算（a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>）。

# --hints--

`isIntegerHypotenuse(3, 4)` 应该返回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` 应该返回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` 应该返回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` 应该返回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` 应该返回 `true`。

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` 应该返回 `false`。

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
