---
id: 694596b0585c11170ac7c7fc
title: "挑战 164：Markdown 内联代码解析器"
challengeType: 28
dashedName: challenge-164
---

# --description--

给定包含一个或多个内联代码块的 Markdown 字串，返回等效的超文本标记语言字串。

Markdown 中的内联代码块使用单个反引号（```）作为代码块文本的开始和结束。

返回给定的字串，将所有代码块转换为超文本标记语言的 `code` 标签。

例如，给定字串 `"Use `let` to declare the variable."`，返回 `"Use <code>let</code> to declare the variable."`。

注意：控制台在记录消息时可能不会显示字串中的超文本标记语言标签。查看浏览器控制台以查看包含标签的日志。

# --hints--

`parseInlineCode("Use `let` to declare the variable.")` 应该返回 `"Use <code>let</code> to declare the variable."`。

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

`parseInlineCode("Use `let` or `const` to declare a variable.")` 应该返回 `"Use <code>let</code> or <code>const</code> to declare a variable."`。

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

`parseInlineCode("Run `npm install` then `npm start`.")` 应该返回 `"Run <code>npm install</code> then <code>npm start</code>."`。

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
