---
id: 694596b0585c11170ac7c7fb
title: "挑战 163：辅音分支"
challengeType: 28
dashedName: challenge-163
---

# --description--

给定一个表现变量名的字串，使用以下规则将其转换为 consonant case：

- 所有辅音字母应转换为大写。
- 所有元音（无论大小写的 `a`、`e`、`i`、`o`、`u`）都应转换为小写。
- 所有连字符（`-`）应转换为下划线（`_`）。

# --hints--

`toConsonantCase("helloworld")` 应该返回 `"HeLLoWoRLD"`。

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` 应该返回 `"HeLLoWoRLD"`。

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` 应该返回 `"_HeLLo_WoRLD_"`。

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` 应该返回 `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`。

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
