---
id: 694596b0585c11170ac7c7fd
title: "挑战 165：类平均值"
challengeType: 28
dashedName: challenge-165
---

# --description--

给定一个考试成绩（数字）数组，按照以下表格以字母等级的形式返回平均成绩：

| Average Score | Letter Grade |
| - | - |
| 97-100 | `"A+"` |
| 93-96 | `"A"` |
| 90-92 | `"A-"` |
| 87-89 | `"B+"` |
| 83-86 | `"B"` |
| 80-82 | `"B-"` |
| 77-79 | `"C+"` |
| 73–76 | `"C"` |
| 70-72 | `"C-"` |
| 67-69 | `"D+"` |
| 63-66 | `"D"` |
| 60–62 | `"D-"` |
| below 60 | `"F"` |

通过将数组中所有分数相加并除以分数总数来计算平均值。

# --hints--

`getAverageGrade([92, 91, 90, 94, 89, 93])` 应该返回 `"A-"`。

```js
assert.equal(getAverageGrade([92, 91, 90, 94, 89, 93]), "A-");
```

`getAverageGrade([84, 89, 85, 100, 91, 88, 79])` 应该返回 `"B+"`。

```js
assert.equal(getAverageGrade([84, 89, 85, 100, 91, 88, 79]), "B+");
```

`getAverageGrade([63, 69, 65, 66, 71, 64, 65])` 应该返回 `"D"`。

```js
assert.equal(getAverageGrade([63, 69, 65, 66, 71, 64, 65]), "D");
```

`getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100])` 应该返回 `"A+"`。

```js
assert.equal(getAverageGrade([97, 98, 99, 100, 96, 97, 98, 99, 100]), "A+");
```

`getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60])` 应该返回 `"C+"`。

```js
assert.equal(getAverageGrade([75, 100, 88, 79, 80, 78, 64, 60]), "C+");
```

`getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59])` 应该返回 `"F"`。

```js
assert.equal(getAverageGrade([45, 48, 50, 52, 100, 54, 56, 58, 59]), "F");
```

# --seed--

## --seed-contents--

```js
function getAverageGrade(scores) {

  return scores;
}
```

# --solutions--

```js
function getAverageGrade(scores) {
  const avg =
    scores.reduce((sum, score) => sum + score, 0) / scores.length;

  if (avg >= 97) return "A+";
  if (avg >= 93) return "A";
  if (avg >= 90) return "A-";

  if (avg >= 87) return "B+";
  if (avg >= 83) return "B";
  if (avg >= 80) return "B-";

  if (avg >= 77) return "C+";
  if (avg >= 73) return "C";
  if (avg >= 70) return "C-";

  if (avg >= 67) return "D+";
  if (avg >= 63) return "D";
  if (avg >= 60) return "D-";

  return "F";
}
```
