---
id: 694596b0585c11170ac7c7f9
title: "挑战 161：免运费"
challengeType: 28
dashedName: challenge-161
---

# --description--

给定一个表现你购物车中项的字串数组和一个表示达到免运费资格的最低订单金额的数字，确定你购物车中的项是否符合免运费资格。

给定的数组将包含以下列表中的项：

| Item | Price |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`getsFreeShipping(["shoes"], 50)` 应该返回 `true`。

```js
assert.isTrue(getsFreeShipping(["shoes"], 50));
```

`getsFreeShipping(["hat", "socks"], 50)` 应该返回 `false`。

```js
assert.isFalse(getsFreeShipping(["hat", "socks"], 50));
```

`getsFreeShipping(["jeans", "shirt", "jacket"], 75)` 应该返回 `true`。

```js
assert.isTrue(getsFreeShipping(["jeans", "shirt", "jacket"], 75));
```

`getsFreeShipping(["socks", "socks", "hat"], 75)` 应该返回 `false`。

```js
assert.isFalse(getsFreeShipping(["socks", "socks", "hat"], 75));
```

`getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100)` 应该返回 `true`。

```js
assert.isTrue(getsFreeShipping(["shirt", "shirt", "jeans", "socks"], 100));
```

`getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` 应该返回 `false`。

```js
assert.isFalse(getsFreeShipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200));
```

# --seed--

## --seed-contents--

```js
function getsFreeShipping(cart, minimum) {

  return cart;
}
```

# --solutions--

```js
function getsFreeShipping(cart, minimum) {
  const prices = {
    shirt: 34.25,
    jeans: 48.50,
    shoes: 75.00,
    hat: 19.95,
    socks: 15.00,
    jacket: 109.95
  };

  let total = 0;

  for (const item of cart) {
    total += prices[item];
  }

  return total >= minimum;
}
```
