---
id: 696655d24b614176d4c9b78a
title: "挑战 167：宾果！字母"
challengeType: 28
dashedName: challenge-167
---

# --description--

给定一个数字，返回与之关联的宾果字母（大写）。宾果数字分组如下：

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)` 应该返回 `"O"`。

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)` 应该返回 `"G"`。

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)` 应该返回 `"I"`。

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)` 应该返回 `"N"`。

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)` 应该返回 `"B"`。

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
