---
id: 696655d24b614176d4c9b78d
title: "挑战 170：奇数天还是偶数天"
challengeType: 28
dashedName: challenge-170
---

# --description--

给定一个时间戳（自 Unix 纪元以来的毫秒数），返回：

- 如果该时间戳的月份日期为“odd”。
- 如果该时间戳的日期是偶数，则为 `"even"`。

例如，给定 `1769472000000`，这是 2026 年 1 月 27 日的时间戳，返回 `"odd"`，因为日期数字（`27`）是奇数。

注意：时间戳以毫秒为单位，你应该使用 UTC 时区的日期，而不是你的局部的时间。

# --hints--

`oddOrEvenDay(1769472000000)` 应该返回 `"odd"`。

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` 应该返回 `"even"`。

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` 应该返回 `"odd"`。

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` 应该返回 `"odd"`。

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` 应该返回 `"even"`。

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
