---
id: 696655d24b614176d4c9b78b
title: "挑战 168：缩放图像"
challengeType: 28
dashedName: challenge-168
---

# --description--

给定一个表示图像宽度和高度的字串，以及一个用于缩放图像的数字，返回缩放后的宽度和高度。

- 输入字串的形式为 `"WxH"`。例如，`"800x600"`。
- 比例是用于乘以宽度和高度的数字。

以相同的 `"WxH"` 形式返回缩放后的维度。

# --hints--

`scaleImage("800x600", 2)` 应该返回 `"1600x1200"`。

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` 应该返回 `"1000x1000"`。

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` 应该返回 `"512x384"`。

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` 应该返回 `"450x300"`。

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
