---
id: 696655d24b614176d4c9b78c
title: "挑战 169：FizzBuzz 迷你"
challengeType: 28
dashedName: challenge-169
---

# --description--

给定一个整数，返回一个基于以下规则的字串：

- 如果数字能被 3 整除，返回 `"Fizz"`。
- 如果数字能被 5 整除，返回 `"Buzz"`。
- 如果数字能被 3 和 5 同时整除，返回 `"FizzBuzz"`。
- 否则，返回给定的数字作为字串。

# --hints--

`fizzBuzzMini(3)` 应该返回 `"Fizz"`。

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` 应该返回 `"4"`。

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` 应该返回 `"Buzz"`。

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` 应该返回 `"FizzBuzz"`。

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` 应该返回 `"98"`。

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
