---
id: 696655d24b614176d4c9b789
title: "挑战 166：Hex 验证器"
challengeType: 28
dashedName: challenge-166
---

# --description--

给定一个字串，判断它是否是有效的 CSS 十六进制颜色。有效的 CSS 十六进制颜色必须：

- 以 `#` 开头，且
- 后面跟随 3 个或 6 个十六进制字符。

十六进制字符是数字 `0` 到 `9` 和字母 `a` 到 `f`（不区分分支）。

# --hints--

`isValidHex("#123")` 应该返回 `true`。

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` 应该返回 `true`。

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` 应该返回 `true`。

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` 应该返回 `true`。

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` 应该返回 `false`。

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` 应该返回 `false`。

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` 应该返回 `false`。

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` 应该返回 `false`。

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
