---
id: 69738771fb5a7b8b24cca2a2
title: "挑战 176：土拨鼠日"
challengeType: 28
dashedName: challenge-176
---

# --description--

今天是土拨鼠日，土拨鼠根据是否看到它的影子来预测天气。

给定一个表现土拨鼠外观的值，返回正确的预测：

- 如果给定值是布尔值 `true`（土拨鼠看到了它的影子），则返回 `"Looks like we'll have six more weeks of winter."`。
- 如果值是布尔类型的 `false`（土拨鼠没有看到它的影子），则返回 `"It's going to be an early spring."`。
- 如果值是其他任何内容（旱獭没有出现），返回 `"No prediction this year."`。

# --hints--

`groundhogDayPrediction(true)` 应该返回 `"Looks like we'll have six more weeks of winter."`。

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` 应该返回 `"It's going to be an early spring."`。

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` 应返回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` 应返回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` 应返回 `"No prediction this year."`。

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
