---
id: 69738771fb5a7b8b24cca2a5
title: "挑战 179：零钱"
challengeType: 28
dashedName: challenge-179
---

# --description--

给定一个表现你口袋中硬币的整数数组，每个整数表示硬币的分值，返回以 `"$D.CC"` 形式表示的总金额。

- 100 分相等 1 美元。
- 在返回值中，对于小于一美元的金额包含前导零，并且始终为分部分保留两位数字。

# --hints--

`countChange([25, 10, 5, 1])` 应该返回 `"$0.41"`。

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` 应该返回 `"$1.43"`。

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` 应该返回 `"$37.55"`。

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` 应该返回 `"$0.70"`。

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` 应该返回 `"$0.01"`。

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` 应该返回 `"$1.00"`。

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
