---
id: 69738771fb5a7b8b24cca2a3
title: "挑战 177：字串 镜射"
challengeType: 28
dashedName: challenge-177
---

# --description--

给定一个字串，返回一个新建的字串，该字串由给定字串和其反转拷贝追加到末尾组成。

# --hints--

`mirror("freeCodeCamp")` 应该返回 `"freeCodeCamppmaCedoCeerf"`。

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` 应该返回 `"RaceCarraCecaR"`。

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` 应该返回 `"helloworlddlrowolleh"`。

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` 应该返回 `"The quick brown fox......xof nworb kciuq ehT"`。

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
