---
id: 69738771fb5a7b8b24cca29d
title: "挑战 171：扁平化数组"
challengeType: 28
dashedName: challenge-171
---

# --description--

给定一个包含嵌套数组的数组，返回一个新建数组，将所有值展平成一个单一的一维数组。保持数组中项的原始顺序。

# --hints--

`flatten([1, [2, 3], 4])` 应该返回 `[1, 2, 3, 4]`。

```js
assert.deepEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4]);
```

`flatten([5, [4, [3, 2]], 1])` 应该返回 `[5, 4, 3, 2, 1]`。

```js
assert.deepEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1]);
```

`flatten(["A", [[[["B"]]]], "C"])` 应该返回 `["A", "B", "C"]`。

```js
assert.deepEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"]);
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` 应该返回 `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`。

```js
assert.deepEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]);
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` 应该返回 `["red","blue","green","yellow","purple","orange","pink","brown"]`。

```js
assert.deepEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"]);
```

# --seed--

## --seed-contents--

```js
function flatten(arr) {

  return arr;
}
```

# --solutions--

```js
function flatten(arr) {
  let result = [];

  for (const item of arr) {
    if (Array.isArray(item)) {
      result = result.concat(flatten(item));
    } else {
      result.push(item);
    }
  }

  return result;
}
```
