---
id: 69738771fb5a7b8b24cca2a4
title: "挑战 178：截断 `text`"
challengeType: 28
dashedName: challenge-178
---

# --description--

给定一个字串，如果它的长度不超过 20 个字符，则原样返回。如果它超过 20 个字符，则截断为前 17 个字符，并在末尾追加 `"..."`（总共 20 个字符），然后返回结果。

# --hints--

`truncateText("Hello, world!")` 应该返回 `"Hello, world!"`。

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` 应该返回 `"This string shoul..."`。

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` 应该返回 `"Exactly twenty chars"`。

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` 应返回 `"...................."`。

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
