---
id: 697a49e9860d24853adef67e
title: "挑战 193：2026 冬季游戏第 14 天：滑雪登山"
challengeType: 28
dashedName: challenge-193
---

# --description--

给定山上的雪深和坡度，确定是否存在雪崩风险。

- 雪深值为 `"Shallow"`、`"Moderate"` 或 `"Deep"`。
- 坡度值为 `"Gentle"`、`"Steep"` 或 `"Very Steep"`。

根据此表格返回 `"Safe"` 或 `"Risky"`：

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` 应该返回 `"Safe"`。

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` 应该返回 `"Safe"`。

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` 应该返回 `"Safe"`。

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` 应该返回 `"Safe"`。

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` 应该返回 `"Risky"`。

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` 应该返回 `"Risky"`。

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` 应该返回 `"Safe"`。

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` 应该返回 `"Risky"`。

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` 应该返回 `"Risky"`。

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
