---
id: 697a49e6ff50d756c9b6935d
title: "æŒ‘æˆ˜ 180ï¼š2026 å†¬å­£æ¸¸æˆç¬¬ 1 å¤©ï¼šå¼€å¹•æ—¥"
challengeType: 28
dashedName: challenge-180
---

# --description--

ä»Šå¤©æ ‡å¿—ç€ 2026 å†¬å­£æ¸¸æˆçš„å¼€å§‹ã€‚æŽ¥ä¸‹æ¥çš„ 17 å¤©å°†ä¸ºä½ å¸¦æ¥å—å…¶å¯å‘çš„ç¼–ç æŒ‘æˆ˜ã€‚

å¯¹äºŽç¬¬ä¸€ä¸ªï¼Œä½ ä¼šå¾—åˆ°ä¸€ä¸ªä¸¤å­—æ¯çš„å›½å®¶ä»£ç ï¼Œéœ€è¦è¿”å›žè¯¥å›½å®¶çš„æ ‡è®°è¡¨æƒ…ç¬¦å·ã€‚

ä½¿ç”¨æ­¤åˆ—è¡¨ï¼š

| Country | Code | Flag |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"ðŸ‡¦ðŸ‡±"` |
| Andorra                    |            `"AD"` | `"ðŸ‡¦ðŸ‡©"` |
| Argentina                  |            `"AR"` | `"ðŸ‡¦ðŸ‡·"` |
| Armenia                    |            `"AM"` | `"ðŸ‡¦ðŸ‡²"` |
| Australia                  |            `"AU"` | `"ðŸ‡¦ðŸ‡º"` |
| Austria                    |            `"AT"` | `"ðŸ‡¦ðŸ‡¹"` |
| Azerbaijan                 |            `"AZ"` | `"ðŸ‡¦ðŸ‡¿"` |
| Belgium                    |            `"BE"` | `"ðŸ‡§ðŸ‡ª"` |
| Benin                      |            `"BJ"` | `"ðŸ‡§ðŸ‡¯"` |
| Bolivia                    |            `"BO"` | `"ðŸ‡§ðŸ‡´"` |
| Bosnia and Herzegovina     |            `"BA"` | `"ðŸ‡§ðŸ‡¦"` |
| Brazil                     |            `"BR"` | `"ðŸ‡§ðŸ‡·"` |
| Bulgaria                   |            `"BG"` | `"ðŸ‡§ðŸ‡¬"` |
| Canada                     |            `"CA"` | `"ðŸ‡¨ðŸ‡¦"` |
| Chile                      |            `"CL"` | `"ðŸ‡¨ðŸ‡±"` |
| China                      |            `"CN"` | `"ðŸ‡¨ðŸ‡³"` |
| Colombia                   |            `"CO"` | `"ðŸ‡¨ðŸ‡´"` |
| Croatia                    |            `"HR"` | `"ðŸ‡­ðŸ‡·"` |
| Cyprus                     |            `"CY"` | `"ðŸ‡¨ðŸ‡¾"` |
| Czech Republic             |            `"CZ"` | `"ðŸ‡¨ðŸ‡¿"` |
| Denmark                    |            `"DK"` | `"ðŸ‡©ðŸ‡°"` |
| Ecuador                    |            `"EC"` | `"ðŸ‡ªðŸ‡¨"` |
| Eritrea                    |            `"ER"` | `"ðŸ‡ªðŸ‡·"` |
| Estonia                    |            `"EE"` | `"ðŸ‡ªðŸ‡ª"` |
| Finland                    |            `"FI"` | `"ðŸ‡«ðŸ‡®"` |
| France                     |            `"FR"` | `"ðŸ‡«ðŸ‡·"` |
| Georgia                    |            `"GE"` | `"ðŸ‡¬ðŸ‡ª"` |
| Germany                    |            `"DE"` | `"ðŸ‡©ðŸ‡ª"` |
| Great Britain              |            `"GB"` | `"ðŸ‡¬ðŸ‡§"` |
| Greece                     |            `"GR"` | `"ðŸ‡¬ðŸ‡·"` |
| Guinea-Bissau              |            `"GW"` | `"ðŸ‡¬ðŸ‡¼"` |
| Haiti                      |            `"HT"` | `"ðŸ‡­ðŸ‡¹"` |
| Hong Kong                  |            `"HK"` | `"ðŸ‡­ðŸ‡°"` |
| Hungary                    |            `"HU"` | `"ðŸ‡­ðŸ‡º"` |
| Iceland                    |            `"IS"` | `"ðŸ‡®ðŸ‡¸"` |
| India                      |            `"IN"` | `"ðŸ‡®ðŸ‡³"` |
| Iran                       |            `"IR"` | `"ðŸ‡®ðŸ‡·"` |
| Ireland                    |            `"IE"` | `"ðŸ‡®ðŸ‡ª"` |
| Israel                     |            `"IL"` | `"ðŸ‡®ðŸ‡±"` |
| Italy                      |            `"IT"` | `"ðŸ‡®ðŸ‡¹"` |
| Jamaica                    |            `"JM"` | `"ðŸ‡¯ðŸ‡²"` |
| Japan                      |            `"JP"` | `"ðŸ‡¯ðŸ‡µ"` |
| Kazakhstan                 |            `"KZ"` | `"ðŸ‡°ðŸ‡¿"` |
| Kenya                      |            `"KE"` | `"ðŸ‡°ðŸ‡ª"` |
| Kosovo                     |            `"XK"` | `"ðŸ‡½ðŸ‡°"` |
| Kyrgyzstan                 |            `"KG"` | `"ðŸ‡°ðŸ‡¬"` |
| Latvia                     |            `"LV"` | `"ðŸ‡±ðŸ‡»"` |
| Lebanon                    |            `"LB"` | `"ðŸ‡±ðŸ‡§"` |
| Liechtenstein              |            `"LI"` | `"ðŸ‡±ðŸ‡®"` |
| Lithuania                  |            `"LT"` | `"ðŸ‡±ðŸ‡¹"` |
| Luxembourg                 |            `"LU"` | `"ðŸ‡±ðŸ‡º"` |
| Madagascar                 |            `"MG"` | `"ðŸ‡²ðŸ‡¬"` |
| Malaysia                   |            `"MY"` | `"ðŸ‡²ðŸ‡¾"` |
| Malta                      |            `"MT"` | `"ðŸ‡²ðŸ‡¹"` |
| Mexico                     |            `"MX"` | `"ðŸ‡²ðŸ‡½"` |
| Moldova                    |            `"MD"` | `"ðŸ‡²ðŸ‡©"` |
| Monaco                     |            `"MC"` | `"ðŸ‡²ðŸ‡¨"` |
| Mongolia                   |            `"MN"` | `"ðŸ‡²ðŸ‡³"` |
| Montenegro                 |            `"ME"` | `"ðŸ‡²ðŸ‡ª"` |
| Morocco                    |            `"MA"` | `"ðŸ‡²ðŸ‡¦"` |
| Netherlands                |            `"NL"` | `"ðŸ‡³ðŸ‡±"` |
| New Zealand                |            `"NZ"` | `"ðŸ‡³ðŸ‡¿"` |
| Nigeria                    |            `"NG"` | `"ðŸ‡³ðŸ‡¬"` |
| North Macedonia            |            `"MK"` | `"ðŸ‡²ðŸ‡°"` |
| Norway                     |            `"NO"` | `"ðŸ‡³ðŸ‡´"` |
| Pakistan                   |            `"PK"` | `"ðŸ‡µðŸ‡°"` |
| Philippines                |            `"PH"` | `"ðŸ‡µðŸ‡­"` |
| Poland                     |            `"PL"` | `"ðŸ‡µðŸ‡±"` |
| Portugal                   |            `"PT"` | `"ðŸ‡µðŸ‡¹"` |
| Puerto Rico                |            `"PR"` | `"ðŸ‡µðŸ‡·"` |
| Romania                    |            `"RO"` | `"ðŸ‡·ðŸ‡´"` |
| San Marino                 |            `"SM"` | `"ðŸ‡¸ðŸ‡²"` |
| Saudi Arabia               |            `"SA"` | `"ðŸ‡¸ðŸ‡¦"` |
| Serbia                     |            `"RS"` | `"ðŸ‡·ðŸ‡¸"` |
| Singapore                  |            `"SG"` | `"ðŸ‡¸ðŸ‡¬"` |
| Slovakia                   |            `"SK"` | `"ðŸ‡¸ðŸ‡°"` |
| Slovenia                   |            `"SI"` | `"ðŸ‡¸ðŸ‡®"` |
| South Africa               |            `"ZA"` | `"ðŸ‡¿ðŸ‡¦"` |
| South Korea                |            `"KR"` | `"ðŸ‡°ðŸ‡·"` |
| Spain                      |            `"ES"` | `"ðŸ‡ªðŸ‡¸"` |
| Sweden                     |            `"SE"` | `"ðŸ‡¸ðŸ‡ª"` |
| Switzerland                |            `"CH"` | `"ðŸ‡¨ðŸ‡­"` |
| Thailand                   |            `"TH"` | `"ðŸ‡¹ðŸ‡­"` |
| Trinidad & Tobago          |            `"TT"` | `"ðŸ‡¹ðŸ‡¹"` |
| Turkey                     |            `"TR"` | `"ðŸ‡¹ðŸ‡·"` |
| Ukraine                    |            `"UA"` | `"ðŸ‡ºðŸ‡¦"` |
| United Arab Emirates       |            `"AE"` | `"ðŸ‡¦ðŸ‡ª"` |
| United States              |            `"US"` | `"ðŸ‡ºðŸ‡¸"` |
| Uruguay                    |            `"UY"` | `"ðŸ‡ºðŸ‡¾"` |
| Uzbekistan                 |            `"UZ"` | `"ðŸ‡ºðŸ‡¿"` |
| Venezuela                  |            `"VE"` | `"ðŸ‡»ðŸ‡ª"` |

# --hints--

`getFlag("AL")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("AL"), "ðŸ‡¦ðŸ‡±");
```

`getFlag("AD")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡©"`ã€‚

```js
assert.equal(getFlag("AD"), "ðŸ‡¦ðŸ‡©");
```

`getFlag("AR")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("AR"), "ðŸ‡¦ðŸ‡·")
```

`getFlag("AM")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("AM"), "ðŸ‡¦ðŸ‡²")
```

`getFlag("AU")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("AU"), "ðŸ‡¦ðŸ‡º")
```

`getFlag("AT")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("AT"), "ðŸ‡¦ðŸ‡¹")
```

`getFlag("AZ")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("AZ"), "ðŸ‡¦ðŸ‡¿")
```

`getFlag("BE")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("BE"), "ðŸ‡§ðŸ‡ª")
```

`getFlag("BJ")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡¯"`ã€‚

```js
assert.equal(getFlag("BJ"), "ðŸ‡§ðŸ‡¯")
```

`getFlag("BO")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("BO"), "ðŸ‡§ðŸ‡´")
```

`getFlag("BA")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("BA"), "ðŸ‡§ðŸ‡¦")
```

`getFlag("BR")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("BR"), "ðŸ‡§ðŸ‡·")
```

`getFlag("BG")` åº”è¯¥è¿”å›ž `"ðŸ‡§ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("BG"), "ðŸ‡§ðŸ‡¬")
```

`getFlag("CA")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("CA"), "ðŸ‡¨ðŸ‡¦")
```

`getFlag("CL")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("CL"), "ðŸ‡¨ðŸ‡±")
```

`getFlag("CN")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("CN"), "ðŸ‡¨ðŸ‡³")
```

`getFlag("CO")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("CO"), "ðŸ‡¨ðŸ‡´")
```

`getFlag("HR")` åº”è¯¥è¿”å›ž `"ðŸ‡­ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("HR"), "ðŸ‡­ðŸ‡·")
```

`getFlag("CY")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("CY"), "ðŸ‡¨ðŸ‡¾")
```

`getFlag("CZ")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("CZ"), "ðŸ‡¨ðŸ‡¿")
```

`getFlag("DK")` åº”è¯¥è¿”å›ž `"ðŸ‡©ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("DK"), "ðŸ‡©ðŸ‡°")
```

`getFlag("EC")` åº”è¯¥è¿”å›ž `"ðŸ‡ªðŸ‡¨"`ã€‚

```js
assert.equal(getFlag("EC"), "ðŸ‡ªðŸ‡¨")
```

`getFlag("ER")` åº”è¯¥è¿”å›ž `"ðŸ‡ªðŸ‡·"`ã€‚

```js
assert.equal(getFlag("ER"), "ðŸ‡ªðŸ‡·")
```

`getFlag("EE")` åº”è¯¥è¿”å›ž `"ðŸ‡ªðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("EE"), "ðŸ‡ªðŸ‡ª")
```

`getFlag("FI")` åº”è¯¥è¿”å›ž `"ðŸ‡«ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("FI"), "ðŸ‡«ðŸ‡®")
```

`getFlag("FR")` åº”è¯¥è¿”å›ž `"ðŸ‡«ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("FR"), "ðŸ‡«ðŸ‡·")
```

`getFlag("GE")` åº”è¯¥è¿”å›ž `"ðŸ‡¬ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("GE"), "ðŸ‡¬ðŸ‡ª")
```

`getFlag("DE")` åº”è¯¥è¿”å›ž `"ðŸ‡©ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("DE"), "ðŸ‡©ðŸ‡ª")
```

`getFlag("GB")` åº”è¯¥è¿”å›ž `"ðŸ‡¬ðŸ‡§"`ã€‚

```js
assert.equal(getFlag("GB"), "ðŸ‡¬ðŸ‡§")
```

`getFlag("GR")` åº”è¯¥è¿”å›ž `"ðŸ‡¬ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("GR"), "ðŸ‡¬ðŸ‡·")
```

`getFlag("GW")` åº”è¯¥è¿”å›ž `"ðŸ‡¬ðŸ‡¼"`ã€‚

```js
assert.equal(getFlag("GW"), "ðŸ‡¬ðŸ‡¼")
```

`getFlag("HT")` åº”è¯¥è¿”å›ž `"ðŸ‡­ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("HT"), "ðŸ‡­ðŸ‡¹")
```

`getFlag("HK")` åº”è¯¥è¿”å›ž `"ðŸ‡­ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("HK"), "ðŸ‡­ðŸ‡°")
```

`getFlag("HU")` åº”è¯¥è¿”å›ž `"ðŸ‡­ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("HU"), "ðŸ‡­ðŸ‡º")
```

`getFlag("IS")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("IS"), "ðŸ‡®ðŸ‡¸")
```

`getFlag("IN")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("IN"), "ðŸ‡®ðŸ‡³")
```

`getFlag("IR")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("IR"), "ðŸ‡®ðŸ‡·")
```

`getFlag("IE")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("IE"), "ðŸ‡®ðŸ‡ª")
```

`getFlag("IL")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("IL"), "ðŸ‡®ðŸ‡±")
```

`getFlag("IT")` åº”è¯¥è¿”å›ž `"ðŸ‡®ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("IT"), "ðŸ‡®ðŸ‡¹")
```

`getFlag("JM")` åº”è¯¥è¿”å›ž `"ðŸ‡¯ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("JM"), "ðŸ‡¯ðŸ‡²")
```

`getFlag("JP")` åº”è¯¥è¿”å›ž `"ðŸ‡¯ðŸ‡µ"`ã€‚

```js
assert.equal(getFlag("JP"), "ðŸ‡¯ðŸ‡µ")
```

`getFlag("KZ")` åº”è¯¥è¿”å›ž `"ðŸ‡°ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("KZ"), "ðŸ‡°ðŸ‡¿")
```

`getFlag("KE")` åº”è¯¥è¿”å›ž `"ðŸ‡°ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("KE"), "ðŸ‡°ðŸ‡ª")
```

`getFlag("XK")` åº”è¯¥è¿”å›ž `"ðŸ‡½ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("XK"), "ðŸ‡½ðŸ‡°")
```

`getFlag("KG")` åº”è¯¥è¿”å›ž `"ðŸ‡°ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("KG"), "ðŸ‡°ðŸ‡¬")
```

`getFlag("LV")` åº”è¯¥è¿”å›ž `"ðŸ‡±ðŸ‡»"`ã€‚

```js
assert.equal(getFlag("LV"), "ðŸ‡±ðŸ‡»")
```

`getFlag("LB")` åº”è¯¥è¿”å›ž `"ðŸ‡±ðŸ‡§"`ã€‚

```js
assert.equal(getFlag("LB"), "ðŸ‡±ðŸ‡§")
```

`getFlag("LI")` åº”è¯¥è¿”å›ž `"ðŸ‡±ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("LI"), "ðŸ‡±ðŸ‡®")
```

`getFlag("LT")` åº”è¯¥è¿”å›ž `"ðŸ‡±ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("LT"), "ðŸ‡±ðŸ‡¹")
```

`getFlag("LU")` åº”è¯¥è¿”å›ž `"ðŸ‡±ðŸ‡º"`ã€‚

```js
assert.equal(getFlag("LU"), "ðŸ‡±ðŸ‡º")
```

`getFlag("MG")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("MG"), "ðŸ‡²ðŸ‡¬")
```

`getFlag("MY")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("MY"), "ðŸ‡²ðŸ‡¾")
```

`getFlag("MT")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("MT"), "ðŸ‡²ðŸ‡¹")
```

`getFlag("MX")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡½"`ã€‚

```js
assert.equal(getFlag("MX"), "ðŸ‡²ðŸ‡½")
```

`getFlag("MD")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡©"`ã€‚

```js
assert.equal(getFlag("MD"), "ðŸ‡²ðŸ‡©")
```

`getFlag("MC")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡¨"`ã€‚

```js
assert.equal(getFlag("MC"), "ðŸ‡²ðŸ‡¨")
```

`getFlag("MN")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡³"`ã€‚

```js
assert.equal(getFlag("MN"), "ðŸ‡²ðŸ‡³")
```

`getFlag("ME")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("ME"), "ðŸ‡²ðŸ‡ª")
```

`getFlag("MA")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("MA"), "ðŸ‡²ðŸ‡¦")
```

`getFlag("NL")` åº”è¯¥è¿”å›ž `"ðŸ‡³ðŸ‡±"`ã€‚

```js
assert.equal(getFlag("NL"), "ðŸ‡³ðŸ‡±")
```

`getFlag("NZ")` åº”è¯¥è¿”å›ž `"ðŸ‡³ðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("NZ"), "ðŸ‡³ðŸ‡¿")
```

`getFlag("NG")` åº”è¯¥è¿”å›ž `"ðŸ‡³ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("NG"), "ðŸ‡³ðŸ‡¬")
```

`getFlag("MK")` åº”è¯¥è¿”å›ž `"ðŸ‡²ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("MK"), "ðŸ‡²ðŸ‡°")
```

`getFlag("NO")` åº”è¯¥è¿”å›ž `"ðŸ‡³ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("NO"), "ðŸ‡³ðŸ‡´")
```

`getFlag("PK")` åº”è¯¥è¿”å›ž `"ðŸ‡µðŸ‡°"`ã€‚

```js
assert.equal(getFlag("PK"), "ðŸ‡µðŸ‡°")
```

`getFlag("PH")` åº”è¯¥è¿”å›ž `"ðŸ‡µðŸ‡­"`ã€‚

```js
assert.equal(getFlag("PH"), "ðŸ‡µðŸ‡­")
```

`getFlag("PL")` åº”è¯¥è¿”å›ž `"ðŸ‡µðŸ‡±"`ã€‚

```js
assert.equal(getFlag("PL"), "ðŸ‡µðŸ‡±")
```

`getFlag("PT")` åº”è¯¥è¿”å›ž `"ðŸ‡µðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("PT"), "ðŸ‡µðŸ‡¹")
```

`getFlag("PR")` åº”è¯¥è¿”å›ž `"ðŸ‡µðŸ‡·"`ã€‚

```js
assert.equal(getFlag("PR"), "ðŸ‡µðŸ‡·")
```

`getFlag("RO")` åº”è¯¥è¿”å›ž `"ðŸ‡·ðŸ‡´"`ã€‚

```js
assert.equal(getFlag("RO"), "ðŸ‡·ðŸ‡´")
```

`getFlag("SM")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡²"`ã€‚

```js
assert.equal(getFlag("SM"), "ðŸ‡¸ðŸ‡²")
```

`getFlag("SA")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("SA"), "ðŸ‡¸ðŸ‡¦")
```

`getFlag("RS")` åº”è¯¥è¿”å›ž `"ðŸ‡·ðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("RS"), "ðŸ‡·ðŸ‡¸")
```

`getFlag("SG")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡¬"`ã€‚

```js
assert.equal(getFlag("SG"), "ðŸ‡¸ðŸ‡¬")
```

`getFlag("SK")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡°"`ã€‚

```js
assert.equal(getFlag("SK"), "ðŸ‡¸ðŸ‡°")
```

`getFlag("SI")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡®"`ã€‚

```js
assert.equal(getFlag("SI"), "ðŸ‡¸ðŸ‡®")
```

`getFlag("ZA")` åº”è¯¥è¿”å›ž `"ðŸ‡¿ðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("ZA"), "ðŸ‡¿ðŸ‡¦")
```

`getFlag("KR")` åº”è¯¥è¿”å›ž `"ðŸ‡°ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("KR"), "ðŸ‡°ðŸ‡·")
```

`getFlag("ES")` åº”è¯¥è¿”å›ž `"ðŸ‡ªðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("ES"), "ðŸ‡ªðŸ‡¸")
```

`getFlag("SE")` åº”è¯¥è¿”å›ž `"ðŸ‡¸ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("SE"), "ðŸ‡¸ðŸ‡ª")
```

`getFlag("CH")` åº”è¯¥è¿”å›ž `"ðŸ‡¨ðŸ‡­"`ã€‚

```js
assert.equal(getFlag("CH"), "ðŸ‡¨ðŸ‡­")
```

`getFlag("TH")` åº”è¯¥è¿”å›ž `"ðŸ‡¹ðŸ‡­"`ã€‚

```js
assert.equal(getFlag("TH"), "ðŸ‡¹ðŸ‡­")
```

`getFlag("TT")` åº”è¯¥è¿”å›ž `"ðŸ‡¹ðŸ‡¹"`ã€‚

```js
assert.equal(getFlag("TT"), "ðŸ‡¹ðŸ‡¹")
```

`getFlag("TR")` åº”è¯¥è¿”å›ž `"ðŸ‡¹ðŸ‡·"`ã€‚

```js
assert.equal(getFlag("TR"), "ðŸ‡¹ðŸ‡·")
```

`getFlag("UA")` åº”è¯¥è¿”å›ž `"ðŸ‡ºðŸ‡¦"`ã€‚

```js
assert.equal(getFlag("UA"), "ðŸ‡ºðŸ‡¦")
```

`getFlag("AE")` åº”è¯¥è¿”å›ž `"ðŸ‡¦ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("AE"), "ðŸ‡¦ðŸ‡ª")
```

`getFlag("US")` åº”è¯¥è¿”å›ž `"ðŸ‡ºðŸ‡¸"`ã€‚

```js
assert.equal(getFlag("US"), "ðŸ‡ºðŸ‡¸")
```

`getFlag("UY")` åº”è¯¥è¿”å›ž `"ðŸ‡ºðŸ‡¾"`ã€‚

```js
assert.equal(getFlag("UY"), "ðŸ‡ºðŸ‡¾")
```

`getFlag("UZ")` åº”è¯¥è¿”å›ž `"ðŸ‡ºðŸ‡¿"`ã€‚

```js
assert.equal(getFlag("UZ"), "ðŸ‡ºðŸ‡¿")
```

`getFlag("VE")` åº”è¯¥è¿”å›ž `"ðŸ‡»ðŸ‡ª"`ã€‚

```js
assert.equal(getFlag("VE"), "ðŸ‡»ðŸ‡ª")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
