---
id: 697a49e6ff50d756c9b69360
title: "挑战 183：2026 冬季游戏 第 4 天：跳台滑雪"
challengeType: 28
dashedName: challenge-183
---

# --description--

给定距离点、风力补偿值和 K 点奖励值，计算你的跳台滑雪得分并确定你是否获得奖牌。

- 你的得分是通过将上述四个值相加计算得出的。

其他跳跃者当前的总分是：

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- 如果你的分数是最高的，返回 `"Gold"`
- 如果是第二名，返回 `"Silver"`
- 如果是第三名，返回 `"Bronze"`
- 否则，返回 `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` 应该返回 `"Gold"`。

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` 应该返回 `"Bronze"`。

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` 应返回 `"Silver"`。

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` 应该返回 `"No Medal"`。

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` 应该返回 `"Gold"`。

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` 应该返回 `"No Medal"`。

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
