---
id: 697a49e6ff50d756c9b6935d
title: "挑战 180：冬季运动会第 1 天：开幕日"
challengeType: 28
dashedName: challenge-180
---

# --description--

接下来的 17 天将为你带来受冬季运动会启发的编程挑战。

对于第一个，你会得到一个两字母的国家代码，需要返回该国家的标记表情符号。

使用此列表：

| 国家 | 代码 | 国旗 |
| - | -: | :--: |
| 阿尔巴尼亚                    |            `"AL"` | `"🇦🇱"` |
| 安道尔                    |            `"AD"` | `"🇦🇩"` |
| 阿根廷                  |            `"AR"` | `"🇦🇷"` |
| 亚美尼亚                    |            `"AM"` | `"🇦🇲"` |
| 澳大利亚                  |            `"AU"` | `"🇦🇺"` |
| 奥地利                    |            `"AT"` | `"🇦🇹"` |
| 阿塞拜疆                 |            `"AZ"` | `"🇦🇿"` |
| 比利时                    |            `"BE"` | `"🇧🇪"` |
| 贝宁                      |            `"BJ"` | `"🇧🇯"` |
| 玻利维亚                    |            `"BO"` | `"🇧🇴"` |
| 波斯尼亚和黑塞哥维那     |            `"BA"` | `"🇧🇦"` |
| 巴西                     |            `"BR"` | `"🇧🇷"` |
| 保加利亚                   |            `"BG"` | `"🇧🇬"` |
| 加拿大                     |            `"CA"` | `"🇨🇦"` |
| 智利                      |            `"CL"` | `"🇨🇱"` |
| 中国                      |            `"CN"` | `"🇨🇳"` |
| 哥伦比亚                   |            `"CO"` | `"🇨🇴"` |
| 克罗地亚                    |            `"HR"` | `"🇭🇷"` |
| 塞浦路斯                     |            `"CY"` | `"🇨🇾"` |
| 捷克共和国             |            `"CZ"` | `"🇨🇿"` |
| 丹麦                    |            `"DK"` | `"🇩🇰"` |
| 厄瓜多尔                    |            `"EC"` | `"🇪🇨"` |
| 厄立特里亚                    |            `"ER"` | `"🇪🇷"` |
| 爱沙尼亚                    |            `"EE"` | `"🇪🇪"` |
| 芬兰                    |            `"FI"` | `"🇫🇮"` |
| 法国                     |            `"FR"` | `"🇫🇷"` |
| 格鲁吉亚                    |            `"GE"` | `"🇬🇪"` |
| 德国                    |            `"DE"` | `"🇩🇪"` |
| 大不列颠              |            `"GB"` | `"🇬🇧"` |
| 希腊                     |            `"GR"` | `"🇬🇷"` |
| 几内亚比绍              |            `"GW"` | `"🇬🇼"` |
| 海地                      |            `"HT"` | `"🇭🇹"` |
| 香港                  |            `"HK"` | `"🇭🇰"` |
| 匈牙利                    |            `"HU"` | `"🇭🇺"` |
| 冰岛                    |            `"IS"` | `"🇮🇸"` |
| 印度                      |            `"IN"` | `"🇮🇳"` |
| 伊朗                       |            `"IR"` | `"🇮🇷"` |
| 爱尔兰                    |            `"IE"` | `"🇮🇪"` |
| 以色列                     |            `"IL"` | `"🇮🇱"` |
| 意大利                      |            `"IT"` | `"🇮🇹"` |
| 牙买加                    |            `"JM"` | `"🇯🇲"` |
| 日本                      |            `"JP"` | `"🇯🇵"` |
| 哈萨克斯坦                 |            `"KZ"` | `"🇰🇿"` |
| 肯尼亚                      |            `"KE"` | `"🇰🇪"` |
| 科索沃                     |            `"XK"` | `"🇽🇰"` |
| 吉尔吉斯斯坦                 |            `"KG"` | `"🇰🇬"` |
| 拉脱维亚                     |            `"LV"` | `"🇱🇻"` |
| 黎巴嫩                    |            `"LB"` | `"🇱🇧"` |
| 列支敦士登              |            `"LI"` | `"🇱🇮"` |
| 立陶宛                  |            `"LT"` | `"🇱🇹"` |
| 卢森堡                 |            `"LU"` | `"🇱🇺"` |
| 马达加斯加                 |            `"MG"` | `"🇲🇬"` |
| 马来西亚                   |            `"MY"` | `"🇲🇾"` |
| 马耳他                      |            `"MT"` | `"🇲🇹"` |
| 墨西哥                     |            `"MX"` | `"🇲🇽"` |
| 摩尔多瓦                    |            `"MD"` | `"🇲🇩"` |
| 摩纳哥                     |            `"MC"` | `"🇲🇨"` |
| 蒙古                   |            `"MN"` | `"🇲🇳"` |
| 黑山                 |            `"ME"` | `"🇲🇪"` |
| 摩洛哥                    |            `"MA"` | `"🇲🇦"` |
| 荷兰                |            `"NL"` | `"🇳🇱"` |
| 新西兰                |            `"NZ"` | `"🇳🇿"` |
| 尼日利亚                    |            `"NG"` | `"🇳🇬"` |
| 北马其顿            |            `"MK"` | `"🇲🇰"` |
| 挪威                     |            `"NO"` | `"🇳🇴"` |
| 巴基斯坦                   |            `"PK"` | `"🇵🇰"` |
| 菲律宾                |            `"PH"` | `"🇵🇭"` |
| 波兰                     |            `"PL"` | `"🇵🇱"` |
| 葡萄牙                   |            `"PT"` | `"🇵🇹"` |
| 波多黎各                |            `"PR"` | `"🇵🇷"` |
| 罗马尼亚                    |            `"RO"` | `"🇷🇴"` |
| 圣马力诺                 |            `"SM"` | `"🇸🇲"` |
| 沙特阿拉伯               |            `"SA"` | `"🇸🇦"` |
| 塞尔维亚                     |            `"RS"` | `"🇷🇸"` |
| 新加坡                  |            `"SG"` | `"🇸🇬"` |
| 斯洛伐克                   |            `"SK"` | `"🇸🇰"` |
| 斯洛文尼亚                   |            `"SI"` | `"🇸🇮"` |
| 南非               |            `"ZA"` | `"🇿🇦"` |
| 韩国                |            `"KR"` | `"🇰🇷"` |
| 西班牙                      |            `"ES"` | `"🇪🇸"` |
| 瑞典                     |            `"SE"` | `"🇸🇪"` |
| 瑞士                |            `"CH"` | `"🇨🇭"` |
| 泰国                   |            `"TH"` | `"🇹🇭"` |
| 特立尼达和多巴哥          |            `"TT"` | `"🇹🇹"` |
| 土耳其                     |            `"TR"` | `"🇹🇷"` |
| 乌克兰                    |            `"UA"` | `"🇺🇦"` |
| 阿拉伯联合酋长国       |            `"AE"` | `"🇦🇪"` |
| 美国              |            `"US"` | `"🇺🇸"` |
| 乌拉圭                    |            `"UY"` | `"🇺🇾"` |
| 乌兹别克斯坦                 |            `"UZ"` | `"🇺🇿"` |
| 委内瑞拉                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` 应该返回 `"🇦🇱"`。

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` 应该返回 `"🇦🇩"`。

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` 应该返回 `"🇦🇷"`。

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` 应该返回 `"🇦🇲"`。

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` 应该返回 `"🇦🇺"`。

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` 应该返回 `"🇦🇹"`。

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` 应该返回 `"🇦🇿"`。

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` 应该返回 `"🇧🇪"`。

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` 应该返回 `"🇧🇯"`。

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` 应该返回 `"🇧🇴"`。

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` 应该返回 `"🇧🇦"`。

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` 应该返回 `"🇧🇷"`。

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` 应该返回 `"🇧🇬"`。

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` 应该返回 `"🇨🇦"`。

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` 应该返回 `"🇨🇱"`。

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` 应该返回 `"🇨🇳"`。

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` 应该返回 `"🇨🇴"`。

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` 应该返回 `"🇭🇷"`。

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` 应该返回 `"🇨🇾"`。

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` 应该返回 `"🇨🇿"`。

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` 应该返回 `"🇩🇰"`。

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` 应该返回 `"🇪🇨"`。

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` 应该返回 `"🇪🇷"`。

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` 应该返回 `"🇪🇪"`。

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` 应该返回 `"🇫🇮"`。

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` 应该返回 `"🇫🇷"`。

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` 应该返回 `"🇬🇪"`。

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` 应该返回 `"🇩🇪"`。

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` 应该返回 `"🇬🇧"`。

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` 应该返回 `"🇬🇷"`。

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` 应该返回 `"🇬🇼"`。

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` 应该返回 `"🇭🇹"`。

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` 应该返回 `"🇭🇰"`。

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` 应该返回 `"🇭🇺"`。

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` 应该返回 `"🇮🇸"`。

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` 应该返回 `"🇮🇳"`。

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` 应该返回 `"🇮🇷"`。

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` 应该返回 `"🇮🇪"`。

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` 应该返回 `"🇮🇱"`。

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` 应该返回 `"🇮🇹"`。

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` 应该返回 `"🇯🇲"`。

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` 应该返回 `"🇯🇵"`。

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` 应该返回 `"🇰🇿"`。

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` 应该返回 `"🇰🇪"`。

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` 应该返回 `"🇽🇰"`。

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` 应该返回 `"🇰🇬"`。

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` 应该返回 `"🇱🇻"`。

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` 应该返回 `"🇱🇧"`。

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` 应该返回 `"🇱🇮"`。

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` 应该返回 `"🇱🇹"`。

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` 应该返回 `"🇱🇺"`。

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` 应该返回 `"🇲🇬"`。

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` 应该返回 `"🇲🇾"`。

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` 应该返回 `"🇲🇹"`。

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` 应该返回 `"🇲🇽"`。

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` 应该返回 `"🇲🇩"`。

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` 应该返回 `"🇲🇨"`。

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` 应该返回 `"🇲🇳"`。

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` 应该返回 `"🇲🇪"`。

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` 应该返回 `"🇲🇦"`。

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` 应该返回 `"🇳🇱"`。

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` 应该返回 `"🇳🇿"`。

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` 应该返回 `"🇳🇬"`。

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` 应该返回 `"🇲🇰"`。

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` 应该返回 `"🇳🇴"`。

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` 应该返回 `"🇵🇰"`。

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` 应该返回 `"🇵🇭"`。

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` 应该返回 `"🇵🇱"`。

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` 应该返回 `"🇵🇹"`。

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` 应该返回 `"🇵🇷"`。

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` 应该返回 `"🇷🇴"`。

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` 应该返回 `"🇸🇲"`。

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` 应该返回 `"🇸🇦"`。

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` 应该返回 `"🇷🇸"`。

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` 应该返回 `"🇸🇬"`。

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` 应该返回 `"🇸🇰"`。

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` 应该返回 `"🇸🇮"`。

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` 应该返回 `"🇿🇦"`。

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` 应该返回 `"🇰🇷"`。

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` 应该返回 `"🇪🇸"`。

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` 应该返回 `"🇸🇪"`。

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` 应该返回 `"🇨🇭"`。

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` 应该返回 `"🇹🇭"`。

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` 应该返回 `"🇹🇹"`。

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` 应该返回 `"🇹🇷"`。

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` 应该返回 `"🇺🇦"`。

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` 应该返回 `"🇦🇪"`。

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` 应该返回 `"🇺🇸"`。

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` 应该返回 `"🇺🇾"`。

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` 应该返回 `"🇺🇿"`。

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` 应该返回 `"🇻🇪"`。

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
