---
id: 697a49e6ff50d756c9b6935e
title: "挑战 181：2026 冬季游戏第 2 天：单板滑雪"
challengeType: 28
dashedName: challenge-181
---

# --description--

给定单板滑雪者的起始姿势和旋转角度，确定他们的着陆姿势。

- 滑雪板手的站姿要么是 `"Regular"`，要么是 `"Goofy"`。
- 技巧旋转是 90 度的倍数。正数表示顺时针旋转，负数表示逆时针旋转。
- 着陆姿势每旋转 180 度翻转一次。

例如，给定 `"Regular"` 和 `90`，返回 `"Regular"`。给定 `"Regular"` 和 `180` 度，返回 `"Goofy"`。

# --hints--

`getLandingStance("Regular", 90)` 应该返回 `"Regular"`。

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` 应该返回 `"Goofy"`。

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` 应该返回 `"Regular"`。

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` 应该返回 `"Goofy"`。

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` 应该返回 `"Goofy"`。

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` 应该返回 `"Regular"`。

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
