---
id: 697a49e9860d24853adef67f
title: "挑战 194：2026 冬季游戏第 15 天：自由式滑雪"
challengeType: 28
dashedName: challenge-194
---

# --description--

给定一个由两个单词组成的花式滑雪动作名称，判断它是否是有效的花式滑雪动作名称。

一个技巧是有效的，当且仅当第一个单词在有效第一个单词的列表中，且第二个单词在有效第二个单词的列表中。

- 这两个单词将由一个空格分隔。

有效的第一个单词：

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

有效的第二个单词：

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` 应该返回 `true`。

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` 应该返回 `true`。

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` 应该返回 `true`。

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` 应该返回 `true`。

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` 应该返回 `true`。

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` 应返回 `false`。

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` 应该返回 `false`。

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
