---
id: 697a49e9860d24853adef67c
title: "挑战 191：2026 冬季游戏第 12 天：雪橇"
challengeType: 28
dashedName: challenge-191
---

# --description--

给定一个表现雪橇队运动员权值的数组和一个表现雪橇权值的数字，确定该队是否有资格比赛。

- 数组的长度决定团队规模：1 人、2 人或 4 人团队。
- 所有给定的权值均以千克（kg）为单位。

雪橇（单独的雪橇）必须具有最低权值：

- 1 人团队的 162 kg
- 2 人团队的 170 kg
- 4 人团队的 210 kg

雪橇（运动员加雪橇）的总权值不得超过：

- 1 人团队的 247 kg
- 2 人团队的 390 kg
- 4 人团队的 630 kg

返回 `"Eligible"` 如果团队满足所有要求，或者返回 `"Not Eligible"` 如果团队未能满足一个或多个要求。

# --hints--

`checkEligibility([78], 165)` 应该返回 `"Eligible"`。

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` 应该返回 `"Eligible"`。

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` 应该返回 `"Eligible"`。

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` 应该返回 `"Not Eligible"`。

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
