---
id: 697a49e6ff50d756c9b6935f
title: "挑战 182：2026 冬季游戏 第 3 天：冬季两项"
challengeType: 28
dashedName: challenge-182
---

# --description--

给定一个整数数组，其中每个值表现一次冬季两项比赛中单轮击中目标的数量，返回运动员必须滑行的总罚距离。

- 每轮包含 5 个 `targets`。
- 每错过一个目标，将导致 150 米的惩罚循环。

# --hints--

`calculatePenaltyDistance([4, 4])` 应该返回 `300`。

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` 应该返回 `0`。

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` 应该返回 `450`。

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` 应该返回 `150`。

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` 应该返回 `1500`。

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
