---
id: 697a49e6ff50d756c9b69366
title: "挑战 189：2026 冬季游戏第 10 天：高山滑雪"
challengeType: 28
dashedName: challenge-189
---

# --description--

给定滑雪坡的垂直落差、水平距离和`type`，确定坡的难度等级。

要确定评级：

- 通过将落差除以距离来计算山坡的陡度。
- 然后，根据丘陵 `hill type` 计算调整后的陡度：
  - `"Downhill"` 将陡度乘以 1.2
  - `"Slalom"`：将陡度乘以 0.9
  - `"Giant Slalom"`：将陡度乘以 1.0

返回：

- `"Green"` 如果调整后的陡度小于或相等于 0.1
- `"Blue"` 如果调整后的陡度大于 0.1 且小于或相等于 0.25
- `"Black"` 如果调整后的陡度大于 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` 应该返回 `"Green"`。

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` 应该返回 `"Black"`。

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` 应该返回 `"Blue"`。

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` 应该返回 `"Green"`。

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` 应该返回 `"Blue"`。

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` 应该返回 `"Black"`。

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
