---
id: 697a49e6ff50d756c9b69365
title: "挑战 188：2026 冬季游戏 第 9 天：Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

给定一个表现骨架赛道弯道的字串，确定赛道的难度。

- 给定的字串将仅包含以下字母：
  - `"L"` 表示左转
  - `"R"` 表示右转
  - `"S"` 表示直线段

- 每次方向改变添加 15 分（一个 `"L"` 后跟一个 `"R"`，或反之亦然）。
- 所有其他曲线各添加 5 点（所有其他 `"L"` 或 `"R"` 字符）。
- 直线段添加 0 分。

赛道的难度基于总分。返回：

- `"Easy"` 如果总数在 0 - 100 之间
- `"Medium"` 如果总数在 101-200 之间
- `"Hard"` 如果总数超过 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` 应该返回 `"Easy"`。

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` 应该返回 `"Hard"`。

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` 应该返回 `"Medium"`。

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` 应该返回 `"Hard"`。

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` 应该返回 `"Medium"`。

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` 应该返回 `"Easy"`。

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
