---
id: 698a1a73ade5ac0e19180fa1
title: "挑战 198：工作日计数"
challengeType: 28
dashedName: challenge-198
---

# --description--

给定一个开始日期和结束日期，返回两者之间的工作日天数。

- 给定的日期格式为 `"YYYY-MM-DD"`。
- 工作日是工作日（星期一至星期五）。
- 周末不是工作日（星期六和星期日）。
- 计算时包括开始和结束日期。

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` 应该返回 `3`。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` 应该返回 `4`。

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` 应该返回 `5`。

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` 应该返回 `7`。

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` 应该返回 `262`。

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
