---
id: 698a1a73ade5ac0e19180fa9
title: "挑战 206：扑克牌数值"
challengeType: 28
dashedName: challenge-206
---

# --description--

给定一组扑克牌，`return` 一个包含每张牌数值的 `新建` `数组`。

牌面值：

- 一张 Ace (`"A"`) 的值为 1。
- 数字牌（`"2"` - `"10"`）的点数分别为 2 - 10。
- 人头牌：杰克（`"J"`）、皇后（`"Q"`）和国王（`"K"`）每张价值 10。

适用情况：

- 每张牌都有一个花色：黑桃（`"S"`）、梅花（`"C"`）、方块（`"D"`）或红心（`"H"`）。

卡片形式：

- 每张牌表现为一个字串：`"valueSuit"`。例如：`"AS"` 是黑桃 A，`"10H"` 是红心 10，`"QC"` 是梅花 Q。

# --hints--

`cardValues(["3H", "4D", "5S"])` 应该返回 `[3, 4, 5]`。

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` 应该返回 `[1, 10, 10, 6, 7]`。

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` 应该返回 `[8, 10, 2, 10, 9]`。

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` 应该返回 `[1, 10]`。

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` 应该返回 `[10, 10, 10, 10, 1]`。

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
