---
id: 698a1a73ade5ac0e19180fa8
title: "挑战 205：完美立方计数"
challengeType: 28
dashedName: challenge-205
---

# --description--

给定两个整数，确定在这两个数字之间（包括这两个数字）存在多少个完全立方数。

- 较小的数字不保证是第一个参数。
- 如果存在一个整数（`n`）使得 `n * n * n = number`，那么该数字就是一个完全立方数。例如，27 是一个完全立方数，因为 `3 * 3 * 3 = 27`。

# --hints--

`countPerfectCubes(3, 30)` 应该返回 `2`。

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` 应该返回 `3`。

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` 应该返回 `4`。

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` 应该返回 `9`。

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` 应该返回 `41`。

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
