---
id: 698a1a73ade5ac0e19180fa0
title: "挑战 197：血型兼容性"
challengeType: 28
dashedName: challenge-197
---

# --description--

给定一个供血者的血型和一个受血者的血型，确定供血者是否可以给受血者输血。

每种血型由以下部分组成：

- 一个字母：`"A"`、`"B"`、`"AB"` 或 `"O"`
- 以及 Rh 因子：`"+"` 或 `"-"`

血型将是有效字母之一，后跟 Rh 因子。例如，`"AB+"` 和 `"O-"` 是有效的血型。

字母规则：

- `"O"` 可以捐赠给其他字母 `type`。
- `"A"` 可以捐赠给 `"A"` 和 `"AB"`。
- `"B"` 可以捐赠给 `"B"` 和 `"AB"`。
- `"AB"` 只能捐赠给 `"AB"`。

Rh 规则：

- 负号（`"-"`）既可以赋予 `"-"`，也可以赋予 `"+"`。
- 正电荷（`"+"`）只能捐赠给 `"+"`。

供体能够向受体捐赠，必须同时通过字母规则和 Rh 规则。

# --hints--

`canDonate("B+", "B+")` 应该返回 `true`。

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` 应该返回 `true`。

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` 应该返回 `false`。

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` 应该返回 `true`。

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` 应该返回 `false`。

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` 应该返回 `true`。

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` 应该返回 `false`。

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` 应该返回 `true`。

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` 应该返回 `false`。

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` 应该返回 `false`。

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
