---
id: 698a1a73ade5ac0e19180fa6
title: "挑战 203：扁平化"
challengeType: 28
dashedName: challenge-203
---

# --description--

给定一个数组，判断它是否是扁平的。

- 如果一个数组的所有元素都不是数组，则该数组是扁平的。

# --hints--

`isFlat([1, 2, 3, 4])` 应该返回 `true`。

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` 应该返回 `false`。

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` 应该返回 `true`。

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` 应该返回 `false`。

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` 应该返回 `false`。

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
