---
id: 698a1a863194f1f4e63f6460
title: "挑战 209：元素大小"
challengeType: 28
dashedName: challenge-209
---

# --description--

给定一个窗口大小、一个元素以视口宽度 `"vw"` 单位表示的宽度，以及一个元素以视口高度 `"vh"` 单位表示的高度，确定该元素的像素大小。

- 给定的 `window` 大小和返回的 `element` 大小是格式为 `"width x height"` 的字串，例如 `"1200 x 800"`。

- `"vw"` 单位是窗口宽度的百分比。例如，`"50vw"` 是窗口宽度的 50%。
- `"vh"` 单位是窗口高度的百分比。例如，`"50vh"` 是窗口高度的 50%。

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` 应该返回 `"600 x 400"`。

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` 应该返回 `"80 x 240"`。

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` 应该返回 `"70 x 15"`。

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` 应该返回 `"1824 x 1080"`。

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` 应该返回 `"0 x 0"`。

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` 应该返回 `"1440 x 1026"`。

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
