---
id: 698a1a863194f1f4e63f645e
title: "挑战 207：最小间隙"
challengeType: 28
dashedName: challenge-207
---

# --description--

给定一个 `string`，返回两个具有最小字符间隔（最小间隙）的相同字符之间的子串。

- 总会至少有一对匹配的字符。
- 返回的子串应排除匹配的字符。
- 如果两个或更多空白长度相同，`return` 第一个空白中的字符。

例如，给定 `"ABCDAC"`，返回 `"DA"`，因为：

- 只有 `"A"` 和 `"C"` 在字串中重复。
- 两个 `"A"` 字符之间的字符数是 3，两个 `"C"` 字符之间的字符数是 2。
- 所以返回两个 `"C"` 字符之间的字串。

# --hints--

`smallestGap("ABCDAC")` 应该返回 `"DA"`。

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` 应该返回 `"e"`。

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` 应该返回 `"#q6e&rkf(p"`。

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` 应该返回 `""`。

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` 应该返回 `"fox"`。

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
