---
id: 698a1a863194f1f4e63f645f
title: "挑战 208：路径遍历"
challengeType: 28
dashedName: challenge-208
---

# --description--

给定一个表现你路径地图的字串数组，返回到达目标所需移动的字串。

给定的字串将包含以下值：

- `"C"`：你当前位置
- `"G"`：你的目标
- `"T"`：路径的可通行部分
- `"-"`：地图中不可通行的部分

返回一个包含从你的位置到目标所需路径移动的字串，其中：

- `"R"` 是向右移动
- `"D"` 是向下移动
- `"L"` 是向左移动
- `"U"` 是向上移动

- 从你当前位置到目标，总会有一条连续的路径，没有任何分支。
- 每个路径位置最多有两个可通行的位置与之相邻。

例如，给定：

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

返回 `"RDDRDD"`。

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` 应该返回 `"RDDRDD"`。

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` 应该返回 `"RRUUURR"`。

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` 应该返回 `"DLDDRRRRD"`。

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` 应该返回 `"RRRDDL"`。

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` 应该返回 `"ULLLUUURRRRRRDDDR"`。

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
