---
id: 698a1a73ade5ac0e19180fa3
title: "挑战 200：字母和数字计数"
challengeType: 28
dashedName: challenge-200
---

# --description--

给定一个字串，返回一条包含其中字母和数字数量的消息。

- 字母是 `A-Z` 和 `a-z`。
- 数字是 `0-9`。
- 忽略所有其他字符。

返回 `"The string has X letters and Y numbers."`，其中 `"X"` 是字母的数量，`"Y"` 是数字的数量。如果任一数量为 1，则该项使用单数形式。例如：使用 `"1 letter"` 而不是 `"1 letters"`，使用 `"1 number"` 而不是 `"1 numbers"`。

# --hints--

`countLettersAndNumbers("helloworld123")` 应该返回 `"The string has 10 letters and 3 numbers."`。

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` 应该返回 `"The string has 5 letters and 2 numbers."`。

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` 应该返回 `"The string has 1 letter and 1 number."`。

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` 应该返回 `"The string has 0 letters and 5 numbers."`。

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` 应该返回 `"The string has 8 letters and 0 numbers."`。

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
