---
id: 6994cff2290543b3aec9f510
title: "挑战 211：数组求和"
challengeType: 28
dashedName: challenge-211
---

# --description--

给定一个数字 `array`，`return` 所有数字的和。

# --hints--

`sumArray([1, 2, 3, 4, 5])` 应该返回 `15`。

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` 应该返回 `42`。

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` 应该返回 `7`。

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` 应该返回 `9395`。

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` 应该返回 `0`。

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
