---
id: 6994cff2290543b3aec9f50f
title: "挑战 210：HSL 验证器"
challengeType: 28
dashedName: challenge-210
---

# --description--

给定一个字串，判断它是否是有效的 CSS `hsl()` 颜色值。

- 有效的 HSL 值必须是 `"hsl(h, s%, l%)"` 形式，其中：
  - `h`（色相）必须是 0 到 360（含）之间的数字。
  - `s`（饱和度）必须是介于 0% 和 100% 之间的百分比。
  - `l`（亮度）必须是介于 0% 和 100% 之间的百分比。

- 只允许空格：
  - 在左括号之后
  - 逗号之前和/或之后
  - 在闭合括号之前和/或之后

- 值可以以分号 (`";"`) 结尾，且这是可选的。

例如，`"hsl(240, 50%, 50%)"` 是一个有效的 HSL 值。

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` 应该返回 `true`。

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` 应该返回 `true`。

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` 应该返回 `true`。

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` 应该返回 `true`。

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` 应该返回 `true`。

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` 应该返回 `false`。

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` 应该返回 `false`。

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` 应该返回 `false`。

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` 应该返回 `false`。

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
