---
id: 699c8e045ee7cb94ed2322dd
title: "挑战 222：昼夜平分影子"
challengeType: 28
dashedName: challenge-222
---

# --description--

今天是昼夜平分点，太阳正好位于赤道上空，并且在中午时分完全处于头顶。给定一个时间，确定一根 4 英尺高的垂直杆投下的阴影。

- 时间将是 `"HH:MM"` 24 小时制格式的字串（例如，`"15:00"` 是下午 3 点）。
- 你将只获得以 30 分钟为增量的时间。

规则：

- 太阳在早上 6 点正对 `"east"` 升起，晚上 6 点正对 `"west"` 落下。
- 影子总是指向与太阳相反的方向。
- 影子的长度（以英尺为单位）是距离中午小时数的立方。
- 日出前（6 点前）、日落后（6 点或更晚）或中午时没有阴影。

返回：

- 如果存在阴影，返回 `"(length)ft (direction)"`。例如，`"8ft west"`。
- 否则，返回 `"No shadow"`。

例如，给定 `"10:00"`，返回 `"8ft west"`，因为上午 10 点距离中午还有 2 小时，所以 2<sup>3</sup> = 8 英尺，且阴影指向西方，因为上午 10 点太阳在东方。

# --hints--

`getShadow("10:00")` 应该返回 `"8ft west"`。

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` 应该返回 `"27ft east"`。

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` 应该返回 `"No shadow"`。

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` 应该返回 `"166.375ft east"`。

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` 应该返回 `"No shadow"`。

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` 应该返回 `"216ft west"`。

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` 应该返回 `"No shadow"`。

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` 应该返回 `"91.125ft west"`。

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` 应该返回 `"No shadow"`。

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
