---
id: 699c8e045ee7cb94ed2322d9
title: "挑战 218：均匀可整除"
challengeType: 28
dashedName: challenge-218
---

# --description--

给定两个整数，判断你是否可以将第一个整数整除第二个整数。

# --hints--

`isEvenlyDivisible(4, 2)` 应该返回 `true`。

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` 应该返回 `false`。

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` 应该返回 `false`。

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` 应该返回 `true`。

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` 应该返回 `true`。

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` 应该返回 `false`。

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
