---
id: 699c8e045ee7cb94ed2322d6
title: "挑战 215：停车费计算器"
challengeType: 28
dashedName: challenge-215
---

# --description--

给定两个表现你停车时间和取车时间的字串，计算停车费。

- 给定的字串将采用 24 小时制的 `"HH:MM"` 形式。例如，`"14:00"` 表示下午 2 点。
- 第一个字串将是你停车的时间，第二个将是你取车的时间。
- 如果取车时间早于进入时间，则表示停车时间跨越了午夜进入了第二天。

费用规则：

- 每停放 1 小时收费 3 美元。
- 部分小时将向上取整到下一个整小时。
- 如果停车时间跨越过夜（超过午夜），将收取额外的 10 美元过夜费用。
- 最低费用为 5 美元（仅在总额低于 5 美元时使用）。

以 `"$cost"` 形式返回总费用，例如 `"$5"`。

# --hints--

`calculateParkingFee("09:00", "11:00")` 应该返回 `"$6"`。

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` 应该返回 `"$5"`。

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` 应该返回 `"$9"`。

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` 应该返回 `"$27"`。

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` 应该返回 `"$34"`。

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` 应该返回 `"$82"`。

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
