---
id: 699c8e045ee7cb94ed2322d8
title: "挑战 217：被捕捉的棋子"
challengeType: 28
dashedName: challenge-217
---

# --description--

给定一个表现你仍然在棋盘上的棋子的字串数组，计算对手已经捕捉的棋子的价值。

在国际象棋中，你开始时有 16 个棋子：

| Piece  | Abbreviation   | Quantity | Value |
| ------ | -------------- | -------- | ----- |
| Pawn   | `"P"`          | 8        | 1     |
| Rook   | `"R"`          | 2        | 5     |
| Knight | `"N"`          | 2        | 3     |
| Bishop | `"B"`          | 2        | 3     |
| Queen  | `"Q"`          | 1        | 9     |
| King   | `"K"`          | 1        | 0     |

- 给定的数组只会包含上述缩略语。
- 给定数组中未包含的任意 16 个部分已被捕捉。
- 返回所有捕捉到的棋子的总价值，除非…
- 如果国王被捕捉，返回 `"Checkmate"`。

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` 应该返回 `8`。

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` 应该返回 `26`。

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` 应该返回 `16`。

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` 应该返回 `4`。

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` 应该返回 `38`。

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` 应该返回 `0`。

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` 应该返回 `"Checkmate"`。

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
