---
id: 699c8e045ee7cb94ed2322db
title: "挑战 220：最大数字"
challengeType: 28
dashedName: challenge-220
---

# --description--

给定一个由各种标点分隔的数字字串，返回最大的数字。

- 给定的字串将只包含数字和分隔符。
- 分隔符可以是逗号（`","`）、感叹号（`"!"`）、问号（`"?"`）、冒号（`":"`）或分号（`";"`）。

# --hints--

`largestNumber("1,2")` 应该返回 `2`。

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` 应该返回 `60`。

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` 应该返回 `-402`。

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` 应该返回 `99.9`。

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
