---
id: 699c8e045ee7cb94ed2322da
title: "挑战 219：周年里程碑"
challengeType: 28
dashedName: challenge-219
---

# --description--

给定一个表现夫妻结婚年数的整数，返回他们根据此表最近的周年纪念里程碑：

| Years Married | Milestone    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- 如果他们还没有达到第一个里程碑，返回 `"Newlyweds"`。

# --hints--

`getMilestone(0)` 应该返回 `"Newlyweds"`。

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` 应该返回 `"Paper"`。

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` 应该返回 `"Wood"`。

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` 应该返回 `"Tin"`。

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` 应该返回 `"Silver"`。

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` 应该返回 `"Ruby"`。

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` 应该返回 `"Gold"`。

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` 应该返回 `"Diamond"`。

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` 应该返回 `"Platinum"`。

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
