---
id: 699c8e045ee7cb94ed2322d5
title: "挑战 214：多米诺链验证器"
challengeType: 28
dashedName: challenge-214
---

# --description--

给定一个表现多米诺骨牌序列的二维数组，判断它是否形成有效链。

- 数组中的每个元素表现一个多米诺骨牌，并且将是一个包含两个从 1 到 6（含）数字的数组。
- 为了使链条有效，每个多米诺骨牌的第二个数字必须与下一个多米诺骨牌的第一个数字匹配。
- 第一个多米诺骨牌的第一个数字和最后一个多米诺骨牌的最后一个数字不需要匹配任何内容。

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` 应该返回 `true`。

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` 应该返回 `false`。

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` 应该返回 `false`。

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` 应该返回 `true`。

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` 应该返回 `false`。

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
