---
id: 69a890af247de743333bd4d1
title: "挑战 228：电影之夜"
challengeType: 28
dashedName: challenge-228
---

# --description--

给定一个表示星期几的字串、另一个表示放映时间的字串和一个整数票数，返回该场电影票的总费用。

给定的日期将是以下之一：

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

放映时间将以 `"H:MMam"` 或 `"H:MMpm"` 形式给出。例如 `"10:00am"` 或 `"10:00pm"`。

以 `"$D.CC"` 形式返回总费用，使用以下规则：

- 周末（星期五－星期日）：每张票 12.00 美元。
- 工作日（星期一 - 星期四）：每张票 10.00 美元。
- 下午 5:00 之前的场次：每张票减免 $2.00（星期二除外）。
- 星期二：所有票价均为 $5.00 每张。

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` 应该返回 `"$12.00"`。

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` 应该返回 `"$10.00"`。

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` 应该返回 `"$10.00"`。

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` 应该返回 `"$30.00"`。

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` 应该返回 `"$32.00"`。

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` 应该返回 `"$50.00"`。

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` 应该返回 `"$5.00"`。

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
