---
id: 69a890af247de743333bd4d0
title: "挑战 227：冷却时间"
challengeType: 28
dashedName: challenge-227
---

# --description--

给定两个时间戳，第一个表现用户完成考试的时间，第二个表现当前时间，确定用户是否可以再次参加考试。

- 两个时间戳都将采用格式：`"YYYY-MM-DDTHH:MM:SS"`，例如 `"2026-03-25T14:00:00"`。请注意时间为 24 小时制。
- 用户必须等待至少 48 小时后才能重新参加考试。

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` 应该返回 `true`。

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` 应该返回 `false`。

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` 应该返回 `true`。

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` 应该返回 `false`。

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
