---
id: 69a890af247de743333bd4d2
title: "挑战 229：截断 `text` 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

给定一个 `string`，返回一个被截断的 `string`，使得字符的总宽度不超过 50 个单位。

每个字符都有特定的宽度：

| Letters | Width |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

上表包含所有大写和小写字母。此外：

- 空格（`" "`）的宽度为 2
- 句号（`"."`）的宽度为 1

- 如果给定的字串长度在 50 个单位或以下，返回该字串本身，否则
- 截断字串并在末尾添加三个句点（`"..."`），使其总宽度（包括三个句点）尽可能接近 60 个单位且不超过。

# --hints--

`truncateText("The quick brown fox")` 应该返回 `"The quick brown f..."`。

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` 应该返回 `"The silky smooth s..."`。

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` 应该返回 `"THE LOUD BRIG..."`。

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` 应该返回 `"The fast striped z..."`。

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` 应该返回 `"The big black bear"`。

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```
