---
id: 69a890af247de743333bd4ce
title: "挑战 225：不允许连续重复"
challengeType: 28
dashedName: challenge-225
---

# --description--

给定一个字串，判断它是否没有重复的字符。

- 如果一个字串中没有相同的字符连续出现两次或更多次，则该字串没有重复。

# --hints--

`hasNoRepeats("hi world")` 应该返回 `true`。

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` 应该返回 `false`。

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` 应该返回 `true`。

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` 应该返回 `false`。

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` 应该 `return` `true`。

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` 应该返回 `false`。

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
