---
id: 69a890af247de743333bd4cd
title: "挑战 224：咖啡烘焙检测器"
challengeType: 28
dashedName: challenge-224
---

# --description--

给定一个表现用来制作一杯咖啡的豆子的字串，确定这杯咖啡的烘焙程度。

- 给定的字串将包含以下字符，每个字符表现一种豆类的类型：

  - 撇号（`'`）是一种轻度烘焙豆，每个价值 1 分。
  - 一个破折号（`-`）是一个中度烘焙豆，每个价值 2 分。
  - 句号（`.`）是每个价值 3 分的深度烘焙豆。

- 烘焙程度由所有豆子的平均值决定。

返回：

- `"Light"` 如果平均值小于 1.75。
- `"Medium"` 如果平均值在 1.75 到 2.5 之间。
- `"Dark"` 如果平均值大于 2.5。

# --hints--

`detectRoast("''-''''''-'-''--''''")` 应该返回 `"Light"`。

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` 应该返回 `"Medium"`。

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` 应该返回 `"Medium"`。

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` 应该返回 `"Dark"`。

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` 应该返回 `"Medium"`。

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` 应该返回 `"Dark"`。

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` 应该返回 `"Light"`。

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
