---
id: 69a890af247de743333bd4cc
title: "挑战 223：QR 解码器"
challengeType: 28
dashedName: challenge-223
---

# --description--

给定一个 6x6 的矩阵（数组的数组），表现一个 QR 码，返回代码中的二进制数据字串。

- 二维码可以以任意 90 度增量旋转的方式给出。
- 一个正确定位的代码在左下角、左上角和右上角有一个 2x2 的 `1`（定位标记）组。
- 这三个 2x2 方向标记不是二进制 `data` 的一部分。
- 当二维码正确放置时，二进制数据按从左到右、从上到下的顺序（像书本一样）被读取。
- 一个代码总是只有一个有效的方向。

例如，给定：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

或者给出相同的代码但方向不同：

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

返回 `"000000000000000000000001"`，所有二进制数据，排除三个 2x2 方向标记。

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` 应该返回 `"000000000000000000000001"`。

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` 应该返回 `"000000000000000000000001"`。

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` 应该返回 `"001101000011000000110100"`。

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` 应该返回 `"010001000100010101010110"`。

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` 应该返回 `"010000100100100101001110"`。

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
