---
id: 69b1028d6e265413d0198a2a
title: "挑战 231：ISBN-10 验证器"
challengeType: 28
dashedName: challenge-231
---

# --description--

给定一个字串，判断它是否是有效的 ISBN-10。

ISBN-10 由连字符（`"-"`）和其他 10 个字符组成。去掉连字符（`"-"`）后：

- 前 9 个字符必须是数字，
- 最后一个字符可以是数字或字母 `"X"`，它代表数字 10。

验证方法：

- 将每个数字（或对应值）乘以它的位置（第一个数字乘以 1，第二个乘以 2，依此类推）。
- 将所有结果相加。
- 如果总和能被 11 整除，则该 ISBN-10 有效。

# --hints--

`isValidIsbn10("0-306-40615-2")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` 应该返回 `false`。

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` 应该返回 `false`。

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
