---
id: 69b1028d6e265413d0198a2c
title: "挑战 233：起床闹钟"
challengeType: 28
dashedName: challenge-233
---

# --description--

给定一个字串，表示你设置的闹钟时间，以及一个字串，表示你实际醒来的时间，判断你是早起、准时还是迟到。

- 两个时间都将以 `"HH:MM"` 24 小时制形式给出。

返回：

- 如果你在闹钟时间之前醒来，返回 `"early"`。
- 如果你在闹钟时间或闹钟时间后 10 分钟的延迟窗口内醒来，返回 `"on time"`。
- 如果你在闹钟时间后超过 10 分钟才醒来，返回 `"late"`。

两个时间都在同一天内。

# --hints--

`alarmCheck("07:00", "06:45")` 应该返回 `"early"`。

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` 应该返回 `"on time"`。

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` 应该返回 `"on time"`。

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` 应该返回 `"late"`。

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` 应该返回 `"on time"`。

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` 应该返回 `"late"`。

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` 应该返回 `"early"`。

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
