---
id: 69b1028d6e265413d0198a2f
title: "挑战 236：浏览器历史记录"
challengeType: 28
dashedName: challenge-236
---

# --description--

给定一个浏览器命令数组，返回一个包含两个值的数组：历史记录作为 URL 数组，以及当前页面的索引。

有效的命令有：

- `"URL"` —— 其中 URL 是一个网页地址（例如 `"freecodecamp.org"`）。导航到给定的 URL，将其添加到历史记录的下一个位置，并丢弃任何前进的历史记录。
- `"Back"` —— 移动到历史记录中的上一页，如果没有上一页则停留在当前页面。
- `"Forward"` —— 移动到历史记录中的下一页，如果没有下一页则停留在当前页面。

例如，给定 `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`，返回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` 应该返回 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`。

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` 应该返回 `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` 应该返回 `[["example.com", "example.com/contact", "example.com/blog"], 1]`。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` 应该返回 `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` 应该返回 `[["example.com", "example.com/about"], 0]`。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` 应该返回 `[["example.com", "example.com/about"], 1]`。

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
