---
id: 69b58ce40693f140c84c8559
title: "挑战 240：回文字符"
challengeType: 28
dashedName: challenge-240
---

# --description--

给定一个字串，判断它是否是回文，并返回中间的字符（如果长度是奇数）或中间的两个字符（如果长度是偶数）。

- 回文是指正着读和反着读都相同的字串。
- 如果不是回文，返回 `"none"`。

# --hints--

`palindromeLocator("racecar")` 应该返回 `"e"`。

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` 应该返回 `"v"`。

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` 应该返回 `"none"`。

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` 应该返回 `"oo"`。

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` 应该返回 `"00"`。

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
