---
id: 69b5b2be76ec8135a7fbe973
title: "挑战 244：车和象的攻击"
challengeType: 28
dashedName: challenge-244
---

# --description--

给定一个表示车在棋盘上位置的字串，以及另一个表示象位置的字串，判断其中一个棋子是否能攻击另一个。

标准棋盘是 8x8，列从 `A` 到 `H`（从左到右）标记，行从 `1` 到 `8`（从下到上）标记。棋盘如下所示：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- 车可以在水平方向或垂直方向上移动任意格数。
- 象可以在任意对角线上移动任意格数。
- 如果一个棋子能移动到另一个棋子的位置，则它可以攻击该棋子。

返回：

- `"rook"` 如果车能攻击象。
- `"bishop"` 如果象能攻击车。
- `"neither"` 如果两个棋子都不能攻击对方。

# --hints--

`rookBishopAttack("A1", "A5")` 应该返回 `"rook"`。

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` 应该返回 `"bishop"`。

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` 应该返回 `"rook"`。

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` 应该返回 `"bishop"`。

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` 应该返回 `"neither"`。

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` 应该返回 `"neither"`。

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
