---
id: 69b58ce40693f140c84c855b
title: "挑战 242：下一个宾果号码"
challengeType: 28
dashedName: challenge-242
---

# --description--

给定一个宾果号码，返回下一个循序的宾果号码。

宾果号码是一个字母，后面跟着根据下表对应范围内的数字：

| 字母 | 数字范围 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

例如，给定 `"B10"`，返回下一个宾果号码 `"B11"`。如果给定的是最后一个宾果号码，则返回 `"B1"`。

# --hints--

`getNextBingoNumber("B10")` 应该返回 `"B11"`。

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` 应该返回 `"N34"`。

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` 应该返回 `"N31"`。

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` 应该返回 `"O61"`。

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` 应该返回 `"B1"`。

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
