---
id: 69b5b2be76ec8135a7fbe974
title: "挑战 245：螺旋矩阵"
challengeType: 28
dashedName: challenge-245
---

# --description--

给定一个二维矩阵，返回一个包含所有值的顺时针顺序的扁平数组。

返回的数组应以左上角的值开始，沿着顶行向右移动，然后沿右列向下，再沿底行向左，最后沿左列向上。对剩余的层重复此过程。

例如，给定：

```json
[
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9]
]
```

返回 `[1, 2, 3, 6, 9, 8, 7, 4, 5]`。

# --hints--

`spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` 应该返回 `[1, 2, 3, 6, 9, 8, 7, 4, 5]`。

```js
assert.deepEqual(spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [1, 2, 3, 6, 9, 8, 7, 4, 5]);
```

`spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]])` 应该返回 `["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]`。

```js
assert.deepEqual(spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]]), ["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]);
```

`spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]])` 应该返回 `[true, false, false, true, false, false, true, true, false, false, true, true]`。

```js
assert.deepEqual(spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]]), [true, false, false, true, false, false, true, true, false, false, true, true]);
```

`spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]])` 应该返回 `[25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`。

```js
assert.deepEqual(spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]]), [25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
function spiralMatrix(matrix) {

    return matrix
}
```

# --solutions--

```js
function spiralMatrix(matrix) {
  const result = [];
  let top = 0, bottom = matrix.length - 1;
  let left = 0, right = matrix[0].length - 1;

  while (top <= bottom && left <= right) {
    for (let i = left; i <= right; i++) result.push(matrix[top][i]);
    top++;

    for (let i = top; i <= bottom; i++) result.push(matrix[i][right]);
    right--;

    if (top <= bottom) {
      for (let i = right; i >= left; i--) result.push(matrix[bottom][i]);
      bottom--;
    }

    if (left <= right) {
      for (let i = bottom; i >= top; i--) result.push(matrix[i][left]);
      left++;
    }
  }

  return result;
}
```
