---
id: 69b559d2903b9e4afe9075f9
title: "挑战 239：今天星期几？"
challengeType: 28
dashedName: challenge-239
---

# --description--

给定一个以毫秒为单位的 Unix 时间戳，返回星期几。

有效的返回星期为：

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

请确保忽略时区。

# --hints--

`getDayOfWeek(1775492249000)` 应该返回 `"Monday"`。

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` 应该返回 `"Saturday"`。

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` 应该返回 `"Tuesday"`。

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` 应该返回 `"Sunday"`。

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` 应该返回 `"Thursday"`。

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
