---
id: 69b58ce40693f140c84c855c
title: "挑战 243：车的攻击"
challengeType: 28
dashedName: challenge-243
---

# --description--

给定两个字串，表示国际象棋棋盘上两辆车的位置，判断它们是否能互相攻击。

标准的国际象棋棋盘是 8x8，列从左到右标记为 `A` 到 `H`，行从下到上标记为 `1` 到 `8`。棋盘如下所示：

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

车可以沿水平方向或垂直方向移动任意格数。因此，如果它们在同一行或同一列上，就能互相攻击。

# --hints--

`rookAttack("A1", "A8")` 应该返回 `true`。

```js
assert.isTrue(rookAttack("A1", "A8"));
```

`rookAttack("B4", "F4")` 应该返回 `true`。

```js
assert.isTrue(rookAttack("B4", "F4"));
```

`rookAttack("E3", "D4")` 应该返回 `false`。

```js
assert.isFalse(rookAttack("E3", "D4"));
```

`rookAttack("H7", "F6")` 应该返回 `false`。

```js
assert.isFalse(rookAttack("H7", "F6"));
```

# --seed--

## --seed-contents--

```js
function rookAttack(rook1, rook2) {

  return rook1;
}
```

# --solutions--

```js
function rookAttack(rook1, rook2) {
  return rook1[0] === rook2[0] || rook1[1] === rook2[1];
}
```
