---
id: 69b559d2903b9e4afe9075f8
title: "挑战 238：数字旋转逃逸"
challengeType: 28
dashedName: challenge-238
---

# --description--

给定一个正整数，判断它本身或它的任何旋转数是否能被它的数字位数整除。

旋转指的是将第一个数字移动到末尾。例如，旋转一次后，123 变成 231。

- 首先查看旋转 `0`（给定的数字）。
- 给定的数字不会包含任何零。
- 如果找到符合条件的旋转数，返回第一个旋转数；否则返回 `"none"`。

# --hints--

`getRotation(123)` 应该返回 `0`。

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` 应该返回 `3`。

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` 应该返回 `"none"`。

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` 应该返回 `6`。

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
