---
id: 69bc6cb30c1d112a2e110a08
title: "挑战 251：数组求和查找器"
challengeType: 28
dashedName: challenge-251
---

# --description--

给定一个数字数组和一个目标数字，返回第一个和为目标的两个或更多数字的子集。

- “第一个”子集是指其元素具有最低可能索引的子集，优先考虑最早的索引。
- 数组中的每个数字只能使用一次。
- 如果不存在有效的子集，返回 `"Sum not found"`。

以它们在原数组中出现的顺序，将匹配的数字作为数组返回。

# --hints--

`findSum([1, 3, 5, 7], 6)` 应该返回 `[1, 5]`。

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` 应该返回 `[1, 4]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` 应该返回 `[1, 2, 3]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` 应该返回 `[-1, -2, 4]`。

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` 应该返回 `[3, 1, 4, 2]`。

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` 应该返回 `[1, 2, 3, 5, 9]`。

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` 应该返回 `"Sum not found"`。

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
