---
id: 69bc6cb30c1d112a2e110a07
title: "挑战 250：隐藏的密钥"
challengeType: 28
dashedName: challenge-250
---

# --description--

欢迎来到第 250 个每日挑战！

给定一个加密字串，使用加密密钥将其解码并返回。

查找密钥的方法：

- 查看截至今天所有挑战编号是 25 的倍数的每日挑战（包括本次）。
- 取出这些挑战标题的第一个字母并组合成一个字串。如果标题以非字母开头，找到它的第一个字母。

解码消息时，遍历加密消息中的每个字母：

  1. 查看密钥中对应的字母（如果消息比密钥长，则重复密钥）。
  2. 将密钥字母转换为对应的数字：`"A"` = 1，`"B"` = 2，…，`"Z"` = 26。
  3. 将加密字母在字母表中向后移动该数字个位置。
  4. 如果移动超出 `"A"` 之前，则从 `"Z"` 重新开始。

例如，如果加密消息以 `"Y"` 开头，且第一个密钥字母是 `"V"`（22），则将 `"Y"` 向后移动 22 位得到 `"C"`。对每个字母重复此过程以解码完整消息。

- 只有字母会被移动，空格保持不变。
- 所有给定和返回的字母均为大写。

# --hints--

`decode("YAVJYNXE")` 应该返回 `"CONGRATS"`。

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` 应该返回 `"CODING LEGEND"`。

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` 应该返回 `"YOU ARE AWESOME"`。

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` 应该返回 `"KEEP GOING"`。

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` 应该返回 `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`。

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
