---
id: 69bc6cb30c1d112a2e110a06
title: "挑战 249：字串数学"
challengeType: 28
dashedName: challenge-249
---

# --description--

给定一个包含数字和其他字符的字串，根据数字之间非数字字符的数量对数字执行数学运算。

- 如果分隔两个数字的字符数量是偶数，则使用加法。
- 如果是奇数，则使用减法。
- 连续的数字构成一个单独的数字。
- 运算从左到右依次执行。
- 忽略开头和结尾非数字的字符。

例如，给定 `"3ab10c8"`，返回 `5`。将 3 和 10 相加得到 13，因为它们之间有偶数个字符。然后从 13 中减去 8，因为结果和 8 之间有奇数个字符。

# --hints--

`doMath("3ab10c8")` 应该返回 `5`。

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` 应该返回 `2`。

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` 应该返回 `12`。

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` 应该返回 `15`。

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` 应该返回 `67`。

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
