---
id: 69bc6cb30c1d112a2e110a04
title: "挑战 247：最后一个字母"
challengeType: 28
dashedName: challenge-247
---

# --description--

给定一个字串，返回字串中在字母表中排最后的字母。

- 如果有两个或更多字母并列最后，返回第一个出现的那个。
- 忽略所有非字母字符。

# --hints--

`getLastLetter("world")` 应该返回 `"w"`。

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` 应该返回 `"W"`。

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` 应该返回 `"z"`。

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` 应该返回 `"L"`。

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` 应该返回 `"T"`。

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
