---
id: 69bc6cb30c1d112a2e110a09
title: "挑战 252：独特的爬楼梯方式"
challengeType: 28
dashedName: challenge-252
---

# --description--

给定楼梯的数量，返回一个人每次可以迈 1 步或 2 步，有多少种不同的爬楼梯方式。

# --hints--

`getUniqueClimbs(4)` 应该返回 `5`。

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` 应该返回 `8`。

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` 应该返回 `89`。

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` 应该返回 `4181`。

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` 应该返回 `832040`。

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` 应该返回 `20365011074`。

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```
