---
id: 69bc6cb30c1d112a2e110a03
title: "挑战 246：姓名首字母"
challengeType: 28
dashedName: challenge-246
---

# --description--

给定一个完整的姓名字串，返回他们的首字母。

- 需要初始化的名字由空格分隔。
- 首字母应大写。
- 首字母之间应以点号分隔。

例如，`"Tommy Millwood"` 返回 `"T.M."`。

# --hints--

`getInitials("Tommy Millwood")` 应返回 `"T.M."`。

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` 应返回 `"S.P."`。

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` 应返回 `"F.C.C."`。

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` 应返回 `"D."`。

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` 应返回 `"D.V.C.H.N."`。

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
