---
id: 69c5f3d787b1725d5f00c8b8
title: "挑战 254：奇数字词"
challengeType: 28
dashedName: challenge-254
---

# --description--

给定一个字串，返回其中字母数为奇数的字词。

- 给定字串中的字词将由单个空格分隔。
- 返回的字词也用单个空格分隔。

# --hints--

`getOddWords("This is a super good test")` 应该返回 `"a super"`。

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` 应该返回 `"one two three"`。

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` 应该返回 `"split rainbow sprinkles top"`。

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` 应该返回 `"The quick brown fox the river"`。

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
