---
id: 69c5f3d787b1725d5f00c8bc
title: "挑战 258：字串解压器"
challengeType: 28
dashedName: challenge-258
---

# --description--

给定一个压缩字串，使用以下规则返回解压后的版本：

- 给定的字串由单词和数字组成，单词和数字之间用空格分隔。
- 保持单词不变。
- 用该位置的单词替换数字，其中第一个单词的位置为 1。

例如，给定 `"practice makes perfect and 3 1 2 3"`，返回 `"practice makes perfect and perfect practice makes perfect"`。

# --hints--

`decompress("practice makes perfect and 3 1 2 3")` 应该返回 `"practice makes perfect and perfect practice makes perfect"`。

```js
assert.equal(decompress("practice makes perfect and 3 1 2 3"), "practice makes perfect and perfect practice makes perfect");
```

`decompress("hello 1 1")` 应该返回 `"hello hello hello"`。

```js
assert.equal(decompress("hello 1 1"), "hello hello hello");
```

`decompress("the cat sat on 1 mat 4 which 1 2 3")` 应该返回 `"the cat sat on the mat on which the cat sat"`。

```js
assert.equal(decompress("the cat sat on 1 mat 4 which 1 2 3"), "the cat sat on the mat on which the cat sat");
```

`decompress("the more you know 1 2 3 realize 3 don't 4")` 应该返回 `"the more you know the more you realize you don't know"`。

```js
assert.equal(decompress("the more you know 1 2 3 realize 3 don't 4"), "the more you know the more you realize you don't know");
```

`decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10")` 应该返回 `"lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero"`。

```js
assert.equal(decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10"), "lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero");
```

# --seed--

## --seed-contents--

```js
function decompress(str) {

  return str;
}
```

# --solutions--

```js
function decompress(str) {
  const tokens = str.split(" ");
  const words = [];

  for (const token of tokens) {
    if (isNaN(parseInt(token))) {
      words.push(token)
    } else {
      words.push(tokens[token - 1])
    }
  }

  return words.join(" ")
}
```
