---
id: 69c5f3d787b1725d5f00c8ba
title: "挑战 256：最近的时间方向"
challengeType: 28
dashedName: challenge-256
---

# --description--

给定两个时间，判断从第一个时间到第二个时间，向前移动还是向后移动更快。

- 时间以 24 小时制形式给出（`"HH:MM"`）
- 时钟是循环的（向前移动时，23:59 会到达 00:00，向后移动时，00:00 会到达 23:59）

返回：

- 如果向前移动更短，返回 `"forward"`
- 如果向后移动更短，返回 `"backward"`
- 如果两个方向所需时间相同，返回 `"equal"`

# --hints--

`getDirection("10:00", "12:00")` 应该返回 `"forward"`。

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` 应该返回 `"backward"`。

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` 应该返回 `"equal"`。

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` 应该返回 `"forward"`。

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` 应该返回 `"backward"`。

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` 应该返回 `"equal"`。

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
