---
id: 69c5f3d787b1725d5f00c8b7
title: "挑战 253：缩写词查找器"
challengeType: 28
dashedName: challenge-253
---

# --description--

给定一个表现缩写词的字串，从下面的列表中返回它所属组织的全称：

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

给定缩写词中的每个字母都应与所属组织中每个单词的首字母匹配，且顺序相同。

# --hints--

`findOrg("NASA")` 应该返回 `"National Avocado Storage Authority"`。

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` 应该返回 `"Cats Infiltration Agency"`。

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` 应该返回 `"Fluffy Beanbag Inspectors"`。

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` 应该返回 `"Department Of Jelly"`。

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` 应该返回 `"Wild Honey Organization"`。

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` 应该返回 `"Eating Pancakes Administration"`。

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
