---
id: 69c6ff713a52713463aa7926
title: "挑战 260：单词分数"
challengeType: 28
dashedName: challenge-260
---

# --description--

给定一个单词，使用标准字母-数值表返回它的分数：

| 字母 | 数值 |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- 大写和小写字母的数值相同。

# --hints--

`getWordScore("hi")` 应该返回 `17`。

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` 应该返回 `52`。

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` 应该返回 `169`。

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` 应该返回 `94`。

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
