---
id: 69c6ff713a52713463aa7927
title: "挑战 261：数字单词"
challengeType: 28
dashedName: challenge-261
---

# --description--

给定一个 0 到 99 的整数，返回它的英文单词表现形式。

- `0` 返回 `"zero"`。
- 1 到 19 的数字有独特的名称（`"one"`、`"two"`、…、`"ten"`、`"eleven"`、…、`"eighteen"`、`"nineteen"`）。
- 20 到 90 的 10 的倍数有它们自己的名称（`"twenty"`、`"thirty"`、…、`"eighty"`、`"ninety"`）。
- 21 到 99 中不是 10 的倍数的数字写作两个单词，中间用连字符连接。例如 `"forty-two"` 和 `"fifty-three"`。

# --hints--

`getNumberWords(0)` 应该返回 `"zero"`。

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` 应该返回 `"ten"`。

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` 应该返回 `"nineteen"`。

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` 应该返回 `"thirty"`。

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` 应该返回 `"fifty-three"`。

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` 应该返回 `"seven"`。

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` 应该返回 `"twelve"`。

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` 应该返回 `"sixty"`。

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` 应该返回 `"sixty-seven"`。

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` 应该返回 `"ninety-eight"`。

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
