---
id: 69c6ff713a52713463aa7928
title: "挑战 262：URL 查询解析器"
challengeType: 28
dashedName: challenge-262
---

# --description--

给定一个包含查询字串的 URL，将查询字串解析成一个包含点对的对象（或字典）。

- 查询字串从 `"?"` 之后开始，
- 每个参数由 `"&"` 分隔，
- 每个点对由 `"="` 分隔

例如，给定 `"https://example.com/search?name=Alice&age=30"`，应返回：

```json
{
  "name": "Alice",
  "age": "30"
}
```

所有值都应作为字串返回。

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` 应该返回 `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` 应该返回 `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` 应该返回 `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` 应该返回 `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
