---
id: 69c6ff713a52713463aa7929
title: "挑战 263：二进制填字游戏"
challengeType: 28
dashedName: challenge-263
---

# --description--

给定一个字符，判断它的 8 位二进制表现是否能在下列网格中找到，横向或纵向，任一方向均可：

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

例如，`"A"` 的二进制表现是 `01000001`，它出现在第一行从左到右的位置。

# --hints--

`isInCrossword("I")` 应该返回 `true`。

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` 应该返回 `true`。

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` 应该返回 `true`。

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` 应该返回 `true`。

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` 应该返回 `false`。

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` 应该返回 `false`。

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` 应该返回 `false`。

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` 应该返回 `false`。

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
