---
id: 69cfca90e8a0a6d4d6871c4e
title: "挑战 264：字谜分组"
challengeType: 28
dashedName: challenge-264
---

# --description--

给定一个单词数组，返回一个二维数组，将单词按字谜分组。

- 如果单词包含相同字母且顺序可以不同，则它们是字谜。
- 每个单词只属于一个分组。
- 返回的顺序无关紧要。

例如，给定 `["listen", "silent", "hello", "enlist", "world"]`，返回 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`。

# --hints--

`groupAnagrams(["listen", "silent", "hello", "enlist", "world"])` 应该返回 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`。

```js
const groups = groupAnagrams(["listen", "silent", "hello", "enlist", "world"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["enlist", "listen", "silent"], ["hello"], ["world"]]);
```

`groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` 应该返回 `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`。

```js
const groups = groupAnagrams(["eat", "tea", "tan", "ate", "nat", "bat"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]);
```

`groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` 应该返回 `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`。

```js
const groups = groupAnagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]);

```

`groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` 应该返回 `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`。

```js
const groups = groupAnagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"]);
const sorted = groups.map(g => g.sort()).sort((a, b) => a[0].localeCompare(b[0]));
assert.deepEqual(sorted, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]);
```

# --seed--

## --seed-contents--

```js
function groupAnagrams(words) {

  return words;
}
```

# --solutions--

```js
function groupAnagrams(words) {
  const map = {};
  for (const word of words) {
    const key = word.split('').sort().join('');
    if (!map[key]) map[key] = [];
    map[key].push(word);
  }
  return Object.values(map);
}
```
