---
id: 69cfca90e8a0a6d4d6871c52
title: "挑战 268：自恋数"
challengeType: 28
dashedName: challenge-268
---

# --description--

给定一个正整数，判断它是否是自恋数。

- 如果一个数字的每个数字的幂次方之和等于它本身，这个数字就是自恋数，幂次方的指数是数字的总位数。

例如，153 有 3 位数字，且 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153，所以它是自恋数。

# --hints--

`isNarcissistic(153)` 应该返回 `true`。

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` 应该返回 `false`。

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` 应该返回 `true`。

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` 应该返回 `false`。

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` 应该返回 `true`。

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` 应该返回 `false`。

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` 应该返回 `true`。

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` 应该返回 `false`。

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
