---
id: 69cfca90e8a0a6d4d6871c4f
title: "挑战 265：最深的方括弧"
challengeType: 28
dashedName: challenge-265
---

# --description--

给定一个包含平衡方括弧的字串，返回最深嵌套方括弧中的内容。

- 方括弧可以是三种类型中的任意一种：`()`、`[]` 和 `{}`。
- 输入中总是只有一个最深的组。

例如，给定 `"(hello (world))"`，返回 `"world"`。

# --hints--

`getDeepestBrackets("(hello (world))")` 应该返回 `"world"`。

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` 应该返回 `"inner"`。

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` 应该返回 `"e"`。

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` 应该返回 `"fox"`。

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` 应该返回 `"C"`。

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
