---
id: 69cfca90e8a0a6d4d6871c51
title: "挑战 267：Parsec 转换器"
challengeType: 28
dashedName: challenge-267
---

# --description--

在一个遥远的星系中，parsec 用来测量时间和距离。给定一个整数 parsec 数，返回它对应的时间或距离。

- 如果给定的整数是奇数，它表示时间。如果是偶数，它表示距离。

使用以下转换率：

| Parsecs | 时间/距离 |
| - | - |
| 1 | 2 小时 |
| 2 | 6 光年 |

将转换后的值作为整数返回。

# --hints--

`convertParsecs(1)` 应该返回 `2`。

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` 应该返回 `6`。

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` 应该返回 `62`。

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` 应该返回 `264`。

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` 应该返回 `34`。

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` 应该返回 `42`。

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
