---
id: 69cfca90e8a0a6d4d6871c53
title: "挑战 269：无过敏原友好餐点"
challengeType: 28
dashedName: challenge-269
---

# --description--

给定一个餐点数组和一个需要避免的过敏原数组，返回所有不含指定过敏原的餐点名称。

- 每个餐点的形式是 `[meal, allergens]`，其中 `meal` 是餐点名称，`allergens` 是该餐点所含过敏原的数组。例如，`["pasta", ["wheat", "milk"]]`。
- 需要避免的过敏原将是一个字串数组。

按给定顺序返回安全的餐点名称。如果没有安全的餐点，返回空数组。

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` 应该返回 `["salad"]`。

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` 应该返回 `["fried rice", "chicken parmesan"]`。

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` 应该返回 `["oatmeal", "granola", "toast"]`。

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` 应该返回 `["granola", "yogurt", "eggs"]`。

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
