---
id: 69e2383af7832c8032603b8e
title: "挑战 272：转置矩阵"
challengeType: 28
dashedName: challenge-272
---

# --description--

给定一个矩阵（一个数组的数组），返回它的转置矩阵。

转置矩阵时，交换行和列。例如：索引 `[0, 1]` 处的值应该移动到索引 `[1, 0]` 处。

例如，给定：

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

返回：

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])` 应该返回 `[[1, 4], [2, 5], [3, 6]]`。

```js
assert.deepEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6]])` 应该返回 `[[1, 3, 5], [2, 4, 6]]`。

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])` 应该返回 `[[1, 3, 5, 7], [2, 4, 6, 8]]`。

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]]);
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])` 应该返回 `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]`。

```js
assert.deepEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]);
```

`transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]])` 应该返回 `[[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]`。

```js
assert.deepEqual(transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]]), [[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]);
```

# --seed--

## --seed-contents--

```js
function transpose(matrix) {

  return matrix;
}
```

# --solutions--

```js
function transpose(matrix) {

  return matrix[0].map((_, i) => matrix.map(row => row[i]));
}
```
