---
id: 69e2383af7832c8032603b8f
title: "挑战 273：ISBN-13 验证器"
challengeType: 28
dashedName: challenge-273
---

# --description--

给定一个字串，判断它是否是有效的 ISBN-13 编号。

有效的 ISBN-13：

- 只包含数字和连字符
- 去除连字符后恰好有 13 个数字
- 通过以下验证：
  1. 交替将每个数字乘以 1 或 3（第一个数字乘以 1，第二个乘以 3，第三个乘以 1，依此类推）。
  2. 结果的总和必须能被 10 整除。

# --hints--

`isValidIsbn13("9780306406157")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` 应该返回 `false`。

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` 应该返回 `false`。

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` 应该返回 `true`。

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
