---
id: 69e2383af7832c8032603b92
title: "挑战 276：错误的元素"
challengeType: 28
dashedName: challenge-276
---

# --description--

给定一个按升序排序的整数数组，只有一个元素位置错误，返回该元素的索引。

- 如果有多个元素可能被认为位置错误，返回第一个的索引。

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` 应该返回 `1`。

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` 应该返回 `3`。

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` 应该返回 `0`。

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` 应该返回 `2`。

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` 应该返回 `6`。

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
