---
id: 69f35a5bb823ed620fcb7cb9
title: "挑战 280：Mongo ID 日期"
challengeType: 28
dashedName: challenge-280
---

# --description--

给定一个 MongoDB ID 字串，返回它的创建时间，格式为 ISO 8601 字串。

- MongoDB ID 是一个 24 字符的十六进制字串。前 8 个字符表现为一个 Unix 时间戳（以秒为单位），以 16 进制整数编码。

例如，`"6a094b50bcf86cd799439011"` 的时间戳是十六进制的 `"6a094b50"`，对应十进制的 `1778994000`，表现的创建时间是 `"2026-05-17T05:00:00.000Z"`。

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` 应该返回 `"2026-05-17T05:00:00.000Z"`。

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` 应该返回 `"2025-12-30T03:20:11.000Z"`。

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` 应该返回 `"2000-01-01T07:01:01.000Z"`。

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` 应该返回 `"2026-05-02T03:38:43.000Z"`。

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` 应该返回 `"2025-08-26T15:08:49.000Z"`。

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```
