---
id: 69f35a5bb823ed620fcb7cbd
title: "挑战 284：I 在 E 之前"
challengeType: 28
dashedName: challenge-284
---

# --description--

给定一个单词或句子，返回一个修正版本，其中每个单词都遵循“I 在 E 之前，除非在 C 之后”的规则。

- 如果一个单词包含 `"ei"`，且前面没有 `"c"`，则将其替换为 `"ie"`。
- 如果一个单词包含 `"ie"`，且前面有 `"c"`，则将其替换为 `"ei"`。
- 其他所有单词保持不变。

# --hints--

`iBeforeE("beleive")` 应该返回 `"believe"`。

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` 应该返回 `"receive"`。

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` 应该返回 `"we received a brief"`。

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` 应该返回 `"she believed the friendly niece could perceive the grief"`。

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` 应该返回 `"we received relief after the thief gave us a brief piece of fierce deceit"`。

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
