---
id: 69f35a5bb823ed620fcb7cbc
title: "挑战 283：字串拉链"
challengeType: 28
dashedName: challenge-283
---

# --description--

给定两个字串，返回一个新字串，将它们的字符交错排列。如果一个字串较长，则将剩余的字符追加到末尾。

从第一个字串的第一个字符开始。

# --hints--

`zipStrings("abc", "123")` 应该返回 `"a1b2c3"`。

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` 应该返回 `"abcdefghijklmnopqrstuvwxyz"`。

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` 应该返回 `"dnaiyght"`。

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` 应该返回 `"pjyatvhaosncript"`。

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` 应该返回 `"freeCodeCamp"`。

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
