---
id: 69f35a5bb823ed620fcb7cbb
title: "挑战 282：睡眠债务"
challengeType: 28
dashedName: challenge-282
---

# --description--

给定一个表示直到今天每晚睡眠小时数的数组，以及每晚的目标睡眠小时数，返回你今晚需要睡多少小时才能消除你的睡眠债务。

- 包括今晚的睡眠时间在内，计算需要补上的总时间。
- 如果你已经睡够了今晚的目标时间或更多，返回 `0`。

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` 应该返回 `20`。

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` 应该返回 `10`。

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` 应该返回 `4`。

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` 应该返回 `0`。

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` 应该返回 `0`。

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
