---
id: 69f8c998d78ad3171a0713be
title: "挑战 291：FizzBuzz 计数"
challengeType: 28
dashedName: challenge-291
---

# --description--

给定一个起始数字和结束数字，计算在该范围内（包含端点）fizz 和 buzz 出现的次数。

- 能被 3 整除的数字计为 fizz。
- 能被 5 整除的数字计为 buzz。
- 能同时被 3 和 5 整除的数字既计为 fizz 也计为 buzz。

返回一个对象或字典，格式为 `{ fizz, buzz }`，其中包含计数结果。

# --hints--

`fizzBuzzCount(1, 11)` 应该返回 `{fizz: 3, buzz: 2}`。

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` 应该返回 `{fizz: 9, buzz: 6}`。

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` 应该返回 `{fizz: 26, buzz: 16}`。

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` 应该返回 `{fizz: 207, buzz: 125}`。

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` 应该返回 `{fizz: 4074, buzz: 2444}`。

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
