---
id: 69f8c998d78ad3171a0713ba
title: "挑战 287：罗马数字修正器"
challengeType: 28
dashedName: challenge-287
---

# --description--

给定一个格式错误的罗马数字字串，返回其标准罗马数字表示的数值。

输入只会使用加法形式，因此每个符号的值都会加到总数上。作为提醒，以下是符号和对应的数值：

| 符号 | 数值 |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

重新编码时，每一步都使用尽可能大的符号，必要时使用减法点对（`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`）。

# --hints--

`fixNumerals("XIIIII")` 应该返回 `"XV"`。

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` 应该返回 `"LXIV"`。

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` 应该返回 `"XL"`。

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` 应该返回 `"MDCCXCIX"`。

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` 应该返回 `"MCDLXIV"`。

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` 应该返回 `"MMCMLXXXIV"`。

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
