---
id: 69f8c998d78ad3171a0713b9
title: "挑战 286：未关闭的问题"
challengeType: 28
dashedName: challenge-286
---

# --description--

给定一个问题编号数组和另一个拉取请求（PR）编号数组，返回所有 PR 合并后仍然未关闭的问题数组。

- 如果 PR 的数字是问题编号数字的旋转，则该 PR 会关闭该问题。例如，问题 123 会被 PR 231 或 312 关闭。
- PR 不会关闭编号完全相同的问题。例如，PR 123 不会关闭问题 123。所以所有数字都相同的问题无法被关闭。
- 两个数字都可能去掉前导零。例如，PR 201 会关闭问题 12（012，是 201 的旋转）。同样，问题 201 会被 PR 12 关闭。

返回剩余未关闭的问题，顺序与给定顺序相同。

# --hints--

`getOpenIssues([123, 234], [231])` 应该返回 `[234]`。

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` 应该返回 `[345, 16]`。

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` 应该返回 `[456, 332, 15]`。

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` 应该返回 `[115, 296, 24]`。

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` 应该返回 `[95, 395, 754, 296, 709, 237, 1802]`。

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
