---
id: 69f8c998d78ad3171a0713bc
title: "挑战 289：差值之和"
challengeType: 28
dashedName: challenge-289
---

# --description--

给定一个数字数组，返回每个数字与其后一个数字之间差值的总和。

例如，给定 `[1, 3, 4]`，返回 `3`（2 + 1）。

# --hints--

`sumOfDifferences([1, 3, 4])` 应该返回 `3`。

```js
assert.equal(sumOfDifferences([1, 3, 4]), 3);
```

`sumOfDifferences([5, -3, 3, 9, 10])` 应该返回 `5`。

```js
assert.equal(sumOfDifferences([5, -3, 3, 9, 10]), 5);
```

`sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7])` 应该返回 `-16`。

```js
assert.equal(sumOfDifferences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16);
```

`sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])` 应该返回 `25`。

```js
assert.equal(sumOfDifferences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25);
```

# --seed--

## --seed-contents--

```js
function sumOfDifferences(arr) {

  return arr;
}
```

# --solutions--

```js
function sumOfDifferences(arr) {
  return arr.slice(0, -1).reduce((sum, n, i) => sum + (arr[i + 1] - n), 0);
}
```
