---
id: 69f8c998d78ad3171a0713bd
title: "挑战 290：披萨派对"
challengeType: 28
dashedName: challenge-290
---

# --description--

给定一个数组，表示每个人今天工作的小时数，返回为披萨派对需要订购的披萨数量。

- 将每个人的工作小时数除以 3，得到他们应得的披萨片数。
- 你不能吃半片披萨，所以将每个人的披萨片数向上取整到最接近的整数。
- 每个人至少得到两片披萨。
- 每个披萨有 8 片。将披萨总数向上取整到最接近的整张披萨。

# --hints--

`getPizzasToOrder([8, 8, 8])` 应该返回 `2`。

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` 应该返回 `3`。

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` 应该返回 `2`。

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` 应该返回 `3`。

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` 应该返回 `1`。

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` 应该返回 `5`。

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
