---
id: 69f8c998d78ad3171a0713bb
title: "挑战 288：秘密数字"
challengeType: 28
dashedName: challenge-288
---

# --description--

给定一个秘密数字和一个猜测，判断猜测是否正确。

返回：

- `"higher"` 如果秘密数字比猜测大。
- `"lower"` 如果秘密数字比猜测小。
- `"you got it!"` 如果猜测正确。

# --hints--

`guessNumber(50, 30)` 应该返回 `"higher"`。

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` 应该返回 `"lower"`。

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` 应该返回 `"you got it!"`。

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` 应该返回 `"higher"`。

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` 应该返回 `"lower"`。

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` 应该返回 `"you got it!"`。

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
