---
id: 69f8c998d78ad3171a0713bf
title: "挑战 292：更宽的宽高比"
challengeType: 28
dashedName: challenge-292
---

# --description--

给定两个表示不同图像维度的字串，返回宽高比更大的图像的宽高比。

- 给定的字串格式为 `"WxH"`，例如 `"1920x1080"`。
- 宽高比是宽度与高度的比值，约简为最简整数比。例如，`"1920x1080"` 约简为 `"16:9"`。
- 返回格式为 `"W:H"` 的字串，例如 `"16:9"`。

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` 应该返回 `"16:9"`。

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` 应该返回 `"4:3"`。

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` 应该返回 `"2:1"`。

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` 应该返回 `"9:16"`。

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` 应该返回 `"43:18"`。

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` 应该返回 `"2089:8539"`。

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
