---
id: 69f90c1329a94b37e2a2086c
title: "挑战 293：最佳牌型"
challengeType: 28
dashedName: challenge-293
---

# --description--

给定一个包含五张扑克牌的字串数组，返回最佳牌型的名称。

- 每张牌用一个两个字符的字串表示：先是点数，后是花色，`"2h"` 例如。
  - 点数从低到高依次是：`"2"`、`"3"`、`"4"`、`"5"`、`"6"`、`"7"`、`"8"`、`"9"`、`"T"`、`"J"`、`"Q"`、`"K"` 和 `"A"`。
  - 花色有：`"h"`、`"d"`、`"c"` 和 `"s"`。
- A（`"A"`）在顺子中既可以作为最大点数，也可以作为最小点数。

牌型从差到好依次是：

| 名称 | 描述 |
| - | - |
| `"High Card"` | 无对子或更好牌型 |
| `"Pair"` | 一对 |
| `"Two Pair"` | 两对 |
| `"Three of a Kind"` | 三条 |
| `"Straight"` | 顺子（五张连续点数） |
| `"Flush"` | 同花（五张同花色） |
| `"Full House"` | 葫芦（三条加一对） |
| `"Four of a Kind"` | 铁支（四条） |
| `"Straight Flush"` | 同花顺（五张同花色且点数连续） |
| `"Royal Flush"` | `"A"`、`"K"`、`"Q"`、`"J"`、`"T"` 的同花色 |

返回最佳牌型的名称。

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` 应该返回 `"Three of a Kind"`。

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` 应该返回 `"Full House"`。

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` 应该返回 `"Flush"`。

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` 应该返回 `"Four of a Kind"`。

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` 应该返回 `"Two Pair"`。

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` 应该返回 `"Straight Flush"`。

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` 应该返回 `"High Card"`。

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` 应该返回 `"Straight"`。

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` 应该返回 `"Pair"`。

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` 应该返回 `"Royal Flush"`。

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
