---
id: 69f90c1329a94b37e2a2086d
title: "挑战 294：括号组合"
challengeType: 28
dashedName: challenge-294
---

# --description--

给定一个整数 `n`，返回 `n` 对括号的有效组合数。

- 有效组合是指每个左括号都有对应的右括号，且没有右括号出现在其匹配的左括号之前的字串。

例如，给定 `2`，有 2 种有效组合：

```json
(())
()()
```

# --hints--

`getCombinations(2)` 应该返回 `2`。

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` 应该返回 `5`。

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` 应该返回 `42`。

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` 应该返回 `1430`。

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` 应该返回 `742900`。

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
