---
id: 6a0dcd03ee4e68698080ef66
title: "挑战 302：时差反应"
challengeType: 28
dashedName: challenge-302
---

# --description--

给定出发城市、到达城市、飞行时长（小时）和旅行方向，返回旅客经历的时差小时数。

给定的城市将来自以下包含其 UTC 偏移的列表：

| 城市 | 偏移 |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

计算时差小时数的方法：

1. 查找两个城市之间的时区差（小时）。
2. 确定方向乘数。如果是 `"east"` 方向，乘数为 1.5，否则为 1.0。
3. 使用公式计算时差小时数：时区差 +（飞行时长 * 0.1）* 方向乘数

返回四舍五入到一位小数的时差小时数。

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` 应该返回 `6.5`。

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` 应该返回 `5.8`。

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` 应该返回 `1.6`。

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` 应该返回 `4.7`。

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` 应该返回 `17.5`。

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` 应该返回 `5.9`。

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` 应该返回 `9.5`。

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
