---
id: 6a0dcd03ee4e68698080ef68
title: "挑战 304：行程安排"
challengeType: 28
dashedName: challenge-304
---

# --description--

给定一个包含至少两个可选停靠点的日间旅行数组，返回有效行程安排的数量。

行程总是包括 `"breakfast"`、`"lunch"` 和 `"dinner"`，这些不会作为参数传入。可选停靠点可以放在行程中的任何位置，但必须遵守以下规则：

- `"breakfast"` 总是第一个，且在 `"lunch"` 之前至少有一个停靠点。
- `"lunch"` 必须出现在 `"dinner"` 之前，且两者之间至少有一个停靠点。
- `"dinner"` 之后最多只能有一个可选停靠点。

返回有效安排的数量。

# --hints--

`getItineraryCount(["library", "park"])` 应该返回 `2`。

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` 应该返回 `18`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` 应该返回 `120`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` 应该返回 `840`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` 应该返回 `55440`。

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
