---
id: 6a0dcc730cb92a616f86f0bf
title: "挑战 295：模式验证器第一部分"
challengeType: 28
dashedName: challenge-295
---

# --description--

给定一个对象（JavaScript）或字典（Python），判断它是否符合以下模式：

```json
{
  username: string
}
```

- 允许存在额外的密钥

# --hints--

`isValidSchema({ username: "bob" })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
