---
id: 6a0dcc730cb92a616f86f0c2
title: "挑战 298：模式验证器 第 4 部分"
challengeType: 28
dashedName: challenge-298
---

# --description--

给定一个对象（JavaScript）或字典（Python），判断它是否符合以下模式：

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- 管道符号（`|`）表示“或”。`role` 必须是列出的 `Roles` 之一。
- `?` 后面的问号（`supporter`）表示该字段是可选的，但如果存在，则必须是指定的类型。
- 允许存在额外的密钥

# --hints--

`isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true }));
```

`isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" }));
```

`isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 })` 应该返回 `true`。

```js
assert.isTrue(isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 }));
```

`isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" }));
```

`isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true }));
```

`isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false }));
```

`isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true }));
```

`isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false })` 应该返回 `false`。

```js
assert.isFalse(isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean')
  );
}
```
