---
id: 6a0dcd03ee4e68698080ef6b
title: "挑战 307：分区规定"
challengeType: 28
dashedName: challenge-307
---

# --description--

给定一个二维网格（数组的数组），表现一个城市的建筑版型，返回所有违反分区规定的建筑的坐标。

网格中的每个单元格包含下表中的一个标签。如果一个建筑的（最多）4 个邻居（水平或垂直方向）中有任何一个是它不能相邻的类型，则该建筑违反规定。

| 标签 | 类型 | 不能相邻的类型 |
| - | - | - |
| `"i"` | 工业 | `"R"`, `"I"` |
| `"A"` | 农业 | `"C"` |
| `"R"` | 住宅 | `"i"`, `"C"` |
| `"I"` | 机构 | `"i"` |
| `"C"` | 商业 | `"R"`, `"A"` |
| `""`（空字串） | 未开发 | 无限制 |

将所有违规单元格的坐标作为 `[row, col]` 点对数组返回，顺序不限。如果没有违规，返回空数组。

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` 应该返回 `[[0, 0], [0, 1]]`。

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` 应该返回 `[[0, 1], [1, 1]]`。

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` 应该返回 `[]`。

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` 应该返回 `[[0, 1], [0, 2], [0, 3]]`。

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` 应该返回 `[[2, 3], [2, 4], [3, 1], [3, 2]]`。

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
