---
id: 6a0dcd03ee4e68698080ef6c
title: "挑战 308：信用卡验证器"
challengeType: 28
dashedName: challenge-308
---

# --description--

给定一个信用卡号码的数字字串，使用以下方法判断它是否为有效卡号：

- 从倒数第二个数字开始，向左每隔一个数字翻倍。
- 如果翻倍后的数字大于 9，则减去 9。
- 将所有数字（翻倍和未翻倍的）相加。
- 如果总和能被 10 整除，则该号码有效。

# --hints--

`isValidCard("4532015112830366")` 应该返回 `true`。

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` 应该返回 `true`。

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` 应该返回 `true`。

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` 应该返回 `true`。

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` 应该返回 `false`。

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` 应该返回 `false`。

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` 应该返回 `false`。

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
