---
id: 6a0dcc730cb92a616f86f0c5
title: "挑战 301：剩余装载量"
challengeType: 28
dashedName: challenge-301
---

# --description--

给定你剩余的洗衣粉勺数和一个数组，数组中记录了你过去几天每天使用的勺数，返回你还能使用多少整天的洗衣粉。

根据使用历史计算你的平均每日用量，并假设未来每天都使用这个量。

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` 应该返回 `5`。

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` 应该返回 `7`。

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` 应该返回 `11`。

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` 应该返回 `12`。

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` 应该返回 `1`。

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
