---
id: 6a15cadf5f240d05a264955e
title: "挑战 320：血库"
challengeType: 28
dashedName: challenge-320
---

# --description--

给定一个血库库存的数组和一个患者血型请求的数组，返回一个格式为 `"X of Y patients served"` 的字串。其中 `X` 是能够从血库库存中输血的最大患者数，`Y` 是患者总数。

两个数组中的每个条目都是以下血型之一：`"AB"`、`"A"`、`"B"` 或 `"O"`。

兼容性规则：

- `"AB"` 可以接受任何血型的血液。
- `"A"` 可以接受 `"A"` 和 `"O"` 的血液。
- `"B"` 可以接受 `"B"` 和 `"O"` 的血液。
- `"O"` 只能接受 `"O"` 的血液。

给定数组中的重复条目表示数量。

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` 应该返回 `"4 of 4 patients served"`。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` 应该返回 `"3 of 5 patients served"`。

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` 应该返回 `"4 of 5 patients served"`。

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` 应该返回 `"4 of 4 patients served"`。

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` 应该返回 `"8 of 11 patients served"`。

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` 应该返回 `"13 of 14 patients served"`。

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
