---
id: 6a151d32d772271dd2448c2d
title: "挑战 311：咒语施放者"
challengeType: 28
dashedName: challenge-311
---

# --description--

给定一个你正在施放的咒语代码字串，计算总分。

字串中的每个字符代表一个咒语：

| 代码 | 咒语 | 类别 | 基础分数 |
| - | - | - | - |
| `"f"` | 火焰 | 破坏 | 3 |
| `"l"` | 闪电 | 破坏 | 3 |
| `"i"` | 冰霜 | 控件 | 2 |
| `"w"` | 风 | 控件 | 2 |
| `"h"` | 治疗 | 恢复 | 1 |
| `"s"` | 盾牌 | 恢复 | 1 |

根据连续施放的咒语类别不同，应用连击倍数：

- 第一个咒语总是按基础分数计分。
- 每当连续施放的咒语类别与前一个不同，连击倍数加 1。
- 施放与前一个咒语类别相同的咒语时，连击倍数重置为 1。
- 每个咒语的得分是其基础分数乘以当前连击倍数。

返回这串咒语的总分。

# --hints--

`cast("fihwl")` 应该返回 `33`。

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` 应该返回 `45`。

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` 应该返回 `37`。

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` 应该返回 `50`。

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` 应该返回 `101`。

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
