---
id: 6a15cadf5f240d05a2649556
title: "挑战 312：流媒体费用"
challengeType: 28
dashedName: challenge-312
---

# --description--

给定一个表示你流媒体服务购物车中电影的数组，以及一个表示你的订阅等级的字串，返回电影的总费用。

购物车中的每一项都是一个包含 `"format"`（`"HD"` 或 `"4K"`）和 `"type"`（`"rent"` 或 `"buy"`）的对象。它们的费用如下：

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

应用以下订阅等级折扣：

- `"none"`：全价
- `"basic"`：九折
- `"premium"`：七五折

返回总费用，四舍五入到小数点后两位，格式为 `"$D.CC"`。

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` 应该返回 `"$3.99"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "buy" }], "premium")` 应该返回 `"$9.74"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "buy" }], "premium"), "$9.74");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` 应该返回 `"$18.87"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` 应该返回 `"$44.98"`。

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$44.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` 应该返回 `"$42.96"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic")` 应该返回 `"$38.66"`。

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic"), "$38.66");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
