---
id: 6a15cadf5f240d05a264955c
title: "挑战 318：DNA 突变"
challengeType: 28
dashedName: challenge-318
---

# --description--

给定两个长度相等的 DNA 链，返回一个数组，包含链中不同（突变）的位置索引。

- DNA 链是由字符 `"A"`、`"T"`、`"C"` 和 `"G"` 组成的字串
- 返回的索引按升序排列
- 如果没有突变，返回空数组

# --hints--

`detectMutations("ATCG", "ATGG")` 应该返回 `[2]`。

```js
assert.deepEqual(detectMutations("ATCG", "ATGG"), [2]);
```

`detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` 应该返回 `[4, 9, 11]`。

```js
assert.deepEqual(detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11]);
```

`detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` 应该返回 `[]`。

```js
assert.deepEqual(detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), []);
```

`detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` 应该返回 `[15, 16, 17, 18]`。

```js
assert.deepEqual(detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18]);
```

`detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` 应该返回 `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`。

```js
assert.deepEqual(detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]);
```

# --seed--

## --seed-contents--

```js
function detectMutations(strand1, strand2) {

  return strand1;
}
```

# --solutions--

```js
function detectMutations(strand1, strand2) {
  return strand1
    .split("")
    .reduce((mutations, base, i) => {
      if (base !== strand2[i]) mutations.push(i);
      return mutations;
    }, []);
}
```
