---
id: 6a151d32d772271dd2448c2c
title: "挑战 310：英式英语转美式英语"
challengeType: 28
dashedName: challenge-310
---

# --description--

给定一个句子，使用下表将任何英式英语拼写转换为美式英语对应拼写，并返回更新后的句子：

| 英式 | 美式 |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- 替换应忽略大小写。例如，`"Colour"` 应该变成 `"Color"`。
- 输入中可能包含基于表中根词的单词，这些单词也需要被更改。例如，`"colouring"` 应该变成 `"coloring"`，`"disorganised"` 应该变成 `"disorganized"`。

# --hints--

`britishToAmerican("I love the colour blue.")` 应该返回 `"I love the color blue."`

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` 应该返回 `"The fiber optic cable is new."`

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` 应该返回 `"It's an honor to meet someone with such humor."`

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` 应该返回 `"The unrecognized artist analyzed his color palette at the center."`

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` 应该返回 `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
