---
id: 6a15cadf5f240d05a264955d
title: "挑战 319：Frontmatter 解析器"
challengeType: 28
dashedName: challenge-319
---

# --description--

给定一个表示 frontmatter 块的字串，解析它并返回一个包含密钥和值的对象（JavaScript）或字典（Python）。

Frontmatter 用 `---` 分隔符包裹，里面包含每行一个的 `key: value` 点对。例如：

```md
---
title: My Post
draft: false
views: 100
---
```

应返回：

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 数字、布尔和字串都应以它们各自的类型返回。
- 给定的字串中的换行符用换行字符（`"\n"`）分隔。上面的例子会给出为：`"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`。

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` 应返回 `{ title: "My Post", draft: false, views: 100 }`。

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` 应返回 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`。

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` 应返回 `{ version: "1.0.0", url: "https://example.com", private: true }`。

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` 应返回 `{ rating: 4.5, price: 9.99 }`。

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
