---
id: 6a15cadf5f240d05a264955a
title: "挑战 316：1337 语言"
challengeType: 28
dashedName: challenge-316
---

# --description--

给定一个小写字串，通过将下面的字母替换为它们的 1337 代换，返回转换后的 1337 语言字串：

| 字母 | 1337 |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- 没有代换的字符保持不变。

# --hints--

`makeLeet("cool")` 应该返回 `"c001"`。

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` 应该返回 `"1337"`。

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` 应该返回 `"h4ck3r"`。

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` 应该返回 `"547311173"`。

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` 应该返回 `"4bcd3f9h1jk1mn0pqr57uvwxyz"`。

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
