---
id: 6a15cadf5f240d05a264955b
title: "挑战 317：BMI 计算器"
challengeType: 28
dashedName: challenge-317
---

# --description--

给定以磅为单位的体重和以英寸为单位的身高，返回四舍五入到一位小数的 BMI（体质指数）。

计算 BMI：用体重除以身高的平方，然后将结果乘以 703。

# --hints--

`calculateBmi(180, 70)` 应该返回 `25.8`。

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` 应该返回 `24.0`。

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` 应该返回 `19.5`。

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` 应该返回 `39.1`。

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` 应该返回 `22.8`。

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
