---
id: 6a15cadf5f240d05a2649558
title: "挑战 314：素数分解"
challengeType: 28
dashedName: challenge-314
---

# --description--

给定一个大于 1 的整数，返回它的素数分解，结果是一个按升序排列的数字数组。

素数分解是一组素数，这些素数相乘得到给定的整数。每个数字只有唯一的一组。例如，20 的素数分解是 `[2, 2, 5]`，因为 2 * 2 * 5 = 20。

如果给定的整数本身是素数，则以单元素数组返回它。

# --hints--

`primeFactorization(20)` 应该返回 `[2, 2, 5]`。

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` 应该返回 `[17]`。

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` 应该返回 `[3, 5]`。

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` 应该返回 `[5, 7]`。

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` 应该返回 `[3, 3, 3, 37]`。

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` 应该返回 `[2, 2, 2, 3, 3, 5]`。

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` 应该返回 `[2, 3, 5, 7, 11, 13, 17]`。

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
