---
id: 6a19b1062d1b153d8ac76d72
title: "挑战 324：副本字符计数"
challengeType: 28
dashedName: challenge-324
---

# --description--

给定两个字串，返回第二个字串中可以在第一个字串中找到的字符数量。

- 第二个字串中的副本字符分别计数。

# --hints--

`duplicateCharacterCount("aloha", "hei")` 应该返回 `1`。

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` 应该返回 `4`。

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` 应该返回 `3`。

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` 应该返回 `0`。

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` 应该返回 `5`。

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` 应该返回 `2`。

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` 应该返回 `26`。

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
