---
id: 6a19b1062d1b153d8ac76d70
title: "挑战 322：连接 3"
challengeType: 28
dashedName: challenge-322
---

# --description--

给定一个表现游戏网格上棋子的字串矩阵，判断是否有玩家连成三子。

- 每个格子包含 `"R"`、`"Y"` 或 `""`（空字串）。
- 连成三子指的是横向、纵向或斜向连续三个相同类型且非空的格子。

返回：

- 一个扁平数组，包含获胜者和他们三个获胜格子的坐标，格式为：`["R", [0,2], [1,3], [2,4]]`。坐标按从上到下、再从左到右的顺序返回。
- 如果没有获胜者，则返回空数组。

# --hints--

`connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])` 应该返回 `["R", [3, 1], [3, 2], [3, 3]]`。

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]]);
```

`connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])` 应该返回 `["Y", [1, 1], [2, 1], [3, 1]]`。

```js
assert.deepEqual(connectThree([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]]);
```

`connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])` 应该返回 `["R", [0, 3], [1, 2], [2, 1]]`。

```js
assert.deepEqual(connectThree([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]]);
```

`connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])` 应该返回 `["Y", [0, 1], [1, 2], [2, 3]]`。

```js
assert.deepEqual(connectThree([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]]);
```

`connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])` 应该返回 `[]`。

```js
assert.deepEqual(connectThree([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), []);
```

# --seed--

## --seed-contents--

```js
function connectThree(matrix) {

  return matrix;
}
```

# --solutions--

```js
function connectThree(matrix) {
  const rows = matrix.length;
  const cols = matrix[0].length;

  const directions = [[0, 1], [1, 0], [1, 1], [1, -1]];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const piece = matrix[r][c];
      if (!piece) continue;

      for (const [dr, dc] of directions) {
        const cells = [[r, c]];

        for (let i = 1; i < 3; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols) break;
          if (matrix[nr][nc] !== piece) break;
          cells.push([nr, nc]);
        }

        if (cells.length === 3) {
          cells.sort(([r1, c1], [r2, c2]) => r1 !== r2 ? r1 - r2 : c1 - c2);
          return [piece, ...cells];
        }
      }
    }
  }

  return [];
}
```
