---
id: 6a19b1062d1b153d8ac76d73
title: "挑战 325：幸运数字"
challengeType: 28
dashedName: challenge-325
---

# --description--

给定一个包含某人名字和姓氏的字串，使用以下规则计算他们的幸运数字：

- 名字和姓氏之间用空格分隔
- 查找每个名字中的元音和构造数量
- 将较小的元音和构造数量相乘，再乘以较短名字的长度
- 对较大的两个数量和较长的名字做同样的操作
- 用较大的值减去较小的值，得到他们的幸运数字

如果最终值是零（`0`），则返回 `13`。

# --hints--

`getLuckyNumber("John Doe")` 应该返回 `21`。

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` 应该返回 `52`。

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` 应该返回 `18`。

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` 应该返回 `81`。

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` 应该返回 `32`。

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` 应该返回 `13`。

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
