---
id: 6a19dd69388e5a6e59e4f2ef
title: "挑战 326：最大利润"
challengeType: 28
dashedName: challenge-326
---

# --description--

给定一个每日股票价格的数组和一个预算（以美元计），计算在给定期间内通过买卖股票你能获得的最大利润。

- 你只能在买入后卖出。
- 你只能购买整股。
- 以字串形式返回最大可能利润，向下取整到最接近的分，并格式化为两位小数。

# --hints--

`getMaxProfit([5, 6], 50)` 应该返回 `"10.00"`。

```js
assert.equal(getMaxProfit([5, 6], 50), "10.00");
```

`getMaxProfit([8, 2, 5, 10], 20)` 应该返回 `"80.00"`。

```js
assert.equal(getMaxProfit([8, 2, 5, 10], 20), "80.00");
```

`getMaxProfit([4, 5, 3, 6], 20)` 应该返回 `"18.00"`。

```js
assert.equal(getMaxProfit([4, 5, 3, 6], 20), "18.00");
```

`getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)` 应该返回 `"8.31"`。

```js
assert.equal(getMaxProfit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31");
```

`getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80)` 应该返回 `"0.00"`。

```js
assert.equal(getMaxProfit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00");
```

`getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)` 应该返回 `"73.80"`。

```js
assert.equal(getMaxProfit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80");
```

# --seed--

## --seed-contents--

```js
function getMaxProfit(prices, budget) {

  return prices;
}
```

# --solutions--

```js
function getMaxProfit(prices, budget) {
  let maxProfit = 0;

  for (let buy = 0; buy < prices.length - 1; buy++) {
    for (let sell = buy + 1; sell < prices.length; sell++) {
      const shares = Math.floor(budget / prices[buy]);
      const profit = shares * (prices[sell] - prices[buy]);
      if (profit > maxProfit) maxProfit = profit;
    }
  }

  return (Math.floor(maxProfit * 100) / 100).toFixed(2);
}
```
