---
id: 6a1d9f98e819ed70a0e994d9
title: "挑战 327：数据库迁移"
challengeType: 28
dashedName: challenge-327
---

# --description--

给定两个数据库对象，返回第二个对象，并用第一个对象中缺失的属性填充它。

- 已经存在于记录中的字段不应被覆盖。

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` 应该返回 `{ username: "", posts: 0, verified: true }`。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })` 应该返回 `{ username: "camper", posts: 5 }`。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })` 应该返回 `{ username: "camper", posts: 0, verified: false }`。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })` 应该返回 `{ username: "camper", role: "admin", posts: 0 }`。

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })` 应该返回 `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }`。

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
