---
id: 6a1d9f98e819ed70a0e994e1
title: "挑战 335：五个骰子"
challengeType: 28
dashedName: challenge-335
---

# --description--

给定一个包含五个骰子（点数为 1-6）的数组，返回可能的最佳牌型。

以下是牌型从低到高的排名：

| 牌型 | 描述 |
|------|-------------|
| `"no pair"` | 无对子或更好 |
| `"pair"` | 两个点数相同的骰子 |
| `"two pair"` | 两个不同的对子 |
| `"three of a kind"` | 三个点数相同的骰子 |
| `"small straight"` | 四个连续点数 |
| `"large straight"` | 五个连续点数 |
| `"full house"` | 三个相同点数加一对 |
| `"four of a kind"` | 四个点数相同的骰子 |
| `"five of a kind"` | 五个点数相同的骰子 |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` 应该返回 `"five of a kind"`。

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` 应该返回 `"four of a kind"`。

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` 应该返回 `"large straight"`。

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` 应该返回 `"three of a kind"`。

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` 应该返回 `"pair"`。

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` 应该返回 `"no pair"`。

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` 应该返回 `"small straight"`。

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` 应该返回 `"full house"`。

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` 应该返回 `"two pair"`。

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
