---
id: 6a1d9f98e819ed70a0e994da
title: "挑战 328：Kaprekar 例程"
challengeType: 28
dashedName: challenge-328
---

# --description--

给定一个 4 位数，返回你需要应用 Kaprekar 例程直到达到 6174 的次数。

Kaprekar 例程的步骤如下：

- 将数字按降序排列，组成最大数
- 将数字按升序排列，组成最小数（必要时用前导零补齐）
- 用较大数减去较小数
- 用新数字重复以上步骤

# --hints--

`kaprekar(1234)` 应该返回 `3`。

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` 应该返回 `6`。

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` 应该返回 `4`。

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` 应该返回 `7`。

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` 应该返回 `2`。

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
