---
id: 6a1d9f98e819ed70a0e994e0
title: "挑战 334：精确找零"
challengeType: 28
dashedName: challenge-334
---

# --description--

给定一个以分为单位的整数金额，返回使用便士（1 分）、镍币（5 分）、一角硬币（10 分）和四分之一美元硬币（25 分）组成精确找零的不同方式数。

# --hints--

`exactChange(3)` 应该返回 `1`。

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` 应该返回 `2`。

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` 应该返回 `6`。

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` 应该返回 `24`。

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` 应该返回 `73`。

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` 应该返回 `213`。

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
