---
id: 6a1d9f98e819ed70a0e994dd
title: "挑战 331：最接近的倍数"
challengeType: 28
dashedName: challenge-331
---

# --description--

给定两个整数，将第一个整数四舍五入到最接近的第二个整数的倍数。

# --hints--

`roundToNearestMultiple(5, 3)` 应该返回 `6`。

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` 应该返回 `16`。

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` 应该返回 `45`。

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` 应该返回 `33`。

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` 应该返回 `96`。

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
