---
id: 6a1d9f98e819ed70a0e994df
title: "挑战 333：问题分类 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

给定一个问题标题和当前标签的数组，根据以下规则返回更新后的标签数组：

如果该问题没有任何标签，添加：

- 如果标题包含 `"error"` 或 `"bug"`，添加 `"bug"` 和 `"needs triage"`
- 如果标题包含 `"feature"` 或 `"add"`，添加 `"enhancement"` 和 `"discussing"`

否则，如果给定的标签包含：

- 如果包含 `"needs triage"` 且标题包含 `"simple"` 或 `"easy"`，移除 `"needs triage"` 并添加 `"good first issue"`
- 如果包含 `"discussing"` 且标题包含 `"planned"` 或 `"next"`，移除 `"discussing"` 并添加 `"on the roadmap"`
- 否则，如果存在 `"needs triage"` 或 `"discussing"`，移除它并添加 `"help wanted"`

如果标题包含：

- `"security"`，添加一个 `"critical"` 标签

# --hints--

`triageIssue("app crashes with error", [])` 应该返回 `["bug", "needs triage"]`。

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` 应该返回 `["bug", "help wanted"]`。

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` 应该返回 `["enhancement", "discussing"]`。

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` 应该返回 `["enhancement", "help wanted"]`。

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` 应该返回 `["bug", "needs triage", "critical"]`。

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` 应该返回 `["critical"]`。

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` 应该返回 `["bug", "good first issue"]`。

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` 应该返回 `["enhancement", "on the roadmap"]`。

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` 应该返回 `["enhancement", "help wanted", "critical"]`。

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
