---
id: 6a1d9f98e819ed70a0e994dc
title: "挑战 330：小写字串"
challengeType: 28
dashedName: challenge-330
---

# --description--

给定一个字串，只返回完全是小写的单词，保持它们的原始顺序，并且单词之间用空格分隔。

# --hints--

`getLowercaseWords("hello GOOD world")` 应该返回 `"hello world"`。

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` 应该返回 `"these are all lowercase"`。

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` 应该返回 `"less is more"`。

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` 应该返回 `"eat pizza every day"`。

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` 应该返回 `"the quick brown fox jumped over the lazy dog"`。

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
