---
id: 6a2037a68a0bc2aef0006006
title: "挑战 344：黄金比例"
challengeType: 28
dashedName: challenge-344
---

# --description--

给定两个数字，判断它们的比值是否接近黄金比例。

- 使用黄金比例 `1.618`
- 允许误差为 `0.01`

# --hints--

`isGoldenRatio(21, 34)` 应该返回 `true`。

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` 应该返回 `false`。

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` 应该返回 `true`。

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` 应该返回 `false`。

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` 应该返回 `true`。

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` 应该返回 `false`。

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
