---
id: 6a2037a68a0bc2aef0006008
title: "挑战 346：储蓄罐"
challengeType: 28
dashedName: challenge-346
---

# --description--

给定一个表现储蓄罐的对象，返回总价值，格式为 `"$D.CC"` 的字串。

该对象可能包含以下任意内容：

| 硬币 | 价值 |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 })` 应该返回 `"$1.98"`。

```js
assert.equal(piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 }), "$1.98");
```

`piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 })` 应该返回 `"$0.41"`。

```js
assert.equal(piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 }), "$0.41");
```

`piggyBank({ nickels: 8, dimes: 6, quarters: 5 })` 应该返回 `"$2.25"`。

```js
assert.equal(piggyBank({ nickels: 8, dimes: 6, quarters: 5 }), "$2.25");
```

`piggyBank({})` 应该返回 `"$0.00"`。

```js
assert.equal(piggyBank({}), "$0.00");
```

`piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 })` 应该返回 `"$6.76"`。

```js
assert.equal(piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 }), "$6.76");
```

# --seed--

## --seed-contents--

```js
function piggyBank(coins) {

  return coins;
}
```

# --solutions--

```js
function piggyBank(coins) {
  const values = {
    pennies: 0.01,
    nickels: 0.05,
    dimes: 0.10,
    quarters: 0.25,
  };

  const total = Object.entries(coins).reduce((sum, [coin, count]) => {
    return sum + (values[coin] * count);
  }, 0);

  return "$" + total.toFixed(2);
}
```
