---
id: 6a2037a68a0bc2aef0006004
title: "挑战 342：骰子概率"
challengeType: 28
dashedName: challenge-342
---

# --description--

给定要掷的六面骰数量和一个目标和，返回以字符串形式表示掷出该和的概率，格式为 `"1 in X"`。

- 骰子数量在 1 到 6 之间。
- 给定数量的骰子总能达到目标和。
- 将 `"X"` 四舍五入到最接近的整数。

# --hints--

`getOdds(1, 5)` 应该返回 `"1 in 6"`。

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` 应该返回 `"1 in 12"`。

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` 应该返回 `"1 in 8"`。

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` 应该返回 `"1 in 65"`。

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` 应该返回 `"1 in 111"`。

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` 应该返回 `"1 in 7776"`。

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
