---
id: 6a2037a68a0bc2aef0006000
title: "挑战 338：宠物年龄计算器"
challengeType: 28
dashedName: challenge-338
---

# --description--

给定宠物的类型和以人类年龄表示的年龄，使用以下换算表返回相当的宠物年龄：

| 宠物 | 乘数 |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` 应该返回 `35`。

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` 应该返回 `54`。

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` 应该返回 `24`。

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` 应该返回 `120`。

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` 应该返回 `60`。

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` 应该返回 `12`。

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` 应该返回 `5`。

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
