---
id: 6a2037a68a0bc2aef0005fff
title: "挑战 337：计数器"
challengeType: 28
dashedName: challenge-337
---

# --description--

给定一串计数符号，返回所表现的总计数。

- 每个管道 `"|"` 表示一个计数。
- 每第五个符号用一个转发斜线 `"/"` 表现，组成五个一组（`"||||/"`）。
- 组与组之间用空格分隔。

# --hints--

`getTallyCount("||||")` 应该返回 `4`。

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` 应该返回 `5`。

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` 应该返回 `8`。

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` 应该返回 `17`。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` 应该返回 `41`。

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
