---
id: 6a2037a68a0bc2aef0006002
title: "挑战 340：猪拉丁文转换器"
challengeType: 28
dashedName: challenge-340
---

# --description--

给定一个字串，使用以下规则将其转换为猪拉丁文：

- 如果一个单词以元音（`"a"`、`"e"`、`"i"`、`"o"` 或 `"u"`）开头，添加 `"way"` 到末尾。例如，`"universe"` 转换为 `"universeway"`。
- 如果一个单词以一个或多个辅音开头，将它们移到末尾并添加 `"ay"`。例如，`"hello"` 转换为 `"ellohay"`。
- 保持首字母的大小写。例如，`"Hello"` 转换为 `"Ellohay"`。

# --hints--

`pigLatin("universe")` 应该返回 `"universeway"`。

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` 应该返回 `"ellohay"`。

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` 应该返回 `"ellohay universeway"`。

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` 应该返回 `"Ellohay universeway"`。

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` 应该返回 `"Igpay Atinlay isway unfay"`。

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` 应该返回 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`。

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
