---
id: 6a2037a68a0bc2aef0006007
title: "挑战 345：单词混合器"
challengeType: 28
dashedName: challenge-345
---

# --description--

给定两个单词，通过将第一个单词的前半部分与第二个单词的后半部分组合，返回一个新单词。

- 对于长度为奇数的单词，前半部分较短。

# --hints--

`blendWords("turtle", "toucan")` 应该返回 `"turcan"`。

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` 应该返回 `"chipingo"`。

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` 应该返回 `"falican"`。

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` 应该返回 `"hyana"`。

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` 应该返回 `"scorilla"`。

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` 应该返回 `"platerine"`。

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
