---
id: 6a22d77ddf034bc4e35b1d5b
title: "挑战 354：对比度评级 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

给定两个表示 RGB 值的数组和一个布尔值表示文本是否为大号，使用以下方法返回 WCAG 对比度评级：

首先，将每个 RGB 值转换为相对亮度：

- 将每个通道 `[R, G, B]` 除以 255，得到介于 `0` 和 `1` 之间的值
- 对每个通道应用伽马校正公式：
  - 如果通道值小于或等于 `0.04045`：`channel / 12.92`
  - 否则：`((channel + 0.055) / 1.055) ^ 2.4`
- 计算亮度：`0.2126 * R + 0.7152 * G + 0.0722 * B`

然后，通过给每个亮度值加 0.05，再将较亮的值除以较暗的值来计算对比度比率。较亮的值总是第一个参数。

根据对比度比率使用下表返回评级：

| 评级 | 正常文本 | 大号文本 |
|--------|-------------|------------|
| `"AAA"` | 7.0 及以上 | 4.5 及以上 |
| `"AA"` | 4.5 及以上 | 3.0 及以上 |
| `"Fail"` | 低于 4.5 | 低于 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` 应返回 `"AAA"`。

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` 应返回 `"AA"`。

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` 应返回 `"Fail"`。

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` 应返回 `"AAA"`。

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` 应返回 `"AA"`。

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` 应返回 `"Fail"`。

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
