---
id: 6a22d77ddf034bc4e35b1d5c
title: "挑战 355：摩斯电码"
challengeType: 28
dashedName: challenge-355
---

# --description--

给定一个摩斯电码字串，使用下表返回解码后的消息：

| 代码 | 字母 | 代码 | 字母 |
|------|--------|------|--------|
| `.-` | `A` | `-.` | `N` |
| `-...` | `B` | `---` | `O` |
| `-.-.` | `C` | `.--.` | `P` |
| `-..` | `D` | `--.-` | `Q` |
| `.` | `E` | `.-.` | `R` |
| `..-.` | `F` | `...` | `S` |
| `--.` | `G` | `-` | `T` |
| `....` | `H` | `..-` | `U` |
| `..` | `I` | `...-` | `V` |
| `.---` | `J` | `.--` | `W` |
| `-.-` | `K` | `-..-` | `X` |
| `.-..` | `L` | `-.--` | `Y` |
| `--` | `M` | `--..` | `Z` |

- 字母之间用一个空格分隔
- 单词之间用三个空格分隔

# --hints--

`decodeMorse("--..")` 应该返回 `"Z"`。

```js
assert.equal(decodeMorse("--.."), "Z");
```

`decodeMorse("... --- ...")` 应该返回 `"SOS"`。

```js
assert.equal(decodeMorse("... --- ..."), "SOS");
```

`decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--.")` 应该返回 `"FREECODECAMP"`。

```js
assert.equal(decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--."), "FREECODECAMP");
```

`decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -..")` 应该返回 `"HELLO WORLD"`。

```js
assert.equal(decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -.."), "HELLO WORLD");
```

`decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --.")` 应该返回 `"THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG"`。

```js
assert.equal(decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --."), "THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG");
```

# --seed--

## --seed-contents--

```js
function decodeMorse(code) {

  return code;
}
```

# --solutions--

```js
function decodeMorse(code) {
  const table = {
    '.-': 'A', '-...': 'B', '-.-.': 'C', '-..': 'D', '.': 'E',
    '..-.': 'F', '--.': 'G', '....': 'H', '..': 'I', '.---': 'J',
    '-.-': 'K', '.-..': 'L', '--': 'M', '-.': 'N', '---': 'O',
    '.--.': 'P', '--.-': 'Q', '.-.': 'R', '...': 'S', '-': 'T',
    '..-': 'U', '...-': 'V', '.--': 'W', '-..-': 'X', '-.--': 'Y',
    '--..': 'Z'
  };

  return code.split('   ')
    .map(word => word.split(' ').map(c => table[c]).join(''))
    .join(' ');
}
```
