---
id: 6a22d77ddf034bc4e35b1d54
title: "挑战 347：博弈论"
challengeType: 28
dashedName: challenge-347
---

# --description--

给定两个等长的字串，分别表现两个玩家在一个游戏中的策略，返回得分数组 `[player1, player2]`。

- 给定的字串只包含两个字母之一：`"C"`（合作）或 `"D"`（背叛）。
- 每个字符代表一轮，得分规则如下：
  - 如果双方都合作，每人得 3 分。
  - 如果双方都背叛，每人得 1 分。
  - 如果一方背叛另一方合作，背叛者得 5 分，合作方得 0 分。

# --hints--

`playGame("CCCC", "CCCC")` 应该返回 `[12, 12]`。

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` 应该返回 `[4, 4]`。

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` 应该返回 `[5, 10]`。

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` 应该返回 `[24, 34]`。

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` 应该返回 `[66, 21]`。

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
