---
id: 6a22d77ddf034bc4e35b1d56
title: "挑战 349：手机信号"
challengeType: 28
dashedName: challenge-349
---

# --description--

给定一个包含三个手机信号塔读数的网格，确定手机的位置。

- 网格中的每个单元格要么是 `0`（无信号塔），要么是一个正整数，表示从该单元格到手机的单元格数，测量方式为直线：水平、垂直或对角线。
- 返回距离所有三个信号塔单元格数都正确的单元格的 `[row, col]`。
- 总是恰好有一个方案。

# --hints--

`findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` 应该返回 `[1, 2]`。

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

`findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` 应该返回 `[2, 1]`。

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

`findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` 应该返回 `[2, 2]`。

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

`findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` 应该返回 `[3, 4]`。

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

`findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` 应该返回 `[3, 3]`。

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
