---
id: 6a22d77ddf034bc4e35b1d57
title: "挑战 350：字母距离"
challengeType: 28
dashedName: challenge-350
---

# --description--

给定两个长度相等的字串，返回每对字符之间最短距离的总和。

- 输入只包含小写字母
- 字母表被视为一个圆环，因此 `a` 和 `z` 之间的距离是 1。

# --hints--

`letterDistance("abc", "bcd")` 应该返回 `3`。

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` 应该返回 `9`。

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` 应该返回 `10`。

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` 应该返回 `43`。

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` 应该返回 `47`。

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` 应该返回 `55`。

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
