---
id: 6a22d77ddf034bc4e35b1d59
title: "挑战 352：对比度评级 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

给定一个对比度比率和一个布尔值表示文本是否为大号，使用下表返回 WCAG 评级：

| 评级 | 正常文本 | 大号文本 |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | 低于 4.5 | 低于 3.0 |

# --hints--

`getContrastRating("7.5", false)` 应该返回 `"AAA"`。

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` 应该返回 `"AA"`。

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` 应该返回 `"Fail"`。

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` 应该返回 `"AAA"`。

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` 应该返回 `"AA"`。

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` 应该返回 `"Fail"`。

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
