---
id: 6a22d77ddf034bc4e35b1d58
title: "挑战 351：积数"
challengeType: 28
dashedName: challenge-351
---

# --description--

给定一个数字，判断它是否是积数。

积数是两个连续整数的乘积。例如，6 是积数，因为 2 * 3 = 6。

# --hints--

`isPronic(6)` 应该返回 `true`。

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` 应该返回 `false`。

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` 应该返回 `true`。

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` 应该返回 `true`。

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` 应该返回 `false`。

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` 应该返回 `true`。

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
