---
id: 6a26df95efa55a2524399746
title: "挑战 365：最后的挑战：桶填充 3"
challengeType: 28
dashedName: challenge-365
---

# --description--

今天标志着每日编码挑战的一周年。这是目前最后一个新挑战。祝你好运！

给定一个由单字母颜色字串组成的二维网格和一个目标颜色，返回将整个网格变成该颜色所需的最少洪泛填充“点击”次数。

- 每次点击会改变被点击单元格的颜色以及所有与其相连且颜色相同的单元格区域（四方向相连）。
- 点击可以使用任何颜色作为中间步骤，而不仅限于目标颜色。

# --hints--

`bucketFill([["B", "B"], ["B", "B"]], "R")` 应该返回 `1`。

```js
assert.equal(bucketFill([["B", "B"], ["B", "B"]], "R"), 1);
```

`bucketFill([["G", "G", "G"], ["G", "G", "G"], ["G", "G", "G"]], "G")` 应该返回 `0`。

```js
assert.equal(bucketFill([["G", "G", "G"], ["G", "G", "G"], ["G", "G", "G"]], "G"), 0);
```

`bucketFill([["P", "P", "Y"], ["Y", "P", "Y"], ["Y", "P", "P"]], "O")` 应该返回 `2`。

```js
assert.equal(bucketFill([["P", "P", "Y"], ["Y", "P", "Y"], ["Y", "P", "P"]], "O"), 2);
```

`bucketFill([["G", "Y", "C", "C"], ["Y", "Y", "Y", "B"], ["C", "Y", "B", "B"], ["C", "B", "B", "C"]], "R")` 应该返回 `4`。

```js
assert.equal(bucketFill([["G", "Y", "C", "C"], ["Y", "Y", "Y", "B"], ["C", "Y", "B", "B"], ["C", "B", "B", "C"]], "R"), 4);
```

`bucketFill([["G", "G", "O", "O"], ["G", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["G", "G", "G", "G"]], "P")` 应该返回 `5`。

```js
assert.equal(bucketFill([["G", "G", "O", "O"], ["G", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["G", "G", "G", "G"]], "P"), 5);
```

`bucketFill([["R", "G", "R", "G"], ["R", "G", "R", "G"], ["B", "B", "B", "B"], ["B", "B", "B", "B"], ["R", "G", "R", "G"]], "Y")` 应该返回 `3`。

```js
assert.equal(bucketFill([["R", "G", "R", "G"], ["R", "G", "R", "G"], ["B", "B", "B", "B"], ["B", "B", "B", "B"], ["R", "G", "R", "G"]], "Y"), 3);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, targetColor) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, targetColor) {
  const rows = grid.length;
  const cols = grid[0].length;

  function gridToString(g) {
    return g.map(r => r.join(",")).join("|");
  }

  function getRegion(g, row, col) {
    const color = g[row][col];
    const visited = new Set();
    const stack = [[row, col]];
    while (stack.length) {
      const [r, c] = stack.pop();
      const key = `${r},${c}`;
      if (visited.has(key)) continue;
      if (r < 0 || r >= rows || c < 0 || c >= cols) continue;
      if (g[r][c] !== color) continue;
      visited.add(key);
      stack.push([r+1,c],[r-1,c],[r,c+1],[r,c-1]);
    }
    return visited;
  }

  function floodFill(g, row, col, color) {
    const next = g.map(r => [...r]);
    const region = getRegion(g, row, col);
    for (const key of region) {
      const [r, c] = key.split(",").map(Number);
      next[r][c] = color;
    }
    return next;
  }

  function isComplete(g) {
    return g.every(row => row.every(cell => cell === targetColor));
  }

  function getColors(g) {
    return [...new Set([...g.flat(), targetColor])];
  }

  function getRegionRoots(g) {
    const visited = new Set();
    const roots = [];
    for (let r = 0; r < rows; r++) {
      for (let c = 0; c < cols; c++) {
        const key = `${r},${c}`;
        if (!visited.has(key)) {
          const region = getRegion(g, r, c);
          for (const k of region) visited.add(k);
          roots.push([r, c]);
        }
      }
    }
    return roots;
  }

  const initial = grid.map(r => [...r]);
  if (isComplete(initial)) return 0;

  const queue = [[initial, 0]];
  const seen = new Set([gridToString(initial)]);

  while (queue.length) {
    const [current, clicks] = queue.shift();
    const colors = getColors(current);
    const roots = getRegionRoots(current);

    for (const [r, c] of roots) {
      for (const color of colors) {
        if (color === current[r][c]) continue;
        const next = floodFill(current, r, c, color);
        if (isComplete(next)) return clicks + 1;
        const key = gridToString(next);
        if (!seen.has(key)) {
          seen.add(key);
          queue.push([next, clicks + 1]);
        }
      }
    }
  }
}
```
