---
id: 6a26df95efa55a2524399745
title: "挑战 364：两个桶之间"
challengeType: 28
dashedName: challenge-364
---

# --description--

给定两个油漆桶，每个都有一个 RGB 颜色和一个满度等级，返回混合后的 RGB 颜色，作为一个包含三个整数的数组。

- 每个桶是一个对象（JavaScript）或字典（Python），具有一个 `color` 属性（一个包含三个整数 `[r, g, b]` 的数组）和一个 `fullness` 属性（0–100）。
- 混合颜色是基于满度等级对两个颜色中每个通道的加权平均，每个通道四舍五入到最接近的整数。

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` 应该返回 `[125, 125, 125]`。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` 应该返回 `[200, 200, 200]`。

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` 应该返回 `[100, 150, 200]`。

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` 应该返回 `[114, 184, 170]`。

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` 应该返回 `[68, 163, 123]`。

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
