---
id: 6a26df3e2988bcdded204897
title: "挑战 361：口语时间"
challengeType: 28
dashedName: challenge-361
---

# --description--

给定模拟时钟的时针和分针角度（以 12 点顺时针方向为起点，单位为度），返回用口语表达的时间。

将分针角度转换为分钟（360° = 60 分钟），然后使用以下规则：

| 分钟 | 口语表达 |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29（不包括 15） | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59（不包括 45） | `"X minutes to Z"`（其中 `X` 是 60 - 分钟数） |

`Y` 是当前小时，`Z` 是下一个小时，二者均由时针角度派生（360° = 12 小时）。

注意：指针角度可能不会正好落在某个数字上，考虑对它们进行某种四舍五入。

# --hints--

`getSpokenTime(90, 0)` 应该返回 `"3 o'clock"`。

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` 应该返回 `"20 minutes past 5"`。

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` 应该返回 `"half past 8"`。

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` 应该返回 `"quarter past 2"`。

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` 应该返回 `"20 minutes to 7"`。

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` 应该返回 `"quarter to 11"`。

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` 应该返回 `"5 minutes to 4"`。

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
