---
id: 6a26df95efa55a2524399743
title: "挑战 362：非ogram 验证器"
challengeType: 28
dashedName: challenge-362
---

# --description--

给定一个线索数字数组和一个单元格数组，判断这些单元格是否满足非ogram 线索。

- 线索是一个数字数组，按顺序表现连续填充单元格的长度。例如，线索 `[3, 2]` 表示应该有 3 个连续填充单元格，后面跟着 2 个连续填充单元格，中间至少有一个空单元格分隔。
- 该行是由 1（填充）和 0（空）组成的数组。

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])` 应该返回 `true`。

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])` 应该返回 `false`。

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` 应该返回 `false`。

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` 应该返回 `true`。

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` 应该返回 `true`。

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` 应该返回 `false`。

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
