---
id: 6a26df95efa55a2524399744
title: "挑战 363：桶填充 2"
challengeType: 28
dashedName: challenge-363
---

# --description--

给定一个由单字母颜色字串组成的二维网格和一个目标颜色，返回将整个网格变成目标颜色所需的最少洪水填充“点击”次数。

- 每次点击会将被点击单元格的颜色以及所有与其相连且颜色相同的单元格区域更改为目标颜色。
- 单元格通过水平和垂直方向相连（不包括对角线方向）。

# --hints--

`bucketFill([["R", "R"], ["R", "R"]], "G")` 应该返回 `1`。

```js
assert.equal(bucketFill([["R", "R"], ["R", "R"]], "G"), 1);
```

`bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")` 应该返回 `0`。

```js
assert.equal(bucketFill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0);
```

`bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")` 应该返回 `3`。

```js
assert.equal(bucketFill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3);
```

`bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")` 应该返回 `5`。

```js
assert.equal(bucketFill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5);
```

`bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")` 应该返回 `12`。

```js
assert.equal(bucketFill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12);
```

# --seed--

## --seed-contents--

```js
function bucketFill(grid, targetColor) {

  return grid;
}
```

# --solutions--

```js
function bucketFill(grid, targetColor) {
  const rows = grid.length;
  const cols = grid[0].length;
  const visited = new Set();

  function dfs(r, c, color) {
    const key = `${r},${c}`;
    if (r < 0 || r >= rows || c < 0 || c >= cols) return;
    if (visited.has(key)) return;
    if (grid[r][c] !== color) return;
    visited.add(key);
    dfs(r + 1, c, color);
    dfs(r - 1, c, color);
    dfs(r, c + 1, color);
    dfs(r, c - 1, color);
  }

  let clicks = 0;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const key = `${r},${c}`;
      if (!visited.has(key) && grid[r][c] !== targetColor) {
        dfs(r, c, grid[r][c]);
        clicks++;
      }
    }
  }

  return clicks;
}
```
