---
id: 6a26df3e2988bcdded204895
title: "挑战 359：高尔夫差点计算器"
challengeType: 28
dashedName: challenge-359
---

# --description--

给定一个高尔夫得分数组和一个对应的球场标准杆数组，使用以下方法返回高尔夫球手的差点指数：

- 通过从得分中减去标准杆来计算每轮的差值，然后返回所有差值的平均值，四舍五入到小数点后一位。

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` 应该返回 `0`。

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` 应该返回 `6`。

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` 应该返回 `8.3`。

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` 应该返回 `8.8`。

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` 应该返回 `11.7`。

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
