---
id: 5e44412c903586ffb414c94c
title: Ein arithmetisches Formatierungsprojekt erstellen
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

Schüler in der Grundschule ordnen Rechenaufgaben oft vertikal an, um sie leichter zu lösen. Zum Beispiel wird aus „235 + 52“:

```py
  235
+  52
-----
```

Vervollständigen Sie die `arithmetic_arranger`-Funktion, die eine Liste von Strings erhält, welche arithmetische Aufgaben sind, und die Aufgaben vertikal und nebeneinander angeordnet zurückgibt. Die Funktion sollte optional ein zweites Argument annehmen. Wenn das zweite Argument auf `True` gesetzt ist, sollen die Antworten angezeigt werden.

## Beispiel

Funktionsaufruf:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

Ausgabe:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

Funktionsaufruf:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

Ausgabe:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## Regeln

Die Funktion gibt die korrekte Umwandlung zurück, wenn die übergebenen Aufgaben richtig formatiert sind. Andernfalls wird ein **String** zurückgegeben, der einen für den Benutzer verständlichen Fehler beschreibt.

- Situationen, die einen Fehler zurückgeben:
  - Wenn der Funktion **zu viele Aufgaben** übergeben werden. Das Limit liegt bei **fünf**, alles darüber hinaus führt zu:
    `'Error: Too many problems.'`
  - Die akzeptierten Operatoren sind **Addition** und **Subtraktion**. Multiplikation und Division führen zu einem Fehler. Andere Operatoren, die in diesem Punkt nicht erwähnt sind, müssen nicht getestet werden. Der zurückgegebene Fehler lautet:
    `"Error: Operator must be '+' or '-'."`
  - Jede Zahl (Operand) darf nur Ziffern enthalten. Andernfalls gibt die Funktion zurück:
    `'Error: Numbers must only contain digits.'`
  - Jeder Operand (also die Zahl auf jeder Seite des Operators) darf maximal vier Stellen breit sein. Andernfalls wird folgender Fehlerstring zurückgegeben:
    `'Error: Numbers cannot be more than four digits.'`
- Wenn der Benutzer die korrekte Formatierung der Aufgaben liefert, folgt die von Ihnen zurückgegebene Umwandlung diesen Regeln:
  - Zwischen dem Operator und dem längeren der beiden Operanden soll genau ein Leerzeichen sein. Der Operator steht in derselben Zeile wie der zweite Operand, beide Operanden bleiben in der Reihenfolge, wie sie übergeben wurden (der erste oben, der zweite unten).
  - Zahlen sollen rechtsbündig ausgerichtet sein.
  - Zwischen jeder Aufgabe sollen vier Leerzeichen sein.
  - Am unteren Rand jeder Aufgabe sollen Striche sein. Die Striche sollen sich über die gesamte Länge jeder einzelnen Aufgabe erstrecken. (Das obige Beispiel zeigt, wie das aussehen soll.)

Hinweis: Öffnen Sie die Browser-Konsole mit F12, um eine ausführlichere Ausgabe der Tests zu sehen.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` sollte `  3801      123\n-    2    +  49\n------    -----` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` sollte `  1         1\n+ 2    - 9380\n---    ------` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` sollte `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` sollte `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])` sollte `'Error: Too many problems.'` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])` sollte `"Error: Operator must be '+' or '-'."` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])` sollte `'Error: Numbers cannot be more than four digits.'` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])` sollte `'Error: Numbers must only contain digits.'` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` sollte `    3      988\n+ 855    +  40\n-----    -----\n  858     1028` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` sollte `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028` zurückgeben.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
