---
id: 68de40f53fcaf106f51cc20f
title: HTML 오디오 및 비디오 플레이어 만들기
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 페이지의 메인 제목을 위한 `h1` 요소가 있어야 합니다.
2. `section` 요소가 두 개 있어야 합니다.
3. 첫 번째 `section` 요소 안에는 재생 중인 비디오 제목을 위한 `h2` 요소가 있어야 합니다.
4. `h2` 요소 아래에는 `video`와 `controls` 속성을 가진 `width` 요소가 있어야 합니다. `width` 속성은 `640`로 설정되어야 합니다.
5. `video` 요소 안에는 비디오 파일을 가리키는 `source` 속성과 `src` 속성을 가진 `type` 요소가 있어야 합니다.
   - `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`를 사용할 수 있습니다.
6. 두 번째 `section` 요소 안에는 재생 중인 노래 제목을 위한 `h2` 요소가 있어야 합니다.
7. `h2` 요소 아래에는 `audio`와 `controls` 속성, 그리고 오디오 파일을 가리키는 `loop` 속성을 가진 `src` 요소가 있어야 합니다.
   - `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`를 사용할 수 있습니다.
   - 또는 `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`를 사용할 수 있습니다.

# --hints--

페이지의 메인 제목을 위한 `h1` 요소가 있어야 합니다.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

`h1`는 비어 있으면 안 됩니다.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

`h1` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`section` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

첫 번째 `section` 요소 안에는 재생 중인 비디오 제목을 위한 `h2` 요소가 있어야 합니다.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

`h2` 요소 아래에는 `video` 요소가 있어야 합니다.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

`video` 요소는 `controls` 속성을 가져야 합니다.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

`video` 요소는 `width` 속성이 `640`로 설정되어야 합니다.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

`video` 요소 안에는 `source` 요소가 있어야 합니다.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

`source` 요소는 비디오 파일을 가리키는 `src` 속성을 가져야 합니다.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

`source` 요소는 `type` 속성을 가져야 합니다.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

`type` 속성은 `src` 속성의 파일 확장자에 맞는 미디어 유형과 일치해야 합니다.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

두 번째 `section` 요소 안에는 재생 중인 노래 제목을 위한 `h2` 요소가 있어야 합니다.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

모든 `h2` 요소는 텍스트를 포함해야 합니다.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

`h2` 요소 아래에는 `audio` 요소가 있어야 합니다.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

`audio` 요소는 `controls` 속성을 가져야 합니다.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

`audio` 요소는 `loop` 속성을 가져야 합니다.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

`audio` 요소는 오디오 파일을 가리키는 `src` 속성을 가져야 합니다.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
