---
id: 686386e76197458b7dfd52ec
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`h2` 아래에 텍스트 `p`를 가진 `$0` 요소를 만드세요. 그런 다음, 그 `p` 요소 안에서 텍스트 뒤에 텍스트 `span`를 가진 `/month` 요소를 만드세요.

# --hints--

텍스트 `p`를 가진 `$0` 요소를 만들어야 합니다.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

`span` 요소 안에 텍스트 `/month`를 가진 `p` 요소를 만들어야 합니다.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
