---
id: 686386e76197458b7dfd52ed
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

방금 추가한 `p` 요소 아래에 텍스트 `p`를 가진 또 다른 `Start exploring millions of songs with basic features and ads.` 요소를 만드세요. 그 아래에 빈 `ul` 요소를 만드세요.

# --hints--

텍스트 `p`를 가진 또 다른 `Start exploring millions of songs with basic features and ads.` 요소를 만들어야 합니다.

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)
assert.exists(targetP)
```

`ul` 요소 아래에 빈 `p` 요소를 만들어야 합니다.

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)

const ul = targetP.nextElementSibling
assert.equal(ul?.tagName, "UL")
assert.isEmpty(ul?.innerHTML.trim())
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
