---
id: 686386e76197458b7dfd52ef
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

이제 CTA 버튼을 만드세요. CTA(Call to Action)는 사용자가 특정 행동을 하도록 유도하는 버튼이나 링크를 말하며, 여기서는 듣기를 시작하도록 하는 버튼입니다.

가장 안쪽 `div` 바깥에 `href`가 `#`인 앵커 요소를 만들고 텍스트는 `Start listening`로 하세요.

이것으로 첫 번째 가격 카드의 HTML이 완성됩니다.

# --hints--

앵커 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector("a"))
```

앵커 요소에 `href`가 `#`로 설정되어야 합니다.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

앵커 요소에 텍스트 `Start listening`가 있어야 합니다.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
