---
id: 686383abac6c625a84458ea9
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 Tailwind CSS의 그리드 유틸리티 클래스를 사용해 가격 컴포넌트를 만드는 방법을 배웁니다.

기본 HTML 보일러플레이트가 제공되어 있습니다. `h1` 요소를 만들고 텍스트로 `Choose your listening plan`을 입력하는 것부터 시작하세요.

# --hints--

`h1` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector("h1"));
```

`h1` 요소에는 텍스트 `Choose your listening plan`가 있어야 합니다.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
