---
id: 686386e76197458b7dfd52ee
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`ul` 안에 두 개의 목록 항목을 만드세요. 첫 번째 목록 항목 안에 `span` 요소를 만들고 그 안에 HTML 엔티티 `&#10003;`를 텍스트로 추가하세요. 그 바로 뒤에 `span`이지만 목록 항목 안에 `Ad-supported streaming` 텍스트를 추가하세요.

두 번째 목록 항목 안에 `span` 요소를 만들고 그 안에 엔티티 `&#10003;`를 텍스트로 추가하세요. 그 바로 뒤에 `span`이지만 목록 항목 안에 `Curated playlists` 텍스트를 추가하세요.

# --hints--

`ul` 요소 안에 두 개의 목록 항목이 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

첫 번째 `li` 요소 안에 HTML 문자 엔티티 `span`가 들어 있는 `&#10003;` 요소가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

첫 번째 `li` 요소 안에 `Ad-supported streaming` 요소 바로 뒤에 `span` 텍스트가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

두 번째 `li` 요소 안에 HTML 문자 엔티티 `span`가 들어 있는 `&#10003;` 요소가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

두 번째 `li` 요소 안에 `Curated playlists` 요소 바로 뒤에 `span` 텍스트가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
