---
id: 686386e76197458b7dfd52eb
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`div` 아래에 `h1` 요소를 만드세요. 그 `div` 안에 또 다른 `div` 요소를 만들고, 그 안에 또 다른 `div` 요소를 중첩하세요.

그리고 가장 안쪽 `div` 안에 텍스트 `h2`를 가진 `Listener` 요소를 만드세요. 이것이 첫 번째 가격 계획을 나타냅니다.

# --hints--

`div` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector("div"));
```

바깥쪽 `div` 안에 또 다른 `div` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector("div > div"));
```

안쪽 `div` 안에 또 다른 `div` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector("div > div > div"));
```

가장 안쪽 `h2` 요소 안에 텍스트 `Listener`를 가진 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
