---
id: 6864f81018a0085ae30afa9b
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

목록 항목 시작 부분의 각 `span`에 있는 체크표시는 장식용입니다. 화면 읽기 프로그램에서 이 `aria-hidden` 요소들을 숨기려면 `span`을 추가하세요.

또한, 각 `text-green-700` 요소에 `mr-2`와 `span` 클래스를 할당하세요.

# --hints--

두 `span` 요소에는 `aria-hidden` 속성이 있어야 합니다.

```js
const spanEls = document.querySelectorAll("ul li span")
const firstTwoSpans = [spanEls[0], spanEls[1]]

firstTwoSpans.forEach(span => {
  assert.isTrue(span.hasAttribute("aria-hidden"))
})
```

두 `span` 요소에는 `text-green-700` 클래스가 있어야 합니다.

```js
const spanEls = document.querySelectorAll("ul li span")
const firstTwoSpans = [spanEls[0], spanEls[1]]

firstTwoSpans.forEach(span => {
  assert.isTrue(span.classList.contains("text-green-700"))
})
```

두 `span` 요소에는 `mr-2` 클래스가 있어야 합니다.

```js
const spanEls = document.querySelectorAll("ul li span")
const firstTwoSpans = [spanEls[0], spanEls[1]]

firstTwoSpans.forEach(span => {
  assert.isTrue(span.classList.contains("mr-2"))
})
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] rounded-xl p-8 gap-6">
          <div class="grid grid-rows-[auto_auto_auto_1fr] gap-y-2">
            <h2 class="text-lg font-semibold text-indigo-600">Listener</h2>
            <p class="text-4xl font-bold text-gray-900">$0<span class="text-base font-medium text-gray-500">/month</span></p>
            <p class="text-gray-600">
              Start exploring millions of songs with basic features and ads.
            </p>
--fcc-editable-region--
            <ul class="mt-6 space-y-2 text-sm text-gray-700">
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
--fcc-editable-region--
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
