---
id: 6864ebcd123f7c335e92d6dd
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

이제 그 계획 안에 있는 항목들의 간격을 조정하세요.

유틸리티 클래스 `rounded-xl`를 추가해 링의 날카로운 모서리를 없애고, `p-8`로 패딩을 늘리며, `gap-6`로 콘텐츠와 CTA 버튼 사이에 공간을 만드세요.

# --hints--

`div` 요소에는 `rounded-xl` 클래스를 지정하세요.

```js
const listenerPlanWrapper = document.querySelectorAll("div")[1]

assert.isTrue(listenerPlanWrapper.classList.contains("rounded-xl"))
```

`div` 요소에는 `p-8` 클래스를 지정하세요.

```js
const listenerPlanWrapper = document.querySelectorAll("div")[1]

assert.isTrue(listenerPlanWrapper.classList.contains("p-8"))
```

`div` 요소에는 `gap-6` 클래스를 지정하세요.

```js
const listenerPlanWrapper = document.querySelectorAll("div")[1]

assert.isTrue(listenerPlanWrapper.classList.contains("gap-6"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
--fcc-editable-region--
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] ">
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
--fcc-editable-region--
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
