---
id: 6864f66df54d9d3abcd4fd31
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

텍스트가 `span`인 `p` 요소 안의 `/month`에 `text-base`, `font-medium`, `text-gray-500` 클래스를 부여하세요. 이렇게 하면 가격 자체와 구분되어 텍스트가 더 연하게 표시됩니다.

# --hints--

`span` 요소에 `text-base` 클래스를 지정하세요.

```js
const PEls = document.querySelectorAll("p")
const targetP = Array.from(PEls).find(p => p.textContent.startsWith("$0"))
const durationSpan = targetP.querySelector("span")

assert.isTrue(durationSpan.classList.contains("text-base"))
```

`span` 요소에 `font-medium` 클래스를 지정하세요.

```js
const PEls = document.querySelectorAll("p")
const targetP = Array.from(PEls).find(p => p.textContent.startsWith("$0"))
const durationSpan = targetP.querySelector("span")

assert.isTrue(durationSpan.classList.contains("font-medium"))
```

`span` 요소에 `text-gray-500` 클래스를 지정하세요.

```js
const PEls = document.querySelectorAll("p")
const targetP = Array.from(PEls).find(p => p.textContent.startsWith("$0"))
const durationSpan = targetP.querySelector("span")

assert.isTrue(durationSpan.classList.contains("text-gray-500"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] rounded-xl p-8 gap-6">
          <div class="grid grid-rows-[auto_auto_auto_1fr] gap-y-2">
            <h2 class="text-lg font-semibold text-indigo-600">Listener</h2>
--fcc-editable-region--
            <p class="text-4xl font-bold text-gray-900">$0<span >/month</span></p>
--fcc-editable-region--
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
