---
id: 6864d9999f8a4e3f0d569a3f
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

다음으로, `h1` 텍스트를 구분되게 만드세요. `class`로 더 크게, `text-3xl`로 가운데 정렬하고, `text-center`로 더 굵게, `font-semibold`로 더 어둡게, 그리고 중간 크기 이상 화면에서는 `text-gray-900`로 더 크게 보이도록 `md:text-5xl`을 할당하세요.

# --hints--

`h1` 요소에 `text-3xl` 클래스를 지정하세요.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-3xl"))
```

`h1` 요소에 `text-center` 클래스를 지정하세요.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-center"))
```

`h1` 요소에 `font-semibold` 클래스를 지정하세요.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("font-semibold"))
```

`h1` 요소에 `text-gray-900` 클래스를 지정하세요.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-gray-900"))
```

`h1` 요소에 `md:text-5xl` 클래스를 지정하세요.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("md:text-5xl"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      --fcc-editable-region--
      <h1>Choose your listening plan</h1>
      --fcc-editable-region--
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
