---
id: 6864e31ef17310b9e6aa376e
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

그리드 항목들이 페이지의 왼쪽과 오른쪽 가장자리와 너무 가까우니, 간격을 조절하세요.

`max-w-6xl`에 `div` 클래스를 할당하여 고정 너비 `72rem`를 설정하고 항목들이 페이지 전체에 너무 넓게 퍼지지 않도록 하세요.

이렇게 하면 모든 항목이 왼쪽으로 몰리므로, `mx-auto`에 `div`를 추가해 항목들을 수평으로 가운데 정렬하세요.

# --hints--

`div` 요소에 `max-w-6xl` 클래스를 지정하세요.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("max-w-6xl"))
```

`div` 요소에 `mx-auto` 클래스를 지정하세요.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("mx-auto"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
--fcc-editable-region--
      <div class="grid grid-cols-1 md:grid-cols-3 ">
--fcc-editable-region--
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
