---
id: 6864f28c46eb6f00dccd5c54
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

**Listener plan** 내부의 `h2`로 이동해서, 더 크게 만들기 위해 유틸리티 클래스 `text-lg`을, 더 굵게 만들기 위해 `font-semibold`를, 그리고 아름다운 인디고 색상을 주기 위해 `text-indigo-600`를 추가하세요.

# --hints--

`h2` 요소에 `text-lg` 클래스를 추가해야 합니다.

```js
const h2Els = document.querySelectorAll("h2")
const listenerH2 = Array.from(h2Els).find(h2 => h2.textContent === "Listener")

assert.isTrue(listenerH2.classList.contains("text-lg"))
```

`h2` 요소에 `font-semibold` 클래스를 추가해야 합니다.

```js
const h2Els = document.querySelectorAll("h2")
const listenerH2 = Array.from(h2Els).find(h2 => h2.textContent === "Listener")

assert.isTrue(listenerH2.classList.contains("font-semibold"))
```

`h2` 요소에 `text-indigo-600` 클래스를 추가해야 합니다.

```js
const h2Els = document.querySelectorAll("h2")
const listenerH2 = Array.from(h2Els).find(h2 => h2.textContent === "Listener")

assert.isTrue(listenerH2.classList.contains("text-indigo-600"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] rounded-xl p-8 gap-6">
          <div class="grid grid-rows-[auto_auto_auto_1fr] gap-y-2">
            --fcc-editable-region--
            <h2>Listener</h2>
            --fcc-editable-region--
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
