---
id: 68d275dd800f404d22a07564
title: Implementieren Sie den Tiefensuche-Algorithmus
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

In diesem Labor implementieren Sie einen Graphdurchlauf-Algorithmus namens <dfn>Tiefensuche</dfn>.

Während die <dfn>Breitensuche</dfn> inkrementelle Kantendistanzen vom Startknoten aus durchsucht, folgt die <dfn>Tiefensuche</dfn> zunächst einem Pfad von Kanten so weit wie möglich.

Sobald sie das Ende eines Pfads erreicht, wird die Suche zum letzten Knoten mit einem noch nicht besuchten Kantenpfad zurückverfolgt und die Suche dort fortgesetzt.

Im Gegensatz zur Breitensuche besucht die Tiefensuche bei jedem besuchten Knoten nicht alle seine Nachbarn. Stattdessen besucht sie zuerst einen Nachbarn und folgt diesem Pfad, bis keine weiteren Knoten auf diesem Pfad mehr besucht werden können.

Um diesen Algorithmus zu implementieren, sollten Sie einen Stack verwenden (ein Array, bei dem das zuletzt hinzugefügte Element als erstes entfernt wird, gemäß dem <dfn>Last-In-First-Out</dfn>-Prinzip). Ein Stack ist bei Tiefensuche-Algorithmen hilfreich, weil Sie beim Hinzufügen von Nachbarn zum Stack die zuletzt hinzugefügten Nachbarn zuerst besuchen und vom Stack entfernen möchten.

Eine einfache Ausgabe dieses Algorithmus ist eine Liste von Knoten, die von einem gegebenen Knoten aus erreichbar sind. Daher sollten Sie auch die besuchten Knoten im Auge behalten.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `dfs` haben.
1. Die `dfs`-Funktion sollte zwei Argumente entgegennehmen:
   - Eine ungerichtete Adjazenzmatrix.
   - Eine Knotenbezeichnung, die der numerische Wert des Knotens zwischen `0` und `n - 1` ist, wobei `n` die Gesamtanzahl der Knoten im Graphen ist.
1. Die `dfs`-Funktion sollte den Tiefensuche-Algorithmus implementieren, um eine Liste aller von dem übergebenen Knoten erreichbaren Knoten auszugeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `dfs` haben, die zwei Argumente entgegennimmt.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` sollte eine Liste mit `1`, `2`, `3` und `0` zurückgeben.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` sollte eine Liste mit `1`, `2`, `3` und `0` zurückgeben.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` sollte `[3]` zurückgeben.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` sollte eine Liste mit `3` und `2` zurückgeben.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` sollte eine Liste mit `0` und `1` zurückgeben.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

Die `dfs`-Funktion sollte die korrekten Ergebnisse zurückgeben.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
