---
id: 6551f1e8737dbe6e0cb08ea4
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

이제 새 변수를 출력하세요.

# --hints--

기존 `print(type(text))` 호출을 제거하지 마세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

`print()` 변수를 전달하여 `shift` 함수를 호출해야 합니다. 함수 호출을 줄 맨 앞에 위치시키는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
