---
id: 6551f6f9ec6db270eb83f3db
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`11`에 11개의 문자가 들어 있으므로 `'Hello World'`가 터미널에 출력되는 것을 볼 수 있습니다.

또 다른 유용한 내장 함수는 `type()`로, 변수의 데이터 유형을 반환합니다. `print()`의 데이터 유형을 출력하도록 `text` 호출을 수정하세요.

# --hints--

`type()` 함수를 호출해야 합니다.

```js
assert.match(code, /type\s*\(.*\)/)
```

소괄호 사이에 `text`를 넣어 `type()` 함수에 전달해야 합니다.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

소괄호 사이에 `type(text)`를 넣어 `print()` 함수에 전달해야 합니다. 함수 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

코드에 `print(len(text))`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
