---
id: 6551ef504d91936d2d4e54f8
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

*인수*는 함수가 호출될 때 소괄호 안에 추가되어 전달되는 객체나 표현식입니다.

```py
greet = 'Hello!'
print(greet)
```

위 예제의 코드는 `'Hello!'` 함수에 인수로 전달된 변수 `greet`의 값인 문자열 `print()`를 출력합니다.

`text` 함수에 `text` 변수를 인수로 전달하여 `print()` 변수를 화면에 출력하세요.

# --hints--

`text` 변수를 소괄호 안에 넣어 `print()` 함수에 전달해야 합니다. 함수 호출이 줄 맨 앞에 오도록 주의하세요.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

코드에 잘못된 구문이나 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
