---
id: 6551f3bfed58796ebb1268dc
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

각 문자열 문자는 숫자 인덱스로 참조할 수 있습니다. 인덱스는 0부터 시작합니다. 따라서 문자열의 첫 번째 문자는 인덱스 `0`를 가집니다. 예를 들어, 문자열 `'Hello World'`에서 `'H'`는 인덱스 `0`에 있고, `'e'`는 인덱스 `1`에 있으며, 이런 식으로 계속됩니다.

문자열의 각 문자는 대괄호 표기법을 사용해 접근할 수 있습니다. 변수 이름 다음에 대괄호를 쓰고, 괄호 안에 문자의 인덱스를 넣으세요:

```py
text = 'Hello World'
r = text[8]
```

이제 `text`를 출력하는 대신, 인덱스 `6`에 있는 문자만 출력하세요.

# --hints--

여전히 `print()` 함수를 호출해야 합니다.

```js
assert.match(code, /print\s*\(.*\)/)
```

소괄호 안에 `text[6]`를 넣어 `print()` 함수에 전달해야 합니다. 함수 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
