---
id: 6551fe3b1df7c9740f13f270
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

Python에서 변수 이름 지정의 주요 사항은 다음과 같습니다:

- 일부 단어는 예약된 키워드입니다(예: `for`, `while`, `True`). Python에서 특별한 의미를 가지므로 변수 이름으로 사용할 수 없습니다.
- 변수 이름은 숫자로 시작할 수 없으며, 영숫자 문자나 밑줄(_)만 포함할 수 있습니다.
- 변수 이름은 대소문자를 구분합니다. 즉, `my_var`는 `my_Var`나 `MY_VAR`와 다릅니다.
- 마지막으로, 변수 이름을 작성할 때 각 공백을 밑줄 문자로 교체하고 단어를 소문자로 작성하는 `snake_case` 관례가 일반적입니다.

`print()` 호출 두 개를 모두 제거하고 `alphabet`라는 다른 변수를 선언하세요. 이 변수에 문자열 `'abcdefghijklmnopqrstuvwxyz'`를 할당하세요.

# --hints--

코드에 `print(type(text))`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

코드에 `print(type(shift))`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

`alphabet`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는지 주의하세요.

```js
assert.match(code, /^alphabet\s*=/m)
```

`'abcdefghijklmnopqrstuvwxyz'` 변수에 문자열 `alphabet`를 할당해야 합니다. 문자열을 감쌀 때는 작은따옴표나 큰따옴표 중 하나를 사용하세요.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
