---
id: 6551ecf597038c36637f3759
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

변수는 Python과 모든 프로그래밍 언어에서 필수적인 부분입니다. 변수는 객체를 참조하거나 가리키는 이름입니다. 변수 이름을 대입 연산자 `=` 왼쪽에 작성하고, 대입 연산자 오른쪽에 그 변수에 할당할 값을 지정하여 변수를 선언할 수 있습니다.

```py
variable_name = value
```

`number`라는 변수를 만들고, 새 변수에 `5` 값을 할당하세요.

# --hints--

`number`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^number\s*=/m)
```

`5` 값을 `number` 변수에 할당해야 합니다.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
