---
id: 6551f22653e6ce6e4f1d1370
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

`print(shift)` 호출을 수정하여 `shift` 변수의 데이터 유형을 출력하세요.

# --hints--

기존 `print(type(text))` 호출은 유지하세요.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

`shift`를 `type()` 함수에 전달해야 합니다.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

`type(shift)`를 `print()` 함수에 전달하세요. 함수 호출을 줄 맨 앞에 위치시키는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

코드에 `print(shift)`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
