---
id: 6551f4259d66026ee28cb205
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

문자열 끝에서부터도 문자열 문자를 접근할 수 있습니다. 마지막 문자는 인덱스가 `-1`이고, 그 전 문자는 `-2`이며 계속 이어집니다.

이제 기존 `print()` 호출을 수정하여 문자열의 마지막 문자를 출력하세요.

# --hints--

여전히 `print()` 함수를 호출해야 합니다.

```js
assert.match(code, /print\s*\(.*\)/)
```

소괄호 사이에 `text[-1]`를 넣어 `print()` 함수에 전달하세요. 함수 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
