---
id: 6551fb531ce7f773071a0b84
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

보시다시피 `type(text)`를 출력한 결과는 `<class 'str'>`이며, 이는 변수의 타입이 `str`인 문자열임을 나타냅니다.

이제 새 줄로 이동해 `shift`라는 변수를 만들고 이 변수에 `3` 값을 할당하세요.

# --hints--

`shift`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는지 주의하세요.

```js
assert.match(code, /^shift\s*=/m)
```

`3` 값을 `shift` 변수에 할당해야 합니다.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
