---
id: 6551f5756c20146fc94f8675
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

내장 `len()` 함수를 사용해 문자열의 문자 수에 접근할 수 있습니다.

기존 `print()` 호출을 수정해 `len(text)` 대신 `text[-1]`를 전달하세요.

# --hints--


`len()` 함수를 호출해야 합니다.

```js
assert.match(code, /len\s*\(.*\)/)
```

소괄호 사이에 변수 `text`를 넣어 `len()` 함수에 전달해야 합니다.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

소괄호 사이에 `len(text)`를 넣어 `print()` 함수를 호출해야 합니다. 함수 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

코드에 `print(text[-1])`가 있으면 안 됩니다.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
