---
id: 655251308f31958d06cdf267
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

현재 암호화된 문자는 매 반복마다 업데이트되고 있습니다. 암호화된 문자열을 새 변수에 저장하는 것이 더 좋습니다. `for` 반복문 전에 `encrypted_text`이라는 변수를 선언하고 이 변수에 빈 문자열(`''`)을 할당하세요.

# --hints--

`encrypted_text` 반복문 전에 `for`이라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

`encrypted_text` 변수에 빈 문자열을 할당해야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
