---
id: 65524b3aa6a1938a069a91ab
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

다음으로 `new_char`를 출력하고 결과를 확인하세요.

# --hints--

`new_char` 변수를 출력하세요.

```js
assert.match(code, /print\s*\(\s*new_char\s*\)/)
```

반복문의 본문 끝에서 `new_char` 변수를 출력하세요.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*new_char\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
--fcc-editable-region--
```
