---
id: 655208d59b131e7816f18c96
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`print()` 함수는 콘솔에 출력만 하지만, 함수와 메서드는 코드에서 사용할 수 있는 반환값을 가질 수 있습니다.

이제 `alphabet.find(text[0])`을 `index`라는 변수에 할당하세요. 이렇게 하면 `index`이 `alphabet.find(text[0])`가 반환한 값을 저장합니다.

# --hints--

`index`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는지 확인하세요.

```js
assert.match(code, /^index\s*=/m)
```

`alphabet.find(text[0])`가 반환한 값을 `index` 변수에 할당해야 합니다.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
