---
id: 655220953ba90d80514d7ee2
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

현재 `print()` 함수는 하나의 인수 `char`만 받고 있지만, 여러 인수를 쉼표로 구분하여 받을 수 있습니다.

`print(char)`에 두 번째 인수를 추가하여 문자와 알파벳 내 인덱스를 출력하게 하세요.

# --hints--

기존 `index` 호출에 두 번째 인수로 `print(char)`를 추가해야 합니다. 인수는 쉼표로 구분하는 것을 잊지 마세요.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
