---
id: 655209a4a27dd37873c4cac3
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`.find()`는 문자열 안에서 일치하는 문자의 인덱스를 반환합니다. 문자를 찾지 못하면 `-1`를 반환합니다. 보시다시피, `text`의 첫 번째 문자인 대문자 `'H'`는 `alphabet`가 소문자만 포함하고 있기 때문에 찾을 수 없습니다.

문자열을 소문자 형태로 변환하려면 `.lower()` 메서드를 사용하세요. `print()`를 출력하도록 `text.lower()` 호출을 하나 더 추가해서 결과를 확인하세요.

# --hints--

코드에 여전히 `print(index)`가 있어야 합니다. 함수 호출이 줄 맨 앞에 오도록 주의하세요.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

`text.lower()`를 출력해야 합니다. 함수 호출이 줄 맨 앞에 오도록 주의하세요.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
