---
id: 65520e6f2b9678799977f24d
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

마지막 `print()` 호출을 제거하세요. 그런 다음 `text[0]` 대신 `text[0].lower()`를 `.find()` 호출의 인수로 전달하고 출력 결과를 확인하세요.

# --hints--

코드에 `print(text.lower())`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

여전히 `index` 변수를 출력해야 합니다.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

`alphabet.find(text[0])` 호출을 `text[0].lower()`를 인수로 사용하도록 업데이트하세요. 메서드 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
