---
id: 65520f8282faf57a0db4f7fe
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

출력에서 볼 수 있듯이, `'h'`는 `7` 문자열의 `alphabet` 인덱스에 있습니다. 이제 `7` 인덱스에 `shift` 값을 더한 위치의 문자를 찾아야 합니다. 이를 위해 수학적 덧셈에 사용하는 것과 같은 방식으로 덧셈 연산자 `+`를 사용할 수 있습니다.

`shifted` 변수를 수정하여 `alphabet`의 `index + shift` 인덱스에 있는 값을 저장하게 하세요.

# --hints--

`alphabet[index + shift]`를 `shifted` 변수에 할당해야 합니다.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
