---
id: 65524d2a1a253b8bb5197ae2
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

`find`는 대문자에 대해 다시 `-1`을 반환하고, 공백 문자에 대해서도 마찬가지입니다. 공백은 나중에 처리할 예정입니다.

지금은 `text`를 순회하는 대신 `for` 루프를 `text.lower()`를 순회하도록 변경하세요.

# --hints--

`for` 루프를 `text.lower()` 대신 `text`를 순회하도록 수정해야 합니다.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
