---
id: 65522255d5b9cd80f335c6f2
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

루프 본문 끝에서 `new_index`라는 변수를 선언하고 이 변수에 `index + shift` 값을 할당하세요.

# --hints--

`new_index` 루프 안에 `for`라는 변수를 선언해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

`index + shift` 루프 본문 끝에서 새 변수에 `for`를 할당해야 합니다.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
