---
id: 65521203d7165c7b84b22ad4
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

루프는 수열의 요소를 체계적으로 순회하며 각 요소에 대해 작업을 실행할 수 있게 합니다.

이번에는 `for` 루프를 사용할 것입니다. `text`를 순회하는 방법은 다음과 같습니다:

```py
for i in text:
```

`for`는 루프 유형을 나타내는 키워드입니다. `i`는 `text`의 요소 값을 순서대로 받는 변수입니다. 문장은 콜론 `:`로 끝납니다.

`alphabet`를 선언한 줄 아래에 `for` 루프를 작성하여 `text`를 순회하세요. 루프 변수로 `i`를 사용하세요.

이렇게 하면 터미널에 오류가 발생합니다. 다음 단계에서 이에 대해 배우게 됩니다.

# --hints--

루프를 만들려면 `for` 키워드를 사용하세요. 키워드 시작 부분에 `for` 키워드를 두고, 키워드 뒤에 공백을 넣으세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

`i` 키워드 뒤에 `for` 변수를 작성하세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

`in` 뒤에 `for i ` 키워드를 작성하세요. `in` 키워드 주변에 공백을 반드시 넣으세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

`text` 뒤에 `for i in `를 작성하세요. 마지막 `:`를 잊지 마세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

`for` 코드 줄 아래에 `alphabet = 'abcdefghijklmnopqrstuvwxyz'` 루프를 작성하세요.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
