---
id: 65521fc818947e800bffe48a
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

`for` 반복문 안에서, 현재 문자를 출력하기 전에 `index`이라는 변수를 선언하고 `alphabet.find(char)`가 반환하는 값을 이 변수에 할당하세요.

# --hints--

`index` 반복문 시작 부분에 `for`라는 새 변수를 선언하세요.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

새로운 `alphabet.find(char)` 변수에 `index`를 할당하세요.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
