---
id: 6552487e689f6e895f658717
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

이제 반복문 본문 끝에 `new_char` 변수를 만드세요. 그 값을 `alphabet[new_index]`로 설정하세요.

# --hints--

`new_char` 반복문 안에 `for` 변수를 만들어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

반복문 본문 끝에 `new_char` 변수를 `alphabet[new_index]`로 설정해야 합니다.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
