---
id: 6552008c0d9d9075cbec9772
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

메시지의 각 글자가 알파벳에서 몇 번째 위치에 있는지 찾기 위해 `.find()` 메서드를 사용할 것입니다. 메서드는 함수와 비슷하지만 객체에 속합니다.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

위에서 `.find()` 메서드는 `sentence` (검색할 문자열)에서 *호출되며*, `'r'` (찾을 문자)이 인수로 전달됩니다. `sentence.find('r')` 호출은 `4`를 반환하는데, 이는 `'r'`에서 `sentence`가 처음 나타나는 인덱스입니다.

코드 마지막에 `.find()`를 `alphabet`에서 호출하고 `'z'`을 메서드 인수로 전달하세요.

# --hints--

`.find()` 메서드를 호출해야 합니다.

```js
assert.match(code, /find\s*\(.*\)/)
```

`.find()` 변수에서 `alphabet` 메서드를 호출해야 합니다.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


`.find()` 변수에서 `alphabet` 메서드를 호출하고 `'z'`을 메서드에 전달해야 합니다. 메서드 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


코드에 잘못된 구문과/또는 잘못된 들여쓰기가 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
