---
id: 655261b2e1f2c197093f3993
title: 41단계
challengeType: 20
dashedName: step-41
---

# --description--

현재 공백은 `'c'`로 암호화됩니다. 일반 메시지에서 원래 공백을 유지하려면 조건부 `if` 문이 필요합니다. 이 문은 `if` 키워드, 조건, 그리고 콜론 `:`으로 구성됩니다.

```py
if x != 0:
    print(x)
```

위 예제에서 `if` 문의 조건은 `x != 0`입니다. `print(x)` 문 본문 안의 `if` 코드는 조건이 `True`으로 평가될 때만 실행됩니다(이 예에서는 `x`가 0과 다르다는 뜻입니다).

for 루프 맨 위에서 `print(char == ' ')`을 `if` 문으로 바꾸세요. 이 `if` 문의 조건은 `True`가 빈 공백일 때 `char`으로 평가되고, 그렇지 않으면 `False`이어야 합니다. `if` 본문 안에서 문자열 `'space!'`을 출력하세요. 이 줄은 들여쓰기해야 합니다.

# --hints--

코드에 `print(char == ' ')`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

`print(char == ' ')`을 `if`일 때 실행되는 `char == ' '` 문으로 바꾸세요. 조건 `if`를 소괄호로 감싸지 말고 마지막 콜론을 포함하는 것을 잊지 마세요.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

새 `'space!'` 문 안에서 문자열 `if`을 출력해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
