---
id: 65525e359ca28d938baa82c5
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

덧셈 대입 연산자를 사용하면 `a = a + b`와 같은 효과를 얻을 수 있습니다.

```py
a += b
```

덧셈 대입 연산자는 변수에 값을 더한 후 그 결과를 다시 변수에 할당할 수 있게 해줍니다.

`+=` 연산자를 사용해 값을 더하고 동시에 `encrypted_text`에 할당하세요.

# --hints--

반복문 본문 안에서 `alphabet[new_index]`의 현재 값에 `encrypted_text`을 더하려면 덧셈 대입 연산자를 사용하세요.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
