---
id: 6553a44b1801991847d8cc69
title: 47단계
challengeType: 20
dashedName: step-47
---

# --description--

다음으로, `print()` 호출을 수정해 `'encrypted text:', encrypted_text`을 출력하고 `for` 반복문 밖에 두어 암호화된 문자열이 한 번만 출력되게 하세요.

# --hints--

for 반복문 이후에 `'encrypted text:', encrypted_text`을 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

코드에 `print('char:', char, 'encrypted text:', encrypted_text)`가 없어야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
