---
id: 6553f8c570f9982e013a8886
title: 57단계
challengeType: 20
dashedName: step-57
---

# --description--

현재는 지정된 오프셋에 따라 모든 글자가 항상 같은 글자로 암호화됩니다. 만약 각 글자마다 오프셋이 다르다면 어떨까요? 그렇게 되면 해독하기 훨씬 더 어려워집니다. 이 알고리즘은 비제네르 암호라고 하며, 각 글자의 오프셋은 *키*라고 하는 다른 텍스트에 의해 결정됩니다.

두 개의 함수 호출을 제거하여 시저 암호를 비제네르 암호로 변환하기 시작하세요.

# --hints--

두 개의 `caesar()` 함수 호출을 제거해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
