---
id: 6553a7d8d05cbb1ae335a665
title: 45단계
challengeType: 20
dashedName: step-45
---

# --description--

루프가 문자 `Z`에 도달했을 때, 합계 `index + shift`가 문자열 `alphabet`의 마지막 인덱스를 초과합니다. 따라서 `alphabet[new_index]`가 잘못된 인덱스를 사용하려고 하여 `IndexError`가 발생합니다.

터미널 출력이 `Z` 바로 앞의 공백에서 멈추는 것을 확인할 수 있습니다. 오류가 발생하기 전 마지막 `print`입니다.

이 경우, 나머지 연산자(`%`)를 사용하여 두 수의 나눗셈에서 나머지를 반환할 수 있습니다. 예를 들어: `5 % 2`는 `1`와 같습니다. 5를 2로 나누면 몫이 2이고 나머지가 1이기 때문입니다.

`index + shift`를 소괄호로 감싸고, `26` 길이인 `alphabet`로 표현식을 나머지 연산하세요.

# --hints--

`new_index = (index + shift) % 26` 문장에 `else`가 있어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
