---
id: 6553f6b4ec51112d44d737c8
title: 48단계
challengeType: 20
dashedName: step-48
---

# --description--

`print` 호출 바로 전에 또 하나를 추가하고 `'plain text:', text`을 `print()`에 인수로 전달하세요. 같은 들여쓰기를 사용하세요.

# --hints--

for 루프가 끝난 후 `'plain text:', text`을 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
