---
id: 6553a572f7a65718f1e42e18
title: 49단계
challengeType: 20
dashedName: step-49
---

# --description--

함수는 본질적으로 재사용 가능한 코드 블록입니다. 이미 `print()`, `find()`, `len()` 같은 내장 함수를 만난 적이 있습니다. 하지만 다음과 같이 사용자 정의 함수를 선언할 수도 있습니다:

```py
def function_name():
    <code>
```

함수 선언은 `def` 키워드로 시작하며, 그 뒤에 함수 이름 — 유효한 변수 이름 — 과 소괄호 한 쌍이 옵니다. 선언은 콜론으로 끝납니다.

`shift` 변수를 선언한 바로 다음에 `caesar`라는 함수를 선언하고, 새 함수의 본문을 만들기 위해 다음 모든 줄을 들여쓰기 하세요.

# --hints--

새 함수를 선언하려면 `def` 키워드를 사용하세요.

```js
assert.match(code, /^def\s+/m)
```

`caesar` 키워드 다음에 함수 이름으로 `def`를 작성하세요. `def` 뒤에 공백을 추가하는 것을 잊지 마세요.

```js
assert.match(code, /^def\s+caesar/m)
```

함수 이름 뒤에 소괄호 한 쌍을 추가하세요. 마지막 콜론도 잊지 마세요.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

`shift = 3` 이후 모든 줄을 들여쓰기 하여 새 함수 본문이 되도록 하세요.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
