---
id: 6553efd6ada3f42aa2d75448
title: 53단계
challengeType: 20
dashedName: step-53
---

# --description--

이제 다시 출력 결과를 볼 수 있습니다. 이 방법은 작동하지만 코드 재사용성을 크게 향상시키지는 않습니다. 함수를 반복해서 호출해도 같은 결과가 나옵니다. 하지만 함수는 *매개변수*를 선언하여 다양성과 맞춤화를 도입할 수 있습니다.

```py
def function_name(param_1, param_2):
    <code>
```

매개변수는 함수 내부에서 사용할 수 있는 변수입니다. 함수는 서로 다른 개수의 매개변수를 선언할 수 있습니다. 위 예제에서 `param_1`와 `param_2`은 매개변수입니다.

함수 선언을 수정하여 `message`와 `offset`라는 두 개의 매개변수를 받도록 하세요.

그다음 터미널에 오류가 나타날 것입니다. 다음 단계에서 해결 방법을 확인할 수 있습니다.

# --hints--

`caesar` 함수는 `message`와 `offset`를 매개변수로 받아야 합니다. 매개변수는 쉼표로 구분하는 것을 잊지 마세요.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
