---
id: 6553ed69ece88d29594748aa
title: 52단계
challengeType: 20
dashedName: step-52
---

# --description--

실행하려면 함수 이름 뒤에 소괄호 한 쌍을 붙여 호출(또는 실행)해야 합니다. 예를 들어 다음과 같습니다:

```py
function_name()
```

코드 끝에 `caesar` 함수를 호출하세요. 들여쓰기에 주의하세요.

# --hints--

`caesar` 함수를 호출해야 합니다. 함수 호출을 줄 맨 앞에 작성했는지 확인하세요.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
