---
id: 6553980e0527fa115c705646
title: 42단계
challengeType: 20
dashedName: step-42
---

# --description--

이제 `'space!'`를 출력하는 대신, 덧셈 대입 연산자를 사용해 현재 `char` 값에 `encrypted_text`에 저장된 공백을 더하세요.

# --hints--

코드에 `print('space!')`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

새로운 `+=` 문장 안에서 `char`의 현재 값에 `encrypted_text`를 더하려면 `if` 연산자를 사용하세요.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
