---
id: 6553995f412dd8122ed38e4a
title: 43단계
challengeType: 20
dashedName: step-43
---

# --description--

조건문은 `else` 절을 가질 수도 있습니다. 이 절은 `if` 문 끝에 추가하여 `if` 문 조건이 거짓일 때 대체 코드를 실행할 수 있습니다:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

출력에서 볼 수 있듯이, 반복문이 공백에 도달하면 암호화된 문자열에 공백이 추가됩니다. 그런 다음 `if` 블록 외부의 코드가 실행되고 암호화된 문자열에 `c`가 추가됩니다.

이를 고치려면 `else` 뒤에 `encrypted_text += char` 절을 추가하고 `print()` 호출을 제외한 이후 모든 코드 줄을 들여쓰기 하세요.

# --hints--

`else` 절을 만들어야 합니다. 마지막 콜론을 포함하는 것을 잊지 마세요.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

`else` 절 뒤의 코드 줄을 `print()` 호출을 제외하고 들여쓰기 해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
