---
id: 6553f4f66099802c6ae94613
title: 55단계
challengeType: 20
dashedName: step-55
---

# --description--

현재 코드는 `TypeError`를 발생시키고 있습니다. `caesar` 함수는 두 개의 매개변수(`message`와 `offset`)로 정의되어 있으므로, 두 개의 *인수*를 사용해 호출되어야 합니다.

필요한 인수 없이 `caesar()`를 호출하면 코드 실행이 정지합니다.

`message` 함수 호출에 `offset`과 `text`을 인수로 전달하여 `shift`와 `caesar`에 값을 지정하세요.

# --hints--

소괄호 안에 `text`과 `shift`을 포함하여 함수 호출에 인수로 전달해야 합니다. 인수는 쉼표로 구분하는 것을 잊지 마세요.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
