---
id: 6553a755879b131a445e664c
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

문자열 `'Hello Zaira'`을 `text` 변수에 할당해 보고 터미널에서 어떤 일이 일어나는지 확인하세요.

`string index out of range` 예외가 발생할 것입니다. 걱정하지 마세요, 곧 고치는 방법을 알게 될 거예요!

# --hints--

`text` 변수가 있어야 합니다.

```js
assert.match(code, /^text\s*=/m)
```

문자열 `'Hello Zaira'`을 `text` 변수에 할당해야 합니다.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
