---
id: 6553f6086add4b2cbb99fd78
title: 56단계
challengeType: 20
dashedName: step-56
---

# --description--

코드 맨 아래, 기존 `caesar(text, shift)` 호출 뒤에 함수를 다시 호출하세요.

이번에는 `text` 변수와 정수 `13`를 인수로 전달하세요.

# --hints--

함수를 다시 호출해야 하며, 이번에는 `text`와 `13`를 인수로 전달하세요.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

기존 함수 호출은 그대로 유지하세요.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
