---
id: 6554e0adc7bb193cbfdb36d5
title: 73단계
challengeType: 20
dashedName: step-73
---

# --description--

함수가 문자열을 다시 일반 텍스트로 복호화할 수 있는지 확인하세요.

`decryption`라는 다른 변수를 선언하고 `vigenere(encryption, custom_key, -1)`을 할당하세요.

# --hints--

`vigenere`, `encryption`, `custom_key`를 인수로 전달하여 `-1`를 호출해야 합니다.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

`decryption` 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

`decryption` 변수는 `vigenere(encryption, custom_key, -1)` 값을 가져야 합니다.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
