---
id: 65549561463f0016876e852c
title: 60단계
challengeType: 20
dashedName: step-60
---

# --description--

키가 암호화할 텍스트보다 짧기 때문에 전체 암호화된 텍스트를 생성하려면 키를 반복해야 합니다.
함수 본문 시작 부분에 `key_index` 변수를 선언하고 0으로 설정하세요.

# --hints--

함수 본문 시작 부분에 `key_index`라는 변수를 선언해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

`0`를 `key_index` 변수에 할당해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
