---
id: 655491bd5b98b813fa5bedca
title: 58단계
challengeType: 20
dashedName: step-58
---

# --description--

이제 함수 선언을 수정하세요: 함수 이름을 `vigenere`로, 두 번째 매개변수를 `key`로 바꾸세요.

# --hints--

함수 이름을 `vigenere`로 수정해야 합니다.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

`vigenere` 함수는 `message`와 `key`를 매개변수로 받아야 합니다.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
