---
id: 6554d25dc5ceaa354307a77e
title: 70단계
challengeType: 20
dashedName: step-70
---

# --description--

암호화와 복호화는 반대 과정이며, 함수에 약간의 수정을 하면 둘 다 할 수 있습니다.

함수 정의에 `direction`라는 세 번째 매개변수를 추가하세요. 또한, 콘솔에서 오류가 발생하지 않도록 마지막 두 줄의 코드를 주석 처리하세요.

# --hints--

코드의 마지막 두 줄을 주석 처리해야 합니다. 각 줄의 시작 부분에 `#`를 넣으세요.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

함수의 세 번째 매개변수로 `direction`를 추가해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
