---
id: 6554de295ade563a069936a1
title: 71단계
challengeType: 20
dashedName: step-71
---

# --description--

해야 할 일은 `offset` 할당에서 `direction`에 `new_index`을 곱하는 것입니다. Python에서 곱셈 연산자는 `*`입니다.

# --hints--

`offset` 할당에서 `direction`에 `new_index`를 곱해야 합니다. 다른 소괄호는 추가하지 마세요.

```js
({ test: () => assert.match(code, /new_index\s*=\s*\(\s*index\s*\+\s*offset\s*\*\s*direction\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
--fcc-editable-region--
            new_index = (index + offset) % len(alphabet)
--fcc-editable-region--
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
#encryption = vigenere(text, custom_key)
#print(encryption)
```
