---
id: 6554dfce1683be3c0c9609a6
title: 72단계
challengeType: 20
dashedName: step-72
---

# --description--

이제 마지막 두 줄의 주석을 해제하고 함수 호출에서 세 번째 인수로 `1`를 전달하도록 수정하세요.

# --hints--

코드의 마지막 두 줄을 복원하고 함수 호출에서 세 번째 인수로 `1`를 전달하세요.

```js
({ test: () => assert.match(code, /^encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*,\s*1\s*\)\s*print\s*\(\s*encryption\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
--fcc-editable-region--
#encryption = vigenere(text, custom_key)
#print(encryption)
--fcc-editable-region--
```
