---
id: 6554930320d70414e7b6acc6
title: 59단계
challengeType: 20
dashedName: step-59
---

# --description--

`shift` 변수를 삭제하세요. 그런 다음 `custom_key`이라는 새 변수를 선언하고 문자열 `'python'`의 값을 이 변수에 할당하세요.

# --hints--

`shift` 변수와 그 값을 삭제해야 합니다.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

`custom_key`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

문자열 `'python'`를 `custom_key` 변수에 할당해야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
