---
id: 6554a88d5af937226f4a9121
title: 65단계
challengeType: 20
dashedName: step-65
---

# --description--

`.index()` 메서드는 `.find()` 메서드와 동일하지만, 부분 문자열을 찾지 못하면 `ValueError` 예외를 발생시킵니다.

`ValueError`는 올바른 타입이지만 부적절한 값을 가진 인수가 함수에 전달될 때 발생하는 내장 예외입니다.

`key_index`를 증가시킨 후 `offset`라는 변수를 선언하세요. `key_char`에서 `alphabet`가 가진 인덱스를 찾아 `offset`에 할당하세요. 인덱스를 찾기 위해 `.index()`를 사용하세요.

# --hints--

`offset`라는 변수를 선언해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

`offset` 변수는 `alphabet.index(key_char)` 값을 저장해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
