---
id: 6554a49a4f782f208abcc87e
title: 63단계
challengeType: 20
dashedName: step-63
---

# --description--

다음 반복을 위해 `key_index` 개수를 늘려야 합니다.
방금 추가한 줄 다음에, 같은 코드 블록 안에서 덧셈 대입 연산자를 사용해 `key_index`을 1만큼 증가시키세요.

# --hints--

`+=` 절 본문 안에서 `1`에 `key_index`을 더하기 위해 `else` 연산자를 사용해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]\s*key_index\s*\+=\s*1/));
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            key_char = key[key_index % len(key)]
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
