---
id: 65549f90cf78131c96ebcf28
title: 61단계
challengeType: 20
dashedName: step-61
---

# --description--

코딩할 때 가독성이 가장 중요합니다. 주석은 코드 뒤에 숨은 논리를 설명하는 효과적인 메모 역할을 합니다. 시간이 지나 프로젝트로 돌아왔을 때 유용하며, 동료들이 코드를 이해하는 데도 도움이 됩니다.

Python에서는 `#`를 사용해 주석을 작성할 수 있습니다. `#` 뒤에 오는 모든 내용은 실행되지 않습니다.

`if` 문장 앞에 `Append space to the message`라는 주석을 추가하세요.

# --hints--

`Append space to the message`라는 주석을 추가해야 합니다. 시작 부분에 `#` 문자를 잊지 마세요.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
