---
id: 6554d0332949b133a0b35eaa
title: 68단계
challengeType: 20
dashedName: step-68
---

# --description--

`text`와 `custom_key`을 인수로 전달하여 함수를 호출하세요. 함수 호출의 반환값을 `encryption`라는 변수에 저장하세요.

# --hints--

`vigenere`를 호출할 때 `text`와 `custom_key`를 인수로 전달해야 합니다.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

`encryption` 변수가 있어야 합니다.

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

함수 호출 결과를 `encryption` 변수에 할당해야 합니다.

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
