---
id: 6554ad2463b8892748f8efdd
title: 67단계
challengeType: 20
dashedName: step-67
---

# --description--

현재 함수는 몇몇 문자열을 출력하지만, 이 값들은 코드의 다른 부분에서 어떤 동작을 수행하는 데 사용할 수 없습니다.

그 목적을 위해 `return` 문장을 사용해야 합니다.

```python
def foo():
    return 'spam'
```

함수가 반환해야 하는 값 뒤에 `return`를 쓰고 한 칸 띄우세요. `return` 문장을 만나면 그 값이 반환되고 함수 실행이 정지되어 함수 호출 다음 줄의 코드로 진행합니다. 위 예시에서 `foo` 함수는 문자열 `'spam'`를 반환합니다.

함수에서 두 개의 `print()` 호출을 제거하고 `encrypted_text`를 반환하세요.

# --hints--

함수에서 두 개의 `print()` 호출을 제거해야 합니다.

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

함수는 `encrypted_text`를 반환해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
