---
id: 6554a57ec0a2c52106e7ee50
title: 64단계
challengeType: 20
dashedName: step-64
---

# --description--

`else` 절 안에 `Find the right key character to encode`라는 주석을 작성하세요.

# --hints--

`Find the right key character to encode`라는 주석을 작성해야 합니다. 시작 부분에 `#`를 잊지 마세요.

```js
({ test: () => assert.match(code, /#\s*Find\sthe\sright\skey\scharacter\sto\sencode/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
