---
id: 6554ac937a49be2701af4f2f
title: 66단계
challengeType: 20
dashedName: step-66
---

# --description--

`offset` 변수 위에 `Define the offset and the encrypted letter`라는 주석을 하나 더 만드세요.

# --hints--

`Define the offset and the encrypted letter`라는 주석을 만드세요. 시작 부분에 `#`를 넣는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /#\s*Define\sthe\soffset\sand\sthe\sencrypted\sletter/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
