---
id: 6555db4a7b788e15795674e5
title: 89단계
challengeType: 20
dashedName: step-89
---

# --description--

두 개 이상의 문자열은 `+` 연산자를 사용해 연결할 수 있습니다. 예를 들어, `'Hello' + ' there!'`는 `'Hello there!'`를 결과로 만듭니다.

`print()` 함수를 호출하고 `+` 연산자를 사용해 `text` 변수를 문자열 `'Encrypted text: '`에 연결하세요. 띄어쓰기에 주의하세요.

# --hints--

`print()`를 호출하고 `'Encrypted text: ' + text`를 인수로 전달해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Encrypted\stext:\s\1\s*\+\s*text\s*\)/m)
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--    

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
