---
id: 6555e43e783ed31a0532b1b2
title: 93단계
challengeType: 20
dashedName: step-93
---

# --description--

줄 바꿈 문자 `\n`는 새 줄을 나타내는 특별한 수열입니다. 문자열에서 역슬래시 `\` 다음에 `n`를 일반 문자 수열로 작성할 수 있으며, 프로그램이 실행될 때 출력에서 새 줄로 교체됩니다.

첫 번째 `print` 호출의 시작 부분에 줄 바꿈 문자를 넣고 출력을 확인하세요.

# --hints--

첫 번째 print 호출을 `print(f'\nEncrypted text: {text}')`로 수정하세요.

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')\\nEncrypted\stext:\s\{\s*text\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print(f'Key: {custom_key}')
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
