---
id: 6555df0c81300b175308557d
title: 90단계
challengeType: 20
dashedName: step-90
---

# --description--

방금 추가한 `print()` 호출 아래에, `print()` 문자열과 `Key: python` 변수 값을 연결하여 `'Key: '`를 출력하는 또 다른 `custom_key` 호출을 추가하세요.

# --hints--

`print()` 문자열과 `Key: python` 변수 값을 연결하여 `'Key: '`를 출력하는 `custom_key` 호출이 있어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Key:\s\1\s*\+\s*custom_key\s*\)/m);
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
