---
id: 655520c20cb1e6177b0641d6
title: 81단계
challengeType: 20
dashedName: step-81
---

# --description--

`not` 연산자는 표현식을 부정하는 데 사용됩니다. 참 같은 값, 즉 `True`로 평가되는 값 앞에 놓으면 `False`를 반환하며 그 반대도 마찬가지입니다.

문자가 알파벳이 아닌지 확인하려면 `not` 조건의 맨 앞에 `if` 연산자를 추가하세요.

# --hints--

조건문 `not`의 조건에 `if` 연산자를 사용하세요.

```js
({ test: () => assert.match(code, /if\s+not\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
