---
id: 6555201d9b7fc917399f9f0b
title: 80단계
challengeType: 20
dashedName: step-80
---

# --description--

`.isalpha()` 메서드는 호출된 문자열의 모든 문자가 문자일 경우 `True`을 반환합니다. 예를 들어, 아래 코드는 `True`를 반환합니다:

```py
'freeCodeCamp'.isalpha()
# True
```

조건 `char == ' '`을 삭제하고 `.isalpha()`에서 `char` 메서드를 호출하는 것으로 바꾸세요.

# --hints--

코드에 `if char.isalpha():`가 있어야 합니다.

```js
({ test: () => assert.match(code, /if\s+char\.isalpha\s*\(\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--    
        # Append space to the message
        if char == ' ':
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
