---
id: 6555d6e12c95701172b55709
title: 77단계
challengeType: 20
dashedName: step-77
---

# --description--

함수는 기본 인수를 사용해 호출할 수 있습니다. 기본 인수는 인수가 전달되지 않았을 때 함수가 가져야 할 값을 나타냅니다. 예를 들어, 아래 함수는 세 개의 인수를 받지만 두 개의 인수만 전달해 호출할 수 있습니다. 세 번째 인수는 기본적으로 지정된 값을 갖게 됩니다:

```py
def foo(a, b, c=value):
    <code>
```

`vigenere` 함수의 `direction` 매개변수가 `1`의 기본값을 갖도록 수정하세요.

# --hints--

`direction` 함수의 `vigenere` 매개변수는 `1`의 기본값을 가져야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
