---
id: 6555231eac4d9f19bd3d44b5
title: 83단계
challengeType: 20
dashedName: step-83
---

# --description--

`pass` 키워드는 미래의 코드를 위한 자리 표시자로 사용할 수 있습니다. 코드에 아무런 영향을 주지 않지만, 코드가 완성되지 않아 발생할 수 있는 오류를 방지할 수 있습니다.

```py
def foo():
    pass
```

암호화할 때 `vigenere`을 `1`와 함께 호출하고, 복호화할 때 `-1`을 호출하는 것은 괜찮지만 약간 난해할 수 있습니다. `encrypt`라는 새 함수를 선언하고, `message`와 `key` 매개변수를 받도록 하세요. 함수 본문에는 `pass`을 사용하세요.


# --hints--

`encrypt`라는 새 함수를 선언해야 합니다. 함수 본문 안에서 `pass` 키워드를 사용하는 것을 기억하세요.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

`encrypt` 함수는 `message`과 `key`를 매개변수로 받아야 합니다. 순서가 중요합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
