---
id: 65552a9593755e1fb2f5ab50
title: 85단계
challengeType: 20
dashedName: step-85
---

# --description--

`decrypt`과 같은 매개변수를 가진 `encrypt`라는 이름의 또 다른 함수를 정의하세요. 이번에는 `vigenere(message, key, -1)`를 반환하세요.

# --hints--

`decrypt`라는 새 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(decrypt)
  `))
})
```

`decrypt` 함수는 `message`와 `key`를 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(decrypt)
    str(sig) == "(message, key)"
  `))
})
```

`decrypt` 함수는 `vigenere(message, key, -1)`를 반환해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "decrypt");
    assert(function_body.match(/return\s+vigenere\s*\(\s*message\s*,\s*key\s*,\s*-\s*1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--
def encrypt(message, key):
    return vigenere(message, key)
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
