---
id: 6555e1bfcd374e18c6be8e58
title: 91단계
challengeType: 20
dashedName: step-91
---

# --description--

Python에는 문자열을 쉽게 포맷하는 방법이 있습니다. *f-strings*는 문자열 안에 값을 보간할 수 있게 해줍니다.

보간이란 프로그램이 실행될 때 지정된 값으로 교체될 자리 표시자를 작성하는 것을 의미합니다. 예를 들어, `'Encrypted text: ' + text`와 같은 결과를 `f'Encrypted text: {text}'`로 얻을 수 있습니다. f-string을 만들려면 따옴표 앞에 `f`를 붙이고, 중괄호 사이에 보간하려는 변수나 표현식을 작성해야 합니다.

첫 번째 `print()` 호출을 f-string을 사용하도록 수정하세요.

# --hints--

f-string을 사용하여 같은 결과를 출력하도록 `print('Encrypted text: ' + text)` 호출을 수정해야 합니다.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
