---
id: 65551a628bcb7e121e32d04b
title: 79단계
challengeType: 20
dashedName: step-79
---

# --description--

현재 구두점, 특수 문자 또는 숫자가 올바르게 인코딩/디코딩되지 않고 있습니다.

`text` 문자열 끝에 느낌표를 추가해서 확인하세요.

# --hints--

`text` 변수가 있어야 합니다.

```js
({ test: () => assert(__userGlobals.has("text")) })
```

`text` 변수는 문자열 `'Hello Zaira!'`와 같아야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("text"), "Hello Zaira!") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'
--fcc-editable-region--
def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
