---
id: 655619327c7b364166f8dd6f
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

`alphabet` 문자열에 숫자나 특수 문자 같은 추가 문자를 포함하려면 나머지 연산의 오른쪽 피연산자를 수동으로 수정해야 합니다.

이 문제를 피하려면 `26`을 `len(alphabet)`로 바꾸세요.

# --hints--

`new_index` 값을 수정할 때 `26`를 `len(alphabet)`로 바꾸세요.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
