---
id: 655a2a7210094920069b117c
title: 40단계
challengeType: 20
dashedName: step-40
---

# --description--

비교 연산자는 두 객체를 값에 따라 비교할 수 있게 해줍니다. 비교하려는 객체 사이에 비교 연산자를 넣어 사용하세요.
이들은 표현식의 참/거짓 여부에 따라 불리언 값 — 즉 `True` 또는 `False` — 을 반환합니다.

Python에는 다음과 같은 비교 연산자가 있습니다:

<table>
  <thead>
    <tr>
      <th>연산자</th>
      <th>의미</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>같음</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>같지 않음</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>보다 큼</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>보다 작음</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>보다 크거나 같음</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>보다 작거나 같음</td>
    </tr>
  </tbody>
</table>

루프 본문 시작 부분에서 `char`와 공백(`' '`)을 비교한 결과를 출력하세요. 이때 동등 연산자 `==`를 사용하세요.

# --hints--

`char`를 공백과 동등 연산자로 비교하여 `for` 루프 안에서 사용하세요.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

`char`를 공백과 비교한 결과를 `for` 루프 안에서 출력하세요.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

루프 시작 부분에서 `char`를 공백과 비교한 결과를 출력하세요.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
