---
id: 656886ec4e4d9b96ded3c565
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

반복 변수는 유효한 이름이면 무엇이든 될 수 있지만, 의미 있는 이름을 주는 것이 편리합니다.

`i` 변수를 `char`으로 이름을 바꾸세요.

# --hints--

`i`를 `char`으로 바꾸세요.

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
