---
id: 6568b7a40c7ba0ccbc4b4425
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

보시다시피, 재할당된 문자열의 각 문자가 루프 안에서 출력됩니다.

`text` 재할당을 삭제하여 원래 문자열로 돌아가세요.

# --hints--

코드에 `text = 'Albatross'`이 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

여전히 `text` 변수에 `'Hello World'` 값이 있어야 합니다.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
