---
id: 65688b5a1655a7a6caede847
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

이전 단계에서처럼 문자열의 개별 문자를 변경하려고 하면 `TypeError`가 발생하는데, 이는 코드에서 부적절한 유형의 객체를 사용했을 때 나타납니다.

오류 메시지에서 알 수 있듯이, 문자열은 불변이기 때문에 항목 할당을 지원하지 않습니다. 하지만 변수에는 다른 문자열을 다시 할당할 수 있습니다.

```py
message = 'Hello World'
message = 'Hello there!'
```

`text[0]` 줄을 삭제하고 `text`에 문자열 `'Albatross'`을 다시 할당하세요.

# --hints--

코드에 `text[0] = 'C'`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

`text`에 문자열 `'Albatross'`을 다시 할당해야 합니다. `text = 'Hello World'`을 수정하지 마세요.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
