---
id: 65687d2f8c7ee27b0446cef3
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

각 반복에서 실행할 코드는 `:` 뒤에 위치하며, 이것이 루프 본문을 구성합니다. 이 코드는 들여쓰기를 해야 합니다. Python에서는 들여쓰기 수준마다 4칸의 공백을 사용하는 것이 권장됩니다. 이 들여쓰기된 수준이 코드 블록입니다.

```py
for i in text:
    print(i)
```

Python은 코드 블록을 표시하기 위해 들여쓰기에 의존합니다. 줄 끝에 콜론이 있으면 새로운 들여쓰기된 코드 블록이 뒤따를 것임을 알리는 신호입니다.

따라서 마지막 콜론 뒤에 들여쓰기된 블록이 없으면 코드 실행이 정지되고 `IndentationError`가 발생합니다. 이 코드는 출력을 표시하지 않고 대신 `IndentationError`를 발생시킵니다:

```py
for i in text:
print(i)
```

`for` 루프에 `print(i)` 호출을 추가하여 본문을 만드세요. 루프 본문은 반드시 들여쓰기해야 합니다.

# --hints--

`print(i)` 루프 본문에 `for`를 추가하세요. 들여쓰기에 주의하세요.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

코드에 잘못된 구문 및/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
