---
id: 6568c86dc9193000d11ca5e0
title: 38단계
challengeType: 20
dashedName: step-38
---

# --description--

`alphabet[new_index]`에 `encrypted_text`을 할당하는 대신, 이 변수에 `encrypted_text`와 `alphabet[new_index]`의 현재 값을 더한 값을 할당하세요.

# --hints--

`encrypted_text + alphabet[new_index]` 변수에 `encrypted_text`를 할당해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*=\s*encrypted_text\s*\+\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
