---
id: 65688a50e6c998a21d8e41d3
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

문자열은 불변이므로 한 번 생성되면 변경할 수 없습니다. 예를 들어, 다음 코드는 `my_string`의 값을 문자열 `'train'`로 변경한다고 생각할 수 있지만, 이는 올바르지 않습니다.

```py
my_string = 'brain'
my_string[0] = 't'
```

대괄호 표기법을 사용해 `text`의 첫 글자에 접근하고 원하는 문자로 변경해 보세요. 출력이 사라지고 오류가 발생하는 것을 확인할 수 있습니다.

# --hints--

코드에 여전히 `text = 'Hello World'`가 있어야 합니다.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

`text`의 첫 글자에 `text[0]`를 사용해 접근해야 합니다.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

`=` 연산자를 사용해 원하는 문자를 `text[0]`에 할당하세요. 문자를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
