---
id: 65951b637ddec63611c482c0
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

내장 함수 `print()`를 사용해 터미널에 코드 출력 결과를 출력할 수 있습니다.

함수는 필요할 때 코드를 실행하도록 호출하거나 호출할 수 있는 재사용 가능한 코드 블록입니다. 함수를 호출하려면 이름 옆에 소괄호 한 쌍을 쓰기만 하면 됩니다. 곧 함수에 대해 더 배우게 됩니다.

지금은 새 줄로 가서 `print()` 함수를 빈 호출로 추가하세요. 아직 출력이 나타나지 않아야 합니다.

# --hints--

코드에 `print()`가 있어야 합니다. 함수 호출을 줄 맨 앞에 두도록 주의하세요.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
