---
id: 65952b70924edb449faf1950
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

`shifted`라는 새 변수를 선언하세요. 대괄호 표기법을 사용해 `alphabet`의 `index` 인덱스에 있는 값을 접근하여 새 변수에 할당하세요.

# --hints--

`shifted`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는지 주의하세요.

```js
assert.match(code, /^shifted\s*=/m)
```

`alphabet[index]`를 `shifted` 변수에 할당해야 합니다.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
