---
id: 659530ca823d6c47f61daa8a
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

`shifted` 변수를 출력하세요.

# --hints--

`shifted` 변수를 출력해야 합니다. 함수 호출을 줄 맨 앞에 두는 것에 주의하세요.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
