---
id: 659526f3e5424b3eac4b6aa7
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

다음으로, `index` 변수를 콘솔에 출력하세요.

# --hints--

`index` 변수를 출력해야 합니다. 함수 호출을 줄 맨 앞에 두는지 주의하세요.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

코드에 잘못된 구문이나 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
