---
id: 659db3e3670d3e712be82593
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

알파벳에서 문자를 찾고 `text`의 각 문자에 대해 이동된 문자를 결정하는 과정을 반복하는 것은 지루할 수 있습니다. 다행히도, 루프를 사용해 이 과정을 간단하게 할 수 있습니다.

지금은 `alphabet` 변수를 선언한 아래의 모든 코드를 삭제하세요.

# --hints--

코드에 여전히 `text = 'Hello World'`가 있어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

코드에 여전히 `shift = 3`가 있어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

코드에 여전히 `alphabet = 'abcdefghijklmnopqrstuvwxyz'`가 있어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

`index` 변수와 그 값을 삭제해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

코드에 `print(index)`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

`shifted` 변수와 그 값을 삭제해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

코드에 `print(shifted)`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

코드에 잘못된 구문과/또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
