---
id: 65a450e8fb2c9d75c7378d28
title: 88단계
challengeType: 20
dashedName: step-88
---

# --description--

이번에는 복호화할 암호화된 문자열에서 시작하므로 `encryption` 변수가 더 이상 필요하지 않습니다.

`encryption`과 `print(encryption)` 호출을 삭제하세요. 또한, 코드의 마지막 두 줄을 주석 처리하세요.

# --hints--

`encryption` 변수와 그 값을 삭제해야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/encryption\s*=\s*encrypt\s*\(\s*text\s*,\s*custom_key\s*\)/.test(commentless_code))
```

코드에 `print(encryption)`가 없어야 합니다.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*encryption\s*\)/.test(commentless_code))
```

`decryption = decrypt(encryption, custom_key)`를 주석 처리해야 합니다.

```js
assert(code.match(/#\s*decryption\s*=/));
```

`print(decryption)`를 주석 처리해야 합니다.

```js
assert(code.match(/#\s*print\s*\(\s*decryption\s*\)/));
```

코드에 잘못된 구문 또는 잘못된 들여쓰기가 포함되어 있습니다.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print(encryption)
decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
