---
id: 6624e0ec0bf47017eed84b1f
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

처음 만들 암호는 *Caesar* 암호라고 합니다. 구체적으로, 메시지의 각 글자를 알파벳에서 찾아 그 위치를 확인한 뒤, 알파벳에서 3칸 뒤에 있는 글자를 찾아 원래 글자를 새 글자로 바꾸는 방식입니다.

이를 구현하려면 이전 단계에서 다룬 `.find()` 메서드를 사용하세요. 기존 `.find()` 호출을 수정하여 `text[0]` 대신 `'z'`를 인수로 전달하세요.

# --hints--

기존 `alphabet.find('z')` 호출을 수정하여 `text[0]`를 메서드에 전달하세요.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
