---
id: 666064e915eba7aa1de03f6b
title: 51단계
challengeType: 20
dashedName: step-51
---

# --description--

이제 `alphabet` 함수 밖에서 `caesar` 변수를 출력하려는 줄을 제거하여 오류를 고치세요.

# --hints--

`print(alphabet)` 줄을 제거해야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--
print(alphabet)
--fcc-editable-region--
```
