---
id: 66605a64a9822f01ec5ae8ce
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

Python에서 변수의 *범위*는 그 변수가 어디에서 접근될 수 있는지를 결정합니다.

- 함수 밖에서 정의된 변수는 *전역* 범위를 가지며 코드의 어느 부분에서든 접근할 수 있습니다.

- 함수 안에서 정의된 변수는 그 함수에 *지역적*이며 함수 밖에서는 접근할 수 없습니다.

이를 직접 확인하려면 코드 끝에서 `alphabet` 변수를 출력해 보세요. 그러면 `NameError` 예외가 발생합니다.

`alphabet`가 정의되지 않았다는 오류 메시지가 나타날 것입니다. 이는 `alphabet`가 `caesar` 함수 안에서 정의되어 함수 밖에서는 접근할 수 없기 때문입니다.

# --hints--

caesar 함수 밖에서 `alphabet` 변수를 출력해 보세요.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
