---
id: 6a89c69cbe1bb69845992d8b
title: 建造一個數字範圍產生器
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

在這個實驗中，你將建置一個 `range_of_numbers` 函式，使用遞迴來產生指定範圍內的數字列表。

例如：

- `range_of_numbers(3, 9)` 應該傳回 `[3, 4, 5, 6, 7, 8, 9]`。
- `range_of_numbers(5, 5)` 應該傳回 `[5]`。

完成以下使用者故事並通過所有測試以完成此實驗。

**使用者故事：**

1. 你應該創建一個名為 `range_of_numbers` 的函式，並帶有兩個參數，分別命名為 `start_num` 和 `end_num`。
1. 該函式應傳回一個連續整數的列表，從 `start_num` 開始，到 `end_num` 結束，包含兩者。
1. `start_num` 的值將永遠小於或等於 `end_num`。
1. 你的函式應該使用遞迴，透過呼叫自身來實現。
1. 你的函式不應使用 `for` 或 `while` 迴圈、推導式或 `range()` 函式。
1. 當 `start_num` 等於 `end_num` 時，函式應傳回一個包含 `start_num` 的列表。這是基本情況。
1. 對於遞迴情況，函式應呼叫自身，並帶入一個向基本情況靠近的引數，然後將目前的數字加入傳回的列表中。
1. 函式不應使用全域變數來儲存傳回的列表。

# --hints--

你應該有一個名為 `range_of_numbers` 的函式，帶有名為 `start_num` 和 `end_num` 的參數。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` 應該傳回 `[1, 2, 3, 4, 5]`。

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` 應該傳回 `[6, 7, 8, 9]`。

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` 應該傳回 `[4]`。

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` 應該傳回 `[10, 11, 12, 13, 14, 15]`。

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

你的 `range_of_numbers` 函式應該呼叫自身。

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

你的 `range_of_numbers` 函式不應使用迴圈、推導式或 `range()` 函式。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

呼叫 `range_of_numbers(1, 3)` 不應影響後續的呼叫。該呼叫之後，`range_of_numbers(6, 9)` 仍應傳回 `[6, 7, 8, 9]`。

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
