---
id: 68f890b73bbfc9efa7379e7b
title: Depura errores de coerción de tipo en una aplicación con fallos
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Acabas de unirte a una empresa local de desarrollo web, y tu primera tarea es limpiar algo de código con errores dejado por los desarrolladores anteriores.

Estaban intentando operaciones aritméticas simples, pero algo salió mal, los resultados no tienen sentido. Tu tarea es revisar, depurar y corregir el código para que funcione como se espera y sea más fácil de leer.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio. 

**Historias de usuario:**

1. Debes tener una variable llamada `firstResult` que sume correctamente los números `5` y `10` para producir el valor `15`.
2. Debes tener una variable llamada `secondResult` que reste correctamente `5` de `8` para producir el valor `3`.
3. Debes tener una variable llamada `thirdResult` que produzca el valor `6` sumando dos números. Reemplaza el booleano que se usa actualmente en la expresión con un número.
4. Debes tener una variable llamada `fourthResult` que produzca el valor `8` sumando dos números. Reemplaza el booleano que se usa actualmente en la expresión con un número.
5. Debes tener una variable llamada `fifthResult` que multiplique correctamente dos números para obtener el producto `20`. Reemplaza la cadena que se usa actualmente en la expresión con un número.
6. Debes tener una variable llamada `sixthResult` que sume correctamente dos números para producir el valor `22`. Reemplaza el `null` que se usa actualmente en la expresión con un número.

# --hints--

Debes tener una variable llamada `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` debería tener el valor numérico `15`.

```js
assert.strictEqual(firstResult, 15);
```

Debes tener una variable llamada `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` debería tener el valor numérico `3`.

```js
assert.strictEqual(secondResult, 3);
```

Debes tener una variable llamada `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` debería tener el valor numérico `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` no debe incluir un valor booleano (`true` o `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Debes tener una variable llamada `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` debería tener el valor numérico `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` no debe incluir un valor booleano (`true` o `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Deberías tener una variable llamada `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` debe tener el valor numérico `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` no debe incluir una multiplicación de cadenas (por ejemplo, "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Debes tener una variable llamada `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` debería tener el valor numérico `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` no debe incluir `null` en su expresión.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
