---
id: 686daa7ed79ceacd0b264e80
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

もう少しで完成です。次の例を見てみましょう：

```html
<svg viewBox="0 0 50 50">
</svg>
```

`viewBox` 属性は、SVG内で画像のどの部分が表示されるかを制御します。

- 最初の2つの数字（`0 0`）は、`viewBox` の開始位置、つまり左上の角（x と y）を設定します。 
- 次の2つの数字（`50 50`）は、`viewBox` の幅と高さを定義します。

`viewBox` 属性を `0 0 24 24` に設定してください。

# --hints--

`viewBox` 属性があるはずです。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

`viewBox` の x 位置を 0 に設定してください。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

`viewBox` の y 位置を 0 に設定してください。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

`viewBox` の幅を 24 に設定してください。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

`viewBox` の高さを 24 に設定してください。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
