---
id: 686d749d6f32f708fdf2b7ad
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

前のレッスンでは`svg`要素について学び、プロジェクトでアイコンとしてよく使われることを知りました。実際のコードベースでは通常、アイコンライブラリを利用するため、`svg`要素を一から作成する必要はありません。しかし、このワークショップでは、`svg`要素内で使われる基本的な属性について学ぶために、ハートのアイコンを作成します。

まずはページ上に`svg`要素を作成してください。

# --hints--

ページ上に`svg`要素があるはずです。

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
