---
id: 6650c9a94d6e13d14a043a69
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CamperBot 正在尝试创建一个随机背景色更换器。 但是他们不断遇到问题，需要你的帮助来调试代码。 

CamperBot 已经添加了项目的 HTML 和 CSS。 但他们对为什么页面上没有显示样式和内容感到困惑。

当他们打开控制台时，看到这条消息：

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

当 JavaScript 引擎遇到无法解释的东西时，会抛出语法错误。 在这种情况下，看起来 CamperBot 在 `darkColorsArr` 数组中有语法错误。

修复 `darkColorsArr` 数组中的语法错误，你应该能看到页面上显示内容和样式。

# --hints--

你应该解决 `darkColorsArr` 数组中的语法错误。 记住，数组元素应该用逗号分隔。

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
