---
id: 66f53dc2c5bd6a11d6c3282f
title: "Crea un'app per generare password"
challengeType: 26
dashedName: lab-password-generator
---

# --description--

In questo laboratorio, eserciterai l'uso delle funzioni creando un generatore di password casuali.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `generatePassword` che prende un parametro, che indica la lunghezza della password generata. Puoi chiamare il parametro come preferisci.
2. La tua funzione dovrebbe restituire una stringa che rappresenta una password generata casualmente. Dovresti usare la seguente stringa e diversi metodi `Math` per aiutarti a restituire una nuova stringa con caratteri casuali al suo interno: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Dovresti definire una variabile chiamata `password` e assegnarle il risultato della chiamata della funzione `generatePassword` con un argomento numerico che rappresenta la lunghezza desiderata della password.
4. Dovresti avere un `console.log` che stampa una singola stringa ottenuta concatenando il messaggio `Generated password:` e la variabile `password` separati da uno spazio.

# --hints--

Dovresti avere una funzione `generatePassword` con un parametro. Puoi chiamare il parametro come preferisci.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

La tua funzione `generatePassword` dovrebbe restituire una stringa.

```js
const result = generatePassword(5);
assert.isString(result);
```

La tua funzione `generatePassword` dovrebbe restituire una nuova stringa della lunghezza corretta.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

La tua funzione dovrebbe restituire una password generata casualmente con caratteri validi.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

La tua funzione dovrebbe restituire una password generata casualmente che contenga più di un carattere unico.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

La tua funzione dovrebbe restituire una nuova stringa casuale ogni volta che viene chiamata.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Dovresti avere una variabile `password`.

```js
assert.isDefined(password);
```

La tua variabile `password` dovrebbe essere una stringa.

```js
assert.isString(password);
```

Dovresti chiamare la funzione `generatePassword` con un argomento numerico e memorizzare la password restituita nella variabile `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Dovresti stampare una singola stringa che combina `Generated password:` e `password` separati da uno spazio singolo usando `+` o un template literal.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
