---
id: aa7697ea2477d1316795783b
title: 돼지 영어 번역기 구현하기
challengeType: 26
dashedName: lab-pig-latin
---

# --description--

돼지 영어는 특정 변환 규칙을 따라 영어 단어를 바꾸는 방법입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 하나의 문자열을 인수로 받는 `translatePigLatin` 함수를 만들어야 합니다.
1. 문자열 인수가 자음으로 시작하면, 함수는 첫 자음 또는 자음군을 단어 끝으로 옮기고 `ay`을 추가한 뒤 결과를 반환해야 합니다.
1. 문자열 인수가 모음으로 시작하면, 함수는 단어 끝에 `way`를 추가하고 결과를 반환해야 합니다.
1. 함수는 단어 중간에 첫 모음이 오는 문자열 인수를 처리하고 적절히 변환된 문자열을 반환해야 합니다.
1. 문자열 인수에 모음이 없으면, 함수는 단어 끝에 `ay`을 추가하고 결과를 반환해야 합니다.

참고: 이 실습에서 모음은 `a`, `e`, `i`, `o`, `u`입니다. `y` 문자는 모음으로 간주하지 않습니다.

# --hints--

`translatePigLatin` 함수를 만들어야 합니다.

```js
assert.isFunction(translatePigLatin);
```

`translatePigLatin("california")`은 문자열 `aliforniacay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")`은 문자열 `aragraphspay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")`는 문자열 `oveglay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")`은 문자열 `algorithmway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")`는 문자열 `eightway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

첫 모음이 단어 중간에 오는 단어를 처리해야 합니다. `translatePigLatin("schwartz")`은 문자열 `artzschway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

모음이 없는 단어를 처리해야 합니다. `translatePigLatin("rhythm")`은 문자열 `rhythmay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  let front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
