---
id: a9bd25c716030ec90084d8a1
title: Implementieren Sie den Chunky Monkey Algorithmus
challengeType: 26
dashedName: implement-the-chunky-monkey-algorithm
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Lab abzuschließen.

**User Stories:**

1. Schreiben Sie eine Funktion namens `chunkArrayInGroups`, die ein Array als erstes Argument und eine Zahl als zweites Argument entgegennimmt. Die Funktion soll das Array in kleinere Arrays aufteilen, deren Länge der zweiten Argumentzahl entspricht, und diese als zweidimensionales Array zurückgeben.

# --hints--

`chunkArrayInGroups(["a", "b", "c", "d"], 2)` sollte `[["a", "b"], ["c", "d"]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups(['a', 'b', 'c', 'd'], 2), [
  ['a', 'b'],
  ['c', 'd']
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3)` sollte `[[0, 1, 2], [3, 4, 5]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 3), [
  [0, 1, 2],
  [3, 4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2)` sollte `[[0, 1], [2, 3], [4, 5]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 2), [
  [0, 1],
  [2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4)` sollte `[[0, 1, 2, 3], [4, 5]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5], 4), [
  [0, 1, 2, 3],
  [4, 5]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3)` sollte `[[0, 1, 2], [3, 4, 5], [6]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6], 3), [
  [0, 1, 2],
  [3, 4, 5],
  [6]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4)` sollte `[[0, 1, 2, 3], [4, 5, 6, 7], [8]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 4), [
  [0, 1, 2, 3],
  [4, 5, 6, 7],
  [8]
]);
```

`chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2)` sollte `[[0, 1], [2, 3], [4, 5], [6, 7], [8]]` zurückgeben.

```js
assert.deepEqual(chunkArrayInGroups([0, 1, 2, 3, 4, 5, 6, 7, 8], 2), [
  [0, 1],
  [2, 3],
  [4, 5],
  [6, 7],
  [8]
]);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function chunkArrayInGroups(arr, size) {
  let out = [];

  for (let i = 0; i < arr.length; i += size) {
    out.push(arr.slice(i, i + size));
  }

  return out;
}

chunkArrayInGroups(['a', 'b', 'c', 'd'], 2);
```
