---
id: 614385513d91ae5c251c2052
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`Magazine` というテキストを含む `title` 要素を追加してください。さらに、外部フォントのスタイルシート `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` を参照する `link` 要素、FontAwesome のスタイルシート `https://use.fontawesome.com/releases/v5.8.2/css/all.css` を参照する `link` 要素、あなた自身のスタイルシート `./styles.css` を参照する `link` 要素をそれぞれ追加してください。

フォントスタイルシートは 3 つのフォントを読み込みます。`Anton`, `Baskervville` そして `Raleway` です。

# --hints--

Your code should have three `link` elements.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Your `link` elements should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

3 つの `link` 要素には `rel` 属性と値 `stylesheet` が必要です。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

link 要素の 1 つは `href` 属性を `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` に設定する必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

link 要素の 1 つは `href` 属性を `https://use.fontawesome.com/releases/v5.8.2/css/all.css` に設定する必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

One of your `link` elements should have the `href` set to `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

Your code should have one `title` element.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

Your `title` element should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

プロジェクトには `Magazine` というタイトルが必要です。 タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
