---
id: 685239f0cad43c7a909cbd98
title: Eseguire il debug di una pagina per l’adozione di animali
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally, proprietaria di un negozio per l’adozione di animali, ha creato la sua prima pagina web ma ci sono alcuni problemi.

Il tuo compito è correggere tutti gli errori così Sally potrà continuare a costruire la sua pagina.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Sally vuole usare un’immagine di alcuni gatti ma non viene visualizzata correttamente. Devi correggere quanto segue nell’elemento `img`:
   - Sostituisci l’attributo `href` con l’attributo corretto per la sorgente dell’immagine.
   - Sostituisci l’attributo `att` con l’attributo corretto che rappresenta un testo breve e descrittivo per le immagini.
   - Rimuovi il tag di chiusura `</img>` perché gli elementi `img` sono elementi void e non hanno tag di chiusura.
2. Sally vuole usare alcuni collegamenti per indirizzare gli utenti alle pagine dei cani e dei gatti. Ma i collegamenti non funzionano correttamente. Devi correggere quanto segue negli elementi `a`:
   - Sostituisci entrambi gli attributi `src` con gli attributi corretti usati per specificare gli URL.

# --hints--

Il tuo elemento `img` dovrebbe avere un attributo `src` invece dell’attributo `href`.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Il tuo elemento `img` dovrebbe avere un attributo `alt` invece dell’attributo `att` inesistente.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Il tuo elemento `img` non dovrebbe avere un tag di chiusura `</img>`.

```js
assert.notMatch(code, /<\/img>/);
```

Il tuo elemento `a` con il testo `Visit cats page` deve avere un attributo `href` invece di un attributo `src`.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Il tuo elemento `a` con il testo `Visit dogs page` deve avere un attributo `href` invece di un attributo `src`.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
