---
id: 694c9c186c089f7bdff0f3c6
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Im Allgemeinen ist es gute Praxis, Ihre Eingaben innerhalb einer Funktion zu validieren. Verwenden Sie eine `for`-Schleife, um das `crew`-Array zu durchlaufen und zu prüfen, ob ein Mitglied bereits denselben `id` wie die Eingabe `astronaut` hat. Wenn ein Duplikat gefunden wird, geben Sie die Fehlermeldung `console.log("Duplicate ID: " + astronaut.id)` aus und rufen Sie dann `return` auf, um die Funktion frühzeitig zu beenden.

Hier ist ein Beispiel für diese häufig verwendete Technik:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Ihre Funktion sollte eine `for`-Schleife verwenden, um vorhandene `crew`-Mitglieder zu durchsuchen (jede gültige `for (...)`-Form ist in Ordnung).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Innerhalb der `for`-Schleife sollten Sie nach einem doppelten `id` suchen. Wenn eines gefunden wird, geben Sie `"Duplicate ID: " + astronaut.id"` in der Konsole aus und rufen Sie dann `return` auf, um die Funktion zu beenden.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Bei gültigen Eingaben sollte Ihre `addCrewMember`-Funktion keine `console.log()`-Fehlermeldung oder einen frühen `return`-Abbruch auslösen.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
