---
id: 694c84106e9023f289c03b84
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Ihr Team beginnt mit einem einzigen Kommandanten-Astronauten. Erstellen Sie ein Objekt namens `firstAstronaut` mit den folgenden Eigenschaften:

| Schlüssel | Wert |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Sie sollten eine `firstAstronaut`-Variable erstellen.

```js
assert.isDefined(firstAstronaut);
```

Ihr `firstAstronaut` sollte ein Objekt sein.

```js
assert.isObject(firstAstronaut);
```

Ihr `firstAstronaut`-Objekt sollte die korrekten Eigenschaftswerte enthalten.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
