---
id: 694c40d86fe6a78d20b6559c
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop nutzen Sie JavaScript, um eine Liste von Astronauten für eine Weltraummission zu erstellen.

Beginnen Sie damit, ein leeres Array für Ihre Liste mit dem Namen `squad` zu erstellen.

# --hints--

Sie sollten eine `squad`-Variable erstellen.

```js
assert.isDefined(squad);
```

Ihre `squad`-Variable sollte ein Array sein.

```js
assert.isArray(squad);
```

Ihr `squad`-Array sollte leer sein.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
