---
id: 694ce8609fb805ba21cebc02
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Die Missionskontrolle benötigt eine Methode, um die Positionen von zwei Crewmitgliedern in einer Besetzungsliste zu tauschen. Erstellen Sie eine leere Funktion namens `swapCrewMembers` mit drei Parametern:

- `crew`: ein Array von Astronautenobjekten

- `fromIndex`: der Index des ersten Astronauten, der getauscht werden soll

- `toIndex`: der Index des zweiten Astronauten, der getauscht werden soll

# --hints--

Sie sollten eine Funktion namens `swapCrewMembers` erstellen.

```js
assert.isFunction(swapCrewMembers);
```

Ihre `swapCrewMembers`-Funktion sollte die Parameter `crew`, `fromIndex` und `toIndex` haben.

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex']);
assert.match(__helpers.removeJSComments(swapCrewMembers.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

--fcc-editable-region--

--fcc-editable-region--
```
