---
id: 694c84330cbe604821b9deef
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Die Missionskontrolle benötigt eine Methode, um neue Mitglieder zur Liste hinzuzufügen. Erstellen Sie eine leere Funktion namens `addCrewMember`, die zwei Parameter akzeptiert, `crew` und `astronaut`.

# --hints--

Sie sollten eine Funktion namens `addCrewMember` erstellen.

```js
assert.isFunction(addCrewMember);
```

Ihre `addCrewMember`-Funktion sollte die Parameter `crew` und `astronaut` haben.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
