---
id: 694d8ddd106340237efb0c80
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Die `splice()`-Methode kann Arrays verändern, indem sie Elemente an beliebigen Positionen hinzufügt oder entfernt, auch in der Mitte. Da `splice()` ein Array mit den entfernten Elementen zurückgibt, können Sie sie verwenden, um zwei Elemente in einem Array zu tauschen, ohne das Original zu verändern. Dies lässt sich mit der folgenden Technik in einer Zeile erledigen:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Hier ist ein Beispiel:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

Die Technik, angewendet in der dritten Zeile, funktioniert wie folgt:

- `splice(3, 1, copyArray[1])` entfernt das Element am Index `3` (`55`)

- Es fügt das Element von `copyArray[1]` (`97`) an Index `3` ein

- `splice()` gibt `[55]` zurück, ein Array mit dem entfernten Element

- `[0]` extrahiert `55` aus diesem Array

- Dieser Wert (`55`) wird zurück an `copyArray[1]` zugewiesen, womit der Tausch abgeschlossen ist

Verwenden Sie diese Technik mit dem `updatedCrew`-Array, um die Astronauten an `fromIndex` und `toIndex` zu tauschen.

# --hints--

Sie sollten den Tausch durchführen, indem Sie die `splice()`-Methode aufrufen.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Sie sollten `splice()` mit dem `updatedCrew`-Array aufrufen, um den Tausch durchzuführen. Verwenden Sie die Ein-Zeilen-Technik.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Ihre `swapCrewMembers`-Funktion sollte das ursprüngliche Eingabe-Array `crew` nicht verändern.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
