---
id: 6140c7e645d8e905819f1dd4
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

在 `meta` 元素中添加正确的 `charset`、`title` 元素和链接 `./styles.css` 文件的 `link` 元素。

将 `title` 设置为 `Ferris Wheel`。

# --hints--

你的代码应该包含 `DOCTYPE` 声明。

```js
assert.match(code, /<!DOCTYPE/gi);
```

你应该在 `DOCTYPE` 声明后添加一个空格。

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

您应将文档类型定义为 `html` 。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

在类型后，您应用 `>` 来关闭 `DOCTYPE` 声明。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

您的 `html` 元素应该有一个结束标签。

```js
assert.match(code, /<\/html\s*>/);
```

您的 `DOCTYPE` 声明应位于最前面（HTML的开头）。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

您应有一个 `head` 开始标签。

```js
assert.match(code, /<head\s*>/i);
```

您应有一个 `head` 结束标签。

```js
assert.match(code, /<\/head\s*>/i);
```

您应有一个 `body` 开始标签。

```js
assert.match(code, /<body\s*>/i);
```

您应有一个 `body` 结束标签。

```js
assert.match(code, /<\/body\s*>/i);
```

你的 `body` 元素应该是 `head` 元素之后的第一个元素。

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

你的 `head` 和 `body` 元素应该是兄弟元素。

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

你的 `head` 元素应该是 `html` 元素中的第一个元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

你的 `head` 元素应该在 `html` 元素中。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素应该在 `html` 元素中。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

你的 `meta` 元素应该在 `head` 元素中。

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

你应该有一个 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
