---
id: 60a3e3396c7b40068ad69974
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

每幅畫都需要一個影格。

將 `.canvas` 元素包裹在另一個 `div` 中。為該 `div` 添加 `frame` 類別。
# --hints--

你應該為你的 `div` 元素添加一個新的元素。

```js
assert(document.querySelectorAll('div').length === 2)
```

你的 `.canvas` 元素應該巢狀在新的 `div` 元素中。

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

你的新 `div` 應該有一個 `class`，其值為 `frame`。

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

你的新 `div` 應該位於你的 `body` 元素內。

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
