---
id: 60a3e3396c7b40068ad6996e
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

外邊距是盒子外部的區域，可用來控制其他盒子或元素之間的間距。

這裡底部元素有較大的上邊距，將它推得更往頁面下方。

既然你已經了解了 CSS 盒模型，讓我們開始進行 Rothko 繪畫。

移除 `<img>` 元素。

# --hints--

你的程式碼中不應該有 `img` 元素。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
