---
id: 60a3e3396c7b40068ad69978
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

使用邊距來調整元素外部的間距。

使用 `margin` 屬性，為 `.frame` 元素設定垂直邊距為 `20px`，水平邊距為 `auto`。這將使訊框向下移動 20 像素，並在頁面上水平置中。

# --hints--

你應該將 `margin` 屬性設為 `20px auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

你的 `.frame` 元素應該有一個 `margin` 值為 `20px auto`。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
