---
id: 60a3e3396c7b40068ad6997d
title: 第 20 步
challengeType: 0
dashedName: step-20
---

# --description--

使用邊距將 `.one` 元素定位在畫布上。

為你的樣式添加縮寫 `margin` 屬性，垂直邊距為 `20px`，水平邊距為 `auto`。

# --hints--

你應該將 `margin` 屬性設為 `20px auto`。

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

你的 `.one` 元素應該有一個 `margin` 值為 `20px auto`。

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
