---
id: 60a3e3396c7b40068ad69971
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

是時候學習 CSS 了。

即使你的 `<div>` 沒有文字，它仍然被視為一個有內容的框。  
撰寫一個使用 `.canvas` 類別選擇器的 CSS 規則，並將其 `width` 設為 500 像素。  
這是一個將類別 `card` 的寬度設為 300 像素的 CSS 規則：

```css
.card {
  width: 300px;
}
```

# --hints--

你的程式碼應該有一個 `.canvas` 選擇器。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

你應該將 `width` 屬性設為 `500px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

你的 `.canvas` 元素應該有 `width` 為 `500px`。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
