---
id: 60a3e3396c7b40068ad6996b
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `head` 元素內，為你的 `meta` 頁籤設定 `charset` 為 `UTF-8`，並新增一個 `title` 元素，值為 `Rothko Painting`。

在 `body` 元素中，為你的 `img` 元素添加一個 `src`，其值為 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`。
# --hints--

你的程式碼應該有一個 `meta` 頁籤。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 頁籤應將 `charset` 設定為 `UTF-8`。

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

你的程式碼應該有一個 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 應該是 `Rothko Painting`。

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

你的程式碼應該有一個 `img` 元素。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 元素的 `src` 應該是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
