---
id: 60a3e3396c7b40068ad69970
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在你可以開始為你添加的 `div` 設定樣式之前，你需要將你的 CSS 聯結到你的 HTML。

為你的 `styles.css` 檔案添加一個 `link` 元素。將 `href` 設定為 `styles.css`，並記得將 `rel` 屬性設定為 `stylesheet`。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
assert(/<link/.test(code))
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
