---
id: 60a3e3396c7b40068ad6996d
title: 步驟4
challengeType: 0
dashedName: step-4
---

# --description--

內容周圍有一個稱爲內邊距的空間，類似於氣泡膜將物品與包裝盒分隔開來。

把邊框想象成裝運物品的紙箱。

將 `src` 屬性更改爲 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`。

# --hints--

`img` 元素的 `src` 應該是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
