---
id: 60a3e3396c7b40068ad6996a
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

到現在為止，你應該已經熟悉一個 HTML 頁面應該具備的基本元素。

為你的程式碼設定一個 `DOCTYPE` 宣告、一個語言設定為英文的 `html` 元素、一個 `head` 元素，以及一個 `body` 元素。

# --hints--

你的程式碼應該有一個 `<!DOCTYPE html>` 宣告式。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個開啟的 `<html>` 頁籤，且其 `lang` 屬性為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的程式碼應該在 `html` 元素內有一個 `head` 元素。

```js
assert(document.querySelectorAll('head').length === 1);
```

你的程式碼應該在 `html` 元素內有一個 `body` 元素。

```js
assert(document.querySelectorAll('body').length === 1);
```

你的 `head` 元素應該位於你的 `body` 元素之前。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
