---
id: 60a3e3396c7b40068ad6997f
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

為畫布的上方、下方、左方和右方各添加 1 像素的內邊距，使其尺寸變為 502 像素 x 602 像素。

將 `padding` 屬性取代為設定 `overflow` 為 `hidden`－將畫布恢復為其原始尺寸。

# --hints--

你應該從 `.canvas` 選擇器中移除 `padding` 屬性。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

你應該將 `overflow` 屬性設定為 `hidden`。

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

你的 `.canvas` 元素的 `overflow` 值應該是 `hidden`。

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
