---
id: 60a3e3396c7b40068ad6998b
title: 步驟 33
challengeType: 0
dashedName: step-33
---

# --description--

讓你的邊距往一個方向推進會很有幫助。

在這種情況下，`.one` 元素的底部外距將 `.two` 向下推動 20 像素。

在 `.two` 選擇器中，使用 `margin` 簡寫屬性將上方外距設為 `0`、水平外距設為 `auto`，以及下方外距設為 `20px`。這將移除它的上方外距、使其水平置中，並將下方外距設為 20 像素。

# --hints--

你應該將 `margin` 屬性設為 `0 auto 20px`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

你的 `.two` 元素應該有一個 `margin` 值為 `0 auto 20px`。

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
