---
id: 6823d6244511f252c8300eed
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

웹페이지에 단락을 추가해야 할 때는 `p` 요소를 이렇게 사용할 수 있습니다:

```html
<p>This is a paragraph element.</p>
```

`Learn the skills to become a full-stack developer`을 단락 요소로 바꾸세요.

# --hints--

`p` 요소에는 여는 태그가 있어야 합니다. 여는 태그의 구문은 다음과 같습니다: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

`p` 요소에는 닫는 태그가 있어야 합니다. 닫는 태그의 구문은 다음과 같습니다: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

`p` 요소의 텍스트는 `Learn the skills to become a full-stack developer`이어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

`p` 요소는 `h2` 요소 아래에 있어야 합니다.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
