---
id: 6823e036ea4b71553558c01b
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

미리 보기에서 각 제목이 조금씩 다르게 보이는 것을 확인하세요. 단락도 마찬가지입니다.

다음으로, 다음 텍스트를 표시하는 또 다른 `p` 요소를 추가하세요:

```md
HTML represents the content and structure of a webpage
```

# --hints--

두 번째 `p` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

`p` 요소에는 `HTML represents the content and structure of a webpage` 텍스트가 있어야 합니다. 지침에서 텍스트를 복사할 수 있습니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
