---
id: 6823ac607bfdbc46331b2559
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML은 하이퍼텍스트 마크업 언어를 뜻합니다. 웹페이지의 구조와 내용을 정의하는 코드입니다. 여기는 HTML을 작성할 코드 편집기입니다.

편집기에서 1번째 줄을 찾아 다음 텍스트를 입력하세요:

```md
Welcome to freeCodeCamp
```

입력이 끝나면 "코드 확인" 버튼을 눌러 올바른지 확인하세요.

# --hints--

편집기에 `Welcome to freeCodeCamp` 텍스트가 있어야 합니다. 철자를 다시 확인하세요.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
