---
id: 6823d9ac8bdc3853df65a1ff
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

HTML에는 여섯 개의 제목 요소가 있습니다: `h1`부터 `h6`까지입니다. 이들은 웹페이지의 섹션 중요도를 나타내는 데 사용되며, `h1`가 가장 중요하고 `h6`가 가장 덜 중요합니다.

`p` 요소 아래에 `h3` 제목을 추가하고 다음 텍스트를 입력하세요:

```md
Introduction to HTML
```

# --hints--

`h3` 요소는 여는 `<h3>` 태그를 가져야 합니다.

```js
assert.exists(document.querySelector("h3"));
```

`h3` 요소는 닫는 `</h3>` 태그를 가져야 합니다.

```js
assert.match(code, /<\/h3\s*\>/);
```

`h3` 요소의 텍스트는 `Introduction to HTML`이어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
