---
id: 6823e169fda14755fbf00445
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

감을 잡고 있네요. 다음으로, 편집기 하단에 텍스트가 있는 또 다른 `h3` 요소를 추가하세요:

```md
Introduction to CSS
```

# --hints--

페이지에 두 번째 `h3` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

`h3` 요소의 텍스트는 `Introduction to CSS`이어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```
