---
id: 6823e637c1c0ed56f781b4fc
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

마지막으로, JavaScript는 웹페이지를 인터랙티브하게 만듭니다 — 누군가 버튼을 클릭하거나, 폼을 제출하거나, 그 밖의 여러 행동을 할 때 페이지가 어떻게 반응할지 지시할 수 있습니다.

워크숍의 마지막 단계로, JavaScript를 설명하는 `h3`와 `p` 요소를 페이지에 하나씩 더 추가하세요.

먼저, 다음 텍스트가 들어간 `h3` 요소를 추가하세요:

```md
Introduction to JavaScript
```

그다음, 그 `h3` 요소 아래에 다음 텍스트가 들어간 `p` 요소를 추가하세요:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

페이지에 세 번째 `h3` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

`h3` 요소의 텍스트는 `Introduction to JavaScript`이어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

페이지에 네 번째 `p` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

`p` 요소에는 `JavaScript adds interactivity to a webpage` 텍스트가 있어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
