---
id: 682ba2318000b62f179bdf04
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

HTML은 요소들로 구성되어 있습니다. 처음 사용할 요소는 `h1` 요소입니다:

```html
<h1>Welcome to freeCodeCamp</h1>
```

이 요소는 여는 태그(`<h1>`)로 시작하고, 닫는 태그(`</h1>`)로 끝나며, 태그 사이에 표시할 텍스트가 들어갑니다.

`Welcome to freeCodeCamp` 텍스트 앞에 여는 태그를 추가하고, 뒤에 닫는 태그를 추가하여 `h1` 요소로 만드세요.

# --hints--

`h1` 요소에는 여는 `<h1>` 태그가 있어야 합니다.

```js
assert.exists(document.querySelector("h1"));
```

`h1` 요소에는 닫는 `</h1>` 태그가 있어야 합니다.

```js
assert.match(code, /<\/h1\s*\>/);
```

`h1` 요소는 다음과 같이 보여야 합니다: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
