---
id: 682cd20b883fc7b25eb539c6
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

다른 두 줄 텍스트 아래에 다음을 추가하세요:

```md
Learn the skills to become a full-stack developer
```

# --hints--

편집기에 `Learn the skills to become a full-stack developer` 텍스트가 있어야 합니다. 철자를 다시 확인하세요.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```
