---
id: 682cd206883fc7b25eb539c5
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

편집기에서 작성한 HTML이 미리보기 창에 나타나는 것을 확인하세요. 이 워크숍에서는 부분적인 커리큘럼 웹페이지용 HTML을 작성할 것입니다.

`h1` 요소 아래 빈 줄에 다음 내용을 입력하세요:

```md
Full-Stack Curriculum
```

# --hints--

편집기에 `Full-Stack Curriculum` 텍스트가 있어야 합니다. 철자를 다시 확인하세요.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
