---
id: 683921c4769fd23dbadec2fe
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

HTML이 웹페이지의 구조와 내용을 정의하는 반면, CSS는 색상, 폰트, 간격, 레이아웃 같은 스타일을 추가하는 데 사용됩니다.

방금 추가한 `h3` 아래에, 다음 텍스트가 들어간 또 다른 단락 요소를 추가하세요:

```md
CSS is used to style a webpage
```

# --hints--

페이지에 두 번째 `h3` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

`h3` 요소의 텍스트는 `Introduction to CSS` 이어야 합니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

페이지에 세 번째 `p` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

`p` 요소에는 `CSS is used to style a webpage` 텍스트가 있어야 합니다. 지침에서 텍스트를 복사할 수 있습니다.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
