---
id: 66ebd4ae2812430bb883c787
title: 이벤트 허브 만들기
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

이 실습에서는 의미론적 HTML 요소를 활용해 웹 페이지 구조를 만듭니다. 콘텐츠와 이미지를 추가해 실제 이벤트 허브처럼 보이게 하세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `header` 요소가 있어야 합니다.
1. `header` 요소 안에 `h1` 요소가 있어야 하며, 이 요소는 `Event Hub` 텍스트와 `nav` 요소를 포함해야 합니다.
1. `nav` 요소 안에는 페이지의 다른 섹션으로 연결되는 링크 두 개가 포함된 순서 없는 목록이 있어야 합니다. 첫 번째 항목은 `Upcoming Events` 텍스트를, 두 번째 항목은 `Past Events` 텍스트를 가져야 합니다.
1. 각 링크는 `a` 요소로 나타내야 하며, `href` 속성으로 페이지의 해당 섹션 `#upcoming-events`과 `#past-events`에 연결되어야 합니다.
1. 페이지의 여러 섹션을 포함하는 `main` 요소가 있어야 합니다.
1. `main` 요소 안에는 두 개의 `section` 요소가 있어야 합니다.
1. 첫 번째 `section` 요소는 `id` 속성을 가지고 값은 `upcoming-events`이어야 합니다.
1. `#upcoming-events` 섹션 안에는 다음이 있어야 합니다:
   
   - `h2` 요소가 `Upcoming Events` 텍스트를 포함해야 합니다.
   - 두 개의 `article` 요소가 있어야 합니다. 각 article은 이벤트를 나타내며 다음을 포함해야 합니다:
      - 이벤트 제목을 위한 `h3` 요소
      - 이벤트 설명을 위한 `p` 요소. 원한다면 하단에 날짜를 추가할 수 있습니다.

1. 두 번째 `section` 요소는 `id` 속성을 가지고 값은 `past-events`이어야 합니다.
1. `#past-events` 섹션 안에는 다음이 있어야 합니다:

   - `h2` 요소가 `Past Events` 텍스트를 포함해야 합니다.
   - 두 개의 `article` 요소가 있어야 합니다. 각 article 요소는 과거 이벤트를 나타내며 다음을 포함해야 합니다:
     - 이벤트 제목을 위한 `h3` 요소
     - 이벤트 설명을 위한 `p` 요소. 원한다면 하단에 날짜를 추가할 수 있습니다.
     - `src` 속성이 이미지 파일을 가리키고 `alt` 속성이 이미지 설명을 포함하는 이미지 요소

**참고:** 이벤트 설명과 날짜는 어떤 텍스트든 사용할 수 있습니다. 원한다면 다음 이미지 URL을 이미지에 사용할 수 있습니다:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

`header` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector("header"));
```

`header` 요소는 `body` 여는 태그 뒤에 와야 합니다.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

`header` 요소 안에는 `h1` 요소가 있어야 하며, 이 요소는 `Event Hub` 텍스트를 포함해야 합니다.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

`header` 요소 안에는 `h1` 요소 뒤에 `nav` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

`nav` 요소 안에는 두 개의 항목이 있는 순서 없는 목록이 있어야 합니다.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

순서 없는 목록의 첫 번째 항목은 링크여야 합니다.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

순서 없는 목록의 두 번째 항목도 링크여야 합니다.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

순서 없는 목록 첫 번째 항목의 텍스트는 `Upcoming Events` 이어야 합니다.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

순서 없는 목록 첫 번째 항목은 `href` 속성이 `#upcoming-events` 로 설정되어야 합니다.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

순서 없는 목록 두 번째 항목의 텍스트는 `Past Events` 이어야 합니다.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

순서 없는 목록 두 번째 항목은 `href` 속성이 `#past-events` 로 설정되어야 합니다.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

`main` 요소는 `header` 요소 닫는 태그 뒤에 있어야 합니다.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

`main` 요소 안에는 두 개의 `section` 요소가 있어야 합니다.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

첫 번째 `section` 요소는 `id` 속성을 가지고 값은 `upcoming-events` 이어야 합니다.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

두 번째 `section` 요소는 `id` 속성을 가지고 값은 `past-events` 이어야 합니다.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

`#upcoming-events` 섹션 안에는 `h2` 요소가 `Upcoming Events` 텍스트를 포함해야 합니다.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

`#upcoming-events` 섹션 안에는 `article` 요소 아래에 두 개의 `h2` 요소가 있어야 합니다.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`article` 섹션 안 두 `#upcoming-events` 요소 모두 이벤트 제목을 위한 `h3` 요소를 가져야 합니다.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`article` 섹션 안 두 `#upcoming-events` 요소 모두 이벤트 설명을 위한 단락 요소를 가져야 합니다.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`#past-events` 섹션 안에는 `h2` 요소가 `Past Events` 텍스트를 포함해야 합니다.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

`#past-events` 섹션 안에는 `article` 요소 아래에 두 개의 `h2` 요소가 있어야 합니다.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

`article` 섹션 안 두 `#past-events` 요소 모두 이벤트 제목을 위한 `h3` 요소를 가져야 합니다.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

`article` 섹션 안 두 `#past-events` 요소 모두 이벤트 설명을 위한 단락 요소를 가져야 합니다.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

`article` 섹션 안 두 `#past-events` 요소 모두 이미지 요소를 가져야 합니다.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

`#past-events` 섹션 안 두 이미지 요소 모두 `src` 속성이 이미지 파일을 가리켜야 합니다.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

`#past-events` 섹션 안 두 이미지 요소 모두 `alt` 속성이 이미지 설명을 포함해야 합니다.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

각 `h3` 요소는 이벤트 제목을 가져야 합니다.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

각 `p` 요소는 이벤트 설명을 가져야 합니다.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```
