---
id: 686ccc2c8b967e17ab18d593
title: Implementieren der Bisektionsmethode
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

Die Bisektionsmethode, auch bekannt als binäre Suchmethode, verwendet eine binäre Suche, um die Nullstellen einer reellwertigen Funktion zu finden. Sie funktioniert, indem sie ein Intervall eingrenzt, in dem die Quadratwurzel liegt, bis sie auf einen Wert innerhalb einer vorgegebenen Toleranz konvergiert.

Wenn die Toleranz zum Beispiel `0.01` beträgt, halbiert die Bisektionsmethode das Intervall so lange, bis die Differenz zwischen der oberen und unteren Grenze kleiner oder gleich `0.01` ist.

In diesem Labor implementieren Sie eine Funktion, die die Bisektionsmethode verwendet, um die Quadratwurzel einer Zahl zu finden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `square_root_bisection` definieren, die drei Parameter hat:
    - Die Zahl, für die Sie die Quadratwurzel finden möchten.
    - Die Toleranz, also die akzeptable Fehlermarge für das Ergebnis. Sie sollten einen Standardwert für die Toleranz festlegen.
    - Die maximale Anzahl der durchzuführenden Iterationen. Sie sollten eine Standardanzahl von Iterationen festlegen.

1. Die `square_root_bisection`-Funktion sollte:
    - Eine `ValueError` mit der Nachricht `Square root of negative number is not defined in real numbers` auslösen, wenn die an die Funktion übergebene Zahl negativ ist.
    - Für die Zahlen `0` und `1` die Nachricht `The square root of [number] is [number]` ausgeben und die Zahl selbst als Quadratwurzel zurückgeben.
    - Für jede andere positive Zahl die ungefähre Quadratwurzel mit der Nachricht `The square root of [square_target] is approximately [root]` ausgeben und den berechneten Wurzelwert zurückgeben.
    - Wenn kein Wert die Toleranzbedingung erfüllt, eine Fehlermeldung `Failed to converge within [maximum] iterations` ausgeben und `None` zurückgeben.

**Hinweis**: Sie dürfen für dieses Labor kein Modul importieren.

# --hints--

Sie sollten kein Modul importieren.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

Sie sollten eine Funktion mit dem Namen `square_root_bisection` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

Ihre `square_root_bisection`-Funktion sollte drei Parameter haben.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

Sie sollten einen Standardwert für die Toleranz und die maximale Anzahl der Iterationen festlegen.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

Ihre `square_root_bisection`-Funktion sollte eine `ValueError` mit der Nachricht `Square root of negative number is not defined in real numbers` auslösen, wenn die an die Funktion übergebene Zahl negativ ist.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` sollte `0` zurückgeben.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` sollte `The square root of 0 is 0` ausgeben.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` sollte eine Zahl zwischen `0.03162267660168379` und `0.031622876601683794` zurückgeben.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` sollte `The square root of 0.001 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `0.03162267660168379` und `0.031622876601683794` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` sollte eine Zahl zwischen `0.4999999` und `0.5000001` zurückgeben.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` sollte `The square root of 0.25 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `0.4999999` und `0.5000001` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` sollte `1` zurückgeben.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` sollte `The square root of 1 is 1` ausgeben.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` sollte eine Zahl zwischen `8.999` und `9.001` zurückgeben.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` sollte `The square root of 81 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `8.999` und `9.001` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` sollte eine Zahl zwischen `14.999` und `15.001` zurückgeben.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` sollte `The square root of 225 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `14.999` und `15.001` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` sollte eine Zahl zwischen `14.99999` und `15.00001` zurückgeben.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` sollte `The square root of 225 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `14.99999` und `15.00001` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` sollte eine Zahl zwischen `14.9999999` und `15.0000001` zurückgeben.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` sollte `The square root of 225 is approximately X` ausgeben, wobei `X` eine Zahl zwischen `14.9999999` und `15.0000001` ist.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` sollte `None` zurückgeben.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` sollte `Failed to converge within 10 iterations` ausgeben.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
