---
id: 66325a250690a3612c1db0f6
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot 創建了一個名爲 `bgHexCodeSpanElement` 的新變量來聲明 `id` 爲 `bg-hex-code` 的 `span` 元素。 但是，當他們嘗試將該變量輸出到控制檯時，得到的是 `null`。

`null` 是 JavaScript 中表示缺少值的特殊值。 當你試圖訪問一個不存在的對象屬性時，會發生這種情況。

在這種情況下，CamperBot 沒有傳遞正確的選擇器給 `document.querySelector` 方法。

修復 `document.querySelector("bg-hex-code")` 行，使其能正確選擇 `id` 爲 `bg-hex-code` 的元素。

# --hints--

你應該修正 `"bg-hex-code"` 以便它正確引用 `id` 爲 `bg-hex-code`。 請記住，CSS 中的 `id` 選擇器以井號 (`#`) 開頭。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
