---
id: 56533eb9ac21ba0edf2244cf
title: 레코드 컬렉션 만들기
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

음악 앨범 컬렉션 유지 관리를 돕는 함수를 만들고 있습니다. 이 컬렉션은 여러 앨범을 포함하는 객체로 구성되며, 각 앨범도 객체입니다. 각 앨범은 컬렉션에서 고유한 `id`를 속성 이름으로 하여 나타냅니다. 각 앨범 객체 내에는 앨범에 관한 다양한 정보를 설명하는 여러 속성이 있습니다. 모든 앨범이 완전한 정보를 가진 것은 아닙니다.

`updateRecords` 함수는 다음 함수 매개변수로 나타내는 4개의 인수를 받습니다:

-   `records` - 여러 개별 앨범을 포함하는 객체
-   `id` - `records` 객체 내 특정 앨범을 나타내는 숫자
-   `prop` - 업데이트할 앨범 속성의 이름을 나타내는 문자열
-   `value` - 앨범 속성을 업데이트하는 데 사용되는 정보를 담은 문자열

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 함수는 항상 전체 `records` 객체를 반환해야 합니다.
2. `value`가 빈 문자열이면, 해당 앨범에서 주어진 `prop` 속성을 삭제하세요.
3. `prop`이 `tracks`가 아니고 `value`가 빈 문자열이 아니면, 그 앨범의 `value`에 `prop`을 할당하세요.
4. `prop`가 `tracks`이고 `value`가 빈 문자열이 아니며, 앨범에 `tracks` 속성이 없으면, 빈 배열을 만들고 `value`를 추가하세요.
5. 프로프가 `tracks`이고 `value`가 빈 문자열이 아니면, 앨범의 기존 `value` 배열 끝에 `tracks`를 추가하세요.

**참고:** 테스트에는 `recordCollection` 객체의 복사본이 사용됩니다. 함수는 `recordCollection` 객체를 직접 참조하지 말고 함수 매개변수만 사용하세요.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

`updateRecords` 함수를 만들어야 합니다.

```js
assert.isFunction(updateRecords);
```

`updateRecords(recordCollection, 5439, "artist", "ABBA")` 후에 `artist`는 문자열 `ABBA`여야 합니다.

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

`updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")` 후에 `tracks`는 문자열 `Take a Chance on Me`를 마지막이자 유일한 요소로 가져야 합니다.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

`updateRecords(recordCollection, 2548, "artist", "")` 후에 `artist`는 설정되어 있으면 안 됩니다.

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

`updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")` 후에 `tracks`는 문자열 `Addicted to Love`를 마지막 요소로 가져야 합니다.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

`updateRecords(recordCollection, 2468, "tracks", "Free")` 후에 `tracks`는 문자열 `1999`를 첫 번째 요소로 가져야 합니다.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

`updateRecords(recordCollection, 2548, "tracks", "")` 후에 `tracks`는 설정되어 있으면 안 됩니다.

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

`updateRecords(recordCollection, 1245, "albumTitle", "Riptide")` 후에 `albumTitle`는 문자열 `Riptide`여야 합니다.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
