---
id: 5e44414f903586ffb414c950
title: 확률 계산기 프로젝트 만들기
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

파란 공 5개, 빨간 공 4개, 초록 공 2개가 들어 있는 모자가 있다고 가정해 보세요. 4개의 공을 무작위로 뽑았을 때 적어도 1개의 빨간 공과 2개의 초록 공이 나올 확률은 얼마일까요? 고급 수학을 사용해 확률을 계산할 수도 있지만, 더 쉬운 방법은 프로그램을 작성해 많은 실험을 수행하여 근사 확률을 추정하는 것입니다.

이 프로젝트에서는 모자에서 특정 공을 무작위로 뽑을 때의 근사 확률을 구하는 프로그램을 작성할 것입니다.

먼저 `Hat`에 `main.py` 클래스를 만드세요. 이 클래스는 모자에 들어 있는 각 색깔 공의 개수를 지정하는 가변 인수를 받아야 합니다. 예를 들어, 클래스 객체는 다음과 같은 방법으로 생성할 수 있습니다:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

모자는 항상 적어도 하나의 공으로 생성됩니다. 생성 시 모자 객체에 전달된 인수는 `contents` 인스턴스 변수로 변환되어야 합니다. `contents`는 모자에 있는 각 공을 나타내는 문자열 하나씩을 포함하는 리스트여야 합니다. 예를 들어, 모자가 `{'red': 2, 'blue': 1}`라면 `contents`는 `['red', 'red', 'blue']`이어야 합니다.

`Hat` 클래스는 `draw` 메서드를 가져야 하며, 이 메서드는 모자에서 뽑을 공의 개수를 인수로 받습니다. 이 메서드는 `contents`에서 무작위로 공을 제거하고, 그 공들을 문자열 리스트로 반환해야 합니다. 공은 뽑는 동안 모자에 다시 들어가지 않아야 하며, 이는 교체 없는 추출 실험과 같습니다. 뽑을 공의 개수가 모자에 있는 공의 수보다 많으면 모든 공을 반환하세요.

다음으로 `experiment`에 `main.py` 함수를 만드세요 (`Hat` 클래스 내부가 아닙니다). 이 함수는 다음 인수를 받아야 합니다:

- `hat`: 모자 객체로, 함수 내부에서 복사되어야 합니다.
- `expected_balls`: 실험에서 모자에서 뽑으려는 정확한 공 그룹을 나타내는 객체입니다. 예를 들어, 모자에서 파란 공 2개와 빨간 공 1개를 뽑을 확률을 구하려면 `expected_balls`를 `{'blue':2, 'red':1}`으로 설정하세요.
- `num_balls_drawn`: 각 실험에서 모자에서 뽑을 공의 개수입니다.
- `num_experiments`: 수행할 실험 횟수입니다. (실험 횟수가 많을수록 근사 확률이 더 정확해집니다.)

`experiment` 함수는 확률을 반환해야 합니다.

예를 들어, 검은 공 6개, 빨간 공 4개, 초록 공 3개가 들어 있는 모자에서 5개의 공을 뽑을 때 적어도 빨간 공 2개와 초록 공 1개가 나올 확률을 구하려면, `N` 실험을 수행하고, `M`만큼 적어도 빨간 공 2개와 초록 공 1개가 나온 횟수를 세어 확률을 `M/N`로 추정합니다. 각 실험은 지정된 공이 들어 있는 모자에서 시작해 여러 개의 공을 뽑고, 뽑으려던 공이 나왔는지 확인하는 과정으로 이루어집니다.

아래는 위 예시를 바탕으로 2000번 실험을 수행하는 `experiment` 함수 호출 방법입니다:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

출력은 다음과 비슷할 것입니다:

```bash
0.356
```

무작위 추출에 기반하므로 코드를 실행할 때마다 확률이 약간 다를 수 있습니다.

_힌트: 상단에 이미 가져온 모듈을 사용하는 것을 고려하세요. 파일 내에서 random 시드를 초기화하지 마세요._

참고: F12를 눌러 브라우저 콘솔을 열면 테스트의 자세한 출력을 볼 수 있습니다.

# --hints--

`hat` 객체 생성 시 올바른 내용이 추가되어야 합니다.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

`draw` 클래스의 `hat` 메서드는 contents의 항목 수를 줄여야 합니다.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

모자에 있는 공의 수보다 뽑을 공의 수가 많을 때 `draw` 메서드는 올바르게 동작해야 합니다.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

`experiment` 메서드는 다른 확률을 반환해야 합니다.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
