---
id: 67fccc2463253b213a000142
title: 멀티미디어 플레이어 만들기
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

이전 수업에서는 `audio`와 `video` 요소를 다루는 방법을 배웠습니다. 이번 실습에서는 `audio` 트랙과 `video` 대본을 표시하는 멀티미디어 플레이어를 만들어 보세요.

`audio` 요소에는 사용되는 미디어를 지정하는 데 쓰이는 `source` 요소를 포함해야 합니다.

예시는 다음과 같습니다:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

`source` 요소는 `video` 요소 안에서도 이렇게 사용할 수 있습니다:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 페이지의 주요 제목을 위한 `h1` 요소가 있어야 합니다.
2. `section` 요소가 세 개 있어야 합니다.
3. 첫 번째 `section` 요소 안에는 재생 중인 노래 제목을 위한 `h2` 요소가 있어야 합니다.
4. `h2` 요소 아래에는 `audio` 속성과 `controls` 속성을 가진 `aria-label` 요소가 있어야 합니다.
5. `audio` 요소 안에는 오디오 파일을 가리키는 `source` 속성과 `src` 속성을 가진 `type` 요소가 있어야 합니다. 원한다면 이 오디오 URL을 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. 두 번째 `section` 요소 안에는 재생 중인 비디오 제목을 위한 `h2` 요소가 있어야 합니다.
7. `h2` 요소 아래에는 `video`, `controls` 속성과 `width` 속성을 가진 `aria-label` 요소가 있어야 합니다.
8. `video` 요소 안에는 비디오 파일을 가리키는 `source` 속성과 `src` 속성을 가진 `type` 요소가 있어야 합니다. 원한다면 이 비디오 URL을 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. `source` 요소 아래에는 자막 파일을 가리키는 `track` 속성과 `src`, `kind`, `srclang` 속성을 가진 `label` 요소가 있어야 합니다. 원한다면 이 자막 URL을 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. 세 번째 `section` 요소 안에는 "Transcript" 같은 섹션 제목을 위한 `h2` 요소가 있어야 합니다.
11. `h2` 요소 아래에는 비디오 대본이 들어 있는 `p` 요소가 있어야 합니다.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector("h1"));
```

`h1`에는 페이지의 주요 제목을 나타내는 텍스트가 있어야 합니다.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

`section` 요소가 세 개 있어야 합니다.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

첫 번째 `section`에는 `h2` 요소가 포함되어야 합니다.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

`h2`에는 재생 중인 노래 제목을 나타내는 텍스트가 있어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

첫 번째 `section`에는 `audio` 요소가 포함되어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

`audio` 요소에는 `controls` 속성이 있어야 합니다.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

`audio` 요소에는 오디오를 설명하는 텍스트가 포함된 `aria-label` 속성이 있어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

`audio` 요소에는 `source` 요소가 있어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

`source` 요소에는 오디오 파일을 가리키는 `src` 속성이 있어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

`source` 요소에는 오디오 유형을 지정하는 값을 가진 `type` 속성이 있어야 합니다.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

두 번째 `section`에는 `h2` 요소가 포함되어야 합니다.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

`h2`에는 재생 중인 비디오 제목을 나타내는 텍스트가 있어야 합니다.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

두 번째 `section`에는 `video` 요소가 포함되어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

`video` 요소에는 `controls` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

`video` 요소에는 비디오 너비를 지정하는 값을 가진 `width` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

`video` 요소에는 `source` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

`source` 요소에는 비디오 파일을 가리키는 `src` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

`source` 요소에는 비디오 유형을 지정하는 값을 가진 `type` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

`video` 요소에는 `track` 요소가 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

`track` 요소에는 자막 파일을 가리키는 `src` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

`track` 요소에는 파일 종류를 설명하는 텍스트가 포함된 `kind` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

`track` 요소에는 자막 언어를 설명하는 텍스트가 포함된 `srclang` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

`track` 요소에는 자막 언어를 설명하는 텍스트가 포함된 `label` 속성이 있어야 합니다.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

세 번째 `section`에는 `h2` 요소가 포함되어야 합니다.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

`h2`에는 섹션 제목을 나타내는 텍스트가 있어야 합니다.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

세 번째 `section`에는 대본을 위한 `p` 요소가 포함되어야 합니다.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
