---
id: 69bac067cd15c4685d07906d
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Agora que você tem funções para processar ticks individuais e lidar com overflow, é hora de criar uma função controladora para um único portão. Essa função vai gerenciar todas as ações de um portão durante um único tick da simulação.

Crie uma função vazia chamada `handleGateAtTick` com os parâmetros:

- `gates`: O array completo de objetos portão.
- `gate`: O portão atual sendo processado.
- `tickIndex`: O índice do tick atual da simulação.
- `throughputSummary`: Um objeto que rastreia o total de participantes processados por portão.

# --hints--

Você deve criar uma função chamada `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

Sua função `handleGateAtTick` deve ter os parâmetros `gates`, `gate`, `tickIndex` e `throughputSummary`.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
