---
id: 69babea99725b43a01adb6d1
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Primeiro, você precisa obter o número de participantes que chegam durante o tick atual. Acesse o valor no array `queue` da entrada `tickIndex` e atribua-o a uma variável chamada `currentTickQueue`.

# --hints--

Você deve atribuir o elemento em `gate.queue[tickIndex]` a uma variável chamada `currentTickQueue`.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const { currentTickQueue } = processGateFlow?.variables ?? {};
assert.isTrue(currentTickQueue?.value.matches("gate.queue[tickIndex]"));
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
