---
id: 69babfd65c4c4632a0c76f63
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Quando um portão não consegue processar todos os participantes durante um tick, alguns participantes permanecem na fila. Para lidar com esse overflow, você vai criar uma função que redirecione os participantes em excesso para outro portão.

Crie uma função vazia chamada `rerouteOverflow` com os seguintes parâmetros:

- `gates`: O array completo de objetos de portão.
- `currentGate`: O portão onde ocorreu o overflow.
- `tickIndex`: A posição atual do tick na fila.
- `overflowAmount`: O número de participantes que não puderam ser processados.

# --hints--

Você deve criar uma função chamada `rerouteOverflow`.

```js
assert.isFunction(rerouteOverflow);
```

Sua função `rerouteOverflow` deve ter os parâmetros `gates`, `currentGate`, `tickIndex` e `overflowAmount`.

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
