---
id: 69babdf3f896711c11516373
title: passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Durante cada tick da simulação:

- Um certo número de participantes chega ao portão (a partir do seu `queue`).
- O portão processa os participantes com base no seu `capacity`.
- Se mais participantes chegarem do que o portão pode atender, alguns permanecerão (overflow).

Agora você vai criar uma função que lida com essa lógica para um portão em um único tick.

Crie uma função vazia chamada `processGateFlow` que aceite dois parâmetros:

- `gate`: O objeto do portão que está sendo processado.
- `tickIndex`: A posição atual do tick da simulação no array `queue`.

# --hints--

Você deve criar uma função chamada `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Sua função `processGateFlow` deve ter os parâmetros `gate` e `tickIndex`.

```js
const regex = __helpers.functionRegex('processGateFlow', ['gate', 'tickIndex']);
assert.match(__helpers.removeJSComments(processGateFlow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
