---
id: 69bac013466bd9bad0b809b1
title: passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Primeiro, você vai buscar a posição de um portão específico dentro do array `gates`.

Use o método `indexOf()` em `gates` para localizar o elemento `currentGate` e atribua o resultado a uma variável chamada `currentIndex`.

Aqui está um lembrete de como isso pode ser feito:

```js
const index = array.indexOf(element);
```

# --hints--

Você deve usar `indexOf()` no array `gates` para acessar o índice de `currentGate` e atribuí-lo a uma variável chamada `currentIndex`. Isso deve ser feito em uma linha.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+currentIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'currentIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)currentIndex=gates\.indexOf\(currentGate\);?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use `indexOf()` on the `gates` array to access the index of `currentGate` and assign it to a variable named `currentIndex`."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
--fcc-editable-region--

--fcc-editable-region--
}
```
