---
id: 69baafd6d683dd9921b75db3
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai usar Javascript para simular o fluxo dos participantes em um festival de música.

Você vai trabalhar com dois conjuntos de dados que descrevem os portões do festival durante os turnos da manhã e da noite:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Cada objeto portão contém as seguintes propriedades:

- `id`: Uma string que identifica o portão.
- `capacity`: O número de participantes que o portão pode processar por tick.
- `queue`: Um array de números que representam quantos participantes chegam ao portão durante um tick específico.

Para começar, crie uma função vazia chamada `initializeThroughput` que aceita um parâmetro chamado `gates`.

# --hints--

Você deve criar uma função chamada `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Sua função `initializeThroughput` deve ter um parâmetro `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
