---
id: 69bac027df64d4433c708d44
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Ao redirecionar overflow, você deve enviar os participantes para o próximo portão no array `gates`. Para isso, é preciso buscar o índice do próximo portão. Normalmente, você poderia adicionar 1 ao índice atual. Porém, se o portão atual for o último do array, é necessário voltar para o primeiro portão.

Você pode usar o operador módulo (`%`) para lidar com isso:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

Isso funciona porque quando `currentIndex + 1` é igual a `array.length`, o resultado se torna 0.

Usando essa técnica, crie uma variável chamada `nextGateIndex` que armazene o índice do próximo portão no array `gates`.

# --hints--

Você deve usar a técnica do módulo para acessar e atribuir o próximo índice de `currentIndex` no array `gates` a uma variável chamada `nextGateIndex`. Siga o mesmo padrão de uma linha do exemplo.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
