---
id: a5229172f011153519423690
title: 奇数のフィボナッチ和計算機を作成する
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

このラボでは、与えられた正の整数以下のすべての奇数のフィボナッチ数の合計を計算する奇数のフィボナッチ和計算機を作成します。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー**

1. `sumFibs` 関数が数値を引数として受け取る必要があります。
1. `sumFibs` 関数は、与えられた数値以下のすべての奇数のフィボナッチ数の合計を返す必要があります。
1. フィボナッチシーケンスは `0` と `1` から始まり、その後の各数は直前の2つの数の合計です。
1. 合計に加算されるのは奇数のフィボナッチ数のみです。

# --hints--

`sumFibs` 関数を用意してください。

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` は数値を返す必要があります。

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` は `1785` を返す必要があります。

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` は `4613732` を返す必要があります。

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` は `5` を返す必要があります。

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` は `60696` を返す必要があります。

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` は `135721` を返す必要があります。

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
