---
id: 673b3d6b7ef7318eef926d5a
title: 무드 보드 만들기
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

무드 보드는 특정 주제에 대한 일반적인 아이디어, 목표 또는 느낌을 전달하는 이미지와 텍스트의 콜라주입니다.

이 실습에서는 재사용 가능한 React 컴포넌트를 사용해 무드 보드를 만듭니다. CSS는 이미 제공되어 있습니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 세 개의 프로프 `MoodBoardItem`, `color`, `image`를 받는 `description` 컴포넌트를 생성하고 내보내야 합니다.
2. `MoodBoardItem` 컴포넌트는 최상위 요소로 클래스 `div`를 가진 `mood-board-item`를 반환해야 합니다.
3. `.mood-board-item` 요소의 배경색을 인라인 스타일로 `color` 프로프 값으로 설정해야 합니다.
4. `img` 요소 안에 클래스 `mood-board-image`를 가진 `src` 요소를 렌더링하고, `image` 속성은 `.mood-board-item` 프로프 값으로 설정해야 합니다.
5. `h3` 요소 안에 클래스 `mood-board-text`를 가진 `description` 요소를 렌더링하고, 텍스트는 `.mood-board-item` 프로프 값이어야 합니다.
6. `MoodBoard` 컴포넌트를 생성하고 내보내야 합니다.
7. `MoodBoard` 컴포넌트는 최상위 요소로 `div`를 반환해야 합니다.
8. `MoodBoard` 컴포넌트는 클래스 `h1`를 가진 `mood-board-heading` 요소를 렌더링하고 텍스트는 `Destination Mood Board`이어야 합니다.
9. `MoodBoard` 컴포넌트는 클래스 `div`를 가진 `mood-board`를 렌더링해야 합니다.
10. `MoodBoard` 컴포넌트는 `MoodBoardItem` 요소 안에 최소 세 개의 `.mood-board` 컴포넌트를 렌더링해야 하며, 각각은 유효한 값의 `color`, `image`, `description` 프로프스를 전달해야 합니다.

원한다면 무드 보드에 다음 이미지를 사용할 수 있습니다:

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

`MoodBoardItem` 컴포넌트를 내보내야 합니다.

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

`MoodBoardItem` 컴포넌트는 최상위 요소로 클래스 `div`를 가진 `mood-board-item`를 반환해야 합니다.

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

`.mood-board-item` 요소의 배경색은 인라인 스타일로 `color` 프로프 값으로 설정해야 합니다.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

`MoodBoardItem` 컴포넌트는 클래스 `img`를 가진 `mood-board-image` 요소를 렌더링하고, `src` 속성은 `image` 프로프 값으로 설정해야 합니다.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

`MoodBoardItem` 컴포넌트는 클래스 `h3`를 가진 `mood-board-text` 요소를 렌더링하고, 텍스트는 `description` 프로프 값이어야 합니다.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

`MoodBoard` 컴포넌트를 내보내야 합니다.

```js
assert.isFunction(window.index.MoodBoard, 1);
```

`MoodBoard` 컴포넌트는 최상위 요소로 `div`를 반환해야 합니다.

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

`MoodBoard` 컴포넌트는 클래스 `h1`를 가진 `mood-board-heading` 요소를 렌더링하고, 텍스트는 `Destination Mood Board`이어야 합니다.

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

`MoodBoard` 컴포넌트는 최소 세 개의 `MoodBoardItem` 컴포넌트를 렌더링해야 하며, 각각은 유효한 값의 `color`, `image`, `description` 프로프스를 전달해야 합니다.

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");
  
  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

`MoodBoard` 컴포넌트는 페이지의 `#root` 요소에 렌더링되어야 합니다.

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
