---
id: 65ef1cb031de30c2cde64c12
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

如果差值在指定的 `tolerance` 范围内，则将 `root` 的值设置为 `mid`，并跳出循环。

# --hints--

你应该删除现有的 `pass` 语句。

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_pass()`))
    }

})

```

你应该将 `mid` 的值分配给 `root` 并跳出循环。

```js
({
    test: () => 
    {
          assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[0].is_equivalent("root = mid\\nbreak")`))
    }

})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

--fcc-editable-region--       
            if abs(square_mid - square_target) < tolerance:
                pass

--fcc-editable-region--
```
