---
id: 65ef1d5e3d2927c5e0f4997b
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 Python 中，`is` 关键字检查对象身份。 它用于确定两个变量是否指向内存中的同一个对象。 与 `is` 相反，相等运算符 (`==`) 确定两个对象的值是否相同，而不管它们在内存中是否是同一个对象。

在 for 循环外部，创建一个 `if` 语句来检查是否为 `root is None`。 如果是，则打印消息 `'Failed to converge within {max_iterations} iterations.'`。 记得使用 f 字符串来格式化消息。

# --hints--

你应该创建一个 `if` 语句来检查是否 `root is None`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

你应该将参数 `f'Failed to converge within {max_iterations} iterations.'` 传递给 `print` 函数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
