---
id: 65ef1a50049cf9bada13266f
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

创建一个 `elif` 语句来检查 `square_target` 是否等于 `0`。 如果是，则将值 `0` 分配给 `root` 变量。 另外，打印消息 `'The square root of {square_target} is 0'`。 记得使用 f 字符串来格式化消息。

# --hints--

你应该有一个 `elif` 语句来检查条件 `square_target == 0`。


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

你应该将值 `0` 分配给 `root` 变量，并将参数 `f'The square root of {square_target} is 0'` 传递给 `print` 调用。


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
