---
id: 65ef190c6b51e9b5a5f7ed29
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

为 `square_root_bisection` 函数提供以下参数：

- `square_target`：你要求平方根的数字。
- `tolerance`（可选的）：近似根值的平方与实际目标值之间可接受的差异（默认值为 `1e-7`）。容差 `1e-7` 意味着解的精度将在真实值的 `0.0000001` 以内，是在准确性和性能之间取得平衡的良好默认选择。
- `max_iterations` (可选): 要执行的最大迭代次数（默认为 `100`）。 如果该方法没有收敛到这个限制内，你将认为没有找到解决方案。
  
# --hints--

你的函数应该具有以下参数：`square_target`、`tolerance = 1e-7` 和 `max_iterations = 100`。 记住参数传递的顺序是很重要的。

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

