---
id: 65ef1d104c2be9c4d8e22075
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

如果 `if` 和 `elif` 条件均不满足，则 `mid` 的值将大于 `square_target`。 在本例中，创建一个 `else` 子句并将 `mid` 的值分配给 `high`。

# --hints--

你应该有一个 `else` 子句。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[2]== _Node()`));
    }
})

```

你应该在 `else` 子句中将 `mid` 的值分配给 `high`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[2].is_equivalent("high = mid")`));
    }
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

--fcc-editable-region--
            elif square_mid < square_target:
                low = mid

--fcc-editable-region--
```
