---
id: 65ef181dad3d96b14a21a03e
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个项目中，你将使用二分法找到给定数字的近似平方根。

二分法是一种寻找实值函数根的技术。 它的工作原理是缩小平方根所在的区间，直到它收敛到指定公差范围内的值。

首先创建一个名为 `square_root_bisection` 的函数。 暂时将参数留空。

# --hints--

你应该定义不带参数的函数 `square_root_bisection`。

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
