---
id: 65ef1be2b958cabf712296cd
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在，你将通过找到当前间隔的中点并将中点的平方与目标值进行比较来反复缩小间隔。

为此，在 `else` 块内，创建一个 `for` 循环，最多运行 `max_iterations` 次。

对于循环，使用 `range` 函数，它会生成一个你可以迭代的数字序列。 语法为 `range(start, stop, step)`，其中 `start` 是起始整数（含），`stop` 是最后一个整数（不含），`step` 是序列中某个数字与前一个数字之间的差值。

另外，使用 `_` 作为循环变量。 `_` 充当占位符，当你需要使用变量但实际上不需要其值时很有用。

# --hints--

你应该创建一个 `for` 循环来迭代 `range(max_iterations)`。 使用 `_` 作为循环变量。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
