---
id: 65ef1c34a541afc12dbb849d
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

`abs()` 函数返回一个数字的绝对值，该值始终为正，无论数字符号如何。 你将使用它来检查估计的平方根是否足够接近实际值。

现在，创建一个 `if` 语句来检查 `square_mid` 和 `square_target` 之间的差的绝对值是否在指定的 `tolerance` 范围内。

# --hints--

你应该在 `for` 循环的体内有一个条件为 `abs(square_mid - square_target) < tolerance:` 的 `if` 语句。

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
