---
id: 65ef198fde24dfb7ff675b42
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

`raise` 语句允许你强制发生特定的异常。 它由 `raise` 关键字和后跟的异常类型组成，并使你能够提供自定义错误消息：

```py
raise ValueError("Invalid value")
```

当上述代码运行时，将引发 `ValueError` 并向用户显示消息 `"Invalid value"`。

如果 `square_target` 小于 `0`，则无法计算实值平方根。 因此，引发 `ValueError` 并显示消息 `'Square root of negative number is not defined in real numbers'`。 不要忘记删除 `pass` 关键字。

# --hints--

你应该删除 `pass` 关键字。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

你应该在 `if` 主体内引发一个带有消息 `'Square root of negative number is not defined in real numbers'` 的 `ValueError`。

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
