---
id: 65ef19425d1b27b6c930bae6
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

如果你要查找平方根的数字是负数，则代码应该引发错误，因为负数的平方根在实数中没有定义。

删除 `pass` 语句并创建 `if` 语句来检查 `square_target` 是否小于 `0`。

# --hints--

你应该删除 `pass` 关键字。

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

你应该创建一个 `if` 语句来查看 `square_target < 0`。

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
