---
id: 65ef1a720d6e65bbb0d22463
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--
接下来，你将处理除 `1` 或 `0` 之外的 `square_target` 为正数的情况。

创建一个 `else` 子句来处理这些情况。 

# --hints--

创建一个 `else` 子句，并且不要忘记添加 `pass` 关键字。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[2] == _Node()`))
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')
--fcc-editable-region--


--fcc-editable-region--
```
