---
id: 68a23004e96a105df3a21fc4
title: بناء مساعد عد البطاقات
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

في لعبة الكازينو بلاك جاك، يمكن للاعب تحديد ما إذا كان لديه ميزة في اليد التالية على الكازينو من خلال تتبع العدد النسبي للبطاقات العالية والمنخفضة المتبقية في الرزمة. يُسمى هذا عد البطاقات.

وجود المزيد من البطاقات العالية المتبقية في الرزمة يُفضل اللاعب. عندما يكون العد إيجابيًا، يجب على اللاعب المراهنة بمبلغ كبير. عندما يكون العد صفرًا أو سالبًا، يجب المراهنة بمبلغ صغير.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب عليك استخدام `let` لإعلان متغير عام باسم `count` وتعيينه إلى `0`.
1. يجب أن يكون لديك دالة تسمى `cardCounter`.
1. يجب أن تستقبل دالة `cardCounter` معلمة `card` التي يمكن أن تكون رقمًا أو سلسلة نصية.
    - للقيم بين `2` و `10`، ستكون معلمة `card` رقمًا.
    - لجميع القيم الأخرى، ستكون معلمة `card` سلسلة نصية.
1. يجب أن تعدل دالة `cardCounter` المتغير العام `count` بناءً على معايير معينة.
1. يجب زيادة المتغير العام `count` بمقدار `1` للبطاقات `2`، `3`، `4`، `5`، أو `6`
1. يجب أن يبقى المتغير العام `count` دون تغيير للبطاقات `7`، `8`، `9`.
1. يجب أن ينخفض المتغير العام `count` بمقدار `1` للبطاقات `10`، `"J"`، `"Q"`، `"K"`، `"A"`
1. يجب أن تُرجع دالة `cardCounter` سلسلة نصية تحتوي على العد الحالي وكلمة `Bet` إذا كان العد إيجابيًا.
1. يجب أن تُرجع دالة `cardCounter` سلسلة نصية تحتوي على العد الحالي وكلمة `Hold` إذا كان العد أقل من أو يساوي `0`.
1. في ناتج الدالة، يجب أن يفصل بين العد الحالي وقرار اللاعب (`Bet` أو `Hold`) مسافة واحدة. على سبيل المثال، `-3 Hold`.

# --hints--

يجب عليك استخدام `let` لإعلان متغير عام باسم `count` وتعيينه إلى `0`.

```js
const explorer = await __helpers.Explorer(code);
const { count } = explorer.variables;
assert.isTrue(count?.matches('let count = 0'));
```

يجب أن يكون لديك دالة باسم `cardCounter`.

```js
assert.isFunction(cardCounter);
```

يجب أن تُرجع دالتك قيمة `count` والنص (`Bet` أو `Hold`) مع وجود مسافة واحدة بينهما.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

بعد البطاقات `2`، `3`، `4`، `5`، يجب أن تُرجع استدعاء `cardCounter(6)` السلسلة النصية `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

بعد البطاقات `7`، `8`، يجب أن تُرجع استدعاء `cardCounter(9)` السلسلة النصية `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

بعد البطاقات `10`، `"J"`، `"Q"`، `"K"`، يجب أن تُرجع استدعاء `cardCounter("A")` السلسلة النصية `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

بعد البطاقات `3`، `7`، `"Q"`، `8`، يجب أن تُرجع استدعاء `cardCounter("A")` السلسلة النصية `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

بعد البطاقات `2`، `"J"`، `9`، `2`، يجب أن تُرجع استدعاء `cardCounter(7)` السلسلة النصية `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

بعد البطاقات `2`، `2`، يجب أن تُرجع استدعاء `cardCounter(10)` السلسلة النصية `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

بعد البطاقات `3`، `2`، `"A"`، `10`، يجب أن تُرجع استدعاء `cardCounter("K")` السلسلة النصية `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
