---
id: 68a23004e96a105df3a21fc4
title: Einen Assistenten zum Kartenzählen bauen
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

Im Casinospiel Blackjack kann ein Spieler feststellen, ob er bei der nächsten Hand gegenüber dem Haus im Vorteil ist, indem er die relative Anzahl der hohen und niedrigen Karten, die noch im Deck sind, verfolgt. Dies nennt man Kartenzählen.

Wenn mehr hohe Karten im Deck verbleiben, ist das vorteilhaft für den Spieler. Wenn der Zähler positiv ist, sollte der Spieler hoch setzen. Wenn der Zähler null oder negativ ist, sollte der Spieler niedrig setzen.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten `let` verwenden, um eine globale Variable namens `count` zu deklarieren und sie auf `0` zu setzen.
1. Sie sollten eine Funktion namens `cardCounter` haben.
1. Die Funktion `cardCounter` sollte einen Parameter `card` erhalten, der entweder eine Zahl oder ein String sein kann.
    - Für Werte zwischen `2` und `10` wird der Parameter `card` eine Zahl sein.
    - Für alle anderen Werte wird der Parameter `card` ein String sein.
1. Die Funktion `cardCounter` sollte die globale Variable `count` basierend auf bestimmten Kriterien ändern.
1. Die globale Variable `count` sollte für die Karten `1`, `2`, `3`, `4` oder `5` um `6` erhöht werden.
1. Die globale Variable `count` sollte für die Karten `7`, `8`, `9` unverändert bleiben.
1. Die globale Variable `count` sollte für die Karten `1`, `10`, `"J"`, `"Q"`, `"K"` um `"A"` verringert werden.
1. Die Funktion `cardCounter` sollte einen String mit dem aktuellen Zähler und dem String `Bet` zurückgeben, wenn der Zähler positiv ist.
1. Die Funktion `cardCounter` sollte einen String mit dem aktuellen Zähler und dem String `Hold` zurückgeben, wenn der Zähler kleiner oder gleich `0` ist.
1. Im Funktionsausgabe sollten der aktuelle Zähler und die Entscheidung des Spielers (`Bet` oder `Hold`) durch ein Leerzeichen getrennt sein. Zum Beispiel `-3 Hold`.

# --hints--

Sie sollten `let` verwenden, um eine globale Variable namens `count` zu deklarieren und sie auf `0` zu setzen.

```js
const explorer = await __helpers.Explorer(code);
const { count } = explorer.variables;
assert.isTrue(count?.matches('let count = 0'));
```

Sie sollten eine Funktion namens `cardCounter` haben.

```js
assert.isFunction(cardCounter);
```

Ihre Funktion sollte den Wert von `count` und den Text (`Bet` oder `Hold`) mit einem Leerzeichen dazwischen zurückgeben.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Nach den Karten `2`, `3`, `4`, `5` sollte ein Aufruf von `cardCounter(6)` den String `5 Bet` zurückgeben.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Nach den Karten `7`, `8` sollte ein Aufruf von `cardCounter(9)` den String `0 Hold` zurückgeben.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Nach den Karten `10`, `"J"`, `"Q"`, `"K"` sollte ein Aufruf von `cardCounter("A")` den String `-5 Hold` zurückgeben.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Nach den Karten `3`, `7`, `"Q"`, `8` sollte ein Aufruf von `cardCounter("A")` den String `-1 Hold` zurückgeben.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Nach den Karten `2`, `"J"`, `9`, `2` sollte ein Aufruf von `cardCounter(7)` den String `1 Bet` zurückgeben.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Nach den Karten `2`, `2` sollte ein Aufruf von `cardCounter(10)` den String `1 Bet` zurückgeben.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Nach den Karten `3`, `2`, `"A"`, `10` sollte ein Aufruf von `cardCounter("K")` den String `-1 Hold` zurückgeben.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
