---
id: 681cb1b0dab50c87ddb2e51b
title: "Desafío 10: 3 Strikes"
challengeType: 29
dashedName: challenge-10
---

# --description--

Dado un entero entre 1 y 10,000, devuelve un conteo de cuántos números desde 1 hasta ese entero tienen un cuadrado que contiene al menos un dígito 3.

# --hints--

`squares_with_three(1)` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1), 0)`)
}})
```

`squares_with_three(10)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10), 1)`)
}})
```

`squares_with_three(100)` debería devolver `19`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(100), 19)`)
}})
```

`squares_with_three(1000)` debería devolver `326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(1000), 326)`)
}})
```

`squares_with_three(10000)` debería devolver `4531`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(squares_with_three(10000), 4531)`)
}})
```

# --seed--

## --seed-contents--

```py
def squares_with_three(n):

    return n
```

# --solutions--

```py
def squares_with_three(n):
    count = 0
    for i in range(1, n + 1):
        square = i * i
        if '3' in str(square):
            count += 1
    return count
```
