---
id: 681cb1b0dab50c87ddb2e51a
title: "Desafío 9: Suma de Cuadrados"
challengeType: 29
dashedName: challenge-9
---

# --description--

Dado un entero positivo hasta 1,000, devuelve la suma de todos los enteros al cuadrado desde 1 hasta el número.

# --hints--

`sum_of_squares(5)` debería devolver `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(5), 55)`)
}})
```

`sum_of_squares(10)` debería devolver `385`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(10), 385)`)
}})
```

`sum_of_squares(25)` debería devolver `5525`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(25), 5525)`)
}})
```

`sum_of_squares(500)` debería devolver `41791750`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(500), 41791750)`)
}})
```

`sum_of_squares(1000)` debería devolver `333833500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(1000), 333833500)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_squares(n):

    return n
```

# --solutions--

```py
def sum_of_squares(n):
    sum = ((n)*(n+1)*(2*n+1))//6
    return sum
```
