---
id: 681cb1afdab50c87ddb2e515
title: "Desafío 4: S  P  A  C  E  J  A  M"
challengeType: 29
dashedName: challenge-4
---

# --description--

Dada una cadena, elimina todos los espacios de la cadena, inserta dos espacios entre cada carácter, convierte todas las letras alfabéticas a mayúsculas y devuelve el resultado.

- Los caracteres no alfabéticos deben permanecer sin cambios (excepto los espacios).

# --hints--

`space_jam("freeCodeCamp")` debería devolver `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("   free   Code   Camp   ")` debería devolver `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P")`)
}})
```

`space_jam("Hello World?!")` debería devolver `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !")`)
}})
```

`space_jam("C@t$ & D0g$")` debería devolver `"C  @  T  $  &  D  0  G  $"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $")`)
}})
```

`space_jam("allyourbase")` debería devolver `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(space_jam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E")`)
}})
```

# --seed--

## --seed-contents--

```py
def space_jam(s):

    return s
```

# --solutions--

```py
def space_jam(s):
    s = s.replace(" ", "")
    s = "  ".join(s)
    return s.upper()
```
