---
id: 681cb1a2dab50c87ddb2e514
title: "Desafío 3: Secuencia de Fibonacci"
challengeType: 29
dashedName: challenge-3
---

# --description--

La secuencia de Fibonacci es una serie de números donde cada número es la suma de los dos anteriores. Al comenzar con `0` y `1`, los primeros 10 números en la secuencia son `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Dado un arreglo que contiene los dos primeros números de una secuencia de Fibonacci, y un entero que representa la longitud de la secuencia, devuelve un arreglo que contiene la secuencia de la longitud dada.

- Tu función debe manejar secuencias de cualquier longitud mayor o igual a cero.
- Si la longitud es cero, devuelve un arreglo vacío.
- Ten en cuenta que los números iniciales son parte de la secuencia.

# --hints--


`fibonacci_sequence([0, 1], 20)` debería devolver `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181])`)
}})
```

`fibonacci_sequence([21, 32], 1)` debería devolver `[21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([21, 32], 1), [21])`)
}})
```

`fibonacci_sequence([0, 1], 0)` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([0, 1], 0), [])`)
}})
```

`fibonacci_sequence([10, 20], 2)` debería devolver `[10, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([10, 20], 2), [10, 20])`)
}})
```

`fibonacci_sequence([123456789, 987654321], 5)` debería devolver `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fibonacci_sequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541])`)
}})
```

# --seed--

## --seed-contents--

```py
def fibonacci_sequence(start_sequence, length):

    return length
```

# --solutions--

```py
def fibonacci_sequence(start_sequence, length):
    if length == 0:
        return []
    if length == 1:
        return [start_sequence[0]]
    if length == 2:
        return start_sequence[:]

    sequence = start_sequence[:]
    while len(sequence) < length:
        next_value = sequence[-1] + sequence[-2]
        sequence.append(next_value)

    return sequence
```
