---
id: 68216ef80f957572e7c340c5
title: "Desafío 12: Decodificador de Mensajes"
challengeType: 29
dashedName: challenge-12
---

# --description--

Dada una cadena de mensaje secreto y un entero que representa el número de letras que se usaron para desplazar el mensaje y codificarlo, devuelve la cadena decodificada.

- Un número positivo significa que el mensaje fue desplazado hacia adelante en el alfabeto.
- Un número negativo significa que el mensaje fue desplazado hacia atrás en el alfabeto.
- Las mayúsculas y minúsculas importan, los caracteres decodificados deben conservar la estructura de mayúsculas y minúsculas de sus contrapartes codificadas.
- Los caracteres no alfabéticos no deben ser decodificados.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` debería devolver `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` debería devolver `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` debería devolver `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` debería devolver `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
