---
id: 6821ebce237de8297eaee790
title: "Desafío 14: Batalla de Personajes"
challengeType: 29
dashedName: challenge-14
---

# --description--

Dadas dos cadenas que representan tu ejército y un ejército opuesto, cada carácter de tu ejército combate contra el carácter en la misma posición del ejército opuesto usando las siguientes reglas:

- Los caracteres `a-z` tienen una fuerza de `1-26`, respectivamente.
- Los caracteres `A-Z` tienen una fuerza de `27-52`, respectivamente.
- Los dígitos `0-9` tienen una fuerza igual a su valor facial.
- Todos los demás caracteres tienen un valor de cero.
- Cada carácter solo puede luchar en una batalla.

Para cada batalla, el personaje más fuerte gana. El ejército con más victorias gana la guerra. Devuelve los siguientes valores:

- `"Opponent retreated"` si tu ejército tiene más personajes que el ejército contrario.
- "Nos retiramos" si el ejército opuesto tiene más personajes que tú.
- "Ganamos" si tu ejército ganó más batallas.
- "Perdimos" si el ejército contrario ganó más batallas.
- "Fue un empate" si ambos ejércitos ganaron el mismo número de batallas.

# --hints--

`battle("Hello", "World")` debería devolver `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` debería devolver `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` debería devolver `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` debería devolver `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` debería devolver `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` debería devolver `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` debería devolver `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
