---
id: 6821ebc9237de8297eaee78f
title: "Desafío 13: Primo No Natural"
challengeType: 29
dashedName: challenge-13
---

# --description--

Dado un entero, determina si ese número es un número primo o un número primo negativo.

- Un número primo es un entero positivo mayor que 1 que solo es divisible por 1 y por sí mismo.
- Un número primo negativo es la versión negativa de un número primo positivo.
- `1` y `0` no se consideran números primos.

# --hints--

`is_unnatural_prime(1)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(1), False)`)
}})
```

`is_unnatural_prime(-1)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-1), False)`)
}})
```

`is_unnatural_prime(19)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(19), True)`)
}})
```

`is_unnatural_prime(-23)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-23), True)`)
}})
```

`is_unnatural_prime(0)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(0), False)`)
}})
```

`is_unnatural_prime(97)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(97), True)`)
}})
```

`is_unnatural_prime(-61)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-61), True)`)
}})
```

`is_unnatural_prime(99)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(99), False)`)
}})
```

`is_unnatural_prime(-44)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-44), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_unnatural_prime(n):

    return n
```

# --solutions--

```py
def is_unnatural_prime(n):
    abs_n = abs(n)

    if abs_n <= 1:
        return False

    for i in range(2, int(abs_n ** 0.5) + 1):
        if abs_n % i == 0:
            return False

    return True
```
