---
id: 6821ebea237de8297eaee795
title: "Desafío 19: Candlelight"
challengeType: 29
dashedName: challenge-19
---

# --description--

Dado un entero que representa la cantidad de velas con las que comienzas, y un entero que representa cuántas velas quemadas se necesitan para crear una nueva, devuelve la cantidad de velas que habrás usado después de crear y quemar tantas como puedas.

Por ejemplo, si te dan 7 velas y se necesitan 2 velas quemadas para hacer una nueva:

1. Quema 7 velas para obtener 7 sobras,
2. Recicla 6 restos en 3 nuevas velas (queda 1 resto),
3. Quema 3 velas para obtener 3 sobras más (4 en total),
4. Recicla 4 restos en 2 nuevas velas,
5. Quema 2 velas para obtener 2 sobras,
6. Recicla 2 restos en 1 vela nueva,
7. Quema 1 vela.

Habrás quemado 13 velas en total en el ejemplo.

# --hints--

`burn_candles(7, 2)` debería devolver `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` debería devolver `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` debería devolver `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` debería devolver `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` debería devolver `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
