---
id: 6821ebdf237de8297eaee793
title: "Desafío 17: Orden de las Operaciones"
challengeType: 29
dashedName: challenge-17
---

# --description--

Dado un arreglo de enteros y un arreglo de operadores de cadena, aplica las operaciones a los números secuencialmente de izquierda a derecha. Repite las operaciones según sea necesario hasta que se usen todos los números. Devuelve el resultado final.

Por ejemplo, dado `[1, 2, 3, 4, 5]` y `['+', '*']`, devuelve el resultado de evaluar `1 + 2 * 3 + 4 * 5` de izquierda a derecha ignorando el orden estándar de las operaciones.

- Los operadores válidos son `+`, `-`, `*`, `/` y `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` debería devolver `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` debería devolver `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` debería devolver `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` debería devolver `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` debería devolver `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
