---
id: 6821ebda237de8297eaee792
title: "Desafío 16: Paréntesis Invertidos"
challengeType: 29
dashedName: challenge-16
---

# --description--

Dada una cadena que contiene paréntesis correctamente anidados, devuelve la versión decodificada de la cadena usando las siguientes reglas:

- Todos los caracteres dentro de cada paréntesis deben invertirse.
- Los paréntesis deben eliminarse del resultado final.
- Si los paréntesis están anidados, el par más interno debe invertirse primero, y luego su resultado debe incluirse en la inversión del par externo.
- Asume que todos los paréntesis están equilibrados y correctamente anidados.

# --hints--

`decode("(f(b(dc)e)a)")` debería devolver `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` debería devolver `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` debería devolver `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
