---
id: 68216eb60f957572e7c340c4
title: "Desafío 11: Ritmo por Milla"
challengeType: 29
dashedName: challenge-11
---

# --description--

Dado un número de millas corridas y un tiempo en `"MM:SS"` (minutos:segundos) que tomó correr esas millas, devuelve una cadena con el tiempo promedio que tomó correr cada milla en el formato `"MM:SS"`.

- Agrega ceros a la izquierda cuando sea necesario.

# --hints--

`mile_pace(3, "24:00")` debería devolver `"08:00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(3, "24:00"), "08:00")`)
}})
```

`mile_pace(1, "06:45")` debería devolver `"06:45"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(1, "06:45"), "06:45")`)
}})
```

`mile_pace(2, "07:00")` debería devolver `"03:30"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(2, "07:00"), "03:30")`)
}})
```

`mile_pace(26.2, "120:35")` debería devolver `"04:36"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mile_pace(26.2, "120:35"), "04:36")`)
}})
```

# --seed--

## --seed-contents--

```py
def mile_pace(miles, duration):

    return miles
```

# --solutions--

```py
def mile_pace(miles, duration):
    minutes, seconds = map(int, duration.split(":"))
    total_seconds = minutes * 60 + seconds
    avg_seconds_per_mile = total_seconds / miles

    avg_minutes = int(avg_seconds_per_mile // 60)
    avg_seconds = round(avg_seconds_per_mile % 60)

    return f"{avg_minutes:02d}:{avg_seconds:02d}"
```
