---
id: 6821ebee237de8297eaee796
title: "Desafío 20: Duplicados en Arrays"
challengeType: 29
dashedName: challenge-20
---

# --description--

Dado un arreglo de enteros, devuelve un arreglo de enteros que aparecen más de una vez en el arreglo inicial, ordenados en orden ascendente. Si ningún valor aparece más de una vez, devuelve un arreglo vacío.

- Solo incluye una instancia de cada valor en el array devuelto.

# --hints--

`find_duplicates([1, 2, 3, 4, 5])` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 5]), [])`)
}})
```

`find_duplicates([1, 2, 3, 4, 1, 2])` debería devolver `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 1, 2]), [1, 2])`)
}})
```

`find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` debería devolver `[-6, 0, 2, 4, 5, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_duplicates(arr):

    return arr
```

# --solutions--

```py
def find_duplicates(arr):
    seen = set()
    duplicates = set()

    for num in arr:
        if num in seen:
            duplicates.add(num)
        else:
            seen.add(num)

    return sorted(duplicates)
```
