---
id: 68adce01c0e1144d0a90295a
title: "Desafío 27: Rotar Matriz"
challengeType: 29
dashedName: challenge-27
---

# --description--

Dada una matriz (un arreglo de arreglos), rota la matriz 90 grados en el sentido de las agujas del reloj y devuélvela. Por ejemplo, dada `[[1, 2], [3, 4]]`, que se ve así:

12
|---|---|
34

Debes devolver `[[3, 1], [4, 2]]`, que se ve así:

31
|---|---|
42

# --hints--

`rotate([[1]])` debería devolver `[[1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1]]), [[1]])`)
}})
```

`rotate([[1, 2], [3, 4]])` debería devolver `[[3, 1], [4, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]])`)
}})
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` debería devolver `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]])`)
}})
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` debería devolver `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def rotate(matrix):

    return matrix
```

# --solutions--

```py
def rotate(matrix):
    n = len(matrix)
    result = [[0] * n for _ in range(n)]

    for i in range(n):
        for j in range(n):
            result[j][n - 1 - i] = matrix[i][j]

    return result
```
