---
id: 68adce01c0e1144d0a902956
title: "Desafío 25: Repetidor de Vocales"
challengeType: 29
dashedName: challenge-25
---

# --description--

Dada una cadena, devuelve una nueva versión de la cadena donde cada vocal se duplica una vez más que la vocal anterior que encontraste. Por ejemplo, la primera vocal en la oración debe permanecer sin cambios. La segunda vocal debe aparecer dos veces seguidas. La tercera vocal debe aparecer tres veces seguidas, y así sucesivamente.

- Las letras `a`, `e`, `i`, `o` y `u`, ya sea en mayúsculas o minúsculas, se consideran vocales.
- La vocal original debe mantener su mayúscula o minúscula.
- Las vocales repetidas deben estar en minúsculas.
- Todos los caracteres que no sean vocales deben mantener su caso original.

# --hints--

`repeat_vowels("hello world")` debería devolver `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` debería devolver `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` debería devolver `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` debería devolver `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
