---
id: 68af0687ef34c76c28ffa54d
title: "Desafío 33: Tiempo frente a la pantalla"
challengeType: 29
dashedName: challenge-33
---

# --description--

Dado un arreglo de entrada de siete enteros, que representan el tiempo de una semana, donde cada entero es la cantidad de horas que pasaste en tu teléfono ese día, determina si es demasiado tiempo frente a la pantalla basado en estas restricciones:

- Si cualquier día tiene 10 horas o más, es demasiado.
- Si el promedio de cualquier tres días consecutivos es mayor o igual a 8 horas, es demasiado.
- Si el promedio de los siete días es mayor o igual a 6 horas, es demasiado.

# --hints--

`too_much_screen_time([1, 2, 3, 4, 5, 6, 7])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 4, 5, 6, 7]), False)`)
}})
```

`too_much_screen_time([7, 8, 8, 4, 2, 2, 3])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([7, 8, 8, 4, 2, 2, 3]), False)`)
}})
```

`too_much_screen_time([5, 6, 6, 6, 6, 6, 6])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([5, 6, 6, 6, 6, 6, 6]), False)`)
}})
```

`too_much_screen_time([1, 2, 3, 11, 1, 3, 4])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 11, 1, 3, 4]), True)`)
}})
```

`too_much_screen_time([1, 2, 3, 10, 2, 1, 0])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 10, 2, 1, 0]), True)`)
}})
```

`too_much_screen_time([3, 3, 5, 8, 8, 9, 4])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 3, 5, 8, 8, 9, 4]), True)`)
}})
```

`too_much_screen_time([3, 9, 4, 8, 5, 7, 6])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 9, 4, 8, 5, 7, 6]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def too_much_screen_time(hours):

    return hours
```

# --solutions--

```py
def too_much_screen_time(hours):
    for h in hours:
        if h >= 10:
            return True

    for i in range(5):
        avg3 = sum(hours[i:i+3]) / 3
        if avg3 >= 8:
            return True

    weekly_avg = sum(hours) / 7
    if weekly_avg >= 6:
        return True

    return False
```
