---
id: 68af0687ef34c76c28ffa54f
title: "Desafío 34: Números faltantes"
challengeType: 29
dashedName: challenge-34
---

# --description--

Dado un arreglo de enteros del 1 a `n`, inclusive, devuelve un arreglo con todos los enteros faltantes entre 1 y `n` (donde `n` es el número más grande en el arreglo dado).

- El array dado puede no estar ordenado y puede contener duplicados.
- El array devuelto debe estar en orden ascendente.
- Si no falta ningún entero, devuelve un arreglo vacío.

# --hints--

`find_missing_numbers([1, 3, 5])` debería devolver `[2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 3, 5]), [2, 4])`)
}})
```

`find_missing_numbers([1, 2, 3, 4, 5])` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5]), [])`)
}})
```

`find_missing_numbers([1, 10])` debería devolver `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

`find_missing_numbers([10, 1, 10, 1, 10, 1])` debería devolver `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9])`)
}})
```

`find_missing_numbers([3, 1, 4, 1, 5, 9])` debería devolver `[2, 6, 7, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([3, 1, 4, 1, 5, 9]), [2, 6, 7, 8])`)
}})
```

`find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` debería devolver `[11]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_missing_numbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_missing_numbers(arr):

    return arr
```

# --solutions--

```py
def find_missing_numbers(arr):
    if not arr:
        return []
    
    n = max(arr)
    seen = set(arr)
    result = [i for i in range(1, n + 1) if i not in seen]
    return result
```
