---
id: 68af0687ef34c76c28ffa54b
title: "Desafío 32: Oración Invertida"
challengeType: 29
dashedName: challenge-32
---

# --description--

Dada una cadena de palabras, devuelve una nueva cadena con las palabras en orden inverso. Por ejemplo, la primera palabra debe estar al final de la cadena devuelta, y la última palabra debe estar al principio de la cadena devuelta.

- En la cadena dada, las palabras pueden estar separadas por uno o más espacios.
- La cadena devuelta debe tener solo un espacio entre palabras.

# --hints--

`reverse_sentence("world hello")` debería devolver `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` debería devolver `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` debería devolver `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` debería devolver `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
