---
id: 68b06e589bf2273243814771
title: "Desafío 36: Ajustador de Termostato"
challengeType: 29
dashedName: challenge-36
---

# --description--

Dada la temperatura actual de una habitación y una temperatura objetivo, devuelve una cadena que indique cómo ajustar la temperatura de la habitación según estas restricciones:

- Devuelve `"heat"` si la temperatura actual está por debajo del objetivo.
- Devuelve `"cool"` si la temperatura actual está por encima del objetivo.
- Devuelve `"hold"` si la temperatura actual es igual a la objetivo.

# --hints--

`adjust_thermostat(68, 72)` debería devolver `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` debería devolver `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` debería devolver `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` debería devolver `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` debería devolver `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` debería devolver `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
