---
id: 68b06e589bf227324381476f
title: "Desafío 35: Frecuencia de Palabras"
challengeType: 29
dashedName: challenge-35
---

# --description--

Dado un párrafo, devuelve un arreglo con las tres palabras que ocurren con mayor frecuencia.

- Las palabras en el párrafo estarán separadas por espacios.
- Ignora mayúsculas y minúsculas en el párrafo dado. Por ejemplo, trata `Hello` y `hello` como la misma palabra.
- Ignora la puntuación en el párrafo dado. La puntuación consiste en comas (`,`), puntos (`.`) y signos de exclamación (`!`).
- El array devuelto debe tener todas las palabras en minúsculas.
- El array devuelto debe estar en orden descendente con la palabra que ocurre con más frecuencia primero.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` debería devolver `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` debería devolver `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` debería devolver `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
