---
id: 68b06e589bf2273243814777
title: "Desafío 39: Llena el tanque"
challengeType: 29
dashedName: challenge-39
---

# --description--

Dado el tamaño de un tanque de combustible, el nivel actual de combustible y el precio por galón, devuelve el costo para llenar el tanque completamente.

- `tankSize` es la capacidad total del tanque en galones.
- `fuelLevel` es la cantidad actual de combustible en el tanque en galones.
- `pricePerGallon` es el costo de un galón de combustible.
- El valor devuelto debe redondearse a dos decimales en el formato: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` debería devolver `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` debería devolver `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` debería devolver `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` debería devolver `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` debería devolver `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
