---
id: 68b06e589bf2273243814775
title: "Desafío 38: Generador de Slugs"
challengeType: 29
dashedName: challenge-38
---

# --description--

Dada una cadena, devuelve una versión amigable para URL de la cadena usando las siguientes restricciones:

- Todas las letras deben estar en minúsculas.
- Todos los caracteres que no sean letras, números o espacios deben ser eliminados.
- Todos los espacios deben ser reemplazados con el código de espacio codificado para URL `%20`.
- Los espacios consecutivos deben ser reemplazados por un solo `%20`.
- La cadena devuelta no debe tener `%20` al principio ni al final.

# --hints--

`generate_slug("helloWorld")` debería devolver `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` debería devolver `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` debería devolver `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` debería devolver `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` debería devolver `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
