---
id: 68b1f72371a5ac895ac70a02
title: "Desafío 40: Almacenamiento de Fotos"
challengeType: 29
dashedName: challenge-40
---

# --description--

Dado un tamaño de foto en megabytes (MB) y la capacidad del disco duro en gigabytes (GB), devuelve el número de fotos que el disco duro puede almacenar usando las siguientes restricciones:

- 1 gigabyte equivale a 1000 megabytes.
- Devuelve el número de fotos completas que puede almacenar el disco.

# --hints--

`number_of_photos(1, 1)` debería devolver `1000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(1, 1), 1000)`)
}})
```

`number_of_photos(2, 1)` debería devolver `500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(2, 1), 500)`)
}})
```

`number_of_photos(4, 256)` debería devolver `64000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(4, 256), 64000)`)
}})
```

`number_of_photos(3.5, 750)` debería devolver `214285`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 750), 214285)`)
}})
```

`number_of_photos(3.5, 5.5)` debería devolver `1571`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_photos(3.5, 5.5), 1571)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_photos(photo_size_mb, drive_size_gb):

    return photo_size_mb
```

# --solutions--

```py
def number_of_photos(photo_size_mb, drive_size_gb):
    drive_size_mb = drive_size_gb * 1000
    return drive_size_mb // photo_size_mb
```
