---
id: 68b1f72371a5ac895ac70a04
title: "Desafío 41: Almacenamiento de Archivos"
challengeType: 29
dashedName: challenge-41
---

# --description--

Dado un tamaño de archivo, una unidad para el tamaño del archivo y la capacidad del disco duro en gigabytes (GB), devuelve el número de archivos que el disco duro puede almacenar usando las siguientes restricciones:

- La unidad para el tamaño del archivo puede ser bytes (`"B"`), kilobytes (`"KB"`), o megabytes (`"MB"`).
- Devuelve el número de archivos completos que el disco puede contener.
- Usa las siguientes conversiones:

| Unidad | Equivalente |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

Por ejemplo, dado `500`, `"KB"` y `1` como argumentos, determina cuántos archivos de 500 KB pueden caber en un disco duro de 1 GB.

# --hints--

`number_of_files(500, "KB", 1)` debería devolver `2000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(500, "KB", 1), 2000)`)
}})
```

`number_of_files(50000, "B", 1)` debería devolver `20000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(50000, "B", 1), 20000)`)
}})
```

`number_of_files(5, "MB", 1)` debería devolver `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(5, "MB", 1), 200)`)
}})
```

`number_of_files(4096, "B", 1.5)` debería devolver `366210`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4096, "B", 1.5), 366210)`)
}})
```

`number_of_files(220.5, "KB", 100)` debería devolver `453514`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(220.5, "KB", 100), 453514)`)
}})
```

`number_of_files(4.5, "MB", 750)` debería devolver `166666`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_files(4.5, "MB", 750), 166666)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_files(file_size, file_unit, drive_size_gb):

    return file_size
```

# --solutions--

```py
def number_of_files(file_size, file_unit, drive_size_gb):
    drive_size_bytes = drive_size_gb * 1000 * 1000 * 1000

    if file_unit == "B":
        file_size_bytes = file_size
    elif file_unit == "KB":
        file_size_bytes = file_size * 1000
    else:
        file_size_bytes = file_size * 1000 * 1000

    return int(drive_size_bytes // file_size_bytes)
```
