---
id: 68b1f72371a5ac895ac70a06
title: "Desafío 42: Almacenamiento de Video"
challengeType: 29
dashedName: challenge-42
---

# --description--

Dado un tamaño de video, una unidad para el tamaño del video, una capacidad de disco duro y una unidad para el disco duro, devuelve el número de videos que el disco duro puede almacenar usando las siguientes restricciones:

- La unidad para el tamaño del video puede ser bytes (`"B"`), kilobytes (`"KB"`), megabytes (`"MB"`), o gigabytes (`"GB"`).
- Si no se proporciona una de las unidades de video anteriores, devuelve `"Invalid video unit"`.
- La unidad de la capacidad del disco duro puede ser gigabytes (`"GB"`) o terabytes (`"TB"`).
- Si no se proporciona una de las unidades de disco duro mencionadas arriba, devuelve `"Invalid drive unit"`.
- Devuelve el número de videos completos que el disco puede almacenar.
- Usa las siguientes conversiones:

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |
| 1 TB |   1000 GB  |

Por ejemplo, dado `500`, `"MB"`, `100` y `"GB"` como argumentos, determina cuántos videos de 500 MB pueden caber en un disco duro de 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` debería devolver `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` debería devolver `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` debería devolver `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` debería devolver `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` debería devolver `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
