---
id: 68b7687dded630607aceccb1
title: "Desafío 48: Detector de Spam"
challengeType: 29
dashedName: challenge-48
---

# --description--

Dado un número de teléfono en el formato `"+A (BBB) CCC-DDDD"`, donde cada letra representa un dígito de la siguiente manera:

- `A` representa el código del país y puede ser cualquier cantidad de dígitos.
- `BBB` representa el código de área y siempre tendrá tres dígitos.
- `CCC` y `DDDD` representan el número local y siempre tendrán tres y cuatro dígitos, respectivamente.

Determina si es un número de spam basado en los siguientes criterios:

- El código del país tiene más de 2 dígitos o no comienza con un cero (`0`).
- El código de área es mayor que 900 o menor que 200.
- La suma de los primeros tres dígitos del número local aparece dentro de los últimos cuatro dígitos del número local.
- El número tiene el mismo dígito cuatro o más veces seguidas (ignorando los caracteres de formato).

# --hints--

`is_spam("+0 (200) 234-0182")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
