---
id: 68b7cadffed0e75a517da677
title: "Desafío 54: ¡P@ssw0rd Str3ngth!"
challengeType: 29
dashedName: challenge-54
---

# --description--

Dada una cadena de contraseña, devuelve `"weak"`, `"medium"` o `"strong"` según la fortaleza de la contraseña.

Una contraseña se evalúa según las siguientes reglas:

- Tiene al menos 8 caracteres de longitud.
- Contiene tanto letras mayúsculas como minúsculas.
- Contiene al menos un número.
- Contiene al menos un carácter especial de este conjunto: `!`, `@`, `#`, `$`, `%`, `^`, `&` o `*`.

Devuelve `"weak"` si la contraseña cumple con menos de dos de las reglas.  
Devuelve `"medium"` si la contraseña cumple con 2 o 3 de las reglas.  
Devuelve `"strong"` si la contraseña cumple con las 4 reglas.

# --hints--

`check_strength("123456")` debería devolver `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` debería devolver `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` debería devolver `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` debería devolver `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` debería devolver `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` debería devolver `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` debería devolver `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` debería devolver `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` debería devolver `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
