---
id: 68b7687dded630607aceccb3
title: "Desafío 49: Analizador de encabezados CSV"
challengeType: 29
dashedName: challenge-49
---

# --description--

Dada la primera línea de un archivo de valores separados por comas (CSV), devuelve un arreglo que contiene los encabezados.

- La primera línea de un archivo CSV contiene encabezados separados por comas.
- Elimina cualquier espacio en blanco al principio o al final de cada encabezado.

# --hints--

`get_headings("name,age,city")` debería devolver `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` debería devolver `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` debería devolver `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
