---
id: 68b7cadffed0e75a517da675
title: "Desafío 53: Decimal a Binario"
challengeType: 29
dashedName: challenge-53
---

# --description--

Dado un entero no negativo, devuelve su representación binaria como una cadena.

Un número binario usa solo los dígitos `0` y `1` para representar cualquier número. Para convertir un número decimal a binario, divide repetidamente el número entre `2` y registra el residuo. Repite hasta que el número sea cero. Lee los residuos desde el último registrado hasta el primero. Por ejemplo, para convertir `12` a binario:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` en binario es `1100`.

# --hints--

`to_binary(5)` debería devolver `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` debería devolver `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` debería devolver `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` debería devolver `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```
