---
id: 68b7cadffed0e75a517da671
title: "Desafío 51: Formateador de Números Telefónicos"
challengeType: 29
dashedName: challenge-51
---

# --description--

Dada una cadena de once dígitos, devuelve la cadena como un número de teléfono en este formato: `"+D (DDD) DDD-DDDD"`.

# --hints--

`format_number("05552340182")` debería devolver `"+0 (555) 234-0182"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_number("05552340182"), "+0 (555) 234-0182")`)
}})
```

`format_number("15554354792")` debería devolver `"+1 (555) 435-4792"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_number("15554354792"), "+1 (555) 435-4792")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_number(number):

    return number
```

# --solutions--

```py
def format_number(number):
    country = number[0]
    area = number[1:4]
    prefix = number[4:7]
    line = number[7:]

    return f"+{country} ({area}) {prefix}-{line}"
```
