---
id: 68b7cadffed0e75a517da673
title: "Desafío 52: Binario a Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

Dada una cadena que representa un número binario, devuelve su equivalente decimal como un número.

Un número binario usa solo los dígitos `0` y `1` para representar cualquier número. Para convertir binario a decimal, multiplica cada dígito por una potencia de `2` y súmalos. Comienza multiplicando el dígito más a la derecha por `2^0`, el siguiente dígito a la izquierda por `2^1`, y así sucesivamente. Una vez que todos los dígitos hayan sido multiplicados por una potencia de `2`, suma el resultado.

Por ejemplo, el número binario `101` equivale a `5` en decimal porque:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` debería devolver `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` debería devolver `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` debería devolver `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
