---
id: 68b7687dded630607aceccaf
title: "Desafío 47: Capturado por Exceso de Velocidad"
challengeType: 29
dashedName: challenge-47
---

# --description--

Dado un arreglo de números que representan la velocidad a la que se observaron vehículos viajando, y un número que representa el límite de velocidad, devuelve un arreglo con dos elementos: el número de vehículos que excedieron la velocidad, seguido por el promedio de cuánto sobrepasaron el límite esos vehículos.

- Si no hay vehículos acelerando, devuelve `[0, 0]`.

# --hints--

`speeding([50, 60, 55], 60)` debería devolver `[0, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([50, 60, 55], 60), [0, 0])`)
}})
```

`speeding([58, 50, 60, 55], 55)` debería devolver `[2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([58, 50, 60, 55], 55), [2, 4])`)
}})
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` debería devolver `[4, 8.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5])`)
}})
```

`speeding([100, 105, 95, 102], 100)` debería devolver `[2, 3.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([100, 105, 95, 102], 100), [2, 3.5])`)
}})
```

`speeding([40, 45, 44, 50, 112, 39], 55)` debería devolver `[1, 57]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57])`)
}})
```

# --seed--

## --seed-contents--

```py
def speeding(speeds, limit):

    return speeds
```

# --solutions--

```py
def speeding(speeds, limit):
    speeding = 0
    total_over = 0

    for speed in speeds:
        if speed > limit:
            speeding += 1
            total_over += (speed - limit)
    
    if speeding == 0:
        return [0, 0]
    
    return [speeding, total_over / speeding]
```
