---
id: 68c1a929005bf54d342aa8d5
title: "Desafío 58: Semana del Espacio Día 4: Lugar de Aterrizaje"
challengeType: 29
dashedName: challenge-58
---

# --description--

En el cuarto día de la Semana Espacial, te dan una matriz de números (un arreglo de arreglos), que representa posibles lugares de aterrizaje para tu rover. Encuentra el lugar de aterrizaje más seguro basado en las siguientes reglas:

- Cada posición en la matriz contendrá un número de `0-9`, inclusive.
- Cualquier `0` representa un posible punto de aterrizaje.
- Cualquier número distinto de `0` es demasiado peligroso para aterrizar. Cuanto mayor sea el número, más peligroso.
- El lugar más seguro se define como la celda `0` cuyas celdas circundantes (hasta 4 vecinas, sin contar diagonales) tienen el menor peligro total.
- Ignora a los vecinos fuera de los límites (las esquinas y los bordes simplemente tienen menos vecinos).
- Devuelve los índices del lugar de aterrizaje más seguro. Siempre habrá un único lugar más seguro.

Por ejemplo, dado:

```js
[
  [1, 0],
  [2, 0]
]
```

Devuelve `[0, 1]`, los índices para el `0` en el primer arreglo.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` debería devolver `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` debería devolver `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` debería devolver `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` debería devolver `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
