---
id: 68c1a929005bf54d342aa8d3
title: "Desafío 56: Semana del Espacio Día 2: Búsqueda de Exoplanetas"
challengeType: 29
dashedName: challenge-56
---

# --description--

Para el segundo día de la Semana del Espacio, se te da una cadena donde cada carácter representa la lectura de luminosidad de una estrella. Determina si las lecturas han detectado un exoplaneta usando el método de tránsito. El método de tránsito es cuando un planeta pasa frente a una estrella, reduciendo su luminosidad observada.

- Las lecturas de luminosidad solo comprenden los caracteres `0-9` y `A-Z`, donde cada lectura corresponde a los siguientes valores numéricos:
- Los caracteres `0-9` corresponden a niveles de luminosidad `0-9`.
- Los caracteres `A-Z` corresponden a niveles de luminosidad `10-35`.

Se considera que una estrella tiene un exoplaneta si alguna lectura individual es menor o igual al 80% del promedio de todas las lecturas. Por ejemplo, si la luminosidad promedio de una estrella es 10, se consideraría que tiene un exoplaneta si alguna lectura individual es 8 o menos.

# --hints--

`has_exoplanet("665544554")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
