---
id: 68c1a929005bf54d342aa8d4
title: "Desafío 57: Semana Espacial Día 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

Para el tercer día de la Semana Espacial, se te da un arreglo de números que representan las distancias (en kilómetros) entre tú, los satélites y tu planeta de origen en una ruta de comunicación. Determina cuánto tiempo tomará un mensaje enviado a través de la ruta para llegar a su planeta destino usando las siguientes restricciones:

- El primer valor en el array es la distancia desde tu ubicación hasta el primer satélite.
- Cada valor subsecuente, excepto el último, es la distancia al siguiente satélite.
- El último valor en el array es la distancia desde el satélite anterior hasta tu planeta natal.
- El mensaje viaja a 300,000 km/s.
- Cada satélite por el que **pasa** el mensaje añade un retraso de transmisión de 0.5 segundos.
- Devuelve un número redondeado a 4 decimales, con ceros finales eliminados.

# --hints--

`send_message([300000, 300000])` debería devolver `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` debería devolver `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` debería devolver `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` debería devolver `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` debería devolver `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` debería devolver `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
