---
id: 68c1a929005bf54d342aa8d6
title: "Desafío 59: Semana Espacial Día 5: Zona Ricitos de Oro"
challengeType: 29
dashedName: challenge-59
---

# --description--

Para el quinto día de la Semana del Espacio, calcularás la "zona de Ricitos de Oro" de una estrella: la región alrededor de una estrella donde las condiciones son "justas" para que exista agua líquida.

Dada la masa de una estrella, devuelve un arreglo con las distancias de inicio y fin de su Zona Ricitos de Oro en Unidades Astronómicas.

Para calcular la Zona de Ricitos de Oro:

1. Encuentra la luminosidad de la estrella elevando su masa a la potencia de 3.5.
2. El inicio de la zona es 0.95 veces la raíz cuadrada de su luminosidad.
3. El final de la zona es 1.37 veces la raíz cuadrada de su luminosidad.

- Devuelve las distancias redondeadas a dos decimales.

Por ejemplo, dado `1` como masa, devuelve `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` debería devolver `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` debería devolver `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` debería devolver `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` debería devolver `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` debería devolver `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
