---
id: 68c497f3aaefc9fd9f1b0e26
title: "Desafío 62: Hexadecimal a Decimal"
challengeType: 29
dashedName: challenge-62
---

# --description--

Dada una cadena que representa un número hexadecimal (base 16), devuelve su valor decimal (base 10) como un entero.

Hexadecimal es un sistema numérico que usa 16 dígitos:

- `0-9` representan valores de `0` a `9`.
- `A-F` representan valores de `10` a `15`.

Aquí tienes una tabla de conversión parcial:

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- La cadena solo contendrá caracteres `0–9` y `A–F`.

# --hints--

`hex_to_decimal("A")` debería devolver `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` debería devolver `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` debería devolver `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` debería devolver `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` debería devolver `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
