---
id: 68c497f3aaefc9fd9f1b0e25
title: "Desafío 61: Semana del Espacio Día 7: Combustible de Lanzamiento"
challengeType: 29
dashedName: challenge-61
---

# --description--

Para el último día de la Semana Espacial, se te dará la masa en kilogramos (kg) de una carga útil que quieres enviar a órbita. Determina la cantidad de combustible necesaria para enviar tu carga útil a órbita usando las siguientes reglas:

- Los cohetes requieren 1 kg de combustible por cada 5 kg de masa que deben levantar.
- El combustible en sí tiene masa. Así que cuando agregas combustible, la masa que hay que levantar aumenta, lo que requiere más combustible, lo que incrementa la masa, y así sucesivamente.
- Para calcular el combustible total necesario: comienza con la masa de la carga útil, calcula el combustible necesario para esa masa, añade ese combustible a la masa total y calcula de nuevo. Repite este proceso hasta que el combustible adicional requerido sea menor a 1 kg, luego detente.
- Ignora la masa del cohete en sí. Solo calcula el combustible necesario para levantar la carga útil y su propio combustible.

Por ejemplo, dado un masa de carga útil de 50 kg, necesitarías 10 kg de combustible para levantarla (payload / 5), lo que aumenta la masa total a 60 kg, que necesita 12 kg para levantarla (2 kg adicionales), lo que aumenta la masa total a 62 kg, que necesita 12.4 kg para levantar - 0.4 kg adicionales - que es menos de 1 kg adicional, así que nos detenemos aquí. La masa total a levantar es de 62.4 kg, 50 de los cuales es la carga útil inicial y 12.4 de combustible.

- Devuelve la cantidad de combustible necesaria redondeada a un decimal.

# --hints--

`launch_fuel(50)` debería devolver `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` debería devolver `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` debería devolver `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` debería devolver `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` debería devolver `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```
