---
id: 68c497f3aaefc9fd9f1b0e24
title: "Desafío 60: Semana del Espacio Día 6: Fase Lunar"
challengeType: 29
dashedName: challenge-60
---

# --description--

Para el día seis de la Semana Espacial, se te dará una fecha en el formato `"YYYY-MM-DD"` y necesitarás determinar la fase de la luna para ese día usando las siguientes reglas:

Usa un ciclo lunar simplificado de 28 días, dividido en cuatro fases iguales:

- "Nuevo": días 1 - 7
- "Luna creciente": días 8 - 14
- "Full": días 15 - 21
- "Waning": días 22 - 28

Después del día 28, el ciclo se repite con el día 1, una luna nueva.

- Usa `"2000-01-06"` como referencia de luna nueva (día 1 del ciclo) para determinar la fase del día dado.
- No se te proporcionarán fechas anteriores a la fecha de referencia.
- Devuelve la fase correcta como una cadena.

**Nota:** El Día 1 representa el día de la luna nueva, lo que significa que han pasado 0 días desde la última luna nueva.

# --hints--

`moon_phase("2000-01-12")` debería devolver `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` debería devolver `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` debería devolver `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` debería devolver `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` debería devolver `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
