---
id: 68cae5b538ff798bbd4da006
title: "Desafío 68: Enmascarador de Tarjetas de Crédito"
challengeType: 29
dashedName: challenge-68
---

# --description--

Dada una cadena de números de tarjeta de crédito, devuelve una versión enmascarada de la misma usando las siguientes restricciones:

- La cadena contendrá cuatro grupos de cuatro dígitos (`0-9`), con todos los grupos separados por un solo espacio o un solo guion (`-`).
- Reemplaza todos los números, excepto los últimos cuatro, con un asterisco (`*`).
- Deja los caracteres restantes sin cambios.

Por ejemplo, dado `"4012-8888-8888-1881"` devuelve `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` debería devolver `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` debería devolver `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` debería devolver `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` debería devolver `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
