---
id: 68cae5b538ff798bbd4da003
title: "Desafío 65: Conteo de Cadenas"
challengeType: 29
dashedName: challenge-65
---

# --description--

Dadas dos cadenas, determina cuántas veces aparece la segunda cadena en la primera.

- La cadena patrón puede superponerse en la primera cadena. Por ejemplo, `"aaa"` contiene `"aa"` dos veces. Las primeras dos `a` y las segundas dos.

# --hints--

`count('abcdefg', 'def')` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('abcdefg', 'def'), 1)`)
}})
```

`count('hello', 'world')` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('hello', 'world'), 0)`)
}})
```

`count('mississippi', 'iss')` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('mississippi', 'iss'), 2)`)
}})
```

`count('she sells seashells by the seashore', 'sh')` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('she sells seashells by the seashore', 'sh'), 3)`)
}})
```

`count('101010101010101010101', '101')` debería devolver `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count('101010101010101010101', '101'), 10)`)
}})
```

# --seed--

## --seed-contents--

```py
def count(text, parameter):

    return text
```

# --solutions--

```py
def count(text, pattern):
    if not pattern:
        return 0
    occurrences = 0

    for i in range(len(text) - len(pattern) + 1):
        if text[i:i+len(pattern)] == pattern:
            occurrences += 1

    return occurrences
```
