---
id: 68cae5b538ff798bbd4da004
title: "Desafío 66: Eliminador de Etiquetas HTML"
challengeType: 29
dashedName: challenge-66
---

# --description--

Dado un string de código HTML, elimina las etiquetas y devuelve el contenido de texto plano.

- La cadena de entrada contendrá solo HTML válido.
- Las etiquetas HTML pueden estar anidadas.
- Elimina las etiquetas y cualquier atributo.

Por ejemplo, `'<a href="#">Click here</a>'` debería devolver `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` debería devolver `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` debería devolver `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` debería devolver una cadena vacía (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` debería devolver `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
