---
id: 68cae5b538ff798bbd4da002
title: "Desafío 64: de 24 a 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

Dada una cadena que representa una hora del día en el formato de 24 horas `"HHMM"`, devuelve la hora en su equivalente formato de 12 horas `"H:MM AM"` o `"H:MM PM"`.

- La entrada dada siempre será una cadena de cuatro dígitos en formato de hora de 24 horas, desde `"0000"` hasta `"2359"`.

# --hints--

`to_12("1124")` debería devolver `"11:24 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` debería devolver `"9:00 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` debería devolver `"2:55 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` debería devolver `"11:46 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` debería devolver `"12:30 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
