---
id: 68cae5b538ff798bbd4da008
title: "Desafío 70: Extractor de Atributos HTML"
challengeType: 29
dashedName: challenge-70
---

# --description--

Dada una cadena de un elemento HTML válido, devuelve los atributos del elemento usando los siguientes criterios:

- Solo se te dará un elemento.
- Los atributos estarán en el formato: `attribute="value"`.
- Devuelve un arreglo de cadenas con cada propiedad y valor del atributo, separados por una coma, en este formato: `["attribute1, value1", "attribute2, value2"]`.
- Devuelve los atributos en el orden en que se proporcionan.
- Si no se encuentran atributos, devuelve un arreglo vacío.

# --hints--

`extract_attributes('<span class="red"></span>')` debería devolver `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` debería devolver `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` debería devolver `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` debería devolver `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
