---
id: 68cae5b538ff798bbd4da001
title: "Desafío 63: Batalla de Palabras"
challengeType: 29
dashedName: challenge-63
---

# --description--

Dadas dos oraciones que representan a tu equipo y a un equipo contrario, donde cada palabra de tu equipo batalla contra la palabra correspondiente del equipo contrario, determina qué equipo gana usando las siguientes reglas:

- Las oraciones dadas siempre contendrán el mismo número de palabras.
- Las palabras están separadas por un solo espacio y solo contienen letras.
- El valor de cada palabra es la suma de sus letras.
- Las letras `a` a `z` corresponden a los valores del `1` al `26`. Por ejemplo, `a` es `1` y `z` es `26`.
- Una letra mayúscula duplica el valor de la letra. Por ejemplo, `A` es `2`, y `Z` es `52`.
- Las palabras luchan en orden: la primera palabra de tu equipo lucha contra la primera palabra del equipo contrario, y así sucesivamente.
- Una palabra gana si su valor es mayor que el valor de la palabra opuesta.
- El equipo con más palabras ganadoras es el ganador.

Devuelve `"We win"` si tu equipo es el ganador, `"We lose"` si tu equipo pierde, y `"Draw"` si ambos equipos tienen el mismo número de victorias.

# --hints--

`battle("hello world", "hello word")` debería devolver `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` debería devolver `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` debería devolver `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` debería devolver `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` debería devolver `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` debería devolver `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` debería devolver `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
