---
id: 68cae5b538ff798bbd4da009
title: "Desafío 71: Calculadora de propinas"
challengeType: 29
dashedName: challenge-71
---

# --description--

Dado el precio de tu comida y un porcentaje de propina personalizado, devuelve un arreglo con tres valores de propina; 15%, 20% y la cantidad personalizada.

- Los precios se darán en el formato: `"$N.NN"`.
- Los porcentajes personalizados de propina se darán en este formato: `"25%"`.
- Devuelve los montos en el mismo formato `"$N.NN"`, redondeados a dos decimales.

Por ejemplo, dado un precio de comida de `"$10.00"` y un valor de propina personalizada de `"25%"`, devuelve `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculate_tips("$10.00", "25%")` debería devolver `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` debería devolver `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` debería devolver `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
