---
id: 68cae5b538ff798bbd4da00a
title: "Desafío 72: Ajustador de Termostato 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Dada la temperatura actual de una habitación en Fahrenheit y una temperatura objetivo en Celsius, devuelve una cadena que indique cómo ajustar la temperatura de la habitación según estas restricciones:

- Devuelve `"Heat: X degrees Fahrenheit"` si la temperatura actual está por debajo del objetivo. Con `X` siendo el número de grados en Fahrenheit para calentar la habitación y alcanzar el objetivo, redondeado a 1 decimal.
- Devuelve `"Cool: X degrees Fahrenheit"` si la temperatura actual está por encima del objetivo. Con `X` siendo el número de grados en Fahrenheit para enfriar la habitación y alcanzar el objetivo, redondeado a 1 decimal.
- Devuelve `"Hold"` si la temperatura actual es igual a la objetivo.

Para convertir Celsius a Fahrenheit, multiplica la temperatura en Celsius por 1.8 y suma 32 al resultado (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` debería devolver `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` debería devolver `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` debería devolver `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` debería devolver `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` debería devolver `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
