---
id: 68d2ba1468508398389487d0
title: "Desafío 75: Tesoro Oculto"
challengeType: 29
dashedName: challenge-75
---

# --description--

Dado un arreglo 2D que representa un mapa del fondo oceánico que incluye un tesoro oculto, y un arreglo con las coordenadas (`[row, column]`) para la próxima inmersión de tu búsqueda del tesoro, devuelve `"Empty"`, `"Found"` o `"Recovered"` usando las siguientes reglas:

- El arreglo 2D dado contendrá exactamente un tesoro no recuperado, que ocupará múltiples celdas.
- Each cell in the 2D array will contain one of the following values:
  - `"-"`: Sin tesoro.
  - `"O"`: Una parte del tesoro que no ha sido encontrada.
  - `"X"`: Una parte del tesoro que ya ha sido encontrada.

- Si la ubicación de la inmersión no tiene tesoro, devuelve `"Empty"`.
- Si la ubicación de la inmersión encuentra tesoro, pero al menos otra parte del tesoro permanece sin encontrar, devuelve `"Found"`.
- Si la ubicación de la inmersión encuentra la última parte del tesoro que no se había encontrado, devuelve `"Recovered"`.

Por ejemplo, dado:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

Y `[2, 1]` para las coordenadas del lugar de la inmersión, devuelve `"Recovered"` porque la inmersión encontró la última parte del tesoro que no se había encontrado.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` debería devolver `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` debería devolver `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` debería devolver `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` debería devolver `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` debería devolver `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` debería devolver `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```
