---
id: 68d2ba1468508398389487cf
title: "Desafío 74: Canciones Favoritas"
challengeType: 29
dashedName: challenge-74
---

# --description--

¿Recuerdas los iPods? El primer modelo salió hace 24 años hoy, el 23 de octubre de 2001.

Dado un arreglo de objetos de canciones que representan tu lista de reproducción de iPod, devuelve un arreglo con los títulos de las dos canciones más reproducidas, con la canción más reproducida primero.

- Cada objeto tendrá una propiedad `"title"` (cadena) y una propiedad `"plays"` (entero).

# --hints--

`favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` debería devolver `["Sync or Swim", "Earbud Blues"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"])`)
}})
```

`favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` debería devolver `["Clickwheel Love", "99 Downloads"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"])`)
}})
```

`favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` debería devolver `["Song B", "Song C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"])`)
}})
```

# --seed--

## --seed-contents--

```py
def favorite_songs(playlist):

    return playlist
```

# --solutions--

```py
def favorite_songs(playlist):
    sorted_songs = sorted(playlist, key=lambda x: x["plays"], reverse=True)
    return [song["title"] for song in sorted_songs[:2]]
```
