---
id: 68d2ba1468508398389487ce
title: "Desafío 73: Habla sabiamente, debes"
challengeType: 29
dashedName: challenge-73
---

# --description--

Dada una oración, devuelve una versión de ella que suene como un consejo de un maestro sabio usando las siguientes reglas:

- Las palabras están separadas por un solo espacio.
- Encuentra la primera ocurrencia de una de las siguientes palabras en la oración: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Move all words before and including that word to the end of the sentence and:
  - Conserva el orden de las palabras cuando las muevas.
  - Ponlos todos en minúsculas.
  - Y añade una coma y un espacio antes de ellos.
- Pon en mayúscula la primera letra de la nueva primera palabra de la oración.
- Todas las oraciones dadas terminarán con un solo signo de puntuación. Mantén la puntuación original al final de la nueva oración.
- Devuelve la nueva oración, asegúrate de que haya un solo espacio entre cada palabra y no haya espacios al principio ni al final de la oración.

Por ejemplo, dado `"You must speak wisely."` devuelve `"Speak wisely, you must."`

# --hints--

`wise_speak("You must speak wisely.")` debería devolver `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` debería devolver `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` debería devolver `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` debería devolver `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` debería devolver `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
