---
id: 68d30845cc08266018fc46bd
title: "Desafío 77: Formateador de Duración"
challengeType: 29
dashedName: challenge-77
---

# --description--

Dado un número entero de segundos, devuelve una cadena que represente la misma duración en el formato `"H:MM:SS"`, donde `"H"` es el número de horas, `"MM"` es el número de minutos y `"SS"` es el número de segundos. Devuelve el tiempo usando las siguientes reglas:

- Segundos: Siempre deben ser dos dígitos.
- Minutos: Deben omitirse los ceros a la izquierda cuando no sean necesarios. Usa `"0"` si la duración es menor a un minuto.
- Horas: Deben incluirse solo si son mayores que cero.

# --hints--

`format(500)` debería devolver `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` debería devolver `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` debería devolver `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` debería devolver `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` debería devolver `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
