---
id: 68d30845cc08266018fc46be
title: "Desafío 78: Secuencia de enteros"
challengeType: 29
dashedName: challenge-78
---

# --description--

Dado un entero positivo, devuelve una cadena con todos los enteros desde `1` hasta, e incluyendo, el número dado, en orden numérico.

Por ejemplo, dado `5`, devuelve `"12345"`.

# --hints--

`sequence(5)` debería devolver `"12345"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` debería devolver `"12345678910"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` debería devolver `"1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` debería devolver `"123456789101112131415161718192021222324252627"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```
