---
id: 68d30845cc08266018fc46bc
title: "Desafío 76: ADN complementario"
challengeType: 29
dashedName: challenge-76
---

# --description--

Dada una cadena que representa una secuencia de ADN, devuelve su hebra complementaria usando las siguientes reglas:

- El ADN consiste en las letras `"A"`, `"C"`, `"G"` y `"T"`.
- Las letras `"A"` y `"T"` se complementan entre sí.
- Las letras `"C"` y `"G"` se complementan entre sí.

Por ejemplo, dado `"ACGT"`, devuelve `"TGCA"`.

# --hints--

`complementary_dna("ACGT")` debería devolver `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` debería devolver `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` debería devolver `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` debería devolver `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
