---
id: 68e39ed6106dac2f0a98fd62
title: "Desafío 80: Clasificador de Emails"
challengeType: 29
dashedName: challenge-80
---

# --description--

El 29 de octubre de 1971, se envió el primer correo electrónico, introduciendo el formato `username@domain` que todavía usamos. Ahora, hay miles de millones de direcciones de correo electrónico.

En este desafío, se te da una lista de direcciones de correo electrónico y necesitas ordenarlas alfabéticamente primero por el nombre de dominio (la parte después del `@`), y segundo por el nombre de usuario (la parte antes del `@`).

- El ordenamiento debe ser insensible a mayúsculas y minúsculas.
- Si más de un correo electrónico tiene el mismo dominio, ordénalos por su nombre de usuario.
- Devuelve un arreglo de las direcciones ordenadas.
- Las direcciones devueltas deben conservar su mayúscula y minúscula original.

Por ejemplo, dado `["jill@mail.com", "john@example.com", "jane@example.com"]`, devuelve `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` debería devolver `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"])`)
}})
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` debería devolver `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"])`)
}})
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` debería devolver `["user@x.com", "user@y.com", "user@z.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"])`)
}})
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` debería devolver `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"])`)
}})
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` debería devolver `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort(emails):

    return emails
```

# --solutions--

```py
def sort(emails):
    return sorted(
        emails,
        key=lambda email: (email.split('@')[1].lower(), email.split('@')[0].lower())
    )
```
