---
id: 68e39ed6106dac2f0a98fd63
title: "Desafío 81: Nth Prime"
challengeType: 29
dashedName: challenge-81
---

# --description--

Un número primo es un entero positivo mayor que 1 que es divisible solo por 1 y por sí mismo. Los primeros cinco números primos son `2`, `3`, `5`, `7` y `11`.

Dado un entero positivo `n`, devuelve el número primo `n`-ésimo. Por ejemplo, dado `5` devuelve el 5º número primo: `11`.

# --hints--

`nth_prime(5)` debería devolver `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(5), 11)`)
}})
```

`nth_prime(10)` debería devolver `29`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(10), 29)`)
}})
```

`nth_prime(16)` debería devolver `53`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(16), 53)`)
}})
```

`nth_prime(99)` debería devolver `523`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(99), 523)`)
}})
```

`nth_prime(1000)` debería devolver `7919`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_prime(1000), 7919)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_prime(n):

    return n
```

# --solutions--

```py
def nth_prime(n):
    primes = []
    num = 2

    while len(primes) < n:
        is_prime = True
        for i in range(2, int(num**0.5) + 1):
            if num % i == 0:
                is_prime = False
                break
        if is_prime:
            primes.append(num)
        num += 1

    return primes[n - 1]
```
