---
id: 68ee9e3066cfd4eb2328e8a7
title: "Desafío 88: Buscador de días de la semana"
challengeType: 29
dashedName: challenge-88
---

# --description--

Dada una cadena de fecha en el formato `YYYY-MM-DD`, devuelve el día de la semana.

Los días válidos para la devolución son:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Asegúrate de ignorar las zonas horarias.

# --hints--

`get_weekday("2025-11-06")` debería devolver `Thursday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2025-11-06"), "Thursday")`)
}})
```

`get_weekday("1999-12-31")` debería devolver `Friday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1999-12-31"), "Friday")`)
}})
```

`get_weekday("1111-11-11")` debería devolver `Saturday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("1111-11-11"), "Saturday")`)
}})
```

`get_weekday("2112-12-21")` debería devolver `Wednesday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2112-12-21"), "Wednesday")`)
}})
```

`get_weekday("2345-10-01")` debería devolver `Monday`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_weekday("2345-10-01"), "Monday")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_weekday(date_string):

    return date_string
```

# --solutions--

```py
import datetime
def get_weekday(date_string):
    days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"]
    year, month, day = map(int, date_string.split("-"))
    date = datetime.date(year, month, day)

    return days[date.weekday() % 7 + 1 if date.weekday() != 6 else 0]
```
