---
id: 68ee9e3066cfd4eb2328e8a8
title: "Desafío 89: Contar cartas"
challengeType: 29
dashedName: challenge-89
---

# --description--

Una baraja estándar de cartas tiene 13 cartas únicas en cada palo. Dado un entero que representa el número de cartas para elegir de la baraja, devuelve el número de combinaciones únicas de cartas que puedes elegir.

- El orden no importa. Elegir la carta A y luego la carta B es lo mismo que elegir la carta B y luego la carta A.

Por ejemplo, dado `52`, devuelve `1`. Solo hay una combinación de 52 cartas para elegir de una baraja de 52 cartas. Y dado `2`, devuelve `1326`. Hay 1326 combinaciones de cartas que puedes obtener al elegir 2 cartas de la baraja.

# --hints--

`combinations(52)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` debería devolver `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` debería devolver `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` debería devolver `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` debería devolver `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` debería devolver `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
