---
id: 68f6587287ad1f4ad39b0c83
title: "Desafío 97: MCD"
challengeType: 29
dashedName: challenge-97
---

# --description--

Dado dos enteros positivos, devuelve su máximo común divisor (GCD).

- El MCD de dos enteros es el número más grande que divide de manera exacta a ambos números sin dejar residuo.

Por ejemplo, los divisores de `4` son `1`, `2` y `4`. Los divisores de `6` son `1`, `2`, `3` y `6`. Entonces, dado `4` y `6`, devuelve `2`, el número más grande que aparece en ambos conjuntos de divisores.

# --hints--

`gcd(4, 6)` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` debería devolver `864`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
