---
id: 68f6587287ad1f4ad39b0c7f
title: "Desafío 93: Vocales y Consonantes"
challengeType: 29
dashedName: challenge-93
---

# --description--

Dada una cadena, devuelve un arreglo con el número de vocales y el número de consonantes en la cadena.

- Las vocales consisten en `a`, `e`, `i`, `o`, `u` en cualquier caso.
- Las consonantes consisten en todas las demás letras en cualquier caso.
- Ignora cualquier carácter que no sea una letra.

Por ejemplo, dado `"Hello World"`, devuelve `[3, 7]`.

# --hints--

`count("Hello World")` debería devolver `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` debería devolver `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` debería devolver `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` debería devolver `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` debería devolver `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` debería devolver `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```
