---
id: 68f6587287ad1f4ad39b0c7c
title: "Desafío 90: Límite de Caracteres"
challengeType: 29
dashedName: challenge-90
---

# --description--

En este desafío, se te da una cadena y necesitas determinar si cabe en una publicación de redes sociales. Devuelve las siguientes cadenas según las reglas dadas:

- "publicación corta" si cabe dentro de un límite de 40 caracteres.
- "long post" si es mayor a 40 caracteres y cabe dentro de un límite de 80 caracteres.
- "publicación inválida" si es demasiado largo para caber dentro de cualquiera de los límites.

# --hints--

`can_post("Hello world")` debería devolver `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` debería devolver `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` debería devolver `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
