---
id: 68f6587287ad1f4ad39b0c7e
title: "Desafío 92: Extractor de Extensiones"
challengeType: 29
dashedName: challenge-92
---

# --description--

Dada una cadena que representa un nombre de archivo, devuelve la extensión del archivo.

- La extensión es la parte del nombre del archivo que viene después del último punto (`.`).
- Si el nombre del archivo no contiene un punto o termina con un punto, devuelve `"none"`.
- La extensión debe devolverse tal cual, preservando mayúsculas y minúsculas.

# --hints--

`get_extension("document.txt")` debería devolver `"txt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` debería devolver `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` debería devolver `"PNG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` debería devolver `"gitignore"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` debería devolver `"gz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` debería devolver `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
