---
id: 68f6587287ad1f4ad39b0c82
title: "Desafío 96: ¿Es fin de semana?"
challengeType: 29
dashedName: challenge-96
---

# --description--

Dada una fecha en el formato `"YYYY-MM-DD"`, devuelve el número de días que faltan hasta el fin de semana.

- El fin de semana comienza el sábado.
- Si la fecha dada es sábado o domingo, devuelve `"It's the weekend!"`.
- De lo contrario, devuelve `"X days until the weekend."`, donde `X` es el número de días hasta el sábado.
- Si `X` es `1`, usa `"day"` (singular) en lugar de `"days"` (plural).
- Asegúrate de que el cálculo ignore tu zona horaria local.

# --hints--

`days_until_weekend("2025-11-14")` debería devolver `"1 day until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-11-14"), "1 day until the weekend.")`)
}})
```

`days_until_weekend("2025-01-01")` debería devolver `"3 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-01-01"), "3 days until the weekend.")`)
}})
```

`days_until_weekend("2025-12-06")` debería devolver `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2025-12-06"), "It's the weekend!")`)
}})
```

`days_until_weekend("2026-01-27")` debería devolver `"4 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-01-27"), "4 days until the weekend.")`)
}})
```

`days_until_weekend("2026-09-07")` debería devolver `"5 days until the weekend."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-09-07"), "5 days until the weekend.")`)
}})
```

`days_until_weekend("2026-11-29")` debería devolver `"It's the weekend!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_weekend("2026-11-29"), "It's the weekend!")`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_weekend(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime

def days_until_weekend(date_string):
    date = datetime.strptime(date_string, "%Y-%m-%d").date()
    day_of_week = date.weekday()

    if day_of_week == 5 or day_of_week == 6:
        return "It's the weekend!"

    days_until_saturday = 5 - day_of_week
    return f"{days_until_saturday} day{'s' if days_until_saturday != 1 else ''} until the weekend."
```
