---
id: 68f6587287ad1f4ad39b0c84
title: "Desafío 98: Conteo de Rectángulos"
challengeType: 29
dashedName: challenge-98
---

# --description--

Dado dos enteros positivos que representan el ancho y la altura de un rectángulo, determina cuántos rectángulos pueden caber dentro del dado.

- Cuenta solo los rectángulos con ancho y alto enteros.

Por ejemplo, dado `1` y `3`, devuelve `6`. Tres rectángulos de 1x1, dos rectángulos de 1x2 y un rectángulo de 1x3.

# --hints--

`count_rectangles(1, 3)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 3), 6)`)
}})
```

`count_rectangles(3, 2)` debería devolver `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(3, 2), 18)`)
}})
```

`count_rectangles(1, 2)` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(1, 2), 3)`)
}})
```

`count_rectangles(5, 4)` debería devolver `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(5, 4), 150)`)
}})
```

`count_rectangles(11, 19)` debería devolver `12540`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_rectangles(11, 19), 12540)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_rectangles(width, height):

    return width
```

# --solutions--

```py
def count_rectangles(width, height):
    total = 0
    for w in range(1, width + 1):
        for h in range(1, height + 1):
            total += (width - w + 1) * (height - h + 1)
    return total
```
