---
id: 68ffb91507a5b645769328c6
title: "Desafío 103: MCM"
challengeType: 29
dashedName: challenge-103
---

# --description--

Dado dos enteros, devuelve el mínimo común múltiplo (LCM) de los dos números.

El MCM de dos números es el entero positivo más pequeño que es múltiplo de ambos números. Por ejemplo, dado `4` y `6`, devuelve `12` porque:

- Los múltiplos de `4` son `4`, `8`, `12` y así sucesivamente.
- Múltiplos de `6` son `6`, `12`, `18` y así sucesivamente.
- `12` es el número más pequeño que es múltiplo de ambos.

# --hints--

`lcm(4, 6)` debería devolver `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` debería devolver `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` debería devolver `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` debería devolver `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` debería devolver `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
