---
id: 68ffb91507a5b645769328c8
title: "Desafío 105: Conteo de Caracteres"
challengeType: 29
dashedName: challenge-105
---

# --description--

Dada una cadena de texto, devuelve un arreglo con el conteo de cada carácter en orden alfabético.

- Trata las letras mayúsculas y minúsculas como la misma letra al contar.
- Ignora números, espacios, puntuación, etc.
- Devuelve el conteo y la letra en el formato `"letter count"`. Por ejemplo, `"a 3"`.
- Todas las letras devueltas deben estar en minúsculas.
- No devuelvas un conteo de letras que no estén en la cadena dada.

# --hints--

`count_characters("hello world")` debería devolver `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_characters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"])`)
}})
```

`count_characters("I love coding challenges!")` debería devolver `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_characters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"])`)
}})
```

`count_characters("// TODO: Complete this challenge ASAP!")` debería devolver `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_characters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"])`)
}})
```

# --seed--

## --seed-contents--

```py
def count_characters(sentence):

    return sentence
```

# --solutions--

```py
def count_characters(sentence):
    from collections import Counter
    
    letters = [c.lower() for c in sentence if c.isalpha()]
    counts = Counter(letters)
    
    return [f"{letter} {counts[letter]}" for letter in sorted(counts)]
```
