---
id: 68ffb91507a5b645769328c4
title: "Desafío 101: Convertidor de Encabezados Markdown"
challengeType: 29
dashedName: challenge-101
---

# --description--

Dada una cadena que representa un encabezado Markdown, devuelve el encabezado HTML equivalente.

Un encabezado válido de Markdown debe:

- Comienza con cero o más espacios, seguido de
- De 1 a 6 caracteres hash (`#`) seguidos, luego
- Al menos un espacio. Y finalmente,
- El texto del encabezado.

El número de símbolos de almohadilla determina el nivel del encabezado. Por ejemplo, un símbolo de almohadilla corresponde a una etiqueta `h1`, y seis símbolos de almohadilla corresponden a una etiqueta `h6`.

Si la cadena dada no tiene el formato exacto anterior, devuelve `"Invalid format"`.

Por ejemplo, dado `"# My level 1 heading"`, devuelve `"<h1>My level 1 heading</h1>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`convert("# My level 1 heading")` debería devolver `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` debería devolver `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` debería devolver `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` debería devolver `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
