---
id: 68ffb91507a5b645769328cc
title: "Desafío 109: ¿Cuál es mi edad otra vez?"
challengeType: 29
dashedName: challenge-109
---

# --description--

Dada la fecha de cumpleaños de alguien en el formato `YYYY-MM-DD`, devuelve la edad de la persona al 27 de noviembre de 2025.

- Asume que todos los cumpleaños son fechas válidas antes del 27 de noviembre de 2025.
- Devuelve la edad como un entero.
- Asegúrate de considerar si la persona ya ha cumplido años en 2025.

# --hints--

`calculate_age("2000-11-20")` debería devolver `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` debería devolver `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` debería devolver `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` debería devolver `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` debería devolver `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
