---
id: 68ffb91507a5b645769328c3
title: "Desafío 100: 100 Caracteres"
challengeType: 29
dashedName: challenge-100
---

# --description--

¡Bienvenido al desafío diario de codificación número 100!

Dada una cadena, repite sus caracteres hasta que el resultado tenga exactamente 100 caracteres. Si tus repeticiones superan los 100 caracteres, recorta el exceso para que tenga exactamente 100.

# --hints--

`one_hundred("One hundred ")` debería devolver `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ")`)
}})
```

`one_hundred("freeCodeCamp ")` debería devolver `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC")`)
}})
```

`one_hundred("daily challenges ")` debería devolver `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge")`)
}})
```

`one_hundred("!")` debería devolver `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(one_hundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!")`)
}})
```



# --seed--

## --seed-contents--

```py
def one_hundred(chars):

    return chars
```

# --solutions--

```py
def one_hundred(chars):
    result = ""
    i = 0

    while len(result) < 100:
        result += chars[i % len(chars)]
        i += 1

    return result[:100]
```
