---
id: 68ffb91507a5b645769328c5
title: "Desafío 102: Palabra más larga"
challengeType: 29
dashedName: challenge-102
---

# --description--

Dada una cadena de oración, devuelve la palabra más larga en la oración.

- Las palabras están separadas por un solo espacio.
- Solo las letras (`a-z`, sin distinguir mayúsculas o minúsculas) cuentan para la longitud de la palabra.
- Si hay varias palabras con la misma longitud, devuelve la primera que aparece.
- Devuelve la palabra tal como aparece en la cadena dada, sin puntuación.

# --hints--

`longest_word("The quick red fox")` debería devolver `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` debería devolver `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` debería devolver `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` debería devolver `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` debería devolver `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` debería devolver `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
