---
id: 68ffb91507a5b645769328cb
title: "Desafío 108: BuzzFizz"
challengeType: 29
dashedName: challenge-108
---

# --description--

Dado un arreglo, determina si es una secuencia correcta de FizzBuzz desde 1 hasta el último elemento del arreglo. Una secuencia es correcta si:

- Los números que son múltiplos de 3 se reemplazan con "Fizz"
- Los números que son múltiplos de 5 se reemplazan con "Buzz"
- Los números que son múltiplos de 3 y 5 se reemplazan con "FizzBuzz"
- Todos los demás números permanecen como enteros en orden ascendente, comenzando desde 1.
- El array debe comenzar en 1 y no tener elementos faltantes ni extras.

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4]), True)`)
}})
```

`is_fizz_buzz([1, 2, 3, 4])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, 3, 4]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", 7])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", 7]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]), False)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(sequence):

    return sequence
```

# --solutions--

```py
def is_fizz_buzz(sequence):
    for i, val in enumerate(sequence):
        n = i + 1
        if n % 15 == 0 and val != "FizzBuzz":
            return False
        if n % 3 == 0 and n % 5 != 0 and val != "Fizz":
            return False
        if n % 5 == 0 and n % 3 != 0 and val != "Buzz":
            return False
        if n % 3 != 0 and n % 5 != 0 and val != n:
            return False
    return True
```
