---
id: 69162d64f96574d9bb629eff
title: "Desafío 114: Camel a Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

Dada una cadena en camel case, devuelve la versión en snake case de la cadena usando las siguientes reglas:

- La cadena de entrada solo contendrá letras (`A-Z` y `a-z`) y siempre comenzará con una letra minúscula.
- Cada letra mayúscula en la cadena camel case inicia una nueva palabra.
- Convierte todas las letras a minúsculas.
- Separa las palabras con un guion bajo (`_`).

# --hints--

`to_snake("helloWorld")` debería devolver `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` debería devolver `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` debería devolver `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
