---
id: 69162d64f96574d9bb629efd
title: "Desafío 112: Detector de IA"
challengeType: 29
dashedName: challenge-112
---

# --description--

El desafío de hoy está inspirado en el lanzamiento de ChatGPT el 30 de noviembre de 2022.

Dada una cadena de una o más oraciones, determina si probablemente fue generada por IA usando las siguientes reglas:

- Contiene dos o más guiones (`-`).
- Contiene dos o más conjuntos de paréntesis (`()`). El texto puede estar dentro de los paréntesis.
- Contiene tres o más palabras con 7 o más letras.

- Las palabras están separadas por un solo espacio y solo consisten en letras (`A-Z`). No incluyas puntuación ni otros caracteres que no sean letras como parte de una palabra.

Si la oración dada cumple con alguna de las reglas anteriores, devuelve `"AI"`, de lo contrario, devuelve `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` debería devolver `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` debería devolver `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` debería devolver `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` debería devolver `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` debería devolver `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
