---
id: 69162d64f96574d9bb629efb
title: "Desafío 110: Adivinador de Palabras"
challengeType: 29
dashedName: challenge-110
---

# --description--

Dadas dos cadenas de la misma longitud, una palabra secreta y una suposición, compara la suposición con la palabra secreta usando las siguientes reglas:

- La palabra secreta y la conjetura solo consistirán en letras mayúsculas (`"A"` a `"Z"`);
- For each letter in the guess, replace it with a number according to how it matches the secret word:
  - "2" si la letra está en la palabra secreta y en la posición correcta.
  - "1" si la letra está en la palabra secreta pero en la posición incorrecta.
  - "0" si la letra no está en la palabra secreta.
- Cada letra en la palabra secreta puede usarse como máximo una vez.
- Las coincidencias exactas (`"2"`) se asignan primero, luego las coincidencias parciales (`"1"`) se asignan de izquierda a derecha para las letras restantes.
- Si una letra aparece varias veces en la suposición, solo puede coincidir tantas veces como aparece en la palabra secreta.

Por ejemplo, dado una palabra secreta de `"APPLE"` y una suposición de `"POPPA"`, devuelve `"10201"`:

La primera `"P"` no está en la ubicación correcta (`"1"`), la `"O"` no está en la palabra secreta (`"0"`), la segunda `"P"` está en la ubicación correcta (`"2"`), la tercera `"P"` es un cero (`"0"`) porque las dos `"P"` en la palabra secreta ya han sido usadas, y la `"A"` no está en la ubicación correcta (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` debería devolver `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` debería devolver `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` debería devolver `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` debería devolver `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` debería devolver `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` debería devolver `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
