---
id: 69162d64f96574d9bb629f01
title: "Desafío 116: Conteo de Permutaciones"
challengeType: 29
dashedName: challenge-116
---

# --description--

Dada una cadena, devuelve el número de permutaciones distintas que se pueden formar con sus caracteres.

- Una permutación es cualquier reordenamiento de los caracteres en la cadena.
- No cuentes permutaciones duplicadas.
- Si la cadena contiene caracteres repetidos, los arreglos repetidos solo deben contarse una vez.
- La cadena contendrá solo letras (`A-Z`, `a-z`).

Por ejemplo, dado `"abb"`, devuelve `3` porque hay tres formas únicas de estructurar las letras: `"abb"`, `"bab"` y `"bba"`.

# --hints--

`count_permutations("abb")` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` debería devolver `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` debería devolver `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
