---
id: 69162d64f96574d9bb629efe
title: "Desafío 113: Millas a Kilómetros"
challengeType: 29
dashedName: challenge-113
---

# --description--

Dada una distancia en millas como un número, devuelve la distancia equivalente en kilómetros.

- La entrada siempre será un número no negativo.
- 1 milla equivale a 1.60934 kilómetros.
- Redondea el resultado a dos decimales.
- Elimina los ceros finales innecesarios del resultado redondeado.

# --hints--

`convert_to_km(1)` debería devolver `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` debería devolver `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` debería devolver `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
