---
id: 69162d64f96574d9bb629f00
title: "Desafío 115: Convertidor de Ítems de Lista Ordenada en Markdown"
challengeType: 29
dashedName: challenge-115
---

# --description--

Dada una cadena que representa un elemento de lista ordenada en Markdown, devuelve la cadena HTML equivalente.

Un elemento válido de lista ordenada en Markdown debe:

- Comienza con cero o más espacios, seguido de
- Un número (1 o mayor) y un punto (`.`), seguido de
- Al menos un espacio, y luego
- El texto del elemento de la lista.

Si la cadena no tiene el formato exacto anterior, devuelve `"Invalid format"`. De lo contrario, envuelve el texto del elemento de la lista en etiquetas `li` y devuelve la cadena.

Por ejemplo, dado `"1. My item"`, devuelve `"<li>My item</li>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`convert_list_item("1. My item")` debería devolver `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` debería devolver `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` debería devolver `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` debería devolver `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
