---
id: 691b559495c5cb5a37b9b489
title: "Desafío 129: Analizador de Citas en Bloque de Markdown"
challengeType: 29
dashedName: challenge-129
---

# --description--

Dada una cadena que incluye un blockquote en Markdown, devuelve la cadena HTML equivalente.

Una cita en bloque en Markdown es cualquier línea que:

- Comienza con cero o más espacios
- >
- Luego, uno o más espacios
- Y finalmente, el texto de la cita en bloque.

Devuelve el texto de la cita rodeado por las etiquetas HTML de apertura y cierre `blockquote`.

Por ejemplo, dado `"> This is a quote"`, devuelve `<blockquote>This is a quote</blockquote>`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parse_blockquote("> This is a quote")` debería devolver `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` debería devolver `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` debería devolver `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
