---
id: 691b559495c5cb5a37b9b488
title: "Desafío 128: Conteo de Consonantes"
challengeType: 29
dashedName: challenge-128
---

# --description--

Dada una cadena y un número objetivo, determina si la cadena contiene exactamente el número objetivo de consonantes.

- Las consonantes son todos los caracteres alfabéticos excepto `"a"`, `"e"`, `"i"`, `"o"` y `"u"` en cualquier caso.
- Ignora dígitos, puntuación, espacios y otros caracteres que no sean letras al contar.

# --hints--

`has_consonant_count("helloworld", 7)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("helloworld", 7), True)`)
}})
```

`has_consonant_count("eieio", 5)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("eieio", 5), False)`)
}})
```

`has_consonant_count("freeCodeCamp Rocks!", 11)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("freeCodeCamp Rocks!", 11), True)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24), False)`)
}})
```

`has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_consonant_count("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_consonant_count(text, target):

    return text
```

# --solutions--

```py
def has_consonant_count(text, target):
    vowels = {"a", "e", "i", "o", "u"}
    count = 0

    for ch in text.lower():
        if ch.isalpha() and ch not in vowels:
            count += 1

    return count == target
```
