---
id: 691b559495c5cb5a37b9b487
title: "Desafío 127: Verificación de Velocidad"
challengeType: 29
dashedName: challenge-127
---

# --description--

Dada la velocidad a la que viajas en millas por hora (MPH), y un límite de velocidad en kilómetros por hora (KPH), determina si estás excediendo el límite y si recibirás una advertencia o una multa.

- 1 milla equivale a 1.60934 kilómetros.
- Si viajas a una velocidad menor o igual al límite, devuelve `"Not Speeding"`.
- Si viajas a 5 KPH o menos sobre el límite de velocidad, devuelve `"Warning"`.
- Si viajas más de 5 KPH sobre el límite de velocidad, devuelve `"Ticket"`.

# --hints--

`speed_check(30, 70)` debería devolver `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` debería devolver `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` debería devolver `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` debería devolver `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` debería devolver `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` debería devolver `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
