---
id: 691b559495c5cb5a37b9b483
title: "Desafío 123: Constructor de Números Romanos"
challengeType: 29
dashedName: challenge-123
---

# --description--

Dado un entero, devuelve su valor equivalente en números romanos.

Los números romanos usan estos símbolos:

SímboloValor
| ------ | ----- |
I1
V5
X10
L50
C100
D500
M1000

Los números romanos se escriben de mayor a menor, de izquierda a derecha, usando las siguientes reglas:

- La adición se usa cuando un símbolo es seguido por otro de igual o menor valor. Por ejemplo, `18` se escribe como `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- La resta se usa cuando un símbolo menor aparece antes de uno mayor, para representar 4 o 9 en cualquier valor posicional. Por ejemplo, 19 se escribe como `XIX` (10 + (10 - 1)).
- Ningún símbolo puede repetirse más de tres veces seguidas. La sustracción se usa cuando de otro modo necesitarías escribir un símbolo más de tres veces seguidas. Así que el número más grande que puedes escribir es 3999.

Aquí tienes un ejemplo más: dado `1464`, devuelve `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` debería devolver `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` debería devolver `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` debería devolver `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` debería devolver `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` debería devolver `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
