---
id: 691b559495c5cb5a37b9b482
title: "Desafío 122: Analizador de Negritas en Markdown"
challengeType: 29
dashedName: challenge-122
---

# --description--

Dada una cadena que puede incluir texto en **negrita** en Markdown, devuelve la cadena HTML equivalente.

- El texto en negrita en Markdown es cualquier texto que comienza y termina con dos asteriscos (`**`) o dos guiones bajos (`__`).
- No puede haber espacios entre el texto y los asteriscos o guiones bajos, pero sí puede haber espacios dentro del texto mismo.
- Convierte todas las ocurrencias en **negrita** a etiquetas HTML `b` y devuelve la cadena.

Por ejemplo, dado `"**This is bold**"`, devuelve `"<b>This is bold</b>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parse_bold("**This is bold**")` debería devolver `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` debería devolver `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` debería devolver `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` debería devolver `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` debería devolver `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
