---
id: 691b559495c5cb5a37b9b480
title: "Desafío 120: Libras a Kilogramos"
challengeType: 29
dashedName: challenge-120
---

# --description--

Dado un peso en libras como número, devuelve la cadena `"(lbs) pounds equals (kgs) kilograms."`.

- Reemplaza `"(lbs)"` con el número de entrada.
- Reemplaza `"(kgs)"` con la entrada convertida a kilogramos, redondeada a dos decimales y siempre incluye dos decimales en el valor.
- 1 libra equivale a 0.453592 kilogramos.
- Si la entrada es `1`, usa `"pound"` en lugar de `"pounds"`.
- Si el valor convertido es `1`, usa `"kilogram"` en lugar de `"kilograms"`.

# --hints--

`convert_to_kgs(1)` debería devolver `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` debería devolver `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` debería devolver `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` debería devolver `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` debería devolver `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
