---
id: 691b559495c5cb5a37b9b485
title: "Desafío 125: Game of Life"
challengeType: 29
dashedName: challenge-125
---

# --description--

Dada una matriz (arreglo de arreglos) que representa el estado actual en el Juego de la Vida de Conway, devuelve el siguiente estado de la matriz usando estas reglas:

- Cada celda es `1` (viva) o `0` (muerta).
- Los vecinos de una celda son hasta ocho celdas circundantes (vertical, horizontal y diagonalmente).
- Las celdas en los bordes tienen menos de ocho vecinos.

Reglas para actualizar cada celda:

- Cualquier célula viva con menos de dos vecinos vivos muere (subpoblación).
- Cualquier célula viva con dos o tres vecinos vivos continúa viviendo.
- Cualquier célula viva con más de tres vecinos vivos muere (sobrepoblación).
- Cualquier célula muerta con exactamente tres vecinos vivos se vuelve viva (reproducción).

Por ejemplo, dado:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

retornar:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Cada celda se actualiza según el número de vecinos vivos. Por ejemplo, `[0][0]` permanece muerta (2 vecinos vivos), `[0][1]` permanece viva (2 vecinos vivos), `[0][2]` muere (3 vecinos vivos), y así sucesivamente.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` debería devolver `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` debería devolver `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` debería devolver `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` debería devolver `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
