---
id: 691b559495c5cb5a37b9b486
title: "Desafío 126: Pon en mayúscula la primera letra"
challengeType: 29
dashedName: challenge-126
---

# --description--

Dada una cadena title, devuelve una nueva cadena formateada en mayúsculas y minúsculas siguiendo las siguientes reglas:

- Pon en mayúscula la primera letra de cada palabra.
- Haz que todas las demás letras de cada palabra estén en minúsculas.
- Las palabras siempre están separadas por un solo espacio.

# --hints--

`title_case("hello world")` debería devolver `"Hello World"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("hello world"), "Hello World")`)
}})
```

`title_case("the quick brown fox")` debería devolver `"The Quick Brown Fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("the quick brown fox"), "The Quick Brown Fox")`)
}})
```

`title_case("JAVASCRIPT AND PYTHON")` debería devolver `"Javascript And Python"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("JAVASCRIPT AND PYTHON"), "Javascript And Python")`)
}})
```

`title_case("AvOcAdO tOAst fOr brEAkfAst")` debería devolver `"Avocado Toast For Breakfast"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(title_case("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast")`)
}})
```

# --seed--

## --seed-contents--

```py
def title_case(title):

    return title
```

# --solutions--

```py
def title_case(title):
    return " ".join(
        w[:1].upper() + w[1:].lower()
        for w in title.split(" ")
    )
```
