---
id: 6925e2068081f40f549ced1c
title: "Desafío 138: Suma de Divisores"
challengeType: 29
dashedName: challenge-138
---

# --description--

Dado un entero positivo, devuelve la suma de todos sus divisores.

- Un divisor es cualquier entero que divide el número de manera exacta (el residuo es `0`).
- Cuenta cada divisor solo una vez.

Por ejemplo, dado `6`, devuelve `12` porque los divisores de `6` son `1`, `2`, `3` y `6`, y la suma de esos es `12`.

# --hints--

`sum_divisors(6)` debería devolver `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` debería devolver `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` debería devolver `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` debería devolver `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` debería devolver `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` debería devolver `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
