---
id: 6925e2068081f40f549ced1a
title: "Desafío 136: Analizador de Imágenes en Markdown"
challengeType: 29
dashedName: challenge-136
---

# --description--

Dada una cadena de una imagen en Markdown, devuelve la cadena HTML equivalente.

Una imagen en Markdown tiene el siguiente formato: `"![alt text](image_url)"`. Donde:

- `alt text` es la descripción de la imagen (el valor del atributo `alt`).
- `image_url` es la URL de origen de la imagen (el valor del atributo `src`).

Devuelve una cadena con la etiqueta HTML `img` con el `src` establecido en la URL de la imagen y el `alt` establecido en el texto alternativo.

Por ejemplo, dado `"![Cute cat](cat.png)"` devuelve `'<img src="cat.png" alt="Cute cat">'`;

- Asegúrate de que la etiqueta, el orden de los atributos, el espaciado y el uso de comillas sean iguales a los anteriores.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes — revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parse_image("![Cute cat](cat.png)")` debería devolver `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` debería devolver `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` debería devolver `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
