---
id: 6925e2068081f40f549ced1d
title: "Desafío 139: Piedra, Papel o Tijeras"
challengeType: 29
dashedName: challenge-139
---

# --description--

Dadas dos cadenas, la primera representando a Player 1 y la segunda representando a Player 2, determina el ganador de un partido de Rock, Paper, Scissors.

- Las cadenas de entrada siempre serán `"Rock"`, `"Paper"` o `"Scissors"`.
- "Rock" vence a "Scissors".
- `"Paper"` vence a `"Rock"`.
- `"Scissors"` vence a `"Paper"`.

Devuelve:

- "Player 1 wins" si Player 1 gana.
- "Player 2 wins" si Player 2 gana.
- "Empate" si ambos jugadores eligen la misma opción.

# --hints--

`rock_paper_scissors("Rock", "Rock")` debería devolver `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Rock"), "Tie")`)
}})
```

`rock_paper_scissors("Rock", "Paper")` debería devolver `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Paper"), "Player 2 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Paper")` debería devolver `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Paper"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Rock", "Scissors")` debería devolver `"Player 1 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Rock", "Scissors"), "Player 1 wins")`)
}})
```

`rock_paper_scissors("Scissors", "Scissors")` debería devolver `"Tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Scissors"), "Tie")`)
}})
```

`rock_paper_scissors("Scissors", "Rock")` debería devolver `"Player 2 wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rock_paper_scissors("Scissors", "Rock"), "Player 2 wins")`)
}})
```

# --seed--

## --seed-contents--

```py
def rock_paper_scissors(player1, player2):

    return player1
```

# --solutions--

```py
def rock_paper_scissors(player1, player2):
    if player1 == player2:
        return "Tie"
    
    if (player1 == "Rock" and player2 == "Scissors") or \
       (player1 == "Paper" and player2 == "Rock") or \
       (player1 == "Scissors" and player2 == "Paper"):
        return "Player 1 wins"
    else:
        return "Player 2 wins"
```
