---
id: 6925e2068081f40f549ced1b
title: "Desafío 137: Generador de Snowflake"
challengeType: 29
dashedName: challenge-137
---

# --description--

Dada una cadena de texto multilínea que usa caracteres de nueva línea (`\n`) para representar un salto de línea, devuelve una nueva cadena donde cada línea está reflejada horizontalmente y se adjunta al final de la línea original.

- Refleja una línea invirtiendo todos sus caracteres, incluidos los espacios.

Por ejemplo, dado `"* \n *\n* "`, que se registra en la consola como:

```sh
* 
 *
* 
```

Devuelve `"*  *\n ** \n*  *"`, que se muestra en la consola como:

```sh
*  *
 ** 
*  *
```

Toma nota cuidadosa de los espacios en blanco en las cadenas dadas y devueltas. Asegúrate de no eliminar ninguno de ellos.

# --hints--

`generate_snowflake("* \n *\n* ")` debería devolver `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` debería devolver `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` debería devolver `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` debería devolver `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` debería devolver `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
