---
id: 69272dcf1c24b44fd79137c4
title: "Desafío 141: Combustible para el despegue"
challengeType: 29
dashedName: challenge-141
---

# --description--

Dado el número de galones de combustible que tienes actualmente en tu avión, y el número requerido de litros de combustible para llegar a tu destino, determina cuántos galones adicionales de combustible debes agregar.

- 1 galón equivale a 3.78541 litros.
- Si el avión ya tiene suficiente combustible, devuelve `0`.
- Solo puedes agregar galones completos.
- No incluyas decimales en el número de retorno.

# --hints--

`fuel_to_add(0, 1)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` debería devolver `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` debería devolver `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
