---
id: 69272dcf1c24b44fd79137c5
title: "Desafío 142: Suma la Cadena"
challengeType: 29
dashedName: challenge-142
---

# --description--

Dada una cadena que contiene dígitos y otros caracteres, devuelve la suma de todos los números en la cadena.

- Trata los dígitos consecutivos como un solo número. Por ejemplo, `"13"` cuenta como 13, no como 1 + 3.
- Ignora cualquier carácter que no sea un dígito.

# --hints--

`string_sum("3apples2bananas")` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` debería devolver `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` debería devolver `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` debería devolver `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` debería devolver `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
