---
id: 69272dcf1c24b44fd79137c3
title: "Desafío 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

Dada una cadena que representa un nombre de variable, devuelve el nombre de variable convertido a SCREAMING_SNAKE_CASE.

Los nombres de variables dados se escribirán en uno de los siguientes formatos:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

En los formatos anteriores, las palabras están separadas por un guion bajo (`_`), un guion (`-`), o una nueva palabra comienza con una letra mayúscula.

Para convertir a SCREAMING_SNAKE_CASE:

- Haz que todas las letras estén en mayúsculas
- Separa las palabras con un guion bajo (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` debería devolver `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` debería devolver `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` debería devolver `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` debería devolver `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` debería devolver `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
