---
id: 69306364df283fcaff2e1ad7
title: "Desafío 146: Asiento para zurdos en la mesa"
challengeType: 29
dashedName: challenge-146
---

# --description--

Dada una matriz 4x2 (arreglo de arreglos) que representa la disposición de los asientos para una comida, determina en cuántos asientos puede sentarse una persona zurda.

- Una persona zurda no puede sentarse donde una persona diestra estaría en el asiento inmediatamente a su izquierda.

En la matriz dada:

- Un `"R"` es un asiento ocupado por una persona diestra.
- Una `"L"` es un asiento ocupado por una persona zurda.
- Una `"U"` es un asiento desocupado.
- Solo los asientos desocupados están disponibles para sentarse.
- Los asientos en la fila superior están orientados hacia "abajo", y los asientos en la fila inferior están orientados hacia "arriba" (como una mesa), por lo que izquierda y derecha son relativos a la orientación del asiento.
- Los asientos de esquina solo tienen un asiento junto a ellos.

Por ejemplo, en la matriz dada:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

El asiento superior izquierdo no se puede ocupar porque hay una persona diestra a la izquierda. Los otros dos asientos libres se pueden ocupar porque no hay una persona diestra a la izquierda.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` debería devolver `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
