---
id: 69306364df283fcaff2e1ad5
title: "Desafío 144: Racha de Resoluciones"
challengeType: 29
dashedName: challenge-144
---

# --description--

Dado un arreglo de arreglos, donde cada subarreglo representa un día de tus actividades de resolución y contiene tres enteros: el número de pasos caminados ese día, los minutos de tiempo frente a la pantalla ese día y el número de páginas leídas ese día; determina si estás cumpliendo tus resoluciones.

- El primer subarreglo es el día 1, el segundo el día 2, y así sucesivamente.

Un día se considera exitoso si se cumplen los tres siguientes objetivos:

- Caminaste al menos 10,000 pasos.
- No tuviste más de 120 minutos de tiempo frente a la pantalla.
- Lees al menos cinco páginas.

Si todos los días dados son exitosos, devuelve `"Resolution on track: N day streak."` Donde `N` es el número de días exitosos.

Si uno o más días no son un éxito, devuelve `"Resolution failed on day X: N day streak."`. Donde `X` es el número del día del primer día no exitoso, y `N` es el número de días exitosos antes del primer día no exitoso.

# --hints--

`resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` debería devolver `"Resolución en curso: racha de 3 días."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.")`)
}})
```

`resolution_streak([[10000, 120, 5], [10950, 121, 11]])` debería devolver `"La resolución falló en el día 2: racha de 1 día."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.")`)
}})
```

`resolution_streak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` debería devolver `"La resolución falló en el día 3: racha de 2 días."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.")`)
}})
```

`resolution_streak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` debería devolver `"Resolución fallida en el día 4: racha de 3 días."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.")`)
}})
```

`resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` debería devolver `"Resolución en curso: racha de 10 días."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.")`)
}})
```

# --seed--

## --seed-contents--

```py
def resolution_streak(days):

    return days
```

# --solutions--

```py
def resolution_streak(days):
    streak = 0

    for i, day in enumerate(days):
        steps, screen, pages = day
        if steps >= 10000 and screen <= 120 and pages >= 5:
            streak += 1
        else:
            return f"Resolution failed on day {i+1}: {streak} day streak."

    return f"Resolution on track: {streak} day streak."
```
