---
id: 69306364df283fcaff2e1ad9
title: "Desafío 148: Presión de Neumáticos"
challengeType: 29
dashedName: challenge-148
---

# --description--

Dado un arreglo con cuatro números que representan las presiones de los neumáticos en psi de los cuatro neumáticos de tu vehículo, y otro arreglo de dos números que representan la presión mínima y máxima para tus neumáticos en bar, devuelve un arreglo de cuatro cadenas que describen el estado de cada neumático.

- 1 bar equivale a 14.5038 psi.

Devuelve un arreglo con los siguientes valores para cada neumático:

- "Low" si la presión del neumático está por debajo del mínimo permitido.
- "Bueno" si está entre el mínimo y el máximo permitido.
- "Alto" si está por encima del máximo permitido.

# --hints--

`tire_status([32, 28, 35, 29], [2, 3])` debería devolver `["Good", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"])`)
}})
```

`tire_status([32, 28, 35, 30], [2, 2.3])` debería devolver `["Good", "Low", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"])`)
}})
```

`tire_status([29, 26, 31, 28], [2.1, 2.5])` debería devolver `["Low", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"])`)
}})
```

`tire_status([31, 31, 30, 29], [1.5, 2])` debería devolver `["High", "High", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"])`)
}})
```

`tire_status([30, 28, 30, 29], [1.9, 2.1])` debería devolver `["Good", "Good", "Good", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"])`)
}})
```

# --seed--

## --seed-contents--

```py
def tire_status(pressures_psi, range_bar):

    return pressures_psi
```

# --solutions--

```py
def tire_status(pressures_psi, range_bar):
    psi_to_bar = 1 / 14.5038
    min_bar, max_bar = range_bar
    result = []

    for psi in pressures_psi:
        bar = psi * psi_to_bar
        if bar < min_bar:
            result.append("Low")
        elif bar > max_bar:
            result.append("High")
        else:
            result.append("Good")
    return result
```
