---
id: 69306364df283fcaff2e1ada
title: "Desafío 149: vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

Dada una cadena, devuelve una nueva cadena donde todas las vocales se convierten en mayúsculas y todos los demás caracteres alfabéticos se convierten en minúsculas.

- Las vocales son `"a"`, `"e"`, `"i"`, `"o"` y `"u"` en cualquier caso.
- Los caracteres no alfabéticos deben permanecer sin cambios.

# --hints--

`vowel_case("vowelcase")` debería devolver `"vOwElcAsE"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` debería devolver `"cOdIng Is fUn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` debería devolver `"hEllO, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` debería devolver `"gIt chErry-pIck"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` debería devolver `"hEAd~1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
