---
id: 69373793f5a867f769cde139
title: "Desafío 154: Par para el hoyo"
challengeType: 29
dashedName: challenge-154
---

# --description--

Dado dos enteros, el par para un hoyo de golf y el número de golpes que un golfista hizo en ese hoyo, devuelve la puntuación del golfista usando términos de golf.

Devuelve:

- "¡Hoyo en uno!" si tomó un solo golpe.
- "Eagle" si tomó dos golpes menos que el par.
- "Birdie" si tomó un golpe menos que el par.
- "Par" si tomó el mismo número de golpes que par.
- "Bogey" si tomó un golpe más que el par.
- "Double bogey" si tomaste dos golpes más que el par.

# --hints--

`golf_score(3, 3)` debería devolver `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` debería devolver `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` debería devolver `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` debería devolver `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` debería devolver `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` debería devolver `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
