---
id: 69373793f5a867f769cde13a
title: "Desafío 155: Planta la Cosecha"
challengeType: 29
dashedName: challenge-155
---

# --description--

Dado un entero que representa el tamaño de tu campo de cultivo, y `"acres"` o `"hectares"` que representan la unidad para el tamaño de tu campo, y un tipo de cultivo, determina cuántas plantas de ese tipo puedes colocar en tu campo.

- 1 acre equivale a 4046.86 metros cuadrados.
- 1 hectárea equivale a 10,000 metros cuadrados.

Aquí hay una lista de cultivos que se darán como entrada y cuánto espacio ocupa una sola planta:

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

Devuelve el número de plantas que caben en el campo, redondeado hacia abajo al número entero más cercano.

# --hints--

`get_number_of_plants(1, "acres", "corn")` debería devolver `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` debería devolver `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` debería devolver `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` debería devolver `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` debería devolver `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
