---
id: 69373793f5a867f769cde138
title: "Desafío 153: Tic-Tac-Toe"
challengeType: 29
dashedName: challenge-153
---

# --description--

Dada una matriz 3×3 (un arreglo de arreglos) que representa un juego de Tic-Tac-Toe completado, determina el ganador.

- Cada elemento en la matriz dada es un `"X"` o un `"O"`.

Un jugador gana si tiene tres de sus personajes en fila - horizontalmente, verticalmente o diagonalmente.

Devuelve:

- "X gana" si el jugador X tiene tres en línea.
- `"O wins"` si el jugador O tiene tres en línea.
- "Empate" si ningún jugador tiene tres en línea.

# --hints--

`tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` debería devolver `"X wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins")`)
}})
```

`tic_tac_toe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]])` debería devolver `"O wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]]), "O wins")`)
}})
```

`tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` debería devolver `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw")`)
}})
```

`tic_tac_toe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]])` debería devolver `"O wins"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]]), "O wins")`)
}})
```

`tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` debería devolver `"X gana"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins")`)
}})
```

`tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` debería devolver `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tic_tac_toe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def tic_tac_toe(board):

    return board
```

# --solutions--

```py
def tic_tac_toe(board):
    lines = []

    for row in board:
        lines.append(row)

    for c in range(3):
        lines.append([board[0][c], board[1][c], board[2][c]])

    lines.append([board[0][0], board[1][1], board[2][2]])
    lines.append([board[0][2], board[1][1], board[2][0]])

    for line in lines:
        if all(cell == "X" for cell in line):
            return "X wins"
        if all(cell == "O" for cell in line):
            return "O wins"

    return "Draw"
```
