---
id: 69373793f5a867f769cde137
title: "Desafío 152: Primo Circular"
challengeType: 29
dashedName: challenge-152
---

# --description--

Dado un entero, determina si es un primo circular.

Un primo circular es un entero donde todas las rotaciones de sus dígitos son primos.

Por ejemplo, `197` es un primo circular porque todas las rotaciones de sus dígitos: `197`, `971` y `719`, son números primos.

# --hints--

`is_circular_prime(197)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(197), True)`)
}})
```

`is_circular_prime(23)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(23), False)`)
}})
```

`is_circular_prime(13)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(13), True)`)
}})
```

`is_circular_prime(89)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(89), False)`)
}})
```

`is_circular_prime(1193)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_circular_prime(1193), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_circular_prime(n):

    return n
```

# --solutions--

```py
import math
def is_prime(n):
    if n < 2:
        return False
    for i in range(2, int(math.isqrt(n)) + 1):
        if n % i == 0:
            return False
    return True

def rotations(n):
    s = str(n)
    return [int(s[i:] + s[:i]) for i in range(len(s))]

def is_circular_prime(n):
    return all(is_prime(rot) for rot in rotations(n))
```
