---
id: 6939b873185d8e00d453563e
title: "Desafío 159: Hipotenusa Entera"
challengeType: 29
dashedName: challenge-159
---

# --description--

Dado dos enteros positivos que representan las longitudes de las dos piernas (los dos lados cortos) de un triángulo rectángulo, determina si la hipotenusa es un entero.

La longitud de la hipotenusa se calcula sumando los cuadrados de las longitudes de los dos catetos y luego tomando la raíz cuadrada de ese total (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`is_integer_hypotenuse(3, 4)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(3, 4), True)`)
}})
```

`is_integer_hypotenuse(2, 3)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(2, 3), False)`)
}})
```

`is_integer_hypotenuse(5, 12)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(5, 12), True)`)
}})
```

`is_integer_hypotenuse(10, 10)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(10, 10), False)`)
}})
```

`is_integer_hypotenuse(780, 1040)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(780, 1040), True)`)
}})
```

`is_integer_hypotenuse(250, 333)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_integer_hypotenuse(250, 333), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_integer_hypotenuse(a, b):

    return a
```

# --solutions--

```py
import math
def is_integer_hypotenuse(a, b):
    total = a*a + b*b
    c = math.isqrt(total)
    return c*c == total
```
