---
id: 6939b873185d8e00d453563d
title: "Desafío 158: Intercambio de Array"
challengeType: 29
dashedName: challenge-158
---

# --description--

Dado un arreglo con dos valores, devuelve un arreglo con los valores intercambiados.

Por ejemplo, dado `["A", "B"]` devuelve `["B", "A"]`.

# --hints--

`array_swap(["A", "B"])` debería devolver `["B", "A"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` debería devolver `[20, 25]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` debería devolver `[False, True]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` debería devolver `[1, "1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
