---
id: 6939b873185d8e00d453563c
title: "Desafío 157: Analizador de enlaces Markdown"
challengeType: 29
dashedName: challenge-157
---

# --description--

Dada la cadena de un enlace en Markdown, devuelve la cadena HTML equivalente.

Una imagen en Markdown tiene el siguiente formato: `"[link_text](link_url)"`. Devuelve la cadena de la etiqueta HTML `a` con el `href` establecido en el `link_url` y el `link_text` como contenido de la etiqueta.

Por ejemplo, dado `"[freeCodeCamp](https://freecodecamp.org/)"` devuelve `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes — revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` debería devolver `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` debería devolver `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribuye a nuestro repositorio en https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` debería devolver `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribuye a nuestro repositorio en https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
