---
id: 6939b873185d8e00d453563b
title: "Desafío 156: ¿Impar o Par?"
challengeType: 29
dashedName: challenge-156
---

# --description--

Dado un entero positivo, devuelve `"Odd"` si es un número impar, y `"Even"` si es par.

# --hints--

`odd_or_even(1)` debería devolver `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(1), "Odd")`)
}})
```

`odd_or_even(2)` debería devolver `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(2), "Even")`)
}})
```

`odd_or_even(13)` debería devolver `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(13), "Odd")`)
}})
```

`odd_or_even(196)` debería devolver `"Even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(196), "Even")`)
}})
```

`odd_or_even(123456789)` debería devolver `"Odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even(123456789), "Odd")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even(n):

    return n
```

# --solutions--

```py
def odd_or_even(n):
    return "Even" if n % 2 == 0 else "Odd"
```
