---
id: 6939b873185d8e00d453563f
title: "Desafío 160: Movimientos del Caballero"
challengeType: 29
dashedName: challenge-160
---

# --description--

Dada la posición de un caballo en un tablero de ajedrez, devuelve el número de casillas válidas a las que el caballo puede moverse.

Un tablero de ajedrez estándar es de 8x8, con columnas etiquetadas de `A` a `H` (de izquierda a derecha) y filas etiquetadas del `1` al `8` (de abajo hacia arriba). Se ve así:

**A8****B8****C8****D8****E8****F8****G8****H8**
|-|-|-|-|-|-|-|-|
**A7****B7****C7****D7****E7****F7****G7****H7**
**A6****B6****C6****D6****E6****F6****G6****H6**
**A5****B5****C5****D5****E5****F5****G5****H5**
**A4****B4****C4****D4****E4****F4****G4****H4**
**A3****B3****C3****D3****E3****F3****G3****H3**
**A2****B2****C2****D2****E2****F2****G2****H2**
**A1****B1****C1****D1****E1****F1****G1****H1**

Un caballo se mueve en forma de "L": dos casillas en una dirección (horizontal o vertical), y una casilla en la dirección perpendicular.

Esto significa que un caballo puede moverse a hasta ocho posiciones posibles, pero menos cuando está cerca de los bordes del tablero. Por ejemplo, si un caballo estuviera en `A1`, solo podría moverse a `B3` o `C2`.

# --hints--

`knight_moves("A1")` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("A1"), 2)`)
}})
```

`knight_moves("D4")` debería devolver `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("D4"), 8)`)
}})
```

`knight_moves("G6")` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("G6"), 6)`)
}})
```

`knight_moves("B8")` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("B8"), 3)`)
}})
```

`knight_moves("H3")` debería devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("H3"), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def knight_moves(position):

    return position
```

# --solutions--

```py
def knight_moves(position):
    col = ord(position[0]) - ord('A')
    row = int(position[1]) - 1

    moves = [
        (2, 1), (2, -1),
        (-2, 1), (-2, -1),
        (1, 2), (1, -2),
        (-1, 2), (-1, -2)
    ]

    valid = 0

    for dx, dy in moves:
        new_col = col + dx
        new_row = row + dy

        if 0 <= new_col < 8 and 0 <= new_row < 8:
            valid += 1

    return valid
```
