---
id: 694596b0585c11170ac7c7f9
title: "Desafío 161: Envío Gratis"
challengeType: 29
dashedName: challenge-161
---

# --description--

Dado un arreglo de cadenas que representan los artículos en tu carrito de compras, y un número para la cantidad mínima de pedido para calificar para envío gratis, determina si los artículos en tu carrito de compras califican para envío gratis.

El array dado contendrá elementos de la lista siguiente:

| Item | Price |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`gets_free_shipping(["shoes"], 50)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shoes"], 50), True)`)
}})
```

`gets_free_shipping(["hat", "socks"], 50)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks"], 50), False)`)
}})
```

`gets_free_shipping(["jeans", "shirt", "jacket"], 75)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["jeans", "shirt", "jacket"], 75), True)`)
}})
```

`gets_free_shipping(["socks", "socks", "hat"], 75)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["socks", "socks", "hat"], 75), False)`)
}})
```

`gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100)` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100), True)`)
}})
```

`gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def gets_free_shipping(cart, minimum):

    return cart
```

# --solutions--

```py
def gets_free_shipping(cart, minimum):
    prices = {
        "shirt": 34.25,
        "jeans": 48.50,
        "shoes": 75.00,
        "hat": 19.95,
        "socks": 15.00,
        "jacket": 109.95
    }

    total = 0

    for item in cart:
        total += prices[item]

    return total >= minimum
```
