---
id: 694596b0585c11170ac7c7fc
title: "Desafío 164: Analizador de Código Inline en Markdown"
challengeType: 29
dashedName: challenge-164
---

# --description--

Dada una cadena de Markdown que incluye uno o más bloques de código en línea, devuelve la cadena HTML equivalente.

Los bloques de código en línea en Markdown usan una sola comilla invertida (```) al inicio y al final del texto del bloque de código.

Devuelve la cadena dada con todos los bloques de código convertidos a etiquetas HTML `code`.

Por ejemplo, dada la cadena `"Use `let` to declare the variable."`, devuelve `"Use <code>let</code> to declare the variable."`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parse_inline_code("Usa `let` para declarar la variable.")` debería devolver `"Usa <code>let</code> para declarar la variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` to declare the variable."), "Use <code>let</code> to declare the variable.")`)
}})
```

`parse_inline_code("Usa `let` o `const` para declarar una variable.")` debería devolver `"Usa <code>let</code> o <code>const</code> para declarar una variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

`parse_inline_code("Run `npm install` then `npm start`.")` debería devolver `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Run \`npm install\` then \`npm start\`."), "Run <code>npm install</code> then <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```
