---
id: 694596b0585c11170ac7c7fa
title: "Desafío 162: Consumo de Energía"
challengeType: 29
dashedName: challenge-162
---

# --description--

Dado el número de Calorías quemadas durante un entrenamiento y el número de vatios-hora usados por tus dispositivos electrónicos durante ese entrenamiento, determina cuál usó más energía.

Para compararlos, convierte ambos valores a julios usando las siguientes conversiones:

- 1 Caloría equivale a 4184 julios.
- 1 vatio-hora equivale a 3600 julios.

Devuelve:

- "Entrenamiento" si el entrenamiento usó más energía.
- "Dispositivos" si el dispositivo usó más energía.
- "Igual" si ambos usaron la misma cantidad de energía.

# --hints--

`compare_energy(250, 50)` debería devolver `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` debería devolver `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` debería devolver `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` debería devolver `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` debería devolver `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` debería devolver `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
