---
id: 696655d24b614176d4c9b78b
title: "Desafío 168: Imagen escalada"
challengeType: 29
dashedName: challenge-168
---

# --description--

Dada una cadena que representa el ancho y alto de una imagen, y un número para escalar la imagen, devuelve el ancho y alto escalados.

- La cadena de entrada está en el formato `"WxH"`. Por ejemplo, `"800x600"`.
- La escala es un número para multiplicar el ancho y el alto.

Devuelve las dimensiones escaladas en el mismo formato `"WxH"`.

# --hints--

`scale_image("800x600", 2)` debería devolver `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` debería devolver `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` debería devolver `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` debería devolver `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
