---
id: 696655d24b614176d4c9b78c
title: "Desafío 169: FizzBuzz Mini"
challengeType: 29
dashedName: challenge-169
---

# --description--

Dado un entero, devuelve una cadena basada en las siguientes reglas:

- Si el número es divisible por 3, devuelve `"Fizz"`.
- Si el número es divisible por 5, devuelve `"Buzz"`.
- Si el número es divisible tanto por 3 como por 5, devuelve `"FizzBuzz"`.
- De lo contrario, devuelve el número dado como una cadena.

# --hints--

`fizz_buzz_mini(3)` debería devolver `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` debería devolver `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` debería devolver `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` debería devolver `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` debería devolver `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
