---
id: 696655d24b614176d4c9b78d
title: "Desafío 170: Día impar o par"
challengeType: 29
dashedName: challenge-170
---

# --description--

Dado un timestamp (número de milisegundos desde la época Unix), devuelve:

- "odd" si el día del mes para ese timestamp es impar.
- "even" si el día del mes para ese timestamp es par.

Por ejemplo, dado `1769472000000`, una marca de tiempo para el 27 de enero de 2026, devuelve `"odd"` porque el número del día (`27`) es un número impar.

# --hints--

`odd_or_even_day(1769472000000)` debería devolver `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` debería devolver `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` debería devolver `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` debería devolver `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` debería devolver `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
