---
id: 696655d24b614176d4c9b789
title: "Desafío 166: Validador Hexadecimal"
challengeType: 29
dashedName: challenge-166
---

# --description--

Dada una cadena, determina si es un color hexadecimal CSS válido. Un color hexadecimal CSS válido debe:

- Comienza con un `#`, y
- deben ser seguidos por 3 o 6 caracteres hexadecimales.

Los caracteres hexadecimales son números `0` a `9` y letras `a` a `f` (sin distinguir mayúsculas de minúsculas).

# --hints--

`is_valid_hex("#123")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123"), True)`)
}})
```

`is_valid_hex("#123abc")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#123abc"), True)`)
}})
```

`is_valid_hex("#ABCDEF")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#ABCDEF"), True)`)
}})
```

`is_valid_hex("#0a1B2c")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#0a1B2c"), True)`)
}})
```

`is_valid_hex("#12G")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12G"), False)`)
}})
```

`is_valid_hex("#1234567")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#1234567"), False)`)
}})
```

`is_valid_hex("#12 3")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("#12 3"), False)`)
}})
```

`is_valid_hex("fff")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hex("fff"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hex(s):

    return s
```

# --solutions--

```py
def is_valid_hex(s):
    if not s.startswith("#"):
        return False

    hex_part = s[1:]
    if len(hex_part) not in (3, 6):
        return False

    for char in hex_part:
        if not char.lower() in "0123456789abcdef":
            return False

    return True
```
