---
id: 696655d24b614176d4c9b78a
title: "Desafío 167: ¡Bingo! Letra"
challengeType: 29
dashedName: challenge-167
---

# --description--

Dado un número, devuelve la letra de bingo asociada con él (en mayúsculas). Los números de bingo se agrupan de la siguiente manera:

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`get_bingo_letter(75)` debería devolver `"O"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(75), "O")`)
}})
```

`get_bingo_letter(54)` debería devolver `"G"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(54), "G")`)
}})
```

`get_bingo_letter(25)` debería devolver `"I"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(25), "I")`)
}})
```

`get_bingo_letter(38)` debería devolver `"N"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(38), "N")`)
}})
```

`get_bingo_letter(11)` debería devolver `"B"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_letter(11), "B")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_bingo_letter(n):

    return n
```

# --solutions--

```py
def get_bingo_letter(n):
    if 1 <= n <= 15:
        return "B"
    if 16 <= n <= 30:
        return "I"
    if 31 <= n <= 45:
        return "N"
    if 46 <= n <= 60:
        return "G"
    if 61 <= n <= 75:
        return "O"
    return None
```
