---
id: 69738771fb5a7b8b24cca29d
title: "Desafío 171: Aplanar el Array"
challengeType: 29
dashedName: challenge-171
---

# --description--

Dado un arreglo que contiene arreglos anidados, devuelve un nuevo arreglo con todos los valores aplanados en un solo arreglo unidimensional. Conserva el orden original de los elementos en los arreglos.

# --hints--

`flatten([1, [2, 3], 4])` debería devolver `[1, 2, 3, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4])`)
}})
```

`flatten([5, [4, [3, 2]], 1])` debería devolver `[5, 4, 3, 2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1])`)
}})
```

`flatten(["A", [[[["B"]]]], "C"])` debería devolver `["A", "B", "C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"])`)
}})
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` debería devolver `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"])`)
}})
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` debería devolver `["red","blue","green","yellow","purple","orange","pink","brown"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"])`)
}})
```

# --seed--

## --seed-contents--

```py
def flatten(arr):

    return arr
```

# --solutions--

```py
def flatten(arr):
    result = []
    for item in arr:
        if isinstance(item, list):
            result.extend(flatten(item))
        else:
            result.append(item)
    return result
```
