---
id: 69738771fb5a7b8b24cca2a3
title: "Desafío 177: Espejo de Cadenas"
challengeType: 29
dashedName: challenge-177
---

# --description--

Dada una cadena, devuelve una nueva cadena que consiste en la cadena dada con una copia invertida de sí misma añadida al final.

# --hints--

`mirror("freeCodeCamp")` debería devolver `"freeCodeCamppmaCedoCeerf"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf")`)
}})
```

`mirror("RaceCar")` debería devolver `"RaceCarraCecaR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("RaceCar"), "RaceCarraCecaR")`)
}})
```

`mirror("helloworld")` debería devolver `"helloworlddlrowolleh"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("helloworld"), "helloworlddlrowolleh")`)
}})
```

`mirror("The quick brown fox...")` debería devolver `"The quick brown fox......xof nworb kciuq ehT"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT")`)
}})
```

# --seed--

## --seed-contents--

```py
def mirror(s):

    return s
```

# --solutions--

```py
def mirror(s):
    reversed_s = s[::-1]
    return s + reversed_s
```
