---
id: 69738771fb5a7b8b24cca2a2
title: "Desafío 176: Día de la Marmota"
challengeType: 29
dashedName: challenge-176
---

# --description--

Hoy es el Día de la Marmota, en el que una marmota predice el clima según si ve o no su sombra.

Dado un valor que representa la apariencia de la marmota, devuelve la predicción correcta:

- Si el valor dado es el booleano `true` (la marmota vio su sombra), devuelve `"Parece que tendremos seis semanas más de invierno."`.
- Si el valor es el booleano `false` (la marmota no vio su sombra), devuelve `"It's going to be an early spring."`.
- Si el valor es cualquier otro (la marmota no apareció), devuelve `"No prediction this year."`.

# --hints--

`groundhog_day_prediction(True)` debería devolver `"Parece que tendremos seis semanas más de invierno."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` debería devolver `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` debería devolver `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` debería devolver `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` debería devolver `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
