---
id: 69738771fb5a7b8b24cca2a4
title: "Desafío 178: Trunca el Texto"
challengeType: 29
dashedName: challenge-178
---

# --description--

Dada una cadena, devuélvela tal cual si tiene 20 caracteres o menos. Si es más larga que 20 caracteres, trúncala a los primeros 17 caracteres y añade `"..."` al final (para que tenga un total de 20 caracteres) y devuelve el resultado.

# --hints--

`truncate_text("Hello, world!")` debería devolver `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` debería devolver `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` debería devolver `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` debería devolver `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
