---
id: 69738771fb5a7b8b24cca29f
title: "Desafío 173: Movimientos válidos de peón"
challengeType: 29
dashedName: challenge-173
---

# --description--

Dada la posición de uno de tus peones en un tablero de ajedrez, devuelve un arreglo con todas las casillas válidas a las que puede moverse en orden ascendente.

Un tablero de ajedrez estándar es de 8x8, con columnas etiquetadas de `A` a `H` (de izquierda a derecha) y filas etiquetadas del `1` al `8` (de abajo hacia arriba). Se ve así:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Para este desafío:

- Eres el jugador en la parte inferior del tablero.
- Los peones solo pueden moverse un cuadro "hacia arriba".
- A menos que el peón esté en la fila inicial (fila 2), entonces puede moverse una o dos casillas hacia arriba.

Por ejemplo, dado `"D4"`, devuelve `["D5"]`, la única casilla a la que tu peón puede moverse.  
Dado `"B2"`, devuelve `["B3", "B4"]`, porque está en la fila inicial y debe estar en orden ascendente.

# --hints--

`find_pawn_moves("D4")` debería devolver `["D5"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("D4"), ["D5"])`)
}})
```

`find_pawn_moves("B2")` debería devolver `["B3", "B4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("B2"), ["B3", "B4"])`)
}})
```

`find_pawn_moves("A7")` debería devolver `["A8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("A7"), ["A8"])`)
}})
```

`find_pawn_moves("G2")` debería devolver `["G3", "G4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("G2"), ["G3", "G4"])`)
}})
```

`find_pawn_moves("E3")` debería devolver `["E4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("E3"), ["E4"])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_pawn_moves(position):

    return position
```

# --solutions--

```py
def find_pawn_moves(position):
    column = position[0].upper()
    row = int(position[1])
    moves = []

    if row >= 8:
        return moves

    moves.append(f"{column}{row + 1}")

    if row == 2:
        moves.append(f"{column}{row + 2}")

    return moves
```
