---
id: 697a49e6ff50d756c9b6935d
title: "Desaf칤o 180: Juegos de Invierno 2026 D칤a 1: D칤a de Apertura"
challengeType: 29
dashedName: challenge-180
---

# --description--

Hoy marca el inicio de los Juegos de Invierno 2026. Los pr칩ximos 17 d칤as te traer치n desaf칤os de programaci칩n inspirados en ellos.

Para el primero, se te da un c칩digo de pa칤s de dos letras y necesitas devolver el emoji de la bandera de ese pa칤s.

Usa esta lista:

| Country | Code | Flag |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"游뷣릖"` |
| Andorra                    |            `"AD"` | `"游뷣릖"` |
| Argentina                  |            `"AR"` | `"游뷣릖"` |
| Armenia                    |            `"AM"` | `"游뷣릖"` |
| Australia                  |            `"AU"` | `"游뷣릖"` |
| Austria                    |            `"AT"` | `"游뷣릖"` |
| Azerbaijan                 |            `"AZ"` | `"游뷣릖"` |
| Belgium                    |            `"BE"` | `"游游"` |
| Benin                      |            `"BJ"` | `"游游"` |
| Bolivia                    |            `"BO"` | `"游游"` |
| Bosnia and Herzegovina     |            `"BA"` | `"游游"` |
| Brazil                     |            `"BR"` | `"游游"` |
| Bulgaria                   |            `"BG"` | `"游游"` |
| Canada                     |            `"CA"` | `"游뻟릖"` |
| Chile                      |            `"CL"` | `"游뻟릖"` |
| China                      |            `"CN"` | `"游뻟릖"` |
| Colombia                   |            `"CO"` | `"游뻟릖"` |
| Croatia                    |            `"HR"` | `"游쇓릖"` |
| Cyprus                     |            `"CY"` | `"游뻟릖"` |
| Czech Republic             |            `"CZ"` | `"游뻟릖"` |
| Denmark                    |            `"DK"` | `"游뾇릖"` |
| Ecuador                    |            `"EC"` | `"游쀯릖"` |
| Eritrea                    |            `"ER"` | `"游쀯릖"` |
| Estonia                    |            `"EE"` | `"游쀯릖"` |
| Finland                    |            `"FI"` | `"游游"` |
| France                     |            `"FR"` | `"游游"` |
| Georgia                    |            `"GE"` | `"游섫릖"` |
| Germany                    |            `"DE"` | `"游뾇릖"` |
| Great Britain              |            `"GB"` | `"游섫릖"` |
| Greece                     |            `"GR"` | `"游섫릖"` |
| Guinea-Bissau              |            `"GW"` | `"游섫릖"` |
| Haiti                      |            `"HT"` | `"游쇓릖"` |
| Hong Kong                  |            `"HK"` | `"游쇓릖"` |
| Hungary                    |            `"HU"` | `"游쇓릖"` |
| Iceland                    |            `"IS"` | `"游쉻릖"` |
| India                      |            `"IN"` | `"游쉻릖"` |
| Iran                       |            `"IR"` | `"游쉻릖"` |
| Ireland                    |            `"IE"` | `"游쉻릖"` |
| Israel                     |            `"IL"` | `"游쉻릖"` |
| Italy                      |            `"IT"` | `"游쉻릖"` |
| Jamaica                    |            `"JM"` | `"游游"` |
| Japan                      |            `"JP"` | `"游游"` |
| Kazakhstan                 |            `"KZ"` | `"游썷릖"` |
| Kenya                      |            `"KE"` | `"游썷릖"` |
| Kosovo                     |            `"XK"` | `"游쬃릖"` |
| Kyrgyzstan                 |            `"KG"` | `"游썷릖"` |
| Latvia                     |            `"LV"` | `"游쐟릖"` |
| Lebanon                    |            `"LB"` | `"游쐟릖"` |
| Liechtenstein              |            `"LI"` | `"游쐟릖"` |
| Lithuania                  |            `"LT"` | `"游쐟릖"` |
| Luxembourg                 |            `"LU"` | `"游쐟릖"` |
| Madagascar                 |            `"MG"` | `"游쓇릖"` |
| Malaysia                   |            `"MY"` | `"游쓇릖"` |
| Malta                      |            `"MT"` | `"游쓇릖"` |
| Mexico                     |            `"MX"` | `"游쓇릖"` |
| Moldova                    |            `"MD"` | `"游쓇릖"` |
| Monaco                     |            `"MC"` | `"游쓇릖"` |
| Mongolia                   |            `"MN"` | `"游쓇릖"` |
| Montenegro                 |            `"ME"` | `"游쓇릖"` |
| Morocco                    |            `"MA"` | `"游쓇릖"` |
| Netherlands                |            `"NL"` | `"游游"` |
| New Zealand                |            `"NZ"` | `"游游"` |
| Nigeria                    |            `"NG"` | `"游游"` |
| North Macedonia            |            `"MK"` | `"游쓇릖"` |
| Norway                     |            `"NO"` | `"游游"` |
| Pakistan                   |            `"PK"` | `"游왫릖"` |
| Philippines                |            `"PH"` | `"游왫릖"` |
| Poland                     |            `"PL"` | `"游왫릖"` |
| Portugal                   |            `"PT"` | `"游왫릖"` |
| Puerto Rico                |            `"PR"` | `"游왫릖"` |
| Romania                    |            `"RO"` | `"游游"` |
| San Marino                 |            `"SM"` | `"游젏릖"` |
| Saudi Arabia               |            `"SA"` | `"游젏릖"` |
| Serbia                     |            `"RS"` | `"游游"` |
| Singapore                  |            `"SG"` | `"游젏릖"` |
| Slovakia                   |            `"SK"` | `"游젏릖"` |
| Slovenia                   |            `"SI"` | `"游젏릖"` |
| South Africa               |            `"ZA"` | `"游游"` |
| South Korea                |            `"KR"` | `"游썷릖"` |
| Spain                      |            `"ES"` | `"游쀯릖"` |
| Sweden                     |            `"SE"` | `"游젏릖"` |
| Switzerland                |            `"CH"` | `"游뻟릖"` |
| Thailand                   |            `"TH"` | `"游좷릖"` |
| Trinidad & Tobago          |            `"TT"` | `"游좷릖"` |
| Turkey                     |            `"TR"` | `"游좷릖"` |
| Ukraine                    |            `"UA"` | `"游쥟릖"` |
| United Arab Emirates       |            `"AE"` | `"游뷣릖"` |
| United States              |            `"US"` | `"游쥟릖"` |
| Uruguay                    |            `"UY"` | `"游쥟릖"` |
| Uzbekistan                 |            `"UZ"` | `"游쥟릖"` |
| Venezuela                  |            `"VE"` | `"游游"` |

# --hints--

`get_flag("AL")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "游뷣릖")`)
}})
```

`get_flag("AD")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "游뷣릖")`)
}})
```

`get_flag("AR")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "游뷣릖")`)
}})
```

`get_flag("AM")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "游뷣릖")`)
}})
```

`get_flag("AU")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "游뷣릖")`)
}})
```

`get_flag("AT")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "游뷣릖")`)
}})
```

`get_flag("AZ")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "游뷣릖")`)
}})
```

`get_flag("BE")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "游游")`)
}})
```

`get_flag("BJ")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "游游")`)
}})
```

`get_flag("BO")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "游游")`)
}})
```

`get_flag("BA")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "游游")`)
}})
```

`get_flag("BR")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "游游")`)
}})
```

`get_flag("BG")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "游游")`)
}})
```

`get_flag("CA")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "游뻟릖")`)
}})
```

`get_flag("CL")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "游뻟릖")`)
}})
```

`get_flag("CN")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "游뻟릖")`)
}})
```

`get_flag("CO")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "游뻟릖")`)
}})
```

`get_flag("HR")` deber칤a devolver `"游쇓릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "游쇓릖")`)
}})
```

`get_flag("CY")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "游뻟릖")`)
}})
```

`get_flag("CZ")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "游뻟릖")`)
}})
```

`get_flag("DK")` deber칤a devolver `"游뾇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "游뾇릖")`)
}})
```

`get_flag("EC")` deber칤a devolver `"游쀯릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "游쀯릖")`)
}})
```

`get_flag("ER")` deber칤a devolver `"游쀯릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "游쀯릖")`)
}})
```

`get_flag("EE")` deber칤a devolver `"游쀯릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "游쀯릖")`)
}})
```

`get_flag("FI")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "游游")`)
}})
```

`get_flag("FR")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "游游")`)
}})
```

`get_flag("GE")` deber칤a devolver `"游섫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "游섫릖")`)
}})
```

`get_flag("DE")` deber칤a devolver `"游뾇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "游뾇릖")`)
}})
```

`get_flag("GB")` deber칤a devolver `"游섫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "游섫릖")`)
}})
```

`get_flag("GR")` deber칤a devolver `"游섫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "游섫릖")`)
}})
```

`get_flag("GW")` deber칤a devolver `"游섫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "游섫릖")`)
}})
```

`get_flag("HT")` deber칤a devolver `"游쇓릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "游쇓릖")`)
}})
```

`get_flag("HK")` deber칤a devolver `"游쇓릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "游쇓릖")`)
}})
```

`get_flag("HU")` deber칤a devolver `"游쇓릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "游쇓릖")`)
}})
```

`get_flag("IS")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "游쉻릖")`)
}})
```

`get_flag("IN")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "游쉻릖")`)
}})
```

`get_flag("IR")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "游쉻릖")`)
}})
```

`get_flag("IE")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "游쉻릖")`)
}})
```

`get_flag("IL")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "游쉻릖")`)
}})
```

`get_flag("IT")` deber칤a devolver `"游쉻릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "游쉻릖")`)
}})
```

`get_flag("JM")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "游游")`)
}})
```

`get_flag("JP")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "游游")`)
}})
```

`get_flag("KZ")` deber칤a devolver `"游썷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "游썷릖")`)
}})
```

`get_flag("KE")` deber칤a devolver `"游썷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "游썷릖")`)
}})
```

`get_flag("XK")` deber칤a devolver `"游쬃릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "游쬃릖")`)
}})
```

`get_flag("KG")` deber칤a devolver `"游썷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "游썷릖")`)
}})
```

`get_flag("LV")` deber칤a devolver `"游쐟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "游쐟릖")`)
}})
```

`get_flag("LB")` deber칤a devolver `"游쐟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "游쐟릖")`)
}})
```

`get_flag("LI")` deber칤a devolver `"游쐟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "游쐟릖")`)
}})
```

`get_flag("LT")` deber칤a devolver `"游쐟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "游쐟릖")`)
}})
```

`get_flag("LU")` deber칤a devolver `"游쐟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "游쐟릖")`)
}})
```

`get_flag("MG")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "游쓇릖")`)
}})
```

`get_flag("MY")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "游쓇릖")`)
}})
```

`get_flag("MT")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "游쓇릖")`)
}})
```

`get_flag("MX")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "游쓇릖")`)
}})
```

`get_flag("MD")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "游쓇릖")`)
}})
```

`get_flag("MC")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "游쓇릖")`)
}})
```

`get_flag("MN")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "游쓇릖")`)
}})
```

`get_flag("ME")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "游쓇릖")`)
}})
```

`get_flag("MA")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "游쓇릖")`)
}})
```

`get_flag("NL")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "游游")`)
}})
```

`get_flag("NZ")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "游游")`)
}})
```

`get_flag("NG")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "游游")`)
}})
```

`get_flag("MK")` deber칤a devolver `"游쓇릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "游쓇릖")`)
}})
```

`get_flag("NO")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "游游")`)
}})
```

`get_flag("PK")` deber칤a devolver `"游왫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "游왫릖")`)
}})
```

`get_flag("PH")` deber칤a devolver `"游왫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "游왫릖")`)
}})
```

`get_flag("PL")` deber칤a devolver `"游왫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "游왫릖")`)
}})
```

`get_flag("PT")` deber칤a devolver `"游왫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "游왫릖")`)
}})
```

`get_flag("PR")` deber칤a devolver `"游왫릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "游왫릖")`)
}})
```

`get_flag("RO")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "游游")`)
}})
```

`get_flag("SM")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "游젏릖")`)
}})
```

`get_flag("SA")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "游젏릖")`)
}})
```

`get_flag("RS")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "游游")`)
}})
```

`get_flag("SG")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "游젏릖")`)
}})
```

`get_flag("SK")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "游젏릖")`)
}})
```

`get_flag("SI")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "游젏릖")`)
}})
```

`get_flag("ZA")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "游游")`)
}})
```

`get_flag("KR")` deber칤a devolver `"游썷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "游썷릖")`)
}})
```

`get_flag("ES")` deber칤a devolver `"游쀯릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "游쀯릖")`)
}})
```

`get_flag("SE")` deber칤a devolver `"游젏릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "游젏릖")`)
}})
```

`get_flag("CH")` deber칤a devolver `"游뻟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "游뻟릖")`)
}})
```

`get_flag("TH")` deber칤a devolver `"游좷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "游좷릖")`)
}})
```

`get_flag("TT")` deber칤a devolver `"游좷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "游좷릖")`)
}})
```

`get_flag("TR")` deber칤a devolver `"游좷릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "游좷릖")`)
}})
```

`get_flag("UA")` deber칤a devolver `"游쥟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "游쥟릖")`)
}})
```

`get_flag("AE")` deber칤a devolver `"游뷣릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "游뷣릖")`)
}})
```

`get_flag("US")` deber칤a devolver `"游쥟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "游쥟릖")`)
}})
```

`get_flag("UY")` deber칤a devolver `"游쥟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "游쥟릖")`)
}})
```

`get_flag("UZ")` deber칤a devolver `"游쥟릖"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "游쥟릖")`)
}})
```

`get_flag("VE")` deber칤a devolver `"游游"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "游游")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
