---
id: 697a49e6ff50d756c9b6935e
title: "Desafío 181: Juegos de Invierno 2026 Día 2: Snowboarding"
challengeType: 29
dashedName: challenge-181
---

# --description--

Dada la postura inicial de un snowboarder y una rotación en grados, determina su postura al aterrizar.

- La postura de un snowboarder es `"Regular"` o `"Goofy"`.
- Las rotaciones de trucos son múltiplos de 90 grados. Un valor positivo indica rotación en el sentido de las agujas del reloj, y un valor negativo indica rotación en sentido contrario.
- La postura de aterrizaje gira cada 180 grados de rotación.

Por ejemplo, dado `"Regular"` y `90`, devuelve `"Regular"`. Dado `"Regular"` y `180` grados, devuelve `"Goofy"`.

# --hints--

`get_landing_stance("Regular", 90)` debería devolver `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` debería devolver `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` debería devolver `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` debería devolver `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` debería devolver `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` debería devolver `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
