---
id: 697a49e6ff50d756c9b69364
title: "Desafío 187: Juegos de Invierno 2026 Día 8: Luge"
challengeType: 29
dashedName: challenge-187
---

# --description--

Dado un arreglo de cinco números, cada uno representando el tiempo (en segundos) que le tomó a un luger completar un segmento de una pista, determina qué segmento tuvo la velocidad más rápida y cuál fue esa velocidad.

La pista está dividida en los siguientes segmentos:

- Segmento 1: 320 metros
- Segmento 2: 280 metros
- Segmento 3: 350 metros
- Segmento 4: 300 metros
- Segmento 5: 250 metros

El primer valor en el array dado corresponde al tiempo para el segmento 1, el segundo valor al segmento 2, y así sucesivamente.

Para calcular la velocidad (en metros por segundo) de un segmento, divide la distancia entre el tiempo.

Devuelve `"La velocidad más rápida del luger fue X m/s en el segmento Y."` Donde `X` es la velocidad más rápida, redondeada a dos decimales, y `Y` es el número del segmento donde ocurrió la velocidad más rápida.

# --hints--

`get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128])` debería devolver `"La velocidad más rápida del lugeista fue 35.07 m/s en el segmento 5."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.")`)
}})
```

`get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284])` debería devolver `"La velocidad más rápida del lugeista fue 37.75 m/s en el segmento 2."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.")`)
}})
```

`get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912])` debería devolver `"La velocidad más rápida del lugeista fue 38.49 m/s en el segmento 3."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.")`)
}})
```

`get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840])` debería devolver `"La velocidad más rápida del luger fue 37.69 m/s en el segmento 1."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.")`)
}})
```

`get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508])` debería devolver `"La velocidad más rápida del lugeista fue 39.24 m/s en el segmento 4."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_fastest_speed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_fastest_speed(times):

    return times
```

# --solutions--

```py
def get_fastest_speed(times):
    distances = [320, 280, 350, 300, 250]
    fastest_speed = 0
    fastest_segment = 0

    for i, t in enumerate(times):
        speed = distances[i] / t
        if speed > fastest_speed:
            fastest_speed = speed
            fastest_segment = i + 1  # 1-indexed

    return f"The luger's fastest speed was {fastest_speed:.2f} m/s on segment {fastest_segment}."
```
