---
id: 697a49e6ff50d756c9b69366
title: "Desafío 189: Juegos de Invierno 2026 Día 10: Esquí Alpino"
challengeType: 29
dashedName: challenge-189
---

# --description--

Dado el desnivel vertical, la distancia horizontal y el tipo de una pista de esquí, determina la calificación de dificultad de la pista.

Para determinar la calificación:

- Calcula la inclinación de la colina tomando la caída dividida por la distancia.
- Then, calculate the adjusted steepness based on the hill type:
  - "Downhill" multiplica la pendiente por 1.2
  - "Slalom": multiplica la pendiente por 0.9
  - "Giant Slalom": multiplica la pendiente por 1.0

Devuelve:

- "Green" si la pendiente ajustada es menor o igual a 0.1
- "Blue" si la pendiente ajustada es mayor que 0.1 y menor o igual a 0.25
- "Negro" si la pendiente ajustada es mayor que 0.25


# --hints--

`get_hill_rating(95, 900, "Slalom")` debería devolver `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` debería devolver `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` debería devolver `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` debería devolver `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` debería devolver `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` debería devolver `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
