---
id: 697a49e6ff50d756c9b69365
title: "Desafío 188: Juegos de Invierno 2026 Día 9: Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

Dada una cadena que representa las curvas en una pista esquelética, determina la dificultad de la pista.

- La cadena dada solo consistirá en las letras:
  - "L" para un giro a la izquierda
  - `"R"` para un giro a la derecha
  - "S" para un segmento recto

- Cada cambio de dirección añade 15 puntos (una `"L"` seguida de una `"R"` o viceversa).
- Todas las demás curvas suman 5 puntos cada una (todos los demás caracteres `"L"` o `"R"`).
- Los segmentos rectos añaden 0 puntos.

La dificultad de la pista se basa en la puntuación total. Devuelve:

- "Fácil" si el total es 0 - 100
- "Medio" si el total es 101-200
- "Difícil" si el total es mayor a 200

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` debería devolver `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` debería devolver `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` debería devolver `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` debería devolver `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` debería devolver `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` debería devolver `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
