---
id: 697a49e6ff50d756c9b69363
title: "Desafío 186: Juegos de Invierno 2026 Día 7: Patinaje de Velocidad"
challengeType: 29
dashedName: challenge-186
---

# --description--

Dadas dos matrices que representan los tiempos por vuelta (en segundos) de dos patinadores de velocidad, devuelve el número de vuelta donde la diferencia en los tiempos por vuelta es la mayor.

El primer elemento de cada arreglo corresponde a la vuelta 1, el segundo a la vuelta 2, y así sucesivamente.

# --hints--

`largest_difference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30])` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([26.11, 25.80, 25.92, 26.23, 26.07], [25.93, 25.74, 26.53, 26.11, 26.30]), 3)`)
}})
```

`largest_difference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23])` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([27.04, 25.94, 26.22, 26.07, 26.18], [25.59, 25.80, 26.11, 26.01, 26.23]), 1)`)
}})
```

`largest_difference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95])` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.82, 25.90, 26.05, 26.00, 26.48], [25.85, 25.92, 26.07, 25.98, 25.95]), 5)`)
}})
```

`largest_difference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01])` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.88, 26.10, 25.95, 26.05, 26.00], [25.90, 26.55, 25.92, 26.03, 26.01]), 2)`)
}})
```

`largest_difference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10])` debería devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_difference([25.92, 26.01, 26.05, 25.88, 26.12], [25.95, 26.00, 26.03, 26.45, 26.10]), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def largest_difference(skater1, skater2):

    return skater1
```

# --solutions--

```py
def largest_difference(skater1, skater2):
    max_diff = -1
    lap_number = 1

    for i in range(len(skater1)):
        diff = abs(skater1[i] - skater2[i])
        if diff > max_diff:
            max_diff = diff
            lap_number = i + 1

    return lap_number
```
