---
id: 697a49e6ff50d756c9b6935d
title: "Desafío 180: Juegos de invierno día 1: día de apertura"
challengeType: 29
dashedName: challenge-180
---

# --description--

Los próximos 17 días te traerán desafíos de código inspirados en los Juegos de invierno.

Para el primero, se te da un código de país de dos letras y necesitas devolver el emoji de la bandera de ese país.

Usa esta lista:

| País | Código | Bandera |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaiyán                 |            `"AZ"` | `"🇦🇿"` |
| Bélgica                    |            `"BE"` | `"🇧🇪"` |
| Benín                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia y Herzegovina       |            `"BA"` | `"🇧🇦"` |
| Brasil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canadá                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croacia                    |            `"HR"` | `"🇭🇷"` |
| Chipre                     |            `"CY"` | `"🇨🇾"` |
| República Checa            |            `"CZ"` | `"🇨🇿"` |
| Dinamarca                  |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finlandia                  |            `"FI"` | `"🇫🇮"` |
| Francia                    |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Alemania                   |            `"DE"` | `"🇩🇪"` |
| Gran Bretaña               |            `"GB"` | `"🇬🇧"` |
| Grecia                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bisáu               |            `"GW"` | `"🇬🇼"` |
| Haití                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungría                   |            `"HU"` | `"🇭🇺"` |
| Islandia                   |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Irán                       |            `"IR"` | `"🇮🇷"` |
| Irlanda                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italia                     |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japón                      |            `"JP"` | `"🇯🇵"` |
| Kazajistán                 |            `"KZ"` | `"🇰🇿"` |
| Kenia                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kirguistán                 |            `"KG"` | `"🇰🇬"` |
| Letonia                    |            `"LV"` | `"🇱🇻"` |
| Líbano                     |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lituania                   |            `"LT"` | `"🇱🇹"` |
| Luxemburgo                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malasia                    |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| México                     |            `"MX"` | `"🇲🇽"` |
| Moldavia                   |            `"MD"` | `"🇲🇩"` |
| Mónaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Marruecos                  |            `"MA"` | `"🇲🇦"` |
| Países Bajos               |            `"NL"` | `"🇳🇱"` |
| Nueva Zelanda              |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| Macedonia del Norte        |            `"MK"` | `"🇲🇰"` |
| Noruega                    |            `"NO"` | `"🇳🇴"` |
| Pakistán                   |            `"PK"` | `"🇵🇰"` |
| Filipinas                  |            `"PH"` | `"🇵🇭"` |
| Polonia                    |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Rumania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Arabia Saudita             |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapur                   |            `"SG"` | `"🇸🇬"` |
| Eslovaquia                 |            `"SK"` | `"🇸🇰"` |
| Eslovenia                  |            `"SI"` | `"🇸🇮"` |
| Sudáfrica                  |            `"ZA"` | `"🇿🇦"` |
| Corea del Sur              |            `"KR"` | `"🇰🇷"` |
| España                     |            `"ES"` | `"🇪🇸"` |
| Suecia                     |            `"SE"` | `"🇸🇪"` |
| Suiza                      |            `"CH"` | `"🇨🇭"` |
| Tailandia                  |            `"TH"` | `"🇹🇭"` |
| Trinidad y Tobago          |            `"TT"` | `"🇹🇹"` |
| Turquía                    |            `"TR"` | `"🇹🇷"` |
| Ucrania                    |            `"UA"` | `"🇺🇦"` |
| Emiratos Árabes Unidos     |            `"AE"` | `"🇦🇪"` |
| Estados Unidos             |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistán                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` debería devolver `"🇦🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` debería devolver `"🇦🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` debería devolver `"🇦🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` debería devolver `"🇦🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` debería devolver `"🇦🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` debería devolver `"🇦🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` debería devolver `"🇦🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` debería devolver `"🇧🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` debería devolver `"🇧🇯"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` debería devolver `"🇧🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` debería devolver `"🇧🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` debería devolver `"🇧🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` debería devolver `"🇧🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` debería devolver `"🇨🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` debería devolver `"🇨🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` debería devolver `"🇨🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` debería devolver `"🇨🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` debería devolver `"🇭🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` debería devolver `"🇨🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` debería devolver `"🇨🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` debería devolver `"🇩🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` debería devolver `"🇪🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` debería devolver `"🇪🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` debería devolver `"🇪🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` debería devolver `"🇫🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` debería devolver `"🇫🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` debería devolver `"🇬🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` debería devolver `"🇩🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` debería devolver `"🇬🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` debería devolver `"🇬🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` debería devolver `"🇬🇼"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` debería devolver `"🇭🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` debería devolver `"🇭🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` debería devolver `"🇭🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` debería devolver `"🇮🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` debería devolver `"🇮🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` debería devolver `"🇮🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` debería devolver `"🇮🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` debería devolver `"🇮🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` debería devolver `"🇮🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` debería devolver `"🇯🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` debería devolver `"🇯🇵"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` debería devolver `"🇰🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` debería devolver `"🇰🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` debería devolver `"🇽🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` debería devolver `"🇰🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` debería devolver `"🇱🇻"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` debería devolver `"🇱🇧"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` debería devolver `"🇱🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` debería devolver `"🇱🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` debería devolver `"🇱🇺"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` debería devolver `"🇲🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` debería devolver `"🇲🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` debería devolver `"🇲🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` debería devolver `"🇲🇽"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` debería devolver `"🇲🇩"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` debería devolver `"🇲🇨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` debería devolver `"🇲🇳"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` debería devolver `"🇲🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` debería devolver `"🇲🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` debería devolver `"🇳🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` debería devolver `"🇳🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` debería devolver `"🇳🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` debería devolver `"🇲🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` debería devolver `"🇳🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` debería devolver `"🇵🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` debería devolver `"🇵🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` debería devolver `"🇵🇱"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` debería devolver `"🇵🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` debería devolver `"🇵🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` debería devolver `"🇷🇴"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` debería devolver `"🇸🇲"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` debería devolver `"🇸🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` debería devolver `"🇷🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` debería devolver `"🇸🇬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` debería devolver `"🇸🇰"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` debería devolver `"🇸🇮"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` debería devolver `"🇿🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` debería devolver `"🇰🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` debería devolver `"🇪🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` debería devolver `"🇸🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` debería devolver `"🇨🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` debería devolver `"🇹🇭"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` debería devolver `"🇹🇹"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` debería devolver `"🇹🇷"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` debería devolver `"🇺🇦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` debería devolver `"🇦🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` debería devolver `"🇺🇸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` debería devolver `"🇺🇾"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` debería devolver `"🇺🇿"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` debería devolver `"🇻🇪"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
