---
id: 697a49e6ff50d756c9b6935f
title: "Desafío 182: Juegos de Invierno 2026 Día 3: Biatlón"
challengeType: 29
dashedName: challenge-182
---

# --description--

Dado un arreglo de enteros, donde cada valor representa el número de objetivos alcanzados en una sola ronda de un biatlón, devuelve la distancia total de penalización que el atleta debe esquiar.

- Cada ronda consiste en 5 objetivos.
- Cada objetivo fallado resulta en un bucle de penalización de 150 metros.

# --hints--

`calculate_penalty_distance([4, 4])` debería devolver `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` debería devolver `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` debería devolver `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` debería devolver `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
