---
id: 697a49e9860d24853adef67e
title: "Desafío 193: Juegos de Invierno 2026 Día 14: Esquí de Montaña"
challengeType: 29
dashedName: challenge-193
---

# --description--

Dada la profundidad de la nieve y la pendiente de una montaña, determina si hay riesgo de avalancha.

- Los valores de profundidad de nieve son `"Shallow"`, `"Moderate"` o `"Deep"`.
- Los valores de pendiente son `"Gentle"`, `"Steep"` o `"Very Steep"`.

Devuelve `"Safe"` o `"Risky"` según esta tabla:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalanche_risk("Shallow", "Gentle")` debería devolver `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Steep")` debería devolver `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Steep"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Very Steep")` debería devolver `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Very Steep"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Gentle")` debería devolver `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Steep")` debería devolver `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Moderate", "Very Steep")` debería devolver `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Very Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Gentle")` debería devolver `"Safe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Deep", "Steep")` debería devolver `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Very Steep")` debería devolver `"Risky"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Very Steep"), "Risky")`)
}})
```

# --seed--

## --seed-contents--

```py
def avalanche_risk(snow_depth, slope):

    return snow_depth
```

# --solutions--

```py
def avalanche_risk(snow_depth, slope):
    risk_table = {
        "Gentle": {
            "Shallow": "Safe",
            "Moderate": "Safe",
            "Deep": "Safe",
        },
        "Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
        "Very Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
    }

    return risk_table[slope][snow_depth]
```
