---
id: 697a49e6ff50d756c9b69362
title: "Desafío 185: Juegos de Invierno 2026 Día 6: Patinaje Artístico"
challengeType: 29
dashedName: challenge-185
---

# --description--

Dado un arreglo de puntuaciones de jueces y penalizaciones opcionales, calcula la puntuación final de una rutina de patinaje artístico.

El primer argumento es un arreglo de 10 puntuaciones de jueces, cada una un número del 0 al 10. Elimina las puntuaciones más alta y más baja de los jueces y suma las 8 puntuaciones restantes para obtener la puntuación base.

Cualquier argumento adicional pasado a la función son penalizaciones. Resta todas las penalizaciones de la puntuación base para obtener la puntuación final.

# --hints--

`compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` debería devolver `64`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64)`)
}})
```

`compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` debería devolver `80`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80)`)
}})
```

`compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` debería devolver `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67)`)
}})
```

`compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` debería devolver `67.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5)`)
}})
```

`compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` debería devolver `59`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59)`)
}})
```

# --seed--

## --seed-contents--

```py
def compute_score(judge_scores, *penalties):

    return judge_scores
```

# --solutions--

```py
def compute_score(judge_scores, *penalties):
    sorted_scores = sorted(judge_scores)
    base_score = sum(sorted_scores[1:9])
    total_penalty = sum(penalties)

    return base_score - total_penalty
```
