---
id: 698a1a73ade5ac0e19180fa0
title: "Desafío 197: Compatibilidad de Tipo de Sangre"
challengeType: 29
dashedName: challenge-197
---

# --description--

Dado un tipo de sangre de donante y un tipo de sangre de receptor, determina si el donante puede donar sangre al receptor.

Cada tipo de sangre consiste en:

- Una letra: `"A"`, `"B"`, `"AB"`, o `"O"`
- Y un factor Rh: `"+"` o `"-"`

Los tipos de sangre serán una de las letras válidas seguida de un factor Rh. Por ejemplo, `"AB+"` y `"O-"` son tipos de sangre válidos.

Reglas de letras:

- `"O"` puede donar a otro tipo de letra.
- `"A"` puede donar a `"A"` y `"AB"`.
- `"B"` puede donar a `"B"` y `"AB"`.
- `"AB"` solo puede donar a `"AB"`.

Reglas Rh:

- Negativo (`"-"`) puede donar tanto a `"-"` como a `"+"`.
- Positivo (`"+"`) solo puede donar a `"+"`.

Tanto la regla de letra como la regla Rh deben cumplirse para que un donante pueda donar al receptor.

# --hints--

`can_donate("B+", "B+")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
