---
id: 698a1a863194f1f4e63f6460
title: "Desafío 209: Tamaño del Elemento"
challengeType: 29
dashedName: challenge-209
---

# --description--

Dado un tamaño de ventana, el ancho de un elemento en unidades de ancho de ventana `"vw"` y la altura de un elemento en unidades de alto de ventana `"vh"`, determina el tamaño del elemento en píxeles.

- El tamaño de ventana dado y el tamaño del elemento devuelto son cadenas en el formato `"width x height"`, por ejemplo, `"1200 x 800"`.

- Las unidades `"vw"` son el porcentaje del ancho de la ventana. `"50vw"`, por ejemplo, es el 50% del ancho de la ventana.
- Las unidades `"vh"` son el porcentaje de la altura de la ventana. `"50vh"`, por ejemplo, es el 50% de la altura de la ventana.

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` debería devolver `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` debe devolver `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` debería devolver `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` debería devolver `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` debe devolver `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` debe devolver `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
