---
id: 698a1a73ade5ac0e19180fa2
title: "Desafío 199: Diferencia Secuencial"
challengeType: 29
dashedName: challenge-199
---

# --description--

Dado un arreglo de números, devuelve un nuevo arreglo que contenga el valor necesario para pasar de cada número al siguiente.

- Para el último número, usa `0` ya que no hay un siguiente número.

Por ejemplo, dado `[1, 2, 4, 7]`, devuelve `[1, 2, 3, 0]`.

# --hints--

`find_differences([1, 2, 4, 7])` debería devolver `[1, 2, 3, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 4, 7]), [1, 2, 3, 0])`)
}})
```

`find_differences([10, 15, 19, 22, 24, 25])` debería devolver `[5, 4, 3, 2, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0])`)
}})
```

`find_differences([25, 20, 16, 13, 11, 10])` debería devolver `[-5, -4, -3, -2, -1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0])`)
}})
```

`find_differences([0, 1, 2, 2, 3, 3, 4, 5])` debe devolver `[1, 1, 0, 1, 0, 1, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([0, 1, 2, 2, 3, 3, 4, 5]), [1, 1, 0, 1, 0, 1, 1, 0])`)
}})
```

`find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])` debería devolver `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_differences(arr):

    return arr
```

# --solutions--

```py
def find_differences(arr):
    result = []

    for i in range(len(arr)):
        if i == len(arr) - 1:
            result.append(0)
        else:
            result.append(arr[i + 1] - arr[i])

    return result
```
