---
id: 698a1a73ade5ac0e19180fa1
title: "Desafío 198: Conteo de Días Hábiles"
challengeType: 29
dashedName: challenge-198
---

# --description--

Dada una fecha de inicio y una fecha de fin, devuelve el número de días hábiles entre ambas.

- Las fechas dadas están en el formato `"YYYY-MM-DD"`.
- Los días laborables son días hábiles (de lunes a viernes).
- Los fines de semana no son días hábiles (sábado y domingo).
- Incluye tanto la fecha de inicio como la de finalización al contar.

# --hints--

`count_business_days("2026-02-24", "2026-02-26")` debe devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-26"), 3)`)
}})
```

`count_business_days("2026-02-24", "2026-02-28")` debe devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2026-02-28"), 4)`)
}})
```

`count_business_days("2026-02-21", "2026-03-01")` debe devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-21", "2026-03-01"), 5)`)
}})
```

`count_business_days("2026-03-08", "2026-03-17")` debe devolver `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-03-08", "2026-03-17"), 7)`)
}})
```

`count_business_days("2026-02-24", "2027-02-24")` debe devolver `262`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_business_days("2026-02-24", "2027-02-24"), 262)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_business_days(start, end):

    return start
```

# --solutions--

```py
from datetime import datetime, timedelta

def count_business_days(start, end):
    start_date = datetime.strptime(start, "%Y-%m-%d")
    end_date = datetime.strptime(end, "%Y-%m-%d")
    
    count = 0
    current = start_date
    
    while current <= end_date:
        if current.weekday() < 5:
            count += 1
        current += timedelta(days=1)
    
    return count
```
