---
id: 698a1a73ade5ac0e19180fa6
title: "Desafío 203: Aplanado"
challengeType: 29
dashedName: challenge-203
---

# --description--

Dado un arreglo, determina si es plano.

- Un arreglo es plano si ninguno de sus elementos es un arreglo.

# --hints--

`is_flat([1, 2, 3, 4])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 2, 3, 4]), True)`)
}})
```

`is_flat([1, [2, 3], 4])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, 3], 4]), False)`)
}})
```

`is_flat([1, 0, False, True, "a", "b"])` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, 0, False, True, "a", "b"]), True)`)
}})
```

`is_flat(["a", [0], "b", True])` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat(["a", [0], "b", True]), False)`)
}})
```

`is_flat([1, [2, [3, [4, [5]]]], 6])` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_flat([1, [2, [3, [4, [5]]]], 6]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_flat(arr):

    return arr
```

# --solutions--

```py
def is_flat(arr):
    for el in arr:
        if isinstance(el, list):
            return False
    return True
```
