---
id: 698a1a73ade5ac0e19180fa3
title: "Desafío 200: Conteo de Letras y Números"
challengeType: 29
dashedName: challenge-200
---

# --description--

Dada una cadena, devuelve un mensaje con el conteo de cuántas letras y números contiene.

- Las letras son `A-Z` y `a-z`.
- Los números son `0-9`.
- Ignora todos los demás caracteres.

Devuelve `"The string has X letters and Y numbers."`, donde `"X"` es el conteo de letras y `"Y"`` es el conteo de números. Si alguno de los conteos es 1, usa la forma singular para ese elemento. Por ejemplo: `"1 letter"` en lugar de `"1 letters"` y `"1 number"` en lugar de `"1 numbers"`.

# --hints--

`count_letters_and_numbers("helloworld123")` debería devolver `"The string has 10 letters and 3 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("helloworld123"), "The string has 10 letters and 3 numbers.")`)
}})
```

`count_letters_and_numbers("Catch 22")` debe devolver `"The string has 5 letters and 2 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("Catch 22"), "The string has 5 letters and 2 numbers.")`)
}})
```

`count_letters_and_numbers("A1!")` debe devolver `"The string has 1 letter and 1 number."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("A1!"), "The string has 1 letter and 1 number.")`)
}})
```

`count_letters_and_numbers("12345")` debería devolver `"The string has 0 letters and 5 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("12345"), "The string has 0 letters and 5 numbers.")`)
}})
```

`count_letters_and_numbers("password")` debería devolver `"The string has 8 letters and 0 numbers."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_letters_and_numbers("password"), "The string has 8 letters and 0 numbers.")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_letters_and_numbers(s):

    return s
```

# --solutions--

```py
def count_letters_and_numbers(s):
    letter_count = 0
    number_count = 0

    for char in s:
        if ('A' <= char <= 'Z') or ('a' <= char <= 'z'):
            letter_count += 1
        elif '0' <= char <= '9':
            number_count += 1

    letter_word = "letter" if letter_count == 1 else "letters"
    number_word = "number" if number_count == 1 else "numbers"

    return f"The string has {letter_count} {letter_word} and {number_count} {number_word}."
```
