---
id: 698a1a73ade5ac0e19180fa8
title: "Desafío 205: Conteo de Cubos Perfectos"
challengeType: 29
dashedName: challenge-205
---

# --description--

Dado dos enteros, determina cuántos cubos perfectos existen en el rango entre y que incluyen a los dos números.

- El número menor no está garantizado como el primer argumento.
- Un número es un cubo perfecto si existe un entero (`n`) donde `n * n * n = number`. Por ejemplo, 27 es un cubo perfecto porque `3 * 3 * 3 = 27`.

# --hints--

`count_perfect_cubes(3, 30)` debe devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(3, 30), 2)`)
}})
```

`count_perfect_cubes(1, 30)` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(1, 30), 3)`)
}})
```

`count_perfect_cubes(30, 0)` debe devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(30, 0), 4)`)
}})
```

`count_perfect_cubes(-64, 64)` debe devolver `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(-64, 64), 9)`)
}})
```

`count_perfect_cubes(9214, -8127)` debe devolver `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(9214, -8127), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_perfect_cubes(a, b):

    return a
```

# --solutions--

```py
def count_perfect_cubes(a, b):
    start, end = min(a, b), max(a, b)
    count = 0

    n = 0
    while n ** 3 <= end:
        if n ** 3 >= start:
            count += 1
        n += 1

    n = -1
    while n ** 3 >= start:
        if n ** 3 <= end:
            count += 1
        n -= 1

    return count
```
