---
id: 698a1a863194f1f4e63f645e
title: "Desafío 207: Brecha más pequeña"
challengeType: 29
dashedName: challenge-207
---

# --description--

Dada una cadena, devuelve la subcadena entre los dos caracteres idénticos que tienen el menor número de caracteres entre ellos (brecha más pequeña).

- Siempre habrá al menos un par de caracteres coincidentes.
- La subcadena devuelta debe excluir los caracteres coincidentes.
- Si dos o más espacios tienen la misma longitud, devuelve los caracteres del primero.

Por ejemplo, dado `"ABCDAC"`, devuelve `"DA"`.

- Solo `"A"` y `"C"` se repiten en la cadena.
- El número de caracteres entre los dos caracteres `"A"` es 3, y entre los caracteres `"C"` es 2.
- Entonces devuelve la cadena entre los dos caracteres `"C"`.

# --hints--

`smallest_gap("ABCDAC")` debería devolver `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` debería devolver `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` debe devolver `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` debería devolver `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` debe devolver `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
