---
id: 698a1a73ade5ac0e19180fa9
title: "Desafío 206: Valores de las Cartas de Juego"
challengeType: 29
dashedName: challenge-206
---

# --description--

Dado un arreglo de cartas, devuelve un nuevo arreglo con el valor numérico de cada carta.

Valores de la carta:

- Un As (`"A"`) tiene un valor de 1.
- Las cartas numeradas (`"2"` - `"10"`) tienen su valor facial: 2 - 10, respectivamente.
- Las cartas con figuras: Jack (`"J"`), Queen (`"Q"`), y King (`"K"`) valen 10 cada una.

Trajes:

- Cada carta tiene un palo: Espadas (`"S"`), Tréboles (`"C"`), Diamantes (`"D"`), o Corazones (`"H"`).

Formato de tarjeta:

- Cada carta se representa como una cadena: `"valueSuit"`. Por ejemplo: `"AS"` es el As de Espadas, `"10H"` es el Diez de Corazones, y `"QC"` es la Reina de Tréboles.

# --hints--

`card_values(["3H", "4D", "5S"])` debería devolver `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` debería devolver `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` debería devolver `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` debería devolver `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` debería devolver `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
