---
id: 698a1a73ade5ac0e19180fa5
title: "Desafío 202: Agrega Puntuación"
challengeType: 29
dashedName: challenge-202
---

# --description--

Dada una cadena de oraciones con puntos faltantes, agrega un punto (`"."`) en los siguientes lugares:

- Antes de cada espacio que viene inmediatamente antes de una letra mayúscula
- Y al final de la cadena

Devuelve la cadena resultante.

# --hints--

`add_punctuation("Hello world")` debe devolver `"Hello world."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world"), "Hello world.")`)
}})
```

`add_punctuation("Hello world It's nice today")` debería devolver `"Hello world. It's nice today."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world It's nice today"), "Hello world. It's nice today.")`)
}})
```

`add_punctuation("JavaScript is great Sometimes")` debería devolver `"JavaScript is great. Sometimes."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.")`)
}})
```

`add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` debe devolver `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.")`)
}})
```

`add_punctuation("Wait.. For it")` debe devolver `"Wait... For it."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Wait.. For it"), "Wait... For it.")`)
}})
```

# --seed--

## --seed-contents--

```py
def add_punctuation(sentences):

    return sentences
```

# --solutions--

```py
def add_punctuation(sentences):
    result = ""
    length = len(sentences)

    for i, c in enumerate(sentences):
        if c == " " and i + 1 < length and sentences[i + 1].isupper():
            result += "."
        result += c

    if not result.endswith("."):
        result += "."

    return result
```
