---
id: 698a1a73ade5ac0e19180fa7
title: "Desafío 204: Suma las Letras"
challengeType: 29
dashedName: challenge-204
---

# --description--

Dada una cadena, devuelve la suma de sus letras.

- Las letras son A-Z en mayúsculas o minúsculas
- Los valores de las letras son: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Las letras mayúsculas y minúsculas tienen el mismo valor.
- Ignora todos los caracteres que no sean letras.

# --hints--

`sum_letters("Hello")` debería devolver `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Hello"), 52)`)
}})
```

`sum_letters("freeCodeCamp")` debería devolver `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("freeCodeCamp"), 94)`)
}})
```

`sum_letters("The quick brown fox jumps over the lazy dog.")` debe devolver `473`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("The quick brown fox jumps over the lazy dog."), 473)`)
}})
```

`sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` debe devolver `1681`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681)`)
}})
```

`sum_letters("</404>")` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_letters("</404>"), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_letters(s):

    return s
```

# --solutions--

```py
def sum_letters(s):
    total = 0
    for char in s:
        upper = char.upper()
        if 'A' <= upper <= 'Z':
            total += ord(upper) - ord('A') + 1
    return total
```
