---
id: 6994cff2290543b3aec9f510
title: "Desafío 211: Suma de Array"
challengeType: 29
dashedName: challenge-211
---

# --description--

Dado un arreglo de números, devuelve la suma de todos los números.

# --hints--

`sum_array([1, 2, 3, 4, 5])` debería devolver `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([1, 2, 3, 4, 5]), 15)`)
}})
```

`sum_array([42])` debe devolver `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([42]), 42)`)
}})
```

`sum_array([5, -2, 7, -3])` debe devolver `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([5, -2, 7, -3]), 7)`)
}})
```

`sum_array([203, 145, -129, 6293, 523, -919, 845, 2434])` debe devolver `9395`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395)`)
}})
```

`sum_array([0, 0])` debe devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_array([0, 0]), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_array(numbers):

    return numbers
```

# --solutions--

```py
def sum_array(numbers):

    return sum(numbers)
```
