---
id: 699c8e045ee7cb94ed2322d8
title: "Desafío 217: Piezas de ajedrez capturadas"
challengeType: 29
dashedName: challenge-217
---

# --description--

Dado un arreglo de cadenas que representan las piezas de ajedrez que aún tienes en el tablero, calcula el valor de las piezas que tu oponente ha capturado.

En ajedrez, comienzas con 16 piezas:

| Piece  | Abbreviation   | Quantity | Value |
| ------ | -------------- | -------- | ----- |
| Pawn   | `"P"`          | 8        | 1     |
| Rook   | `"R"`          | 2        | 5     |
| Knight | `"N"`          | 2        | 3     |
| Bishop | `"B"`          | 2        | 3     |
| Queen  | `"Q"`          | 1        | 9     |
| King   | `"K"`          | 1        | 0     |

- El arreglo dado solo contendrá las abreviaturas anteriores.
- Cualquiera de las 16 piezas que no están incluidas en el arreglo dado han sido capturadas.
- Devuelve el valor total de todas las piezas capturadas, a menos que...
- Si el Rey ha sido capturado, devuelve `"Checkmate"`.

# --hints--

`get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` debe devolver `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8)`)
}})
```

`get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"])` debería devolver `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "P", "P", "P", "P", "R", "B", "K"]), 26)`)
}})
```

`get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` debe devolver `16`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16)`)
}})
```

`get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` debe devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4)`)
}})
```

`get_captured_value(["P", "K"])` debería devolver `38`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["P", "K"]), 38)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` debe devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0)`)
}})
```

`get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` debe devolver `"Checkmate"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_captured_value(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_captured_value(pieces):

    return pieces
```

# --solutions--

```py
def get_captured_value(pieces):
    if "K" not in pieces:
        return "Checkmate"

    values = {"P":1, "R":5, "N":3, "B":3, "Q":9, "K":0}
    remaining = "PPPPPPPPRRNNBBQK"

    for p in pieces:
        remaining = remaining.replace(p, "", 1)

    return sum(values[p] for p in remaining)
```
