---
id: 699c8e045ee7cb94ed2322db
title: "Desafío 220: Número más grande"
challengeType: 29
dashedName: challenge-220
---

# --description--

Dada una cadena de números separados por varios signos de puntuación, devuelve el número más grande.

- La cadena dada solo contendrá números y separadores.
- Los separadores pueden ser comas (`","`), signos de exclamación (`"!"`), signos de interrogación (`"?"`), dos puntos (`":"`) o puntos y coma (`";"`).

# --hints--

`largest_number("1,2")` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("1,2"), 2)`)
}})
```

`largest_number("4;15:60,26?52!0")` debe devolver `60`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("4;15:60,26?52!0"), 60)`)
}})
```

`largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` debe devolver `-402`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402)`)
}})
```

`largest_number("12;-50,99.9,49.1!-10.1?88?16")` debe devolver `99.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(largest_number("12;-50,99.9,49.1!-10.1?88?16"), 99.9)`)
}})
```

# --seed--

## --seed-contents--

```py
def largest_number(s):

    return s
```

# --solutions--

```py
import re
def largest_number(s):
    numbers = [float(n) for n in re.split(r'[,!?:;]', s)]
    return max(numbers)
```
