---
id: 699c8e045ee7cb94ed2322d6
title: "Desafío 215: Calculadora de tarifa de estacionamiento"
challengeType: 29
dashedName: challenge-215
---

# --description--

Dadas dos cadenas que representan la hora en que estacionaste tu auto y la hora en que lo recogiste, calcula la tarifa de estacionamiento.

- Las cadenas dadas estarán en el formato `"HH:MM"` usando un reloj de 24 horas. Por ejemplo, `"14:00"` es 2pm.
- La primera cadena será la hora en que estacionaste tu coche, y la segunda será la hora en que lo recogiste.
- Si la hora de recogida es anterior a la hora de entrada, significa que el estacionamiento se extendió más allá de la medianoche hasta el día siguiente.

Reglas de tarifas:

- Cada hora estacionada cuesta $3.
- Las horas parciales se redondean hacia arriba a la siguiente hora completa.
- Si el estacionamiento abarca toda la noche (más allá de la medianoche), se aplica una tarifa adicional de $10 por la noche.
- Hay una tarifa mínima de $5 (solo se usa si el total sería menor a $5).

Devuelve el costo total en el formato `"$cost"`, por ejemplo, `"$5"`.

# --hints--

`calculate_parking_fee("09:00", "11:00")` debe devolver `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` debe devolver `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` debería devolver `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` debe devolver `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` debería devolver `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` debe devolver `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
