---
id: 699c8e045ee7cb94ed2322dd
title: "Desafío 222: Sombras del Equinoccio"
challengeType: 29
dashedName: challenge-222
---

# --description--

Hoy es el equinoccio, cuando el sol está directamente sobre el ecuador y perfectamente en lo alto al mediodía. Dado un tiempo, determina la sombra proyectada por un poste vertical de 4 pies.

- La hora será una cadena en formato de 24 horas `"HH:MM"` (por ejemplo, `"15:00"` es 3pm).
- Solo se te dará una hora en incrementos de 30 minutos.

Reglas:

- El sol sale a las 6am directamente `"east"`, y se pone a las 6pm directamente `"west"`.
- Una sombra siempre apunta en dirección opuesta al sol.
- La longitud de la sombra (en pies) es el número de horas desde el mediodía, elevado al cubo.
- No hay sombra antes del amanecer (antes de las 6am), después del atardecer (6pm o más tarde), o al mediodía.

Devuelve:

- Si existe una sombra, devuelve `"(length)ft (direction)"`. Por ejemplo, `"8ft west"`.
- De lo contrario, devuelve `"No shadow"`.

Por ejemplo, dado `"10:00"`, devuelve `"8ft west"` porque las 10am están a 2 horas del mediodía, así que 2<sup>3</sup> = 8 pies, y la sombra apunta al oeste porque el sol está en el este a las 10am.

# --hints--

`get_shadow("10:00")` debería devolver `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` debe devolver `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` debe devolver `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` debe devolver `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` debería devolver `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` debe devolver `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` debe devolver `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` debe devolver `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` debe devolver `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
