---
id: 699c8e045ee7cb94ed2322d5
title: "Desafío 214: Validador de Cadena de Dominó"
challengeType: 29
dashedName: challenge-214
---

# --description--

Dado un arreglo 2D que representa una secuencia de dominós, determina si forma una cadena válida.

- Cada elemento en el arreglo representa un dominó y será un arreglo de dos números del 1 al 6, (inclusive).
- Para que la cadena sea válida, el segundo número de cada dominó debe coincidir con el primer número del siguiente dominó.
- El primer número del primer dominó y el último número del último dominó no necesitan coincidir con nada.

# --hints--

`is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

`is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

`is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

`is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

`is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
