---
id: 699c8e045ee7cb94ed2322d9
title: "Desafío 218: Divisible Uniformemente"
challengeType: 29
dashedName: challenge-218
---

# --description--

Dado dos enteros, determina si puedes dividir el primero por el segundo de manera exacta.

# --hints--

`is_evenly_divisible(4, 2)` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4, 2), True)`)
}})
```

`is_evenly_divisible(7, 3)` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(7, 3), False)`)
}})
```

`is_evenly_divisible(5, 10)` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(5, 10), False)`)
}})
```

`is_evenly_divisible(48, 6)` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(48, 6), True)`)
}})
```

`is_evenly_divisible(3186, 9)` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(3186, 9), True)`)
}})
```

`is_evenly_divisible(4192, 11)` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_evenly_divisible(4192, 11), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_evenly_divisible(a, b):

    return a
```

# --solutions--

```py
def is_evenly_divisible(a, b):
    return a % b == 0
```
