---
id: 69a890af247de743333bd4d2
title: "Desafío 229: Trunca el Texto 2"
challengeType: 29
dashedName: challenge-229
---

# --description--

Dada una cadena, devuelve una nueva cadena que esté truncada para que el ancho total de los caracteres no exceda 50 unidades.

Cada carácter tiene un ancho específico:

| Letters | Width |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

La tabla anterior incluye todas las letras mayúsculas y minúsculas. Además:

- Los espacios (`" "`) tienen un ancho de 2
- Los puntos (`"."`) tienen un ancho de 1

- Si la cadena dada tiene 50 unidades o menos, devuelve la cadena tal cual, de lo contrario
- Trunca la cadena y añade tres puntos al final (`"..."`) para que su ancho total, incluidos los tres puntos, sea lo más cercano posible a 60 unidades sin pasarse.

# --hints--

`truncate_text("The quick brown fox")` debería devolver `"The quick brown f..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The quick brown fox"), "The quick brown f...")`)
}})
```

`truncate_text("The silky smooth sloth")` debe devolver `truncate_text("The silky smooth sloth")`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The silky smooth sloth"), truncate_text("The silky smooth sloth"))`)
}})
```

`truncate_text("THE LOUD BRIGHT BIRD")` debe devolver `"THE LOUD BRIG..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...")`)
}})
```

`truncate_text("The fast striped zebra")` debería devolver `"The fast striped z..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The fast striped zebra"), "The fast striped z...")`)
}})
```

`truncate_text("The big black bear")` debería devolver `"The big black bear"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The big black bear"), "The big black bear")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(s):

    return s
```

# --solutions--

```py
def truncate_text(s):
    MAX_WIDTH = 50
    ELLIPSIS = "..."
    ELLIPSIS_WIDTH = 3
    TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH

    char_groups = {
        "ilI.": 1,
        "fjrt ": 2,
        "abcdeghkmnopqrstuvwxyzJL": 3,
        "ABCDEFGHKMNOPQRSTUVWXYZ": 4
    }

    def get_char_width(char):
        for key, width in char_groups.items():
            if char in key:
                return width
        return 3

    def string_width(string):
        total = 0
        for char in string:
            total += get_char_width(char)
        return total

    if string_width(s) <= MAX_WIDTH:
        return s

    result = ""
    total_width = 0

    for char in s:
        char_width = get_char_width(char)

        if total_width + char_width > TRUNCATE_LIMIT:
            break

        result += char
        total_width += char_width

    return result + ELLIPSIS
```
