---
id: 69a890af247de743333bd4cd
title: "Desafío 224: Detector de Tostado de Café"
challengeType: 29
dashedName: challenge-224
---

# --description--

Dada una cadena que representa los granos usados para hacer una taza de café, determina el tueste de la taza.

- La cadena dada contendrá los siguientes caracteres, cada uno representando un tipo de frijol:

  - Un apóstrofe (`'`) es un grano de tueste ligero que vale 1 punto cada uno.
  - Un guion (`-`) es un grano de tueste medio que vale 2 puntos cada uno.
  - Un punto (`.`) es un grano de tueste oscuro que vale 3 puntos cada uno.

- El nivel de tueste se determina por el promedio de todos los granos.

Devuelve:

- `"Light"` si el promedio es menor que 1.75.
- `"Medium"` si el promedio está entre 1.75 y 2.5.
- `"Dark"` si el promedio es mayor que 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` debe devolver `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` debe devolver `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` debe devolver `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` debería devolver `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` debe devolver `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` debe devolver `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` debe devolver `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
