---
id: 69a890af247de743333bd4d0
title: "Desafío 227: Tiempo de enfriamiento"
challengeType: 29
dashedName: challenge-227
---

# --description--

Dadas dos marcas de tiempo, la primera que representa cuándo un usuario terminó un examen, y la segunda que representa el tiempo actual, determina si el usuario puede volver a tomar un examen.

- Ambas marcas de tiempo tendrán el formato: `"YYYY-MM-DDTHH:MM:SS"`, por ejemplo `"2026-03-25T14:00:00"`. Ten en cuenta que la hora está en formato de 24 horas.
- Un usuario debe esperar al menos 48 horas antes de volver a presentar un examen.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` debe devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` debe devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
