---
id: 69a890af247de743333bd4cf
title: "Desafío 226: Contar exámenes aprobados"
challengeType: 29
dashedName: challenge-226
---

# --description--

Dado un arreglo de puntajes de exámenes de estudiantes y el puntaje necesario para aprobar, devuelve el número de estudiantes que aprobaron el examen.

# --hints--

`passing_count([90, 85, 75, 60, 50], 70)` debe devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([90, 85, 75, 60, 50], 70), 3)`)
}})
```

`passing_count([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75)` debe devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([100, 80, 75, 88, 72, 74, 79, 71, 60, 92], 75), 6)`)
}})
```

`passing_count([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60)` debe devolver `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([79, 60, 88, 72, 74, 59, 75, 71, 80, 92], 60), 9)`)
}})
```

`passing_count([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85)` debe devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([76, 79, 80, 70, 71, 65, 79, 78, 59, 72], 85), 0)`)
}})
```

`passing_count([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60)` debería devolver `27`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(passing_count([84, 65, 98, 53, 58, 71, 91, 80, 92, 70, 73, 83, 86, 69, 84, 77, 72, 58, 69, 75, 66, 68, 72, 96, 90, 63, 88, 63, 80, 67], 60), 27)`)
}})
```

# --seed--

## --seed-contents--

```py
def passing_count(scores, passing_score):

    return scores
```

# --solutions--

```py
def passing_count(scores, passing_score):
    count = 0

    for score in scores:
        if score >= passing_score:
            count += 1

    return count
```
