---
id: 69b58ce40693f140c84c8559
title: "Desafío 240: caracteres de palíndromo"
challengeType: 29
dashedName: challenge-240
---

# --description--

Dada una cadena, determina si es un palíndromo y devuelve el carácter del medio (si tiene longitud impar) o los dos caracteres del medio (si tiene longitud par).

- Un palíndromo es una cadena que es igual de adelante hacia atrás y de atrás hacia adelante.
- Si no es un palíndromo, devuelve `"none"`.

# --hints--

`palindrome_locator("racecar")` debería devolver `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` debería devolver `"v"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` debería devolver `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` debería devolver `"oo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` debería devolver `"00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
