---
id: 69b58ce40693f140c84c855b
title: "Desafío 242: siguiente número de bingo"
challengeType: 29
dashedName: challenge-242
---

# --description--

Dado un número de bingo, devuelve el siguiente número de bingo de forma secuencial.

Un número de bingo es una letra seguida de un número dentro de su rango según esta tabla:

| Letra | Rango de números |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Por ejemplo, dado `"B10"`, devuelve `"B11"`, el siguiente número de bingo. Si se recibe el último número de bingo, devuelve `"B1"`.

# --hints--

`get_next_bingo_number("B10")` debería devolver `"B11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("B10"), "B11")`)
}})
```

`get_next_bingo_number("N33")` debería devolver `"N34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("N33"), "N34")`)
}})
```

`get_next_bingo_number("I30")` debería devolver `"N31"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("I30"), "N31")`)
}})
```

`get_next_bingo_number("G60")` debería devolver `"O61"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("G60"), "O61")`)
}})
```

`get_next_bingo_number("O75")` debería devolver `"B1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_bingo_number("O75"), "B1")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_bingo_number(n):

    return n
```

# --solutions--

```py
def get_next_bingo_number(n):
    num = int(n[1:])
    next_num = 1 if num == 75 else num + 1

    if next_num <= 15:
        return "B" + str(next_num)
    if next_num <= 30:
        return "I" + str(next_num)
    if next_num <= 45:
        return "N" + str(next_num)
    if next_num <= 60:
        return "G" + str(next_num)
    return "O" + str(next_num)
```
