---
id: 69bc6cb30c1d112a2e110a07
title: "Desafío 250: clave oculta"
challengeType: 29
dashedName: challenge-250
---

# --description--

¡Bienvenido al desafío diario número 250!

Dado un texto codificado, descífralo usando una clave de cifrado y devuélvelo.

Para encontrar la clave:

- Observa todos los desafíos diarios hasta hoy cuyo número sea múltiplo de 25 (incluyendo este).
- Toma la primera letra de cada uno de esos títulos y combínalas en una cadena. Si el título empieza con un carácter que no es letra, busca su primera letra.

Para descifrar el mensaje, recorre cada letra del mensaje codificado y:

  1. Mira la letra correspondiente en la clave (repite la clave si el mensaje es más largo que la clave).
  2. Convierte la letra de la clave a su número correspondiente: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Desplaza la letra codificada hacia atrás en el alfabeto esa cantidad de lugares.
  4. Si el desplazamiento pasa antes de `"A"`, continúa desde `"Z"`.

Por ejemplo, si el mensaje codificado empieza con `"Y"` y la primera letra de la clave es `"V"` (22), desplaza `"Y"` 22 lugares hacia atrás para obtener `"C"`. Repite este proceso para cada letra y descifra el mensaje completo.

- Solo se desplazan letras, los espacios se mantienen igual.
- Todas las letras dadas y devueltas están en mayúsculas.

# --hints--

`decode("YAVJYNXE")` debería devolver `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` debería devolver `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` debería devolver `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` debería devolver `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` debería devolver `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
