---
id: 69bc6cb30c1d112a2e110a08
title: "Challenge 251: Array Sum Finder"
challengeType: 29
dashedName: challenge-251
---

# --description--

Dado un array de números y un número objetivo, devuelve el primer subconjunto de dos o más números que sumen el objetivo.

- El "primer" subconjunto es aquel cuyos elementos tienen los índices más bajos posibles, priorizando primero el índice más temprano.
- Cada número en el array solo puede usarse una vez.
- Si no existe un subconjunto válido, devuelve `"Sum not found"`.

Devuelve los números que coincidan como un array en el orden en que aparecen en el array original.

# --hints--

`find_sum([1, 3, 5, 7], 6)` should return `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 3, 5, 7], 6), [1, 5])`)
}})
```

`find_sum([1, 2, 3, 4, 5], 5)` should return `[1, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5], 5), [1, 4])`)
}})
```

`find_sum([1, 2, 3, 4, 5], 6)` should return `[1, 2, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5], 6), [1, 2, 3])`)
}})
```

`find_sum([-1, -2, 3, 4], 1)` should return `[-1, -2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([-1, -2, 3, 4], 1), [-1, -2, 4])`)
}})
```

`find_sum([3, 1, 4, 1, 5, 9, 2, 6], 10)` should return `[3, 1, 4, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2])`)
}})
```

`find_sum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` should return `[1, 2, 3, 5, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9])`)
}})
```

`find_sum([7, 9, 4, 2, 5], 10)` should return `"Sum not found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_sum([7, 9, 4, 2, 5], 10), "Sum not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_sum(arr, target):

    return arr
```

# --solutions--

```py
def find_sum(arr, target):
    def search(start, remaining, current):
        if remaining == 0 and len(current) >= 2:
            return current
        if start >= len(arr):
            return None
        for i in range(start, len(arr)):
            result = search(i + 1, remaining - arr[i], current + [arr[i]])
            if result is not None:
                return result
        return None

    return search(0, target, []) or "Sum not found"
```
