---
id: 69c5f3d787b1725d5f00c8ba
title: "Desafío 256: dirección de tiempo más cercana"
challengeType: 29
dashedName: challenge-256
---

# --description--

Dadas dos horas, determina si puedes llegar de la primera a la segunda más rápido moviéndote hacia adelante o hacia atrás.

- Las horas se dan en formato de 24 horas (`"HH:MM"`)
- El reloj da la vuelta (las 23:59 pasan a 00:00 al avanzar, y las 00:00 pasan a 23:59 al retroceder)

Devuelve:

- `"forward"` si avanzar es más rápido
- `"backward"` si retroceder es más rápido
- `"equal"` si ambas direcciones toman el mismo tiempo

# --hints--

`get_direction("10:00", "12:00")` debería devolver `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")` debería devolver `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")` debería devolver `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")` debería devolver `"forward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")` debería devolver `"backward"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")` debería devolver `"equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
