---
id: 69c5f3d787b1725d5f00c8bd
title: "Desafío 259: explosión fizzbuzz"
challengeType: 29
dashedName: challenge-259
---

# --description--

Dado un entero, devuelve el número de pasos que se necesitan para convertir la palabra `"fizzbuzz"` en una cadena con al menos el número dado de `"z"` usando las siguientes reglas:

- Comienza con la cadena `"fizzbuzz"`.
- En cada paso, aplica las reglas estándar de FizzBuzz usando la posición de la letra en la cadena (la primera `"f"` está en la posición 1).
  - Si la posición de la letra es divisible por 3, reemplaza la letra con `"fizz"`
  - Si es divisible por 5, reemplaza la letra con `"buzz"`
  - Si es divisible por 3 y 5, reemplaza la letra con `"fizzbuzz"`

Así que después de 1 paso, `"fizzbuzz"` se convierte en `"fifizzzbuzzfizzzz"`, que tiene 9 `"z"`.

# --hints--

`explode_fizzbuzz(9)` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` debería devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` debería devolver `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` debería devolver `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
