---
id: 69c6ff713a52713463aa7927
title: "Desafío 261: palabras numéricas"
challengeType: 29
dashedName: challenge-261
---

# --description--

Dado un número entero de 0 a 99, devuelve su representación en palabras en inglés.

- `0` devuelve `"zero"`.
- Los números del 1 al 19 tienen nombres únicos (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Los múltiplos de 10 del 20 al 90 tienen sus propios nombres (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Los números del 21 al 99 que no son múltiplos de 10 se escriben como dos palabras unidas por un guion. Por ejemplo `"forty-two"` y `"fifty-three"`.

# --hints--

`get_number_words(0)` debería devolver `"zero"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(0), "zero")`)
}})
```

`get_number_words(10)` debería devolver `"ten"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(10), "ten")`)
}})
```

`get_number_words(19)` debería devolver `"nineteen"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(19), "nineteen")`)
}})
```

`get_number_words(30)` debería devolver `"thirty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(30), "thirty")`)
}})
```

`get_number_words(53)` debería devolver `"fifty-three"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(53), "fifty-three")`)
}})
```

`get_number_words(7)` debería devolver `"seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(7), "seven")`)
}})
```

`get_number_words(12)` debería devolver `"twelve"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(12), "twelve")`)
}})
```

`get_number_words(60)` debería devolver `"sixty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(60), "sixty")`)
}})
```

`get_number_words(67)` debería devolver `"sixty-seven"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(67), "sixty-seven")`)
}})
```

`get_number_words(98)` debería devolver `"ninety-eight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_words(98), "ninety-eight")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_words(n):

    return n
```

# --solutions--

```py
def get_number_words(n):
    ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
            "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
            "seventeen","eighteen","nineteen"]
    tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"]

    if n < 20:
        return ones[n]
    if n % 10 == 0:
        return tens[n // 10]
    return tens[n // 10] + "-" + ones[n % 10]
```
