---
id: 69cfca90e8a0a6d4d6871c51
title: "Desafío 267: convertidor de parsecs"
challengeType: 29
dashedName: challenge-267
---

# --description--

En una galaxia lejana, los parsecs se usan para medir tanto el tiempo como la distancia. Dado un número entero de parsecs, devuelve su equivalente en tiempo o distancia.

- Si el número entero es impar, representa tiempo. Si es par, representa distancia.

Usa estas tasas de conversión:

| Parsecs | Tiempo/Distancia |
| - | - |
| 1 | 2 horas |
| 2 | 6 años luz |

Devuelve el valor convertido como un entero.

# --hints--

`convert_parsecs(1)` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` debería devolver `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` debería devolver `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` debería devolver `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` debería devolver `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
