---
id: 69e2383af7832c8032603b93
title: "Desafío 277: imagen en espejo"
challengeType: 29
dashedName: challenge-277
---

# --description--

Dadas dos cadenas, determina si la segunda cadena es una imagen en espejo de la primera.

Una imagen en espejo se forma invirtiendo la cadena y reemplazando cada carácter por su equivalente en espejo.

- Los caracteres simétricos se ven igual en un espejo:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, y el espacio (` `).

- Los pares reflejados se intercambian entre sí en un espejo:

| Carácter | Se intercambia con |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Si alguna de las cadenas incluye un carácter que no está en las listas anteriores, no tiene una imagen en espejo que pueda crearse con esos caracteres.

Por ejemplo, la imagen en espejo de `"[HOW]"` es `"[WOH]"`.

# --hints--

`is_mirror_image("[HOW]", "[WOH]")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("[HOW]", "[WOH]"), True)`)
}})
```

`is_mirror_image("MOM", "MOM")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("MOM", "MOM"), True)`)
}})
```

`is_mirror_image("vow", "wov")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("vow", "wov"), True)`)
}})
```

`is_mirror_image("TIM", "TIM")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("TIM", "TIM"), False)`)
}})
```

`is_mirror_image("{WOW}", "}WOW{")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("{WOW}", "}WOW{"), False)`)
}})
```

`is_mirror_image("XXVII", "IIV%X")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("XXVII", "IIV%X"), False)`)
}})
```

`is_mirror_image("><(((*>", "<*)))><")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("><(((*>", "<*)))><"), True)`)
}})
```

`is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror_image(s1, s2):

    return s1
```

# --solutions--

```py
def is_mirror_image(s1, s2):
    symmetric = set('WTYUIOHA XVMwoxv08=+:|-_*^!.')
    pairs = {'[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<',
             'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '('}
    mirrored = []
    for c in reversed(s1):
        if c in symmetric:
            mirrored.append(c)
        elif c in pairs:
            mirrored.append(pairs[c])
        else:
            return False
    return ''.join(mirrored) == s2
```
