---
id: 69f35a5bb823ed620fcb7cbe
title: "Challenge 285: Meeting Time"
challengeType: 29
dashedName: challenge-285
---

# --description--

Dado un arreglo 3D que representa las ventanas de disponibilidad de varias personas, devuelve la hora más temprana en la que todos tienen una hora libre. Si no existe tal hora, devuelve `"None"`.

- La disponibilidad de cada persona es un arreglo de pares enteros `[start, end]` en formato de 24 horas. Por ejemplo, `[10, 12]` significaría que la persona está disponible de 10 a 12. Los horarios de inicio van de 0 a 23, y los horarios de fin van de 1 a 24.

Por ejemplo, dado:

```json
[
  [[10, 12], [15, 16]], // person 1
  [[11, 14], [15, 16]]  // person 2
]
```

Devuelve `11`, el inicio de su primera hora libre compartida.

# --hints--

`get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11)`)
}})
```

`get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` should return `13`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13)`)
}})
```

`get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` should return `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9)`)
}})
```

`get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` should return `None`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None")`)
}})
```

`get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` should return `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_meeting_time([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_meeting_time(availability):

    return availability
```

# --solutions--

```py
def get_meeting_time(availability):
    for hour in range(24):
        if all(
            any(start <= hour and hour + 1 <= end for start, end in person)
            for person in availability
        ):
            return hour
    return "None"
```
