---
id: 69f8c998d78ad3171a0713bb
title: "Challenge 288: Secret Number"
challengeType: 29
dashedName: challenge-288
---

# --description--

Dado un número secreto y una suposición, determina si la suposición es correcta.

Return:

- `"higher"` si el número secreto es mayor que la suposición.
- `"lower"` si el número secreto es menor que la suposición.
- `"you got it!"` si la suposición es correcta.

# --hints--

`guess_number(50, 30)` should return `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(50, 30), "higher")`)
}})
```

`guess_number(85, 99)` should return `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(85, 99), "lower")`)
}})
```

`guess_number(2026, 2026)` should return `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(2026, 2026), "you got it!")`)
}})
```

`guess_number(92904, 11283)` should return `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(92904, 11283), "higher")`)
}})
```

`guess_number(230495, 423920)` should return `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(230495, 423920), "lower")`)
}})
```

`guess_number(120349, 120349)` should return `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(120349, 120349), "you got it!")`)
}})
```

# --seed--

## --seed-contents--

```py
def guess_number(secret, guess):

    return secret
```

# --solutions--

```py
def guess_number(secret, guess):
    if guess == secret:
        return "you got it!"
    return "higher" if guess < secret else "lower"
```
