---
id: 69f8c998d78ad3171a0713b9
title: "Desafío 286: issues abiertas"
challengeType: 29
dashedName: challenge-286
---

# --description--

Dado un arreglo de números de issues y otro arreglo de números de pull requests (PR), devuelve un arreglo con los issues que permanecen abiertos después de que todos los PR se hayan fusionado.

- Un PR cierra un issue si sus dígitos son una rotación uno del otro. Por ejemplo, el issue 123 sería cerrado por el PR 231 o 312.
- Un PR no cierra un issue con el mismo número exacto. Por ejemplo, el PR 123 no cierra el issue 123. Así que un issue con todos los mismos dígitos no puede cerrarse.
- Cualquiera de los números puede tener ceros a la izquierda eliminados. Por ejemplo, el PR 201 cerraría el issue 12 (012, una rotación de 201). De manera similar, el issue 201 sería cerrado por el PR 12.

Devuelve los issues abiertos restantes en el orden en que fueron dados.

# --hints--

`get_open_issues([123, 234], [231])` debería devolver `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` debería devolver `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` debería devolver `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` debería devolver `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` debería devolver `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
