---
id: 69f8c998d78ad3171a0713bd
title: "Desafío 290: fiesta de pizza"
challengeType: 29
dashedName: challenge-290
---

# --description--

Dado un arreglo con las horas trabajadas hoy por persona, devuelve la cantidad de pizzas que se deben pedir para una fiesta de pizza.

- Divide las horas trabajadas de cada persona entre 3 para obtener la cantidad de porciones.
- No puedes comer una porción parcial, así que redondea hacia arriba la cantidad de porciones de cada persona al número entero más cercano.
- Cada persona recibe un mínimo de dos porciones.
- Cada pizza tiene 8 porciones. Redondea hacia arriba el total de pizzas al número entero más cercano.

# --hints--

`get_pizzas_to_order([8, 8, 8])` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` debería devolver `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
