---
id: 69f90c1329a94b37e2a2086c
title: "Desafío 293: mejor mano"
challengeType: 29
dashedName: challenge-293
---

# --description--

Dado un arreglo de cinco cadenas que representan cartas de juego, devuelve el nombre de la mejor mano.

- Cada carta se representa como una cadena de dos caracteres: el rango seguido del palo, `"2h"` por ejemplo.
  - Los rangos, de menor a mayor, son: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` y `"A"`.
  - Los palos son: `"h"`, `"d"`, `"c"` y `"s"`.
- Los ases (`"A"`) pueden usarse como altos o bajos en una escalera.

Las manos, en orden de peor a mejor, son:

| Nombre | Descripción |
| - | - |
| `"High Card"` | Sin pareja ni mejor combinación |
| `"Pair"` | Dos cartas del mismo rango |
| `"Two Pair"` | Dos cartas de un rango y dos de otro |
| `"Three of a Kind"` | Tres cartas del mismo rango |
| `"Straight"` | Cinco rangos consecutivos |
| `"Flush"` | Cinco cartas del mismo palo |
| `"Full House"` | Tres cartas de un rango y dos de otro |
| `"Four of a Kind"` | Cuatro cartas del mismo rango |
| `"Straight Flush"` | Cinco rangos consecutivos del mismo palo |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` del mismo palo |

Devuelve el nombre de la mejor mano.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` debería devolver `"Three of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` debería devolver `"Full House"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` debería devolver `"Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` debería devolver `"Four of a Kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` debería devolver `"Two Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` debería devolver `"Straight Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` debería devolver `"High Card"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` debería devolver `"Straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` debería devolver `"Pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` debería devolver `"Royal Flush"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
