---
id: 69f90c1329a94b37e2a2086d
title: "Desafío 294: combinaciones de paréntesis"
challengeType: 29
dashedName: challenge-294
---

# --description--

Dado un entero, `n`, devuelve el número de combinaciones válidas de `n` pares de paréntesis.

- Una combinación válida es una cadena donde cada paréntesis de apertura tiene un paréntesis de cierre correspondiente, y ningún paréntesis de cierre aparece antes que su paréntesis de apertura correspondiente.

Por ejemplo, dado `2`, hay 2 combinaciones válidas:

```json
(())
()()
```

# --hints--

`get_combinations(2)` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` debería devolver `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` debería devolver `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` debería devolver `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
