---
id: 6a0dcd03ee4e68698080ef66
title: "Desafío 302: Jet lag"
challengeType: 29
dashedName: challenge-302
---

# --description--

Dadas una ciudad de salida, una ciudad de llegada, una duración del vuelo en horas y una dirección de viaje, devuelve el número de horas de jet lag que experimenta el viajero.

Las ciudades dadas serán de la siguiente lista que incluye su diferencia horaria UTC:

| Ciudad | Diferencia |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Para calcular las horas de jet lag:

1. Encuentra la diferencia de zona horaria en horas entre las dos ciudades.
2. Determina el multiplicador de dirección. Si viajas `"east"`, es 1.5; de lo contrario, es 1.0.
3. Obtén las horas de jet lag con la fórmula: diferencia de zona horaria + (duración del vuelo * 0.1) * multiplicador de dirección

Devuelve las horas de jet lag redondeadas a un decimal.

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` debería devolver `6.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` debería devolver `5.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` debería devolver `1.6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` debería devolver `4.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` debería devolver `17.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` debería devolver `5.9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` debería devolver `9.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
