---
id: 6a0dcd03ee4e68698080ef68
title: "Desafío 304: arreglos de itinerarios"
challengeType: 29
dashedName: challenge-304
---

# --description--

Dado un arreglo de al menos dos paradas opcionales para una excursión de un día, devuelve el número de arreglos válidos del itinerario.

El itinerario siempre incluye `"breakfast"`, `"lunch"` y `"dinner"`, estos no se pasan como argumentos. Las paradas opcionales pueden colocarse en cualquier parte del itinerario, sujetas a las siguientes reglas:

- `"breakfast"` siempre es la primera, con al menos una parada antes de `"lunch"`.
- `"lunch"` debe aparecer antes de `"dinner"`, con al menos una parada entre ellas.
- Como máximo, una parada opcional puede aparecer después de `"dinner"`.

Devuelve el número de arreglos válidos.

# --hints--

`get_itinerary_count(["library", "park"])` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` debería devolver `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` debería devolver `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` debería devolver `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` debería devolver `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
