---
id: 6a0dcd03ee4e68698080ef6b
title: "Desafío 307: regulaciones de zonificación"
challengeType: 29
dashedName: challenge-307
---

# --description--

Dado un diseño 2D (arreglo de arreglos) que representa la distribución de edificios de una ciudad, devuelve las coordenadas de todos los edificios que están violando las reglas de zonificación.

Cada celda en la rejilla contiene una de las etiquetas de la tabla a continuación. Un edificio está en violación si cualquiera de sus (hasta) 4 vecinos, horizontales o verticales, es de un tipo al que no puede estar adyacente.

| Etiqueta | Tipo | No puede estar adyacente a |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | agrícola | `"C"` |
| `"R"` | residencial | `"i"`, `"C"` |
| `"I"` | institucional | `"i"` |
| `"C"` | comercial | `"R"`, `"A"` |
| `""` (cadena vacía) | sin desarrollar | sin restricciones |

Devuelve las coordenadas de todas las celdas que violan las reglas como un arreglo de pares `[row, col]`, en cualquier orden. Si no existen violaciones, devuelve un arreglo vacío.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` debería devolver `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` debería devolver `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` debería devolver `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` debería devolver `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
