---
id: 6a0dcd03ee4e68698080ef6c
title: "Desafío 308: validador de tarjeta de crédito"
challengeType: 29
dashedName: challenge-308
---

# --description--

Dada una cadena de dígitos para un número de tarjeta de crédito, determina si es un número válido usando el siguiente método:

- Empezando desde el penúltimo dígito, duplica cada otro dígito moviéndote hacia la izquierda.
- Si al duplicar un dígito el resultado es un número mayor que 9, réstale 9.
- Suma todos los dígitos (duplicados y no duplicados).
- Si el total es divisible por 10, el número es válido.

# --hints--

`is_valid_card("4532015112830366")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` debería devolver `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` debería devolver `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
