---
id: 6a15cadf5f240d05a2649559
title: "Desafío 315: solsticio de verano"
challengeType: 29
dashedName: challenge-315
---

# --description--

Hoy es el solsticio de verano, el día más largo del año en el hemisferio norte y el más corto en el hemisferio sur. Dada una latitud, devuelve una cadena que representa las horas de día y de noche.

- La latitud estará entre 90 (polo norte) y -90 (polo sur), inclusive
- El número de horas de día = 12 + (latitud / 90) * 12
- Redondea las horas de día al número par más cercano

Devuelve una cadena de 24 caracteres usando `"☀️"` para las horas de día y `"🌑"` para las horas de noche, donde:

- Cada carácter representa una hora, comenzando a medianoche (hora 0)
- El amanecer y el atardecer son simétricos alrededor del mediodía

Por ejemplo, una latitud de `0` (el ecuador) tiene 12 horas de luz diurna, así que el amanecer es a las 6:00 AM y el atardecer a las 6:00 PM. Devuelve: `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

# --hints--

`get_daytime_hours(0)` debería devolver `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(90)` debería devolver `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️")`)
}})
```

`get_daytime_hours(-90)` debería devolver `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(-33)` debería devolver `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(66.5)` debería devolver `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑")`)
}})
```

`get_daytime_hours(40)` debería devolver `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑")`)
}})
```

`get_daytime_hours(-50)` debería devolver `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_daytime_hours(latitude):

    return latitude
```

# --solutions--

```py
def get_daytime_hours(latitude):
    daylight_hours = round((12 + (latitude / 90) * 12) / 2) * 2
    sunrise = 12 - daylight_hours // 2
    sunset = 12 + daylight_hours // 2
    result = ""
    for hour in range(24):
        result += "☀️" if sunrise <= hour < sunset else "🌑"
    return result
```
