---
id: 6a15cadf5f240d05a2649558
title: "Desafío 314: factorización en factores primos"
challengeType: 29
dashedName: challenge-314
---

# --description--

Dado un entero mayor que 1, devuelve su factorización en factores primos como un arreglo de números en orden ascendente.

Una factorización en factores primos es el conjunto de números primos que se multiplican para producir el entero dado. Cada número tiene exactamente un conjunto. Por ejemplo, la factorización en factores primos de 20 es `[2, 2, 5]` porque 2 * 2 * 5 = 20.

Si el entero dado es primo, devuélvelo en un arreglo de un solo elemento.

# --hints--

`prime_factorization(20)` debería devolver `[2, 2, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` debería devolver `[17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` debería devolver `[3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` debería devolver `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` debería devolver `[3, 3, 3, 37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` debería devolver `[2, 2, 2, 3, 3, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` debería devolver `[2, 3, 5, 7, 11, 13, 17]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
