---
id: 6a151d32d772271dd2448c2d
title: "Desafío 311: lanzador de hechizos"
challengeType: 29
dashedName: challenge-311
---

# --description--

Dada una cadena de códigos de hechizos que estás lanzando, calcula la puntuación total.

Cada carácter en la cadena representa un hechizo:

| Código | Hechizo | Categoría | Puntuación base |
| - | - | - | - |
| `"f"` | Fuego | Destrucción | 3 |
| `"l"` | Rayo | Destrucción | 3 |
| `"i"` | Hielo | Control | 2 |
| `"w"` | Viento | Control | 2 |
| `"h"` | Curar | Restauración | 1 |
| `"s"` | Escudo | Restauración | 1 |

Se aplica un multiplicador de combo según cuántos hechizos consecutivos se hayan lanzado de diferentes categorías:

- El primer hechizo siempre puntúa con su valor base.
- Cada hechizo consecutivo de una categoría diferente a la anterior aumenta el multiplicador en 1.
- Lanzar un hechizo de la misma categoría que el anterior reinicia el multiplicador a 1.
- La puntuación de cada hechizo es su puntuación base multiplicada por el multiplicador actual.

Devuelve la puntuación total de la secuencia de hechizos.

# --hints--

`cast("fihwl")` debería devolver `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` debería devolver `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` debería devolver `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` debería devolver `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` debería devolver `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
