---
id: 6a15cadf5f240d05a264955e
title: "Desafío 320: banco de sangre"
challengeType: 29
dashedName: challenge-320
---

# --description--

Dado un arreglo con el inventario de un banco de sangre y un arreglo con las solicitudes de tipo de sangre de los pacientes, devuelve una cadena en el formato `"X of Y patients served"`. Donde `X` es el número máximo de pacientes que pueden recibir sangre del inventario del banco, y `Y` es el número total de pacientes.

Cada entrada en ambos arreglos es uno de los siguientes tipos de sangre: `"AB"`, `"A"`, `"B"` o `"O"`.

Reglas de compatibilidad:

- `"AB"` puede recibir de cualquier tipo de sangre.
- `"A"` puede recibir de `"A"` y `"O"`.
- `"B"` puede recibir de `"B"` y `"O"`.
- `"O"` solo puede recibir de `"O"`.

Las entradas duplicadas en los arreglos dados representan cantidad.

# --hints--

`triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` debería devolver `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` debería devolver `"3 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served")`)
}})
```

`triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` debería devolver `"4 of 5 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served")`)
}})
```

`triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` debería devolver `"4 of 4 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` debería devolver `"8 of 11 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served")`)
}})
```

`triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` debería devolver `"13 of 14 patients served"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_blood(bank, patients):

    return bank
```

# --solutions--

```py
def triage_blood(bank, patients):
    b = {}
    p = {}
    for t in bank:
        b[t] = b.get(t, 0) + 1
    for t in patients:
        p[t] = p.get(t, 0) + 1

    saved = 0

    def serve(patient, donors):
        nonlocal saved
        for donor in donors:
            n = min(p.get(patient, 0), b.get(donor, 0))
            p[patient] = p.get(patient, 0) - n
            b[donor] = b.get(donor, 0) - n
            saved += n

    serve("O", ["O"])
    serve("A", ["A", "O"])
    serve("B", ["B", "O"])
    serve("AB", ["AB", "A", "B", "O"])
    return f"{saved} of {len(patients)} patients served"
```
