---
id: 6a19b1062d1b153d8ac76d70
title: "Desafío 322: Conectar 3"
challengeType: 29
dashedName: challenge-322
---

# --description--

Dada una matriz de cadenas que representan piezas en una rejilla de juego, determina si algún jugador tiene tres en línea.

- Cada celda contiene `"R"`, `"Y"` o `""` (cadena vacía).
- Tres en línea significa tres celdas consecutivas no vacías del mismo tipo horizontal, vertical o diagonalmente.

Devuelve:

- Un arreglo plano con el ganador y las coordenadas de sus tres celdas ganadoras en el formato: `["R", [0,2], [1,3], [2,4]]`. Las coordenadas se devuelven de arriba hacia abajo y luego de izquierda a derecha.
- Un arreglo vacío si no hay ganador.

# --hints--

`connect_three([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]])` debería devolver `["R", [3, 1], [3, 2], [3, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "", ""], ["", "", "", ""], ["", "Y", "", ""], ["Y", "R", "R", "R"]]), ["R", [3, 1], [3, 2], [3, 3]])`)
}})
```

`connect_three([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]])` debería devolver `["Y", [1, 1], [2, 1], [3, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "", ""], ["", "Y", "Y", ""], ["", "Y", "R", "R"], ["", "Y", "R", "R"]]), ["Y", [1, 1], [2, 1], [3, 1]])`)
}})
```

`connect_three([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]])` debería devolver `["R", [0, 3], [1, 2], [2, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "", "Y", "R"], ["", "Y", "R", "Y"], ["", "R", "Y", "R"], ["", "R", "Y", "R"]]), ["R", [0, 3], [1, 2], [2, 1]])`)
}})
```

`connect_three([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]])` debería devolver `["Y", [0, 1], [1, 2], [2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["", "Y", "", ""], ["", "Y", "Y", ""], ["", "R", "R", "Y"], ["R", "R", "Y", "R"]]), ["Y", [0, 1], [1, 2], [2, 3]])`)
}})
```

`connect_three([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]])` debería devolver `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(connect_three([["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"], ["Y", "R", "R", "Y"], ["R", "Y", "Y", "R"]]), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def connect_three(matrix):

    return matrix
```

# --solutions--

```py
def connect_three(matrix):
    rows = len(matrix)
    cols = len(matrix[0])

    directions = [[0, 1], [1, 0], [1, 1], [1, -1]]

    for r in range(rows):
        for c in range(cols):
            piece = matrix[r][c]
            if not piece:
                continue

            for dr, dc in directions:
                cells = [[r, c]]

                for i in range(1, 3):
                    nr = r + dr * i
                    nc = c + dc * i
                    if nr < 0 or nr >= rows or nc < 0 or nc >= cols:
                        break
                    if matrix[nr][nc] != piece:
                        break
                    cells.append([nr, nc])

                if len(cells) == 3:
                    cells.sort(key=lambda x: (x[0], x[1]))
                    return [piece] + cells

    return []
```
