---
id: 6a1d9f98e819ed70a0e994df
title: "Desafío 333: Clasificación de problemas 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

Dado un título de problema y un arreglo de etiquetas actuales, devuelve un arreglo actualizado de etiquetas según las siguientes reglas:

Si el problema no tiene ninguna etiqueta, agrega:

- `"bug"` y `"needs triage"` si el título contiene `"error"` o `"bug"`
- `"enhancement"` y `"discussing"` si el título contiene `"feature"` o `"add"`

De lo contrario, si las etiquetas dadas contienen:

- `"needs triage"` y el título contiene `"simple"` o `"easy"`, elimina `"needs triage"` y agrega `"good first issue"`
- `"discussing"` y el título contiene `"planned"` o `"next"`, elimina `"discussing"` y agrega `"on the roadmap"`
- De lo contrario, si `"needs triage"` o `"discussing"` está presente, elimínalo y agrega `"help wanted"`

Si el título contiene:

- `"security"`, agrega una etiqueta `"critical"`

# --hints--

`triage_issue("app crashes with error", [])` debería devolver `["bug", "needs triage"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` debería devolver `["bug", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` debería devolver `["enhancement", "discussing"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` debería devolver `["enhancement", "help wanted"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` debería devolver `["bug", "needs triage", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` debería devolver `["critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` debería devolver `["bug", "good first issue"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` debería devolver `["enhancement", "on the roadmap"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` debería devolver `["enhancement", "help wanted", "critical"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
