---
id: 6a1d9f98e819ed70a0e994da
title: "Desafío 328: la rutina de Kaprekar"
challengeType: 29
dashedName: challenge-328
---

# --description--

Dado un número de 4 dígitos, devuelve la cantidad de veces que necesitas aplicar la rutina de Kaprekar hasta llegar a 6174.

La rutina de Kaprekar funciona así:

- Ordena los dígitos en orden descendente para formar el número más grande
- Ordena los dígitos en orden ascendente para formar el número más pequeño (rellena con ceros a la izquierda si es necesario)
- Resta el número más pequeño del más grande
- Repite con el nuevo número

# --hints--

`kaprekar(1234)` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` debería devolver `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` debería devolver `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` debería devolver `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
