---
id: 6a1d9f98e819ed70a0e994e1
title: "Desafío 335: cinco dados"
challengeType: 29
dashedName: challenge-335
---

# --description--

Dado un arreglo de cinco dados con valores del 1 al 6, devuelve la mejor mano posible.

Aquí están las manos ordenadas de menor a mayor:

| Mano | Descripción |
|------|-------------|
| `"no pair"` | Sin pares ni mejor |
| `"pair"` | Dos dados con el mismo valor |
| `"two pair"` | Dos pares diferentes |
| `"three of a kind"` | Tres dados con el mismo valor |
| `"small straight"` | Cuatro valores consecutivos |
| `"large straight"` | Cinco valores consecutivos |
| `"full house"` | Tres iguales y un par |
| `"four of a kind"` | Cuatro dados con el mismo valor |
| `"five of a kind"` | Los cinco dados con el mismo valor |

# --hints--

`five_dice([1, 1, 1, 1, 1])` debería devolver `"five of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 1, 1, 1, 1]), "five of a kind")`)
}})
```

`five_dice([5, 5, 5, 6, 5])` debería devolver `"four of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([5, 5, 5, 6, 5]), "four of a kind")`)
}})
```

`five_dice([2, 5, 6, 4, 3])` debería devolver `"large straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 5, 6, 4, 3]), "large straight")`)
}})
```

`five_dice([4, 3, 3, 3, 1])` debería devolver `"three of a kind"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 3, 3, 3, 1]), "three of a kind")`)
}})
```

`five_dice([4, 6, 2, 6, 5])` debería devolver `"pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 6, 2, 6, 5]), "pair")`)
}})
```

`five_dice([1, 4, 5, 6, 2])` debería devolver `"no pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 4, 5, 6, 2]), "no pair")`)
}})
```

`five_dice([1, 3, 4, 6, 2])` debería devolver `"small straight"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 3, 4, 6, 2]), "small straight")`)
}})
```

`five_dice([2, 2, 5, 2, 5])` debería devolver `"full house"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 2, 5, 2, 5]), "full house")`)
}})
```

`five_dice([6, 4, 5, 6, 4])` debería devolver `"two pair"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([6, 4, 5, 6, 4]), "two pair")`)
}})
```

# --seed--

## --seed-contents--

```py
def five_dice(dice):

    return dice
```

# --solutions--

```py
def five_dice(dice):
    from collections import Counter
    counts = Counter(dice)
    values = sorted(counts.values(), reverse=True)
    keys = sorted(counts.keys())

    is_large_straight = len(keys) == 5 and keys[-1] - keys[0] == 4
    is_small_straight = not is_large_straight and any(
        keys[i + 3] - keys[i] == 3 and len(set(keys[i:i + 4])) == 4
        for i in range(len(keys) - 3)
    )

    if values[0] == 5: return "five of a kind"
    if values[0] == 4: return "four of a kind"
    if values[0] == 3 and values[1] == 2: return "full house"
    if is_large_straight: return "large straight"
    if is_small_straight: return "small straight"
    if values[0] == 3: return "three of a kind"
    if values[0] == 2 and values[1] == 2: return "two pair"
    if values[0] == 2: return "pair"
    return "no pair"
```
