---
id: 6a1d9f98e819ed70a0e994dc
title: "Desafío 330: palabras en minúsculas"
challengeType: 29
dashedName: challenge-330
---

# --description--

Dada una cadena, devuelve solo las palabras que están completamente en minúsculas, en su orden original y con un espacio entre cada palabra.

# --hints--

`get_lowercase_words("hello GOOD world")` debería devolver `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("hello GOOD world"), "hello world")`)
}})
```

`get_lowercase_words("these are all lowercase")` debería devolver `"these are all lowercase"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("these are all lowercase"), "these are all lowercase")`)
}})
```

`get_lowercase_words("less is NoT more")` debería devolver `"less is more"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("less is NoT more"), "less is more")`)
}})
```

`get_lowercase_words("DonT eat pizza every OTHER day")` debería devolver `"eat pizza every day"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("DonT eat pizza every OTHER day"), "eat pizza every day")`)
}})
```

`get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` debería devolver `"the quick brown fox jumped over the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lowercase_words(s):

    return s
```

# --solutions--

```py
def get_lowercase_words(s):
    return " ".join(word for word in s.split(" ") if word == word.lower())
```
