---
id: 6a2037a68a0bc2aef0006004
title: "Desafío 342: probabilidades con dados"
challengeType: 29
dashedName: challenge-342
---

# --description--

Dado un número de dados de seis caras para lanzar y una suma objetivo, devuelve las probabilidades de obtener esa suma como una cadena en el formato `"1 in X"`.

- El número de dados estará entre 1 y 6.
- La suma objetivo siempre es alcanzable con el número dado de dados.
- Redondea `"X"` al número entero más cercano.

# --hints--

`get_odds(1, 5)` debería devolver `"1 in 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(1, 5), "1 in 6")`)
}})
```

`get_odds(2, 4)` debería devolver `"1 in 12"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(2, 4), "1 in 12")`)
}})
```

`get_odds(3, 10)` debería devolver `"1 in 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(3, 10), "1 in 8")`)
}})
```

`get_odds(4, 7)` debería devolver `"1 in 65"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(4, 7), "1 in 65")`)
}})
```

`get_odds(5, 26)` debería devolver `"1 in 111"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(5, 26), "1 in 111")`)
}})
```

`get_odds(6, 35)` debería devolver `"1 in 7776"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(6, 35), "1 in 7776")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odds(dice, target):

    return dice
```

# --solutions--

```py
def get_odds(dice, target):
    valid_combos = 0
    total_combos = 6 ** dice

    def roll(dice_left, total):
        nonlocal valid_combos
        if dice_left == 0:
            if total == target:
                valid_combos += 1
            return
        for face in range(1, 7):
            roll(dice_left - 1, total + face)

    roll(dice, 0)
    return f"1 in {round(total_combos / valid_combos)}"
```
