---
id: 6a2037a68a0bc2aef0006002
title: "Desafío 340: convertidor a Pig Latin"
challengeType: 29
dashedName: challenge-340
---

# --description--

Dada una cadena, conviértela a Pig Latin usando las siguientes reglas:

- Si una palabra comienza con una vocal (`"a"`, `"e"`, `"i"`, `"o"` o `"u"`), agrega `"way"` al final. Por ejemplo, `"universe"` se convierte en `"universeway"`.
- Si una palabra comienza con una o más consonantes, muévelas al final y agrega `"ay"`. Por ejemplo, `"hello"` se convierte en `"ellohay"`.
- Conserva la mayúscula de la primera letra. Por ejemplo, `"Hello"` se convierte en `"Ellohay"`.

# --hints--

`pig_latin("universe")` debería devolver `"universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` debería devolver `"ellohay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` debería devolver `"ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` debería devolver `"Ellohay universeway"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` debería devolver `"Igpay Atinlay isway unfay"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` debería devolver `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
