---
id: 6a2037a68a0bc2aef0006000
title: "Desafío 338: calculadora de edad de mascotas"
challengeType: 29
dashedName: challenge-338
---

# --description--

Dado un tipo de mascota y su edad en años humanos, devuelve la edad equivalente en años de mascota usando la siguiente tabla de conversión:

| Mascota | Multiplicador |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`pet_years("dog", 5)` debería devolver `35`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

`pet_years("cat", 9)` debería devolver `54`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

`pet_years("rabbit", 3)` debería devolver `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

`pet_years("hamster", 4)` debería devolver `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

`pet_years("guinea pig", 5)` debería devolver `60`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

`pet_years("goldfish", 2)` debería devolver `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

`pet_years("bird", 1)` debería devolver `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
