---
id: 6a2037a68a0bc2aef0006005
title: "Desafío 343: paradas del elevador"
challengeType: 29
dashedName: challenge-343
---

# --description--

Dado un número para el piso actual de un elevador y un arreglo de pisos solicitados, devuelve un arreglo con el orden en que el elevador debe visitarlos para minimizar el número de pisos recorridos.

- Si hay empate, sube primero
- Los pisos con una solicitud deben ser visitados cuando el elevador los pase por primera vez

# --hints--

`elevator_stops(5, [2, 8, 3, 9])` debería devolver `[3, 2, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(5, [2, 8, 3, 9]), [3, 2, 8, 9])`)
}})
```

`elevator_stops(6, [2, 10, 8, 3, 1, 9])` debería devolver `[8, 9, 10, 3, 2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1])`)
}})
```

`elevator_stops(1, [4, 8, 3, 6, 9])` debería devolver `[3, 4, 6, 8, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9])`)
}})
```

`elevator_stops(12, [6, 10, 7, 3, 1, 4])` debería devolver `[10, 7, 6, 4, 3, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1])`)
}})
```

`elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` debería devolver `[10, 9, 8, 6, 5, 2, 12, 19, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def elevator_stops(current_floor, stops):

    return current_floor
```

# --solutions--

```py
def elevator_stops(current_floor, stops):
    stops = sorted(stops)
    above = [s for s in stops if s > current_floor]
    below = [s for s in stops if s < current_floor][::-1]
    up_distance = above[-1] - current_floor if above else float('inf')
    down_distance = current_floor - below[-1] if below else float('inf')
    return below + above if down_distance < up_distance else above + below
```
