---
id: 6a2037a68a0bc2aef0006001
title: "Desafío 339: Dividir arreglos en partes"
challengeType: 29
dashedName: challenge-339
---

# --description--

Dado un arreglo y un tamaño de parte, devuelve el arreglo dividido en subarreglos de ese tamaño.

- La última parte puede ser más pequeña si el arreglo no se divide de manera uniforme.

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` debería devolver `[[1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` debería devolver `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` debería devolver `[[1, 2, 3], [4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` debería devolver `[["a"], ["b"], ["c"], ["d"], ["e"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` debería devolver `[[1, 2, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
