---
id: 6a22d77ddf034bc4e35b1d59
title: "Desafío 352: clasificación de relación de contraste 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

Dada una relación de contraste y un booleano que indica si el texto es grande, devuelve la clasificación WCAG usando la siguiente tabla:

| Calificación | Texto normal | Texto grande |
|-------------|--------------|--------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | por debajo de 4.5 | por debajo de 3.0 |

# --hints--

`get_contrast_rating("7.5", False)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
