---
id: 6a22d77ddf034bc4e35b1d5a
title: "Desafío 353: clasificación de contraste 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

Dadas dos valores de luminancia relativa y un booleano que indica si el texto es grande, devuelve la clasificación de contraste WCAG usando el siguiente método:

Calcula la relación de contraste agregando 0.05 a cada valor de luminancia, luego dividiendo el más claro entre el más oscuro. El más claro siempre será el primer argumento.

Devuelve la calificación basada en la relación de contraste usando la siguiente tabla:

| Calificación | Texto normal | Texto grande |
|-------------|--------------|--------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | por debajo de 4.5 | por debajo de 3.0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
