---
id: 6a22d77ddf034bc4e35b1d55
title: "Desafío 348: calculadora de préstamos"
challengeType: 29
dashedName: challenge-348
---

# --description--

Dado un monto de préstamo, un porcentaje de tasa de interés anual y un pago mensual fijo, devuelve un arreglo con los saldos restantes después de cada pago mensual hasta que el préstamo esté pagado.

- Cada mes, se calcula el interés sobre el saldo restante usando la tasa de interés mensual: `(annual rate / 100) / 12`, luego se resta el pago mensual.
- Devuelve cada saldo restante redondeado al dólar más cercano.
- Incluye el monto del préstamo en el arreglo devuelto. El primer elemento del arreglo siempre será el monto del préstamo y el último elemento siempre será 0.

# --hints--

`get_loan_schedule(1000, 0, 200)` debería devolver `[1000, 800, 600, 400, 200, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_loan_schedule(1000, 0, 200), [1000, 800, 600, 400, 200, 0])`)
}})
```

`get_loan_schedule(1000, 5, 200)` debería devolver `[1000, 804, 608, 410, 212, 13, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_loan_schedule(1000, 5, 200), [1000, 804, 608, 410, 212, 13, 0])`)
}})
```

`get_loan_schedule(10, 50, 1)` debería devolver `[10, 9, 9, 8, 8, 7, 6, 5, 5, 4, 3, 2, 1, 0, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_loan_schedule(10, 50, 1), [10, 9, 9, 8, 8, 7, 6, 5, 5, 4, 3, 2, 1, 0, 0])`)
}})
```

`get_loan_schedule(5500, 8, 400)` debería devolver `[5500, 5137, 4771, 4403, 4032, 3659, 3283, 2905, 2525, 2141, 1756, 1367, 977, 583, 187, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_loan_schedule(5500, 8, 400), [5500, 5137, 4771, 4403, 4032, 3659, 3283, 2905, 2525, 2141, 1756, 1367, 977, 583, 187, 0])`)
}})
```

`get_loan_schedule(50000, 5.2, 1650)` debería devolver `[50000, 48567, 47127, 45681, 44229, 42771, 41306, 39835, 38358, 36874, 35384, 33887, 32384, 30874, 29358, 27835, 26306, 24770, 23227, 21678, 20122, 18559, 16990, 15413, 13830, 12240, 10643, 9039, 7428, 5810, 4186, 2554, 915, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_loan_schedule(50000, 5.2, 1650), [50000, 48567, 47127, 45681, 44229, 42771, 41306, 39835, 38358, 36874, 35384, 33887, 32384, 30874, 29358, 27835, 26306, 24770, 23227, 21678, 20122, 18559, 16990, 15413, 13830, 12240, 10643, 9039, 7428, 5810, 4186, 2554, 915, 0])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_loan_schedule(loan_amount, annual_rate, monthly_payment):

    return loan_amount
```

# --solutions--

```py
def get_loan_schedule(loan_amount, annual_rate, monthly_payment):
    monthly_rate = (annual_rate / 100) / 12
    balances = [loan_amount]
    balance = loan_amount

    while balance > 0:
        balance += balance * monthly_rate
        balance -= monthly_payment
        balances.append(round(max(balance, 0)))

    return balances
```
