---
id: 6a22d77ddf034bc4e35b1d57
title: "Desafío 350: distancia entre letras"
challengeType: 29
dashedName: challenge-350
---

# --description--

Dadas dos cadenas de igual longitud, devuelve la suma de las distancias más cortas entre cada par de caracteres.

- La entrada solo contendrá letras minúsculas
- El alfabeto se trata como un círculo, por lo que la distancia entre `a` y `z` es 1.

# --hints--

`letter_distance("abc", "bcd")` debería devolver `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "bcd"), 3)`)
}})
```

`letter_distance("abc", "xyz")` debería devolver `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "xyz"), 9)`)
}})
```

`letter_distance("encrypt", "decrypt")` debería devolver `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("encrypt", "decrypt"), 10)`)
}})
```

`letter_distance("algorithm", "codeblock")` debería devolver `43`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("algorithm", "codeblock"), 43)`)
}})
```

`letter_distance("lobster", "penguin")` debería devolver `47`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("lobster", "penguin"), 47)`)
}})
```

`letter_distance("alligator", "crocodile")` debería devolver `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("alligator", "crocodile"), 55)`)
}})
```

# --seed--

## --seed-contents--

```py
def letter_distance(str1, str2):

    return str1
```

# --solutions--

```py
def letter_distance(str1, str2):
    total = 0
    for a, b in zip(str1, str2):
        diff = abs(ord(a) - ord(b))
        total += min(diff, 26 - diff)
    return total
```
