---
id: 6a22d77ddf034bc4e35b1d5b
title: "Desafío 354: Calificación de contraste 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Dado dos arreglos que representan valores RGB y un booleano que indica si el texto es grande, devuelve la calificación de contraste WCAG usando el siguiente método:

Primero, convierte cada valor RGB a luminancia relativa:

- Divide cada canal `[R, G, B]` entre 255 para obtener un valor entre `0` y `1`
- Aplica la fórmula de corrección gamma a cada canal:
  - Si el valor del canal es menor o igual a `0.04045`: `channel / 12.92`
  - De lo contrario: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcula la luminancia: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Luego, calcula la relación de contraste sumando 0.05 a cada valor de luminancia, y luego dividiendo el más claro entre el más oscuro. El más claro siempre será el primer argumento.

Devuelve la calificación basada en la relación de contraste usando la siguiente tabla:

| Calificación | Texto normal | Texto grande |
|-------------|--------------|--------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | por debajo de 4.5 | por debajo de 3.0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` debería devolver `"AAA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` debería devolver `"AA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` debería devolver `"Fail"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
