---
id: 6a22d77ddf034bc4e35b1d54
title: "Desafío 347: teoría de juegos"
challengeType: 29
dashedName: challenge-347
---

# --description--

Dadas dos cadenas de igual longitud que representan las estrategias de dos jugadores para un juego, devuelve las puntuaciones como un arreglo `[player1, player2]`.

- Las cadenas dadas solo contendrán una de dos letras: `"C"` (cooperar) o `"D"` (traicionar).
- Cada carácter representa una ronda, puntuándose de la siguiente manera:
  - Si ambos jugadores cooperan, cada uno obtiene 3 puntos.
  - Si ambos jugadores traicionan, cada uno obtiene 1 punto.
  - Si un jugador traiciona y el otro coopera, el traidor obtiene 5 puntos y el cooperador 0.

# --hints--

`play_game("CCCC", "CCCC")` debería devolver `[12, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` debería devolver `[4, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` debería devolver `[5, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` debería devolver `[24, 34]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` debería devolver `[66, 21]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
