---
id: 6a26df3e2988bcdded204897
title: "Desafío 361: hora hablada"
challengeType: 29
dashedName: challenge-361
---

# --description--

Dadas las posiciones en grados de las manecillas de la hora y los minutos de un reloj analógico (en sentido horario desde el 12), devuelve la hora en inglés hablado.

Convierte el ángulo de la manecilla de los minutos a minutos (360° = 60 minutos), luego usa las siguientes reglas:

| Minutos | Hablado |
|---------|---------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (excluyendo 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (excluyendo 45) | `"X minutes to Z"` (donde `X` es 60 - minutos) |

Donde `Y` es la hora actual y `Z` es la siguiente hora, ambas derivadas del ángulo de la manecilla de la hora (360° = 12 horas).

Nota: Los ángulos de las manecillas pueden no coincidir exactamente con un número, considera redondearlos de alguna manera.

# --hints--

`get_spoken_time(90, 0)` debería devolver `"3 o'clock"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` debería devolver `"20 minutes past 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` debería devolver `"half past 8"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` debería devolver `"quarter past 2"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` debería devolver `"20 minutes to 7"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` debería devolver `"quarter to 11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` debería devolver `"5 minutes to 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
