---
id: 6a26df3e2988bcdded204895
title: "Desafío 359: Calculadora de hándicap de golf"
challengeType: 29
dashedName: challenge-359
---

# --description--

Dado un arreglo de puntuaciones de golf y un arreglo correspondiente de valores par del campo, devuelve el índice de hándicap del golfista usando el siguiente método:

- Calcula la diferencia para cada ronda restando el par a la puntuación, luego devuelve el promedio de todas las diferencias redondeado a un decimal.

# --hints--

`calculate_handicap([72, 72, 72], [72, 72, 72])` debería devolver `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([72, 72, 72], [72, 72, 72]), 0)`)
}})
```

`calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72])` debería devolver `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72]), 6)`)
}})
```

`calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36])` debería devolver `8.3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3)`)
}})
```

`calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` debería devolver `8.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8)`)
}})
```

`calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` debería devolver `11.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_handicap(scores, pars):

    return scores
```

# --solutions--

```py
def calculate_handicap(scores, pars):
    differentials = [s - p for s, p in zip(scores, pars)]
    average = sum(differentials) / len(differentials)
    return int(average * 10 + 0.5) / 10
```
