---
id: 68a23004e96a105df3a21fc4
title: 建置一個計牌助手
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

在賭場遊戲 Blackjack 中，玩家可以透過記錄牌堆中剩餘的高牌和低牌的相對數量，來判斷他們在下一手牌中是否對莊家有優勢。這稱為 Card Counting。

牌堆中剩餘較多高牌對玩家有利。當計數為正時，玩家應該下注較高。當計數為零或負時，玩家應該下注較低。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該使用 `let` 宣告一個名為 `count` 的全域變數並將其設為 `0`。
1. 你應該有一個名為 `cardCounter` 的函式（程式）。
1. `cardCounter` 函式（程式）應該接收一個 `card` 參數，該參數可以是數字或字串。
    - 對於介於 `2` 到 `10` 之間的值，`card` 參數將會是數字。
    - 對於所有其他值，`card` 參數將會是字串。
1. `cardCounter` 函式（程式）應該根據特定條件修改全域的 `count` 變數。
1. 全域的 `count` 變數應該為卡片 `2`、`3`、`4`、`5` 或 `6` 增加 `1`。
1. 全域的 `count` 變數應保持不變，針對卡片 `7`、`8`、`9`。
1. 全域的 `count` 變數應該為卡片 `10`、`"J"`、`"Q"`、`"K"`、`"A"` 減少 `1`。
1. `cardCounter` 函式（程式）應該傳回一個包含目前計數和字串 `Bet` 的字串，若計數為正數。
1. `cardCounter` 函式（程式）應該傳回一個包含目前計數的字串，且如果計數小於或相等於 `0`，則字串中應包含 `Hold`。
1. 在函式（程式）輸出中，當前計數和玩家的決定（`Bet` 或 `Hold`）應該以空格分開。例如，`-3 Hold`。

# --hints--

你應該使用 `let` 宣告一個名為 `count` 的全域變數並將其設為 `0`。 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

你應該有一個名為 `cardCounter` 的函式（程式）。 

```js
assert.isFunction(cardCounter);
```

你的函式應該傳回 `count` 的值和文字（`Bet` 或 `Hold`），中間有一個空白字元。

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

在牌 `2`、`3`、`4`、`5` 之後，呼叫 `cardCounter(6)` 應該傳回字串 `5 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

在牌 `7`、`8` 之後，呼叫 `cardCounter(9)` 應該傳回字串 `0 Hold`。

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

在牌 `10`、`"J"`、`"Q"`、`"K"` 之後，呼叫 `cardCounter("A")` 應該傳回字串 `-5 Hold`。

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

在牌 `3`、`7`、`"Q"`、`8` 之後，呼叫 `cardCounter("A")` 應該傳回字串 `-1 Hold`。

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

在牌組 `2`、`"J"`、`9`、`2` 之後，呼叫 `cardCounter(7)` 應該傳回字串 `1 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

在 `2`、`2` 這兩張牌之後，呼叫 `cardCounter(10)` 應該會傳回字串 `1 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

在牌 `3`、`2`、`"A"`、`10` 之後，呼叫 `cardCounter("K")` 應該傳回字串 `-1 Hold`。

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
