---
id: 616968c2c94c8071b349c146
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

All'interno del `body`, annida un elemento `h1` con il testo `CSS Color Markers`.

# --hints--

Il tuo codice dovrebbe avere un elemento `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Dovresti avere un tag `h1` di apertura.

```js
assert.match(code, /<h1\s*>/i);
```

Dovresti avere un tag `h1` di chiusura.

```js
assert.match(code, /<\/h1\s*>/i);
```

Il tuo elemento `h1` dovrebbe essere all'interno dell'elemento `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Il tuo elemento `h1` dovrebbe avere il testo `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Ricorda, la maiuscola e la minuscola e l'ortografia sono importanti per il testo `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
