---
id: 61696ef7ac756c829f9e4048
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

In questo progetto lavorerai con un file CSS esterno per definire lo stile della pagina. Un file `styles.css` è già stato creato per te. Ma prima di poterlo usare, devi collegarlo alla pagina.

Annida un elemento `link` all'interno dell'elemento `head`. Dagli un attributo `rel` impostato su `"stylesheet"` e un attributo `href` impostato su `"styles.css"`.

# --hints--

Dovresti avere un elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Il tuo elemento `link` dovrebbe essere un elemento vuoto, non dovrebbe avere un tag di chiusura `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Il tuo elemento `link` dovrebbe essere dentro l'elemento `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
