---
id: 616d595270d902f0e7086e18
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Anche se hai tre elementi marker `div` separati, sembrano un unico grande rettangolo. Dovresti aggiungere un po’ di spazio tra di loro per rendere più facile vedere ogni elemento.

Quando la proprietà abbreviata `margin` ha due valori, imposta `margin-top` e `margin-bottom` sul primo valore, e `margin-left` e `margin-right` sul secondo valore.

Nella tua regola CSS `.marker`, imposta la proprietà `margin` su `10px auto`.

# --hints--

La tua regola CSS `.marker` dovrebbe avere una proprietà `margin` impostata su `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
