---
id: 616d55bd160a95e22453a081
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora che hai un marcatore centrato con colore, è il momento di aggiungere gli altri marcatori.

Nell'elemento `.container`, aggiungi altri due elementi `div` e assegna a ciascuno la classe `marker`.

# --hints--

Il tuo primo nuovo elemento `div` dovrebbe avere un tag di apertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

Il tuo primo nuovo elemento `div` dovrebbe avere un tag di chiusura.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Il tuo secondo nuovo elemento `div` dovrebbe avere un tag di apertura.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

Il tuo secondo nuovo elemento `div` dovrebbe avere un tag di chiusura.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

I tuoi nuovi elementi `div` dovrebbero essere all'interno del `div` con la classe `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

I tuoi nuovi elementi `div` dovrebbero entrambi avere la classe `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
