---
id: 616d50b93ba424d6282c99cf
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Il colore di sfondo è stato applicato, ma poiché l'elemento marker `div` non ha contenuto, per impostazione predefinita non ha alcuna altezza.

Nella tua regola CSS `.marker`, imposta la proprietà `height` su `25px` e la proprietà `width` su `200px`

# --hints--

La tua regola CSS `.marker` dovrebbe avere una proprietà `width` impostata su `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.width, '200px');
```

La tua regola CSS `.marker` dovrebbe avere una proprietà `height` impostata su `25px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
