---
id: 617b65579ce424bf5f02ca73
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Come hai imparato nelle lezioni, ci sono due modelli di colore principali: il modello RGB (rosso, verde, blu) <dfn>additivo</dfn>, usato nei dispositivi elettronici, e il modello CMYK (ciano, magenta, giallo, nero) <dfn>sottrattivo</dfn>, usato nella stampa.

In questo progetto lavorerai con il modello RGB. Questo significa che i colori partono dal nero e cambiano man mano che vengono introdotti diversi livelli di rosso, verde e blu. Un modo semplice per vedere questo è con la funzione CSS `rgb`.

Crea una nuova regola CSS che abbia come obiettivo la classe `container` e imposta la sua `background-color` su nero con `rgb(0, 0, 0)`.

# --hints--

Dovresti usare un selettore di classe per prendere di mira la classe `container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

La tua regola CSS `.container` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
