---
id: 617b954d9f4f6217a749380e
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

Ora aggiungi un po’ più di spazio verticale tra i tuoi marker e il bordo dell’elemento `container` in cui si trovano.

Nella regola CSS `.container`, usa la proprietà abbreviata `padding` per aggiungere `10px` di spaziatura interna superiore e inferiore, e imposta la spaziatura interna sinistra e destra a `0`. Funziona in modo simile alla proprietà abbreviata `margin` che hai usato prima.

# --hints--

Non dovresti rimuovere la proprietà `background-color` e il suo valore dalla regola CSS `.container`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

La tua regola CSS `.container` dovrebbe avere una proprietà `padding` impostata su `10px 0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
