---
id: 617b8e0d93a8d10d9a90e720
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

Nota che il `background-color` per il tuo marker è ancora rosso. Questo perché hai impostato il valore rosso della funzione `rgb` al massimo di `255`, cioè 100% rosso, e hai impostato sia il valore verde che quello blu a `0`.

Ora usa la funzione `rgb` per impostare gli altri colori.

Nella regola CSS `.two`, usa la funzione `rgb` per impostare il `background-color` al valore massimo per il verde, e `0` per gli altri valori. E nella regola CSS `.three`, usa la funzione `rgb` per impostare il `background-color` al valore massimo per il blu, e `0` per gli altri valori.

# --hints--

La tua regola CSS `.two` non dovrebbe usare la parola chiave di colore `green` per impostare la proprietà `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

La tua regola CSS `.three` non dovrebbe usare la parola chiave di colore `blue` per impostare la proprietà `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

La tua regola CSS `.three` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
