---
id: 617bc3386dc7d07d6469bf20
title: Passaggio 27
challengeType: 0
dashedName: step-27
---

# --description--

Nota che, per creare l'arancione, hai aumentato l'intensità del rosso e diminuito l'intensità del verde nei valori `rgb`. Questo perché l'arancione si ottiene mescolando il rosso con il giallo.

Per creare il colore terziario verde primavera, combina il ciano con il verde. Aggiorna la funzione `rgb` nella regola CSS `.two` impostando il verde al suo valore massimo e il blu a `127`.

# --hints--

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 255, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
