---
id: 617bd6ec666b1da2587e4e37
title: Passaggio 29
challengeType: 0
dashedName: step-29
---

# --description--

Ci sono altri tre colori terziari: chartreuse green (verde + giallo), azure (blu + ciano) e rose (rosso + magenta).

Per creare chartreuse green, aggiorna la funzione `rgb` nella regola CSS `.one` in modo che il rosso sia a `127` e imposta il verde al valore massimo.

Per azure, aggiorna la funzione `rgb` nella regola CSS `.two` in modo che il verde sia a `127` e il blu sia al valore massimo.

E per rose, che a volte viene chiamato bright pink, aggiorna la funzione `rgb` nella regola CSS `.three` in modo che il blu sia a `127` e il rosso sia al valore massimo.

# --hints--

La tua regola CSS `.one` dovrebbe avere una proprietà `background-color` impostata su `rgb(127, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 127, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

La tua regola CSS `.three` dovrebbe avere una proprietà `background-color` impostata su `rgb(255, 0, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
