---
id: 618a132676346ac9f7fd59dd
title: Passaggio 34
challengeType: 0
dashedName: step-34
---

# --description--

Hai notato come i tuoi occhi siano naturalmente attratti dal colore rosso al centro? Quando progetti un sito, puoi usare questo effetto per attirare l'attenzione su intestazioni, pulsanti o collegamenti importanti.

Ci sono diverse altre combinazioni di colori importanti oltre ai colori complementari, ma le imparerai un po' più avanti.

Per ora, usa la funzione `rgb` nella regola CSS `.two` per impostare la `background-color` su nero.

# --hints--

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 0);
}

```
