---
id: 618a00ed1ca871a2b3aca0cb
title: Passaggio 30
challengeType: 0
dashedName: step-30
---

# --description--

Ora che hai esplorato tutti i colori primari, secondari e terziari sulla ruota dei colori, sarà più facile comprendere altri concetti della teoria del colore e come influenzano il design.

Per prima cosa, nelle regole CSS `.one`, `.two` e `.three`, modifica i valori nella funzione `rgb` in modo che la `background-color` di ogni elemento sia impostata su nero puro. Ricorda che la funzione `rgb` usa il modello di colore additivo, dove i colori partono dal nero e cambiano man mano che aumentano i valori di rosso, verde e blu.

# --hints--

La tua regola CSS `.one` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(0, 0, 0)');
```

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

La tua regola CSS `.three` dovrebbe avere una proprietà `background-color` impostata su `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
