---
id: 618a16873520a8d088ffdf44
title: Passaggio 35
challengeType: 0
dashedName: step-35
---

# --description--

Nella regola CSS `h1`, rimuovi la proprietà `background-color` e il suo valore per tornare al colore bianco predefinito.

# --hints--

La tua regola CSS `h1` non dovrebbe avere una proprietà o un valore `background-color`.

```js
const backgroundColorInstances = code.match(/background-color:.*;/gi);
assert.lengthOf(backgroundColorInstances, 4);
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('h1')?.backgroundColor);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
