---
id: 61a498399534644f59cff05e
title: Passaggio 48
challengeType: 0
dashedName: step-48
---

# --description--

Come puoi vedere, la funzione `linear-gradient` ha prodotto un gradiente fluido dal rosso al verde. Mentre la funzione `linear-gradient` richiede un minimo di due argomenti colore per funzionare, può accettarne molti.

Usa la funzione `rgb` per aggiungere il blu puro come terzo argomento colore alla funzione `linear-gradient`.

# --hints--

La tua regola CSS `.red` dovrebbe avere una proprietà `background` impostata su `linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
