---
id: 61a5d7ef1cfcf45764df07a2
title: Passaggio 65
challengeType: 0
dashedName: step-65
---

# --description--

Ora che i pennarelli hanno i colori corretti, è il momento di costruire le coperture per i pennarelli. Inizia con il pennarello rosso.

All'interno dell'elemento `div` del pennarello rosso, crea un nuovo elemento `div` e assegnagli una classe `sleeve`.

# --hints--

Il tuo nuovo elemento `div` dovrebbe trovarsi all'interno dell'elemento `div` del pennarello rosso.

```js
const redMarkerChildren = [...document.querySelector('.red')?.children];
redMarkerChildren.forEach(child => {
  assert.strictEqual(child?.localName, 'div')
})
assert.lengthOf(redMarkerChildren, 1)
```

Il tuo nuovo elemento `div` dovrebbe avere una classe `sleeve`.

```js
const redMarkerChild = [...document.querySelector('.red')?.children][0];
assert.isTrue(redMarkerChild?.classList?.contains('sleeve'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
