---
id: 61b30995968123ceb6b76167
title: Passaggio 83
challengeType: 0
dashedName: step-83
---

# --description--

Nota che i bordi dell'ombra sono netti. Questo perché c'è un valore `blurRadius` opzionale per la proprietà `box-shadow`:

```css
box-shadow: offsetX offsetY blurRadius color;
```

Se non viene incluso un valore `blurRadius`, viene usato il valore predefinito `0` e si ottengono bordi netti. Più alto è il valore di `blurRadius`, maggiore sarà l'effetto di sfocatura.

Nella regola CSS `.green`, aggiungi la proprietà `box-shadow` con i valori `5px` per `offsetX`, `5px` per `offsetY`, `5px` per `blurRadius` e `green` per `color`.

# --hints--

La tua regola CSS `.green` dovrebbe avere una proprietà abbreviata `box-shadow` con il valore `5px 5px 5px green`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow, 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
--fcc-editable-region--
  
--fcc-editable-region--
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
