---
id: 68b94f23f7b3ee14078c89c5
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você usará o elemento `iframe` para exibir um vídeo. O boilerplate básico de HTML foi preparado para você.

Comece criando um elemento `h1` com o texto `iframe Video Display`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Seu elemento `h1` deve ter o texto `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
