---
id: 68ba9c4f1688914d72876097
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Adicione `encrypted-media`, `gyroscope` e `web-share` aos valores existentes no atributo `allow`.

Esses três permitirão o uso de extensões de mídia criptografadas para proteger o vídeo, conceder acesso aos sensores de movimento e orientação do dispositivo e permitir o compartilhamento do conteúdo do iframe através dos diálogos nativos de compartilhamento do dispositivo. 

# --hints--

O atributo `allow` do seu elemento `iframe` deve ter `encrypted-media` como um dos seus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

O atributo `allow` do seu elemento `iframe` deve ter `gyroscope` como um dos seus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

O atributo `allow` do seu elemento `iframe` deve ter `web-share` como um dos seus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
