---
id: 68ba9c4f1688914d72876093
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Agora, crie um elemento `iframe`. Ainda não coloque nada nele.

# --hints--

Você deve ter um elemento `iframe`.

```js
assert.exists(document.querySelector('iframe'))
```

Seu elemento `iframe` deve estar vazio.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
