---
id: 68ba9c4f1688914d72876098
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

O próximo atributo que você vai adicionar é `referrerpolicy`. É a regra que determina quanta informação você compartilha quando sua página se conecta a outra página.

Adicione o atributo `referrerpolicy` e defina-o como `strict-origin-when-cross-origin`. Isso compartilha o endereço completo no mesmo site, apenas o nome do site em outros sites e nada em sites inseguros.

# --hints--

Seu elemento `iframe` deve ter um atributo `referrerpolicy` definido como `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
