---
id: 68ba9c4f1688914d72876095
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

O elemento `iframe` também aceita um atributo `src` com um valor que indica a URL ou o caminho do recurso a ser exibido.

Adicione um atributo `src` com o valor `https://www.youtube.com/embed/I0_951_MPE0` ao seu elemento `iframe`.

Neste ponto, você deve ver o vídeo sendo exibido na página, mas há mais alguns atributos que você precisa adicionar.

# --hints--

Seu elemento `iframe` deve ter um atributo `src` definido como `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
