---
id: 68ba9c4f1688914d72876094
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Na primeira lição sobre o elemento `iframe`, você aprendeu que ele é um elemento substituído assim como o `img`. Isso significa que ele também pode receber as propriedades `width` e `height` para determinar quão alto e largo ele deve ser.

Dê ao seu elemento `iframe` uma `width` de `560` e uma `height` de `315`.

# --hints--

Seu elemento `iframe` deve ter um atributo `width` definido como `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Seu elemento `iframe` deve ter um atributo `height` definido como `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
