---
id: 68ba9c4f1688914d72876096
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Um dos atributos é `allow`. É como uma lista de permissões que informa ao navegador quais recursos o `iframe` tem permissão para usar.

Aqui está um elemento `iframe` com o atributo `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Adicione o atributo `allow` com o valor `accelerometer`, `autoplay` e `clipboard-write`.

`accelerometer` permite que o `iframe` use sensores de movimento para que ele possa detectar coisas como inclinação e rotação do dispositivo. `autoplay` permite que o vídeo comece a ser reproduzido automaticamente e `clipboard-write` permite que o iframe escreva dados na área de transferência do usuário.

# --hints--

Seu elemento `iframe` deve ter um atributo `allow` definido como `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

O atributo `allow` do seu elemento `iframe` deve ter `autoplay` como um dos seus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

O atributo `allow` do seu elemento `iframe` deve ter `clipboard-write` como um dos seus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"

    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```
