---
id: 69bac067cd15c4685d07906d
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

個々の目盛りを処理し、オーバーフローを扱う関数ができたので、次は単一のゲート用のコントローラー関数を作成しましょう。この関数はシミュレーションの1つの目盛りの間に、1つのゲートに対するすべての操作を処理します。

次のパラメータを持つ空の関数`handleGateAtTick`を作成してください。

- `gates`: ゲートオブジェクトの完全な配列
- `gate`: 現在処理中のゲート
- `tickIndex`: シミュレーションの現在の目盛りインデックス
- `throughputSummary`: ゲートごとの合計処理済み参加者数を追跡するオブジェクト

# --hints--

`handleGateAtTick`という名前の関数を作成してください。

```js
assert.isFunction(handleGateAtTick);
```

`handleGateAtTick`関数は`gates`、`gate`、`tickIndex`、`throughputSummary`のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
