---
id: 69bac027df64d4433c708d44
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

オーバーフローをリルーティングするときは、参加者を`gates`配列の次のゲートに送る必要があります。そのためには、次のゲートのインデックスを見つける必要があります。通常は現在のインデックスに1を加えればよいですが、現在のゲートが配列の最後の場合は、最初のゲートに戻る必要があります。

これを処理するには、モジュラス演算子（`%`）を使うことができます。

```js
const nextIndex = (currentIndex + 1) % array.length;
```

これは、`currentIndex + 1`が`array.length`と等しいとき、結果が0になるため機能します。

この方法を使って、`nextGateIndex`配列の次のゲートのインデックスをストアする`gates`という名前の変数を作成してください。

# --hints--

モジュラスの手法を使って、`currentIndex`配列の`gates`から次のインデックスにアクセスし、`nextGateIndex`という名前の変数に代入してください。例の1行パターンに従ってください。

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
