---
id: 69baafd6d683dd9921b75db3
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptを使って音楽フェスティバルの参加者の流れをシミュレートします。

朝と夜のシフト中のフェスティバルのゲートを表す2つのデータセットを扱います。

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

各ゲートオブジェクトには次のプロパティが含まれています：

- `id`：ゲートを識別する文字列です。
- `capacity`：ゲートが1ティックあたり処理できる参加者の数です。
- `queue`：特定のティックでゲートに到着する参加者数を表す数値の配列です。

まず、`gates`というパラメータを持つ空の関数`initializeThroughput`を作成してください。

# --hints--

`initializeThroughput`という名前の関数を作成してください。

```js
assert.isFunction(initializeThroughput);
```

`initializeThroughput`関数には`gates`というパラメータが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { initializeThroughput } = explorer.allFunctions;
const parameters = initializeThroughput?.parameters;
assert.lengthOf(parameters, 1);
assert.equal(parameters[0].toString(), "gates");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
