---
id: 69bac18fe3729f6761344365
title: ステップ 27
challengeType: 1
dashedName: step-27
---

# --description--

シミュレーションが何目盛り（ticks）実行されるかを知る必要があります。次のことを思い出してください：

- 各ゲートには、各目盛りで何人の参加者が到着するかを示す `queue` 配列があります。
- すべてのゲートが同じ数の目盛りを持つため、最初のゲートの `queue` 配列の長さを使えば十分です。

`maxTicks`という名前の変数を作成し、`gates[0].queue.length`の値を代入してください。

この値はシミュレーションループの制御に使われ、すべてのゲートで各目盛りが処理されることを保証します。

# --hints--

`maxTicks`という名前の変数を作成し、`gates[0].queue.length`の値を代入してください。

```js
const explorer = await __helpers.Explorer(code);
const { simulateFestival } = explorer.allFunctions;
const { maxTicks } = simulateFestival?.variables ?? {};
assert.isTrue(maxTicks?.value.matches("gates[0].queue.length"));
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
  const throughputSummary = initializeThroughput(gates);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
