---
id: 69bac067cd15c4685d07906d
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

個々の目盛りを処理し、オーバーフローを扱う関数ができたので、次は単一のゲート用のコントローラー関数を作成しましょう。この関数はシミュレーションの1つの目盛りで1つのゲートに対するすべての処理を担当します。

`handleGateAtTick` という名前の空の関数を、以下のパラメータで作成してください。

- `gates`: ゲートオブジェクトの完全な配列です。
- `gate`: 現在処理中のゲートです。
- `tickIndex`: `queue` 配列内の現在の目盛り位置です。
- `throughputSummary`: ゲートごとに処理された参加者の合計を追跡するオブジェクトです。

# --hints--

`handleGateAtTick`という名前の関数を作成してください。

```js
assert.isFunction(handleGateAtTick);
```

`handleGateAtTick`関数は`gates`、`gate`、`tickIndex`、`throughputSummary`のパラメータを持つべきです。

```js
const explorer = await __helpers.Explorer(code);
const { handleGateAtTick } = explorer.allFunctions;
const parameters = handleGateAtTick?.parameters;
assert.lengthOf(parameters, 4);
assert.equal(parameters[0].toString(), "gates");
assert.equal(parameters[1].toString(), "gate");
assert.equal(parameters[2].toString(), "tickIndex");
assert.equal(parameters[3].toString(), "throughputSummary");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
