---
id: 69babfd65c4c4632a0c76f63
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

ゲートが1回のティックで全ての参加者を処理できない場合、いくつかの参加者がキューに残ります。このオーバーフローを処理するために、オーバーフローした参加者を別のゲートにルーティングする関数を作成します。

次のパラメータを持つ空の関数 `rerouteOverflow` を作成してください。

- `gates`: ゲートオブジェクトの完全な配列。
- `currentGate`: オーバーフローが発生したゲート。
- `tickIndex`: キュー内の現在のティック位置。
- `overflowAmount`: 処理できなかった参加者の数。

# --hints--

`rerouteOverflow` という名前の関数を作成してください。

```js
assert.isFunction(rerouteOverflow);
```

`rerouteOverflow` 関数は `gates`、`currentGate`、`tickIndex`、`overflowAmount` のパラメータを持つ必要があります。

```js
const regex = __helpers.functionRegex('rerouteOverflow', ['gates', 'currentGate', 'tickIndex', 'overflowAmount']);
assert.match(__helpers.removeJSComments(rerouteOverflow.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
