---
id: 69babfc35c3c58d0993d404d
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

`processGateFlow` 関数を完成させるには、2つのプロパティを含むオブジェクトを返してください。

- `processed`：このティック中にゲートが処理できた参加者の数。
- `overflow`：処理後にキューに残っている参加者の数（ループ後の `currentTickQueue` と同じ）。

オブジェクトは次のように直接返してください。

```js
return {
  firstKey: firstValue,
  secondKey: secondValue
}
```

# --hints--

`processed` と `overflow` のプロパティを持つオブジェクトを返す必要があります。変数に代入せず、オブジェクトを直接返してください。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(processGateFlow.toString()));
const viableSolutions = [
  "return{processed:processed,overflow:currentTickQueue}}",
  "return{processed:processed,overflow:currentTickQueue,}}",
  "return{overflow:currentTickQueue,processed:processed,}}",
  "return{overflow:currentTickQueue,processed:processed}}",
  "return{processed:processed,overflow:currentTickQueue};}",
  "return{processed:processed,overflow:currentTickQueue,};}",
  "return{overflow:currentTickQueue,processed:processed,};}",
  "return{overflow:currentTickQueue,processed:processed};}",
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your processGateFlow function must return the object with processed and overflow properties."
);

const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert(
  typeof result === 'object' && result !== null && !Array.isArray(result),
  'Your function should return an object'
);
assert.hasAllKeys(
  result,
  ['processed', 'overflow'],
  'Your returned object should have properties `processed` and `overflow`'
);
```

返すオブジェクトの `processed` プロパティは `processed` の値を持つべきです。

```js
const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert.strictEqual(
  result.processed,
  3,
  'For the North gate at tick 0 with capacity 5 and queue 3, processed should be 3'
);
const westGate = morningGates[3];
const resultWest = processGateFlow(westGate, 0);
assert.strictEqual(
  resultWest.processed,
  2,
  'For the West gate at tick 0 with capacity 2 and queue 4, processed should be 2'
);
```

返すオブジェクトの `overflow` プロパティは `currentTickQueue` の値を持つべきです。

```js
const northGate = morningGates[0];
const result = processGateFlow(northGate, 0);
assert.strictEqual(
  result.overflow,
  0,
  'For the North gate at tick 0 with capacity 5 and queue 3, overflow should be 0'
);
const westGate = morningGates[3];
const resultWest = processGateFlow(westGate, 0);
assert.strictEqual(
  resultWest.overflow,
  2,
  'For the West gate at tick 0 with capacity 2 and queue 4, overflow should be 2'
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
