---
id: 69babdf3f896711c11516373
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

各シミュレーションのティックごとに：

- 一定数の参加者がゲートに到着します（その`queue`から）。
- ゲートは`capacity`に基づいて参加者を処理します。
- ゲートが処理できる数より多くの参加者が到着した場合、一部は残ります（オーバーフロー）。

これから、1つのゲートで1回のティックの間にこのロジックを処理する関数を作成します。

`processGateFlow` という名前の空の関数を、次のパラメータで作成してください。

- `gate`: 処理されるゲートオブジェクト。
- `tickIndex`: `queue` 配列内の現在のティック位置。

# --hints--

`processGateFlow`という名前の関数を作成してください。

```js
assert.isFunction(processGateFlow);
```

あなたの `processGateFlow` 関数は、`gate` と `tickIndex` のパラメータを持つ必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
