---
id: 69cfdc80443b715a8aa452a7
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`initializeThroughput` 関数を完成させるには：

- 入力の `gates` 配列内の各ゲートを繰り返し処理してください。各ゲートについて、`summary` にプロパティを追加します。ここで：
  - キーはゲートの `id` です。
  - 値は `0` です。
- 繰り返し処理の後、`summary` オブジェクトを返してください。

`initializeThroughput` 関数は、元の入力配列 `gates` を変更してはいけません。

# --hints--

`initializeThroughput()` を使って `[]` を呼び出すと、元の入力を変更せずに `{}` を返すはずです。

```js
const input = [];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  {},
  "Calling `initializeThroughput([])` should return an empty object."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given an empty array."
);
```

`initializeThroughput()` を使って `[{ id: "North", capacity: 5, queue: [1, 2, 3] }]` を呼び出すと、元の入力を変更せずに `{ North: 0 }` を返すはずです。

```js
const input = [{ id: "North", capacity: 5, queue: [1, 2, 3] }];
const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0 },
  "Calling `initializeThroughput([{ id: 'North', ... }])` should return `{ North: 0 }`."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given one gate."
);
```

`initializeThroughput()` を使って `[{ id: "North", capacity: 5, queue: [] }, { id: "East", capacity: 3, queue: [] }, { id: "South", capacity: 4, queue: [] }]` を呼び出すと、元の入力を変更せずに `{ North: 0, East: 0, South: 0}` を返すはずです。

```js
const input = [
  { id: "North", capacity: 5, queue: [] },
  { id: "East", capacity: 3, queue: [] },
  { id: "South", capacity: 4, queue: [] }
];

const original = JSON.parse(JSON.stringify(input));
const result = initializeThroughput(input);

assert.deepEqual(
  result,
  { North: 0, East: 0, South: 0 },
  "Calling `initializeThroughput` with multiple gates should return an object with an entry for each gate id set to 0."
);

assert.deepEqual(
  input,
  original,
  "`initializeThroughput` should not mutate the source array when given multiple gates."
);
```

入力の `for` 配列内の各ゲートオブジェクトを繰り返し処理するには、`gates` ループを使うべきです。

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));
assert.match(cleaned, /for\s*\(/);
```

繰り返し処理の後、`summary` を返すべきです。これは関数の最後の行になります。

```js
const funcStr = __helpers.removeWhiteSpace(__helpers.removeJSComments(initializeThroughput.toString()));

assert.match(
  funcStr,
  /returnsummary;?}$/,
  "Your `initializeThroughput` function should end with `return summary`."
);
```

`initializeThroughput` 関数は、元の入力配列 `gates` を変更してはいけません。

```js
const morningCopy = JSON.parse(JSON.stringify(morningGates));
const nightCopy = JSON.parse(JSON.stringify(nightGates));

initializeThroughput(morningGates);
initializeThroughput(nightGates);

assert.deepEqual(
  morningGates,
  morningCopy,
  "`initializeThroughput` should not mutate the `morningGates` seed array."
);

assert.deepEqual(
  nightGates,
  nightCopy,
  "`initializeThroughput` should not mutate the `nightGates` seed array."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
--fcc-editable-region--

--fcc-editable-region--
}
```
