---
id: a77dbc43c33f39daa4429b4f
title: 建置布林值檢查函式（程式）
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

在本實驗中，你將建置一個函式（程式），用來檢查一個值是否被分類為布林值原始型別。

布林值原始值是 `true` 和 `false`。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `booWho` 的函式（程式），它接收一個引數。
1. 如果接收到的引數是布林值原始型態，該函式應該傳回 `true`。
1. 如果引數是任何其他值，函式（程式）應該傳回 `false`。

# --hints--

你應該有一個 `booWho` 函式（程式）。

```js
assert.isFunction(booWho);
```

`booWho(true)` 應該傳回 `true`。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` 應該傳回 `true`。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` 應該傳回 `false`。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` 應該傳回 `false`。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` 應該傳回 `false`。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` 應該傳回 `false`。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` 應該傳回 `false`。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` 應該傳回 `false`。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` 應該傳回 `false`。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` 應該傳回 `false`。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
