---
id: a26cbbe9ad8655a977e1ceb5
title: Construir um Aplicativo Localizador da Palavra Mais Longa
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

Neste laboratório, você construirá uma função que retorna o comprimento da palavra mais longa na sentença fornecida.

Por exemplo, na frase `"The quick brown fox jumped over the lazy dog"`, a palavra mais longa é `"jumped"`, que tem um comprimento de 6.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `findLongestWordLength` que recebe uma string como argumento.
2. A função deve retornar o comprimento da palavra mais longa na string.

# --hints--

Você deve criar uma função chamada `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` deve ter um único parâmetro.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` deve retornar um número.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` deve retornar `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` deve retornar `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` deve retornar `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` deve retornar `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` deve retornar `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
