---
id: 67fccc2463253b213a000142
title: Створіть мультимедійний плеєр
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

У попередніх відеолекціях ви ознайомились з принципами роботи елементів `audio` та `video`. У цій лабораторній роботі ви створите мультимедійний плеєр, який відтворюватиме `audio` та `video` з текстовою версією.

До елемента `audio` потрібно буде застосувати елемент `source`, який вказує використане медіа.

Ось приклад:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

Елемент `source` також можна застосувати до елемента `video`:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `h1` з головним заголовком сторінки.
2. Ви повинні мати три елементи `section`.
3. В межах першого елемента `section` має бути елемент `h2` з назвою пісні.
4. Під елементом `h2` має бути елемент `audio` з атрибутами `controls` та `aria-label`.
5. В межах елемента `audio` має бути елемент `source` з атрибутом `src`, що вказує на аудіофайл, та атрибутом `type`. Якщо хочете, можете використати це посилання на аудіо: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. В межах другого елемента `section` має бути елемент `h2` з назвою відео.
7. Під елементом `h2` має бути елемент `video` з атрибутами `controls`, `width` та `aria-label`.
8. В межах елемента `video` має бути елемент `source` з атрибутом `src`, що вказує на відеофайл, та атрибутом `type`. Якщо хочете, можете використати це посилання на відео: `https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4`
9. Після елемента `source` має бути елемент `track` з атрибутом `src`, що вказує на файл з субтитрами, а також з атрибутами `kind`, `srclang` та `label`.
10. В межах третього елемента `section` має бути елемент `h2` з заголовком розділу (наприклад, «Transcript»).
11. Під елементом `h2` має бути елемент `p` з текстовою версією відео.

# --hints--

Ви повинні мати елемент `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Елемент `h1` повинен містити головний заголовок сторінки.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

Ви повинні мати три елементи `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

Перший елемент `section` повинен містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Елемент `h2` повинен містити назву пісні, яка виконується.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

Перший елемент `section` повинен містити елемент `audio`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Елемент `audio` повинен мати атрибут `controls`.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Елемент `audio` повинен мати атрибут `aria-label`, що описує аудіо.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Елемент `audio` повинен містити елемент `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Елемент `source` повинен мати атрибут `src`, що веде до аудіофайлу.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Елемент `source` повинен мати атрибут `type`, що вказує тип аудіо.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

Другий елемент `section` повинен містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Елемент `h2` повинен містити назву відео, яке виконується.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Другий елемент `section` повинен містити елемент `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Елемент `video` повинен мати атрибут `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Елемент `video` повинен мати атрибут `width`, що вказує ширину відео.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

Елемент `video` повинен містити елемент `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

Елемент `source` повинен мати атрибут `src`, що веде до відеофайлу.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

Елемент `source` повинен мати атрибут `type`, що вказує тип відео.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Елемент `video` повинен містити елемент `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Елемент `track` повинен мати атрибут `src`, що веде до файлу з субтитрами.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Елемент `track` повинен мати атрибут `kind` з текстом, що описує тип файлу.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Елемент `track` повинен мати атрибут `srclang`, що вказує мову субтитрів.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Елемент `track` повинен мати атрибут `label`, що вказує мову субтитрів.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

Третій елемент `section` повинен містити елемент `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Елемент `h2` повинен містити заголовок розділу.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Третій елемент `section` повинен містити елемент `p` з текстовою версією відео.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4" type="video/mp4">

      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
