---
id: 63ee3fe4381756f9716727f0
title: CSS基礎演習B
challengeType: 25
dashedName: css-foundations-exercise-b
---

# --description--

**目的:** 提供されたHTMLファイルにはいくつかの要素があり、それらにクラスまたはID属性を追加する必要があります。その後、提供されたCSSファイルに正しいセレクター構文を使ってルールを追加してください。

## ユーザーストーリー

1. リストの奇数番目のすべての要素に`yellow`の背景が表示されるはずです。

1. リストの奇数番目のすべての要素に`class`セレクターが使われているはずです。

1. リストの2番目の要素には`blue`のテキストと`font-size`が`36px`であることが見えるはずです。

1. 2番目の要素の`font-size`とテキストの色は`id`属性を使って設定されているはずです。

1. リストの3番目の要素には`font-size`が`24px`であることが見えるはずです。

1. 3番目の要素の`font-size`は`class`属性を使って設定されているはずです。

1. リストの4番目の要素には`red`の背景、`font-size`が`24px`、そして`font-weight`が`bold`であることが見えるはずです。

1. 4番目の要素の`font-size`は`class`属性で設定されているはずです。`font-weight`と色は`id`属性で設定されているはずです。

# --hints--

すべての奇数番目の要素には`class`属性が必要です。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasClass = p?.every((paragraph) => paragraph.classList.length > 0);

assert(everyPHasClass);
```

奇数番目の要素には`background-color`が`yellow`である必要があります。

```js
const p = Array.from(document.querySelectorAll('P'));

const everyPHasBackgroundColor = p?.every((paragraph) => {
  const style = getComputedStyle(paragraph);
  
  return style?.backgroundColor === 'rgb(255, 255, 0)';
})

assert.equal(everyPHasBackgroundColor, true);
```

2番目の要素には`blue`のテキストと`font-size`が`36px`である必要があります。

```js
const secondElementId = document.querySelectorAll('div')?.[0]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${secondElementId}`);

assert.equal(style?.color, 'blue');
assert.equal(style?.fontSize, '36px');
```

3番目の要素にはテキストと`font-size`が`24px`である必要があります。

```js
const thirdElement = document.querySelectorAll('p')?.[1];

assert(thirdElement?.innerText?.length > 0);

const thirdElementClasses = Array.from(thirdElement?.classList?.values());

assert(thirdElementClasses.some(thirdElementClass => {

  const style = new __helpers.CSSHelp(document).getStyle(`.${thirdElementClass}`);

  return style?.fontSize === '24px';

}))

``` 

4番目の要素には`font-size`が`24px`である必要があります。

```js
const fourthElementClass = document.querySelectorAll('div')?.[1]?.classList[0];

const style = new __helpers.CSSHelp(document).getStyle(`.${fourthElementClass}`);

assert(style?.fontSize === '24px');
```

4番目の要素には`red`の`background-color`があります。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.backgroundColor === 'red');
```

4番目の要素には`font-weight`が`bold`である必要があります。

```js
const fourthElement = document.querySelectorAll('div')?.[1]?.id;

const style = new __helpers.CSSHelp(document).getStyle(`#${fourthElement}`);

assert(style?.fontWeight === 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p>Number 1 - I'm a class!</p>
    <div>Number 2 - I'm one ID.</div>
    <p>Number 3 - I'm a class, but cooler!</p>
    <div>Number 4 - I'm another ID.</div>
    <p>Number 5 - I'm a class!</p>
  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta http-equiv="X-UA-Compatible" content="IE=edge">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Class and ID Selectors</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <p class="odd">Number 1 - I'm a class!</p>
    <div id="two">Number 2 - I'm one ID.</div>
    <p class="odd adjust-font-size">Number 3 - I'm a class, but cooler!</p>
    <div id="four" class="adjust-font-size">Number 4 - I'm another ID.</div>
    <p class="odd">Number 5 - I'm a class!</p>
  </body>
</html>
```

```css
.odd {
  background-color: yellow;
  font-family: Verdana, "DejaVu Sans", sans-serif;
}

.adjust-font-size {
  font-size: 24px;
}

#two {
  color: blue;
  font-size: 36px;
}

#four {
  background-color: red;
  font-weight: bold;
}
```
