---
id: 63ee3f71381756f9716727ef
title: CSS基礎演習A
challengeType: 25
dashedName: css-foundations-exercise-a
---

# --description--

**目的:**  
この演習では、外部CSS、内部CSS、インラインCSSの3つの方法すべてを使ってHTMLファイルにCSSを追加する練習をします。外部および内部の方法でスタイルを追加するときは、型セレクターのみを使用してください。また、色には`RGB`や`HEX`の値ではなく、キーワード（例："blue"）を使うことをお勧めします。

## ユーザーストーリー

- `div`要素にテキストが表示されているはずです。  
  - `red`の背景色、`white`のテキスト色、フォントサイズは`32px`、テキストは中央揃えで`bold`であるべきです。  
  - `div`要素のCSSは外部で追加され、型セレクターを使っている必要があります。  
- `p`要素にテキストが表示されているはずです。  
  - `green`の背景色、`white`のテキスト色、フォントサイズは`18px`であるべきです。  
  - `p`要素のCSSは内部で追加され、型セレクターを使っている必要があります。  
- `button`要素にテキストが表示されているはずです。  
  - `button`要素は`orange`の背景色と`18px`のフォントサイズを持つべきです。  
  - `button`要素のCSSはインラインスタイルで追加されている必要があります。

# --hints--

テキストを含む`div`要素が1つあるはずです。

```js
const divElement = document.querySelector('div');

assert.isNotNull(divElement);
assert.isAtLeast(divElement?.innerText.length, 1);
```

`div`要素のスタイルを含む外部スタイルシートがあるはずです。

```js
const styleSheet = new __helpers.CSSHelp(document).getStyleSheet();
const isExternal = styleSheet?.ownerNode.classList.contains('fcc-injected-styles');
const divStyle = new __helpers.CSSHelp(document).getStyle('div');

assert.isTrue(isExternal);
assert.isNotNull(divStyle);
```

`div`要素には内部またはインラインスタイルでCSSを追加してはいけません。

```js
const styleElement = document.querySelector('style:not([class])');

assert.isNotTrue(styleElement?.innerText.includes('div'));
assert.isNotTrue(document.querySelector('div')?.hasAttribute('style'));
```

`div`要素は`background-color`が`red`、`color`が`white`であるべきです。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const divBGColor = divStyle?.getPropertyValue('background-color');
const divColor = divStyle?.getPropertyValue('color');

assert.equal(divBGColor, 'red');
assert.equal(divColor, 'white');
```

`div`要素は`font-weight`が`bold`、`font-size`が`32px`、`text-align`が`center`に設定されているべきです。

```js
const divStyle = new __helpers.CSSHelp(document).getStyle('div');
const textAlign = divStyle?.getPropertyValue('text-align');
const fontSize = divStyle?.getPropertyValue('font-size');
const fontWeight = divStyle?.getPropertyValue('font-weight');

assert.equal(textAlign, 'center');
assert.equal(fontSize, '32px');
assert.equal(fontWeight,'bold');
```

テキストを含む`p`要素が1つあるはずです。

```js
const pElement = document.querySelector('p');

assert.isNotNull(pElement);
assert.isAtLeast(pElement?.innerText.length, 1)
```

`p`要素のスタイルは`style`要素を使って内部で追加されているべきです。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let isStyled = false;

if (rules && rules.selectorText === 'p') {
  isStyled = true;
}

assert.isTrue(isStyled);
```

`p`要素は`font-size`が`18px`で、`color`が`white`に設定されているべきです。

```js
const styleElement = document.querySelector('style:not([class])');
const rules = styleElement?.sheet?.cssRules?.[0] || styleElement?.sheet?.rules?.[0];
let fontSize, color;

if (rules && rules.selectorText === 'p') {
  fontSize = rules.style.fontSize;
  color = rules.style.color;
}

assert.equal(fontSize, "18px");
assert.equal(color, 'white');
```

テキストを含む`button`要素が1つあるはずです。

```js
const btnElement = document.querySelector('button');

assert.isNotNull(btnElement);
assert.isAtLeast(btnElement?.innerText.length, 1);
```

`button`要素はインラインスタイルを持つべきです。

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('style'));
```

`button`要素は`background-color`が`orange`に設定されているべきです。

```js
assert.equal(document.querySelector('button')?.style.backgroundColor, 'orange')
```

`button`要素は`font-size`が`18px`に設定されているべきです。

```js
assert.equal(document.querySelector('button')?.style.fontSize, '18px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>

  </body>
</html>
```

```css
/* styles.css */
```

# --solutions--

```html
<!DOCTYPE html>
<html>
<head>
  <title>My Styling Example</title>
  <link rel="stylesheet" type="text/css" href="styles.css">
  <style>
    p {
      background-color: green;
      color: white;
      font-size: 18px;
    }
  </style>
</head>
<body>
  <div>Hello World!</div>
  <p>This is a paragraph.</p>
  <button style="background-color: orange; font-size: 18px;">Click Me</button>
</body>
</html>
```

```css
div {
  background-color: red;
  color: white;
  font-size: 32px;
  text-align: center;
  font-weight: bold;
}
```
