---
id: 5f3313e74582ad9d063e3a38
title: 步驟2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. The `title` element's text should be `Cafe Menu`.

# --hints--

你需要一個開始標籤 `<head>`

```js
assert(code.match(/<head>/i));
```

需要一個結束標籤 `</head>`

```js
assert(code.match(/<head>/i));
```

需要一個開始標籤 `<title>`

```js
assert(code.match(/<title>/i));
```

需要一個結束標籤 `</title>`

```js
assert(code.match(/<\/title>/i));
```

所有的 `title` 元素都應嵌套在 `head` 元素中。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

你的 `title` 元素內應該有文本 `Cafe Menu`。 您可能需要檢查您的拼寫。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
