---
id: 5f332a88dc25a0fd25c7687a
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

你應該有一個 `<h1>` 開始標籤。

```js
assert.match(code, /<h1>/i);
```

你應該有一個 `</h1>` 結束標籤。

```js
assert.match(code, /<\/h1>/i);
```

你不應該改變你現有的 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你的 `h1` 元素應該嵌套在你的 `main` 元素中。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

你的 `h1` 元素應該有文字 `CAMPER CAFE`。

```js
assert.match(code, /<h1>CAMPER CAFE<\/h1>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

