---
id: 5f33294a6af5e9188dbdb8f3
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

你的代碼應該有一個 `<main>` 開始標籤。

```js
assert(code.match(/<main>/i));
```

你的代碼應該有一個 `</main>` 結束標籤。

```js
assert(code.match(/<\/main>/i));
```

你不應該改變你的 `body` 元素。 確認你沒有不小心刪除結束標籤。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

你的 `main` 元素應該在你的 `body` 元素之後。

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
