---
id: 683ecb010f4c6bafd428802c
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 의료 데이터 집합이 일련의 규칙을 준수하는지 검증할 것입니다.

`medical_records`라는 변수를 선언하고 빈 리스트를 할당하세요. 다음 단계에서 이 변수를 사용해 의료 데이터를 저장할 것입니다.

# --hints--

`medical_records`라는 변수를 선언해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

`medical_records` 변수에 빈 리스트를 할당해야 합니다.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
