---
id: 6842d5b537c64266420faf62
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

정수 `age` 값을 가진 키 `34`와 문자열 `gender` 값을 가진 키 `Female`를 사전에 추가하세요. 키-값 쌍 사이에 쉼표를 잊지 마세요.

# --hints--

`medical_records` 리스트에는 하나의 사전이 들어 있어야 합니다.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

사전에 정수 `age` 값을 가진 `34` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

사전에 문자열 `gender` 값을 가진 `Female` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

사전에 문자열 `patient_id` 값을 가진 `P1001` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
