---
id: 6842d659a4ae7973dabab3b0
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

다음 세 가지 키-값 쌍을 추가하여 딕셔너리를 완성하세요:

- 키 `diagnosis`에 `Hypertension` 값을 할당하세요
- 키 `medications`에 리스트 `['Lisinopril']` 값을 할당하세요
- 키 `last_visit_id`에 `V2301` 값을 할당하세요.

# --hints--

딕셔너리에 문자열 `diagnosis` 값을 가진 `Hypertension` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

딕셔너리에 리스트 `medications` 값을 가진 `['Lisinopril']` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

딕셔너리에 문자열 `last_visit_id` 값을 가진 `V2301` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

딕셔너리에 문자열 `patient_id` 값을 가진 `P1001` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

딕셔너리에 정수 `age` 값을 가진 `34` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

딕셔너리에 문자열 `gender` 값을 가진 `Female` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
