---
id: 6842aee5bcaf4ab8883f8191
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이전 수업에서 배운 것처럼, 딕셔너리는 키-값 쌍을 저장하는 자료 구조이며, 여기서 키는 변경할 수 없는 데이터 유형(예: 문자열)이고 딕셔너리 내에서 고유합니다:

```py
person = {
    'name': 'John',
    'age': 33
}
```

`medical_records` 리스트는 각각 환자를 나타내는 딕셔너리를 저장합니다. `patient_id` 리스트에 키가 `P1001`이고 값이 문자열 `medical_records`인 딕셔너리를 추가하세요.

# --hints--

`medical_records` 리스트에는 딕셔너리 하나만 포함되어야 합니다.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

딕셔너리에는 문자열 `patient_id` 값을 가진 `P1001` 키가 있어야 합니다.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
