---
id: 6842ba0086201089ea268c92
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이전 단계에서 사용한 것과 같은 구조를 따라 `medical_records` 리스트는 다른 환자들의 데이터로 채워져 있습니다. 자유롭게 살펴보세요.

다음으로 데이터 세트를 검증하는 함수를 작성하기 시작할 것입니다. `validate`라는 이름의 함수를 하나의 매개변수 `data`로 선언하세요.

데이터가 리스트나 튜플인지 확인하고 싶습니다. 따라서 `validate` 함수 내에서 `is_sequence`라는 변수를 선언하고 `isinstance`를 호출한 결과를 할당하세요. 첫 번째 인수로 `data`를 전달하고 두 번째 인수로 `list`과 `tuple`을 포함하는 튜플을 전달하세요.

# --hints--

`validate`라는 이름의 함수가 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

`validate` 함수는 하나의 매개변수 `data`을 가져야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

`is_sequence` 함수 내에 `validate`라는 변수가 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

`isinstance(data, (list, tuple))` 변수에 `is_sequence`를 할당해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
