---
id: 6846a33f665a435203da0ad7
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

코드 맨 아래에서 `validate` 함수를 `medical_records` 인수로 호출하세요. 터미널에 `Valid format.`가 출력되는 것을 확인할 수 있습니다.

# --hints--

`validate` 함수를 `medical_records` 인수로 호출해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("validate(medical_records)")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True
            
    if is_invalid:
        return False
    print('Valid format.')
    return True

--fcc-editable-region--

--fcc-editable-region--
```
