---
id: 6846b0048f71fc4efd786b42
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`keys()` 메서드는 딕셔너리의 모든 키를 포함하는 뷰 객체를 반환합니다:

```py
person = {
   'name': 'John',
   'age': 33
}

print(person.keys()) # dict_keys(['name, 'age'])
```

`for` 반복문 안에서, 첫 번째 `if` 문장 뒤에 현재 딕셔너리의 키 집합이 `if`와 다를 때 실행되는 `key_set` 문장을 만드세요. 이는 딕셔너리에 누락되거나 잘못된 키가 없는지 확인하기 위함입니다.

새로운 `if` 문장 안에서 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.`을 출력하세요 (여기서 `<dictionary>`과 `<index>`은 현재 반복의 딕셔너리와 인덱스로 교체되어야 하며) 그리고 `is_invalid`을 `True`으로 설정하세요.

# --hints--

`if` 반복문 안에 두 번째 `for` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1]
`) })
```

새로운 `if` 문장은 `set(dictionary.keys())`가 `key_set`와 다를 때 실행되어야 합니다.

```js
({ test: () => assert(runPython(`
_condition = _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_conditions()[0]

_condition.is_equivalent("set(dictionary.keys()) != key_set") or _condition.is_equivalent("key_set != set(dictionary.keys())")
`)) })
```

새로운 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.` 문장 안에서 `<dictionary>`을 출력해야 합니다 (여기서 `<index>`과 `if`은 현재 반복의 딕셔너리와 인덱스로 교체되어야 합니다).

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
_test_data = [
  {'testing': 'stuff'},
  {
    'patient_id': 'P1003',
    'age': 29,
    'gender': 'Female',
    'diagnosis': 'Asthma',
    'medications': ['Albuterol'],
    'last_visit_id': 'V2303',
  },
  {'keep testing': 'stuff'}
]
validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
len(out) >= 2 and _first in out and _second in out
`)) })
```

새로운 `is_invalid` 문장 안에서 `True`을 `if`로 설정해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

--fcc-editable-region--
        
--fcc-editable-region--

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
