---
id: 6846a9d1f115c1de0d357778
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

두 번째 조건문을 테스트하려면, `medical_records` 리스트 끝에 사전이 아닌 항목 두 개를 추가하세요. 터미널에 두 개의 유효성 검사 메시지가 출력되는 것을 확인할 수 있습니다.

# --hints--

사전이 아닌 항목 두 개를 `medical_records` 리스트 끝에 추가하세요.

```js
({ test: () => runPython(`
assert not isinstance(medical_records[-1], dict)
assert not isinstance(medical_records[-2], dict)
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'P1002',
        'age': 47,
        'gender': 'Male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'V2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'Female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'V2303',
    },
    {
        'patient_id': 'P1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    },
--fcc-editable-region--
    
--fcc-editable-region--
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
