---
id: 6846a1f2805ac73ce671a5d8
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이전 수업에서 배운 것처럼, `enumerate` 함수는 반복 가능한 객체를 순회할 때 인덱스를 추적할 수 있게 해줍니다:

```py
person = {'name': 'John', 'age': 33}

for index, item in enumerate(person):
    print(index, item)

# 0 name
# 1 age
```

`for`를 순회하는 `data` 반복문을 만드세요. 각 반복에서 `enumerate`의 인덱스와 항목을 모두 얻기 위해 `data` 함수를 사용하세요. 반복 변수로 `index`와 `dictionary`를 사용하세요.

지금은 반복문 본문을 채우기 위해 `pass`를 사용하세요.

# --hints--

`for` 함수 안에 `validate` 반복문이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0]
`) })
```

`for` 반복문은 `enumerate(data)`를 순회해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_iter().is_equivalent("enumerate(data)")
`) })
```

`index` 반복문에서 반복 변수로 `dictionary`와 `for`를 사용해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_vars().is_equivalent("index, dictionary")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

--fcc-editable-region--
    
--fcc-editable-region--
```
