---
id: 6846ee2d4d66bda7e22362ab
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

정규 표현식, 또는 정규식은 텍스트에서 문자열 수열과 일치시키는 데 사용하는 패턴입니다. `search` 함수는 `re` 모듈에서 정규식 패턴과 문자열을 인수로 받습니다.

패턴이 일치하면 해당하는 매치 객체를 반환합니다. 그렇지 않으면 `None`를 반환합니다.

```py
import re

greeting = "Hello there!"
print(re.search('Hi', greeting)) # None
print(re.search('Hello', greeting)) # <re.Match object; span=(0, 5), match='Hello'>
```

문자열 `re.search`를 첫 번째 인수로, `p`를 두 번째 인수로 하여 `patient_id`를 호출하세요. `and` 연산자를 사용해 함수 호출을 `patient_id` 키의 값에 두 번째 표현식으로 추가하세요.

# --hints--

`constraints` 딕셔너리에 `patient_id` 키가 있고 그 값이 `isinstance(patient_id, str) and re.search('p', patient_id)`이어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent("constraints = {'patient_id': isinstance(patient_id, str) and re.search('p', patient_id)}")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
