---
id: 6849311be42a0431b8e59109
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

함수에 전달된 데이터가 리스트인지 튜플인지 확인하고 있습니다. 수열의 각 항목이 사전인지도 반드시 확인해야 합니다.

`for` 반복문 안에서, `dictionary`의 항목이 `dict`의 인스턴스가 아니라면 `Invalid format: expected a dictionary at position <index>.`을 출력하세요(`<index>`는 현재 인덱스로 교체되어야 합니다). 그리고 `is_invalid`를 `True`으로 설정하세요.

# --hints--

`if` 반복문 안에 `for` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0]
`) })
```

`if` 문장은 `dictionary`이 `dict`의 인스턴스가 아닌지 확인해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(dictionary, dict)")
`) })
```

`Invalid format: expected a dictionary at position <index>.` 문장 안에서 `<index>`를 출력해야 합니다(`if`는 현재 인덱스로 교체되어야 합니다).

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
validate([{}, "", {}, [], 0])
print = built_in_print

_first = 'Invalid format: expected a dictionary at position 1.'
_second = 'Invalid format: expected a dictionary at position 3.'
_third = 'Invalid format: expected a dictionary at position 4.'
len(out) >= 3 and _first in out and _second in out and _third in out
`)) })
```

`is_invalid` 문장 안에서 `True`를 `if`으로 설정해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
