---
id: 684fe1be09b05ad785832546
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

이제 정규식이 `p` 문자 다음에 하나 이상의 숫자를 일치시키므로, 문자열에 불필요한 문자가 없는지 마지막으로 확인해야 합니다.

이를 위해 `re` 모듈의 또 다른 함수를 사용할 수 있습니다. `fullmatch` 함수는 정규식 패턴이 전체 문자열과 일치하면 매치 객체를 반환하고, 그렇지 않으면 `None`를 반환합니다.

```py
import re

book = "Fahrenheit 451"
print(re.fullmatch('\d+', book)) #None

print(re.fullmatch('Fahrenheit \d+', book))
# <re.Match object; span=(0, 14), match='Fahrenheit 451'>
```

동일한 인수를 유지하면서 `search` 호출을 `fullmatch` 호출로 바꾸세요.

# --hints--

`re.search('p\d+', patient_id, re.IGNORECASE)`를 `re.fullmatch('p\d+', patient_id, re.IGNORECASE)`로 바꾸세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent("constraints = {'patient_id': isinstance(patient_id, str) and re.fullmatch('p\\d+', patient_id, re.IGNORECASE)}")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str) and re.search('p\d+', patient_id, re.IGNORECASE)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
