---
id: 684fdb4386b473631d8e27ca
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

정규 표현식은 역슬래시(`\`)와 그 뒤에 오는 문자로 이루어진 특수 수열을 포함할 수 있습니다. 이 수열들은 특별한 의미를 가집니다. 예를 들어, `\d`은 10진수 숫자 하나와 일치합니다.

Python에서 문자열 리터럴 앞에 `r` 또는 `R`을 붙이면 *로우 문자열*이 생성됩니다. 로우 문자열은 역슬래시를 이스케이프 시퀀스가 아닌 문자 그대로 취급하여 정규 표현식에 적합합니다.

```py
import re

book = "Fahrenheit 451"
print(re.search(r'\d', book))
# <re.Match object; span=(11, 12), match='4'>
```

`p` 문자 다음에 `patient_id`에는 숫자 수열이 와야 합니다. 따라서 정규식 패턴을 로우 문자열 `r'p\d'`로 수정하세요.

# --hints--

정규식 패턴을 `r'p\d'`로 수정하세요.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent(r"constraints = {'patient_id': isinstance(patient_id, str) and re.search(r'p\\d', patient_id, re.IGNORECASE)}")
`) })
```

정규식 패턴에 로우 문자열을 사용하세요.

```js
({ test: () => runPython(`
import ast

_target = _Node(_code).find_function("find_invalid_records").find_variable("constraints").tree.value
_keys = [k.value for k in _target.keys]
_index = _keys.index('patient_id')
_val = _target.values[_index]
_pattern = _val.values[1].args[0]
assert ast.get_source_segment(_code, _pattern).lower().startswith('r')
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str) and re.search('p', patient_id, re.IGNORECASE)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
