---
id: 68529fb26c581f13e901de3a
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

`invalid_records` 변수 바로 다음에 `for` 반복문을 만들어 순회하세요. 각 유효하지 않은 레코드에 대해 `Unexpected format '<key>: <val>' at position <index>.`를 출력하세요. `<key>`, `<val>`, `<index>`는 현재 키, 값, 인덱스로 교체하세요.

`invalid_records`는 현재 `dictionary`에서 유효하지 않은 레코드를 가리키는 키들의 리스트임을 기억하세요. `invalid_records`에서 키를 가져와 `dictionary`에서 값을 찾아야 합니다.

Position 또는 `index`는 함수 내 외부 `medical_records` 반복문으로 정의된 `for`의 현재 딕셔너리를 가리킵니다.

이미 생성한 반복문과 변수를 다시 확인하고 싶으면 코드를 검토하세요.

그다음 `is_invalid`을 `True`로 설정하세요.

유효성 검사 메시지를 확인하려면 `validate` 함수를 유효하지 않은 데이터로 테스트해도 좋습니다.

이로써 의료 유효성 검사 워크숍이 완료됩니다.

# --hints--

`for`를 순회하는 `invalid_records` 반복문이 있어야 합니다.

```js
({ test: () => assert(runPython(`
_Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_for_iter().is_equivalent("invalid_records")
`)) })
```

`Unexpected format '<key>: <val>' at position <index>.` 반복문 안에서 `<key>`를 출력해야 하며, `<val>`, `<index>`, `for`는 현재 키, 값, 인덱스로 교체해야 합니다.

```js
({ test: () => runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
_test_data = [
  {'testing': 'stuff'},
  {
    'patient_id': 'P1003',
    'age': 29,
    'gender': 'emale',
    'diagnosis': 'Asthma',
    'medications': ['Albuterol'],
    'last_visit_id': 'V2303',
  },
  {'keep testing': 'stuff'},
  {
    'patient_id': '1003',
    'age': 29,
    'gender': 'female',
    'diagnosis': 'Asthma',
    'medications': 'Albuterol',
    'last_visit_id': 'V2303',
  },
]

validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = "Unexpected format 'gender: emale' at position 1."
_third = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
_fourth = "Unexpected format 'patient_id: 1003' at position 3."
_fifth = "Unexpected format 'medications: Albuterol' at position 3."

assert len(out) >= 5
assert _first in out
assert _second in out
assert _third in out
assert _fourth in out
assert _fifth in out
`) })
```

새로운 `is_invalid` 반복문 안에서 `True`을 `for`로 설정해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
    return [key for key, value in constraints.items() if not value]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True
            continue

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True
            continue

        invalid_records = find_invalid_records(**dictionary)
--fcc-editable-region--
        
--fcc-editable-region--

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```

# --solutions--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
    return [key for key, value in constraints.items() if not value]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True
            continue

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True
            continue

        invalid_records = find_invalid_records(**dictionary)
        for key in invalid_records:
            print(
                f"Unexpected format '{key}: {dictionary.get(key)}' at position {index}."
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
