---
id: 684ccbdc1ef89e042b744e09
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

La parte successiva del workshop è aggiungere il pulsante di chiusura.

Sotto il tuo elemento `p`, aggiungi un elemento `button` con un insieme `aria-label` impostato su `"Dismiss"`. All’interno di quell’elemento `button`, aggiungi un elemento `span` con un attributo `aria-hidden`. Il testo dell’elemento `span` dovrebbe essere `&times;`.

*NOTA*: Poiché questo workshop si concentrerà solo sul lavoro con gli stili Tailwind CSS, al `button` non verrà aggiunto alcun tipo di interattività.

# --hints--

Dovresti avere un elemento `button`.

```js
assert.exists(document.querySelector("button"));
```
 
Il tuo elemento `button` dovrebbe avere un attributo `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Il tuo elemento `button` dovrebbe avere un attributo `aria-label` impostato su `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Il tuo elemento `button` dovrebbe avere un elemento `span` annidato all’interno.

```js
assert.exists(document.querySelector("button span"));
```

Il tuo elemento `span` dovrebbe avere un attributo `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Il tuo elemento `span` dovrebbe avere il testo `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
