---
id: 684cc59eed2706fa50e8933c
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop imparerai come iniziare a lavorare con le basi di Tailwind CSS costruendo un componente per messaggi di errore.

Nei primi passaggi ti concentrerai sulla creazione del markup HTML iniziale.

All'interno dell'elemento `body` aggiungi un elemento `div`. Poi, all'interno di quell'elemento `div`, aggiungi un elemento `p`.

# --hints--

Dovresti avere un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Dovresti avere un elemento `p`.

```js
assert.exists(document.querySelector("p"));
```

Il tuo elemento `p` dovrebbe essere all'interno dell'elemento `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
