---
id: 684cd027744c2a0520ba4d58
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS usa classi di utilità che vengono applicate direttamente al markup. Queste classi includono stili per caratteri tipografici, colori, layout Flexbox, griglia e altro.

Per applicare un colore di sfondo in Tailwind CSS, puoi usare il prefisso `bg`, seguito dal nome del colore e da un numero. Il numero rappresenta la tonalità del colore.

Ecco un esempio di applicazione di un colore di sfondo blu scuro a un elemento paragrafo.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

Nel tuo tag di apertura `div`, aggiungi una classe con la classe di utilità `bg-red-100` che rappresenta un colore di sfondo rosso chiaro.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere un attributo class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Il tuo attributo class dovrebbe avere la classe di utilità `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
