---
id: 684cd8bf5445de089af9c837
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Ora è il momento di applicare un colore al bordo. Per impostare la proprietà `border-color` usando Tailwind CSS, puoi usare una classe di utilità con il seguente pattern: `border-color-number`, dove `number` indica la tonalità del colore del bordo.

Ecco un esempio di come applicare un colore di bordo viola a un elemento:

```html
<p class="border-purple-400">This is border-color.</p>
```

All’interno del tuo tag di apertura `div`, applica un colore di bordo rosso con una tonalità `300`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `border-red-300`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
