---
id: 684cca09aac0fc035db4099f
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno del tuo elemento `p`, aggiungi il testo `Error! Something went wrong. Please try again.` dove `Error!` deve essere racchiuso all'interno di un elemento `strong`.

# --hints--

Dovresti avere un elemento `strong` all'interno del tuo elemento `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Il tuo elemento `strong` dovrebbe contenere il testo `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Il tuo elemento `p` dovrebbe contenere il testo corretto che termina con `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
