---
id: 684cff0bf0a4ca0ad7e42990
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Il passaggio successivo è arrotondare gli angoli del bordo. Per impostare la proprietà `border-radius` usando Tailwind CSS, puoi usare una classe di utilità con il seguente pattern: `rounded-size`, dove `size` indica il raggio.

Ecco un esempio di applicazione di un raggio del bordo extra-large a un elemento:

```html
<p class="rounded-xl">This is border radius.</p>
```

All’interno del tuo tag di apertura `div`, applica un raggio del bordo di dimensione `md`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
