---
id: 684cd59ec94b5e063bf776f0
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Sarebbe bello avere un bordo intorno a quello sfondo rosso chiaro.

Per impostare il `border-width` in Tailwind CSS, puoi usare la classe di utilità `border-number` dove `number` rappresenterà la luminosità o lo spessore del tuo bordo.

Ecco un esempio di applicazione di un `border-width` di `4px` su tutti i lati:

```html
<p class="border-4">This is how to set the border-width.</p>
```

All’interno del tuo tag di apertura `div`, aggiungi la classe di utilità corretta per un insieme `border-width` impostato su `2px` su tutti i lati.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere una classe `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
