---
id: 684d0eeb9c81e5109e52fe8b
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

L'elemento `div` sta iniziando ad avere un aspetto migliore. Ma ora il `button` per chiudere è un po' troppo vicino al paragrafo.

Per creare spazio tra gli elementi flex, puoi usare il seguente pattern: `gap-number`. `number` rappresenta lo spazio fisso tra righe e colonne in un layout flex.

Ecco un esempio di applicazione di un `gap` a un elemento:

```html
<p class="flex gap-8">This is gap.</p>
```

All'interno del tuo tag di apertura `div`, usa la classe di utilità corretta per impostare il `gap` su `4`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
