---
id: 684d11ffd15a531225d081e9
title: Passaggio 17
challengeType: 0
dashedName: step-17
---

# --description--

Ora che gli stili sono stati applicati all'elemento `div`, è il momento di applicare alcuni stili all'elemento `p`.

Attualmente, il colore del testo per l'elemento paragrafo è nero. Ma sarebbe bello se il colore fosse un rosso scuro invece.

In un passaggio precedente, hai imparato a lavorare con i colori di sfondo in questo modo: `bg-red-100`. Puoi fare qualcosa di simile con il colore del testo seguendo questo pattern: `text-color-shade`.

Ecco un esempio di come applicare un colore del testo blu a un elemento:

```html
<p class="text-blue-500">This has blue text.</p>
```

Inizia aggiungendo un attributo classe al tuo elemento `p`. Poi, applica la classe di utilità corretta per definire un colore del testo rosso con la tonalità impostata su `700`.

# --hints--

Il tuo elemento `p` dovrebbe avere un attributo classe.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Il tuo tag di apertura `p` dovrebbe avere la classe di utilità `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
