---
id: 684d1942a1904a14a3709c6c
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

Lo stile per il tuo componente del messaggio di errore è quasi completo, ma c’è ancora una cosa da sistemare. Poiché hai impostato dimensioni del carattere diverse per il paragrafo e il testo del pulsante, l’allineamento verticale tra quegli elementi flex non è corretto.

Per posizionare gli elementi flex lungo l’asse trasversale di un contenitore usando Tailwind CSS, puoi usare il seguente pattern: `items-value`. `value` può essere qualsiasi cosa tra `start`, `end`, `center`, `baseline` e altro.

All’interno del tuo tag di apertura `div`, usa la classe di utilità corretta per impostare la proprietà `align-items` su `center`.

E con quest’ultima modifica, il tuo componente del messaggio di errore è completo!

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
