---
id: 684d0d369e2b3b0f288d2c20
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Sarebbe utile centrare orizzontalmente l'elemento `div`. Per applicare `margin` su due lati contemporaneamente, puoi usare i seguenti pattern: `mx-number` o `my-number`.

`mx` imposta il margine a sinistra e a destra, mentre `my` imposta il margine in alto e in basso.

Ecco un esempio di applicazione dei margini superiore e inferiore a un elemento:

```html
<p class="my-8">This is margin.</p>
```

All'interno del tuo tag di apertura `div`, usa la classe di utilità corretta per impostare il margine a sinistra e a destra su `auto`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
