---
id: 684d160a2370b01310ae0fda
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

Il nuovo colore rosso è bello, ma sarebbe meglio se il testo fosse più grande. Per impostare la `font-size` di un elemento usando Tailwind CSS, puoi usare il seguente pattern: `text-value`.

Ecco un esempio di come impostare la `font-size` di un elemento su `0.875rem` o (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

All’interno del tuo elemento `p` di apertura, usa la classe di utilità corretta per impostare la `font-size` su `xl` che rappresenta `1.25rem` o (`20px`).

# --hints--

Il tuo tag `p` di apertura dovrebbe avere la classe di utilità `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
