---
id: 684d009994a3840b98c9aa3a
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Al momento, la larghezza dell'elemento `div` occupa tutta la larghezza della pagina. Per applicare larghezze in Tailwind CSS, hai diverse opzioni, come usare il pattern `w-number` per larghezze fisse o il pattern `w-fraction` per percentuali.

Ecco un esempio di applicazione di una larghezza impostata a `66%` del suo contenitore genitore:

```html
<p class="w-2/3">This is a percentage width.</p>
```

All'interno del tag di apertura `div`, applica una larghezza impostata a `50%` del suo contenitore genitore.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
