---
id: 684d17a3b350bb13d3c19c0f
title: Passaggio 19
challengeType: 0
dashedName: step-19
---

# --description--

Ora è il momento di definire lo stile all'elemento `button`. 

Inizia aggiungendo un attributo class al tuo elemento `button`. Poi, applica la classe di utilità corretta per definire il colore del testo in rosso con la tonalità impostata su `700`.

Quindi, all'interno del tag di apertura `button`, usa la classe di utilità corretta per impostare il `font-size` su `2xl` che rappresenta `1.5rem` o (`24px`).

# --hints--

Il tuo elemento `button` dovrebbe avere un attributo class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Il tuo tag di apertura `button` dovrebbe avere la classe di utilità `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Il tuo tag di apertura `button` dovrebbe avere la classe di utilità `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
