---
id: 684d07cbfa43a20d6a002ed7
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Gli stili per l'elemento `div` stanno iniziando a prendere forma. Ma sarebbe bello se ci fosse un po' di spazio intorno al testo.

Per applicare `padding` a un elemento usando Tailwind CSS, puoi usare il seguente pattern: `p-number`, dove `number` indica i pixel che vuoi applicare su tutti e quattro i lati.

Ecco un esempio di applicazione di `padding` di `8px` a un elemento:

```html
<p class="p-8">This is padding.</p>
```

All'interno del tuo tag di apertura `div`, usa la classe di utilità corretta per applicare `padding` di `4px`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
