---
id: 684d09cad4ec900e41db2e74
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Al momento l'elemento `div` è posizionato nell'angolo in alto a sinistra dello schermo di anteprima. Ma sarebbe utile aggiungere un po' di spazio sopra l'elemento. Per applicare `margin` a un elemento usando Tailwind CSS, puoi usare il seguente pattern: `m-number`, dove `number` rappresenta i pixel fissi da applicare a tutti e quattro i lati.

Se vuoi applicare il margine solo a un lato, puoi usare uno dei seguenti pattern: `mt-number` (margine superiore), `mb-number` (margine inferiore), `mr-number` (margine destro) e `ml-number` (margine sinistro).

Ecco un esempio di applicazione di `margin-bottom` a un elemento:

```html
<p class="mb-8">This is margin.</p>
```

All'interno del tag di apertura `div`, usa la classe di utilità corretta per applicare `margin-top` di `4px`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
