---
id: 684d105e488dbd1166a717ae
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Se aumenti la dimensione della finestra di anteprima, noterai che tutti gli elementi sono ancora allineati a sinistra nel contenitore. Sarebbe bello se gli elementi flex fossero centrati.

Per posizionare gli elementi flex lungo l’asse principale di un contenitore con Tailwind CSS, puoi usare il seguente pattern: `justify-value`. `value` può essere `start`, `end`, `center`, `between` e altri ancora. Qualsiasi valore valido usato nel CSS vanilla con `justify-content` può essere usato qui con Tailwind CSS.

All’interno del tuo tag di apertura `div`, usa la classe di utilità corretta per impostare la proprietà `justify-content` su `center`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
