---
id: 684d03df2a0ddf0c735865f0
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora vedi che la larghezza dell'elemento è la metà della dimensione del `body`. Questo funziona bene per dispositivi più grandi ma non altrettanto per quelli più piccoli.

Tailwind CSS segue un approccio mobile-first nel design. Quindi, se vuoi creare design responsivi per altri dispositivi, puoi usare le varianti responsive di Tailwind come `md` e `lg` in questo modo:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

L'esempio seguente imposterà la larghezza a `66%` dell'elemento genitore per dimensioni di dispositivo `1024px` e superiori.

All'interno del tuo tag di apertura `div`, modifica il tuo `w-1/2` in modo che venga applicato solo a dispositivi di dimensione `768px` e superiori. Puoi anteporre questa classe con la variante `md`.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Non dovresti più avere la classe `w-1/2` nel tuo tag di apertura `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
