---
id: 68501d32b69f9d7756334ffc
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Se ingrandisci la finestra di anteprima, dovresti vedere l'elemento `div` centrato orizzontalmente nella pagina. Ma per dispositivi più piccoli, la larghezza dell'elemento `div` occupa tutta la larghezza della pagina.

Aggiorna la classe `mx-auto` in modo che venga applicata solo per dispositivi `md` e superiori.

# --hints--

Il tuo tag di apertura `div` dovrebbe avere la classe di utilità `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Non dovresti più avere la classe `mx-auto` nel tuo tag di apertura `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
