---
id: 65f01969115f933073b6be03
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Вектор — це об’єкт, який має довжину (або величину) і напрямок, і його не можна виразити одним числом. У фізиці вектори часто використовують для представлення сил, швидкостей, прискорень та інших величин.

У цьому проєкті ви створите векторний простір — множину, у якій визначено ряд операцій між елементами цієї множини. Незабаром ви дізнаєтеся всі деталі.

Поки що почніть проєкт з оголошення порожнього класу з ім’ям `Vector`.

# --hints--

Ви повинні оголосити порожній клас з ім’ям `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
