---
id: 65f06a8e5a57673d700c79c3
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Поза межами класу `Vector` створіть екземпляр `Vector`, передаючи цілі числа `2` і `3` як значення `x` і `y`, і призначте його змінній з ім’ям `v1`.

# --hints--

Ви повинні оголосити змінну `v1` і призначити їй `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
