---
id: 65f058e06f34fd1f0ee6e55d
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Далі, у класі `Vector` створіть порожній метод `norm` і додайте йому параметр `self`.

# --hints--

Ваш клас `Vector` має містити метод `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Ваш метод `norm` має мати параметр `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
