---
id: 65f03d9f92eac9183a4d3281
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Вектор можна визначити за двома координатами, `x` та `y`, на евклідовій площині. Відстань між початком координат і точкою `(x, y)` буде його довжиною, або нормою. А напрямок вектора вказуватиме на `(x, y)`.

<img alt="вектор у двовимірному просторі з координатами (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

У класі `Vector` створіть метод `__init__` і надайте йому три параметри: `self`, `x` та `y`.

# --hints--

Ви повинні визначити метод `__init__` всередині класу `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Ваш метод `__init__` має приймати три параметри: `self`, `x` та `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
