---
id: 65f055b9190fc41ca35549b8
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Python пропонує різні методи, які мають подвійний підкреслення на початку і в кінці своїх назв. Ви, можливо, вже знайомі з деякими, як-от `__init__` і `__str__`. Ці методи, що відповідають шаблону іменування `__<name>__`, називають спеціальними методами, магічними методами або *dunder* (що означає подвійне підкреслення).

Визначення спеціальних методів у класі впливає на поведінку цього класу. Вони викликаються "під капотом" у певних ситуаціях (наприклад, `__init__` під час створення екземпляра, `__str__` коли об’єкт виводиться або передається до `str()`). У цьому проєкті ви дізнаєтеся про деякі з найпоширеніших спеціальних методів.

Поки що призначте `x` атрибуту `x` об’єкта `Vector`.

# --hints--

Ви повинні призначити `x` `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
