---
id: 65f3fa097f26b510db6c710b
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Чудово, `norm` працює як очікувалося. Тепер, якщо ви спробуєте вивести `v1`, отримаєте рядкове представлення об’єкта за замовчуванням (щось на кшталт `<__main__.Vector object at 0x11eb778>`).

Всередині класу `Vector` оголосіть порожній метод `__str__` для реалізації читабельного рядкового представлення. Не забудьте надати йому параметр `self`.

Зверніть увагу, що не слід виводити `v1` до того, як `__str__` поверне рядок, інакше отримаєте `TypeError`, бо `__str__` завжди має повертати рядок.

# --hints--

Ви повинні визначити метод `__str__` у класі `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Ваш метод `__str__` має мати параметр `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
