---
id: 65f3fe07cc763212efe91285
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Вектор може мати число `n` вимірів (компонентів). Ось представлення тривимірного вектора:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

До цього ви створили двовимірний вектор. Ви хочете мати можливість представляти вектори з різною кількістю вимірів, не переписуючи необхідний код для кожного конкретного випадку. Для цього ви використаєте успадкування.

Почніть із перейменування класу `Vector` на `R2Vector`, щоб вказати, що цей клас працюватиме з двовимірними векторами. Не забудьте також змінити створення екземпляра `v1`.

# --hints--

Ви повинні перейменувати клас `Vector` на `R2Vector`. Не забудьте також змінити створення екземпляра `v1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
