---
id: 65f4535bbdb28d436ff3ddc9
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Як видно з вихідних даних, `__dict__` містить значення ваших атрибутів екземпляра. Замість того, щоб явно додавати квадрати `self.x` і `self.y`, ви ітеруватимете над значеннями, збереженими в `__dict__`, щоб обчислити значення норми.

У тілі методу `norm` замініть вміст круглих дужок на твірний вираз, який підносить кожне значення `val` у `self.__dict__.values()` до степеня `2`. Також передайте цей твірний вираз як аргумент функції `sum`, щоб усі квадрати додалися разом перед обчисленням квадратного кореня.

# --hints--

Ви повинні повернути `sum(val**2 for val in self.__dict__.values())**0.5` з методу `norm()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
