---
id: 65f412c208c3791fee305acf
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

У Python ви можете примусити використовувати лише ключові аргументи, додавши `*` як додатковий параметр у сигнатуру функції або методу. Змініть обидва методи `__init__`, додавши `*` як другий параметр (після `self`). Кожен параметр, розміщений після нього, вимагатиме використання ключового аргументу під час виклику функції/методу.

Це означає, що вам також потрібно змінити виклик `super().__init__(x, y)`. Зробіть це, призначивши `x` значення `x`, а `y` — значення `y`.

Нарешті, змініть створення екземплярів `v1` і `v2`, використовуючи ключові аргументи.

# --hints--

Ви повинні примусити використання ключових аргументів, додавши `*` як другий параметр в обох ваших методах `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Ви повинні змінити виклик `super.__init__(x, y)`, щоб використовувати ключові аргументи.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Ви повинні змінити присвоєння `v1`, щоб воно було `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Ви повинні змінити присвоєння `v2`, щоб воно було `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
