---
id: 65f412c208c3791fee305acf
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

У Python можна змусити використання аргументів лише за ключовими словами, додавши `*` як додатковий аргумент до сигнатури функції чи методу. Змініть обидва методи `__init__`, додавши `*` як другий параметр (після `self`). Кожен параметр, розміщений після цього, вимагатиме використання аргументу ключового слова у виклику функції/методу.

Це означає, що також потрібно змінити виклик `super().__init__(x, y)`. Зробіть це, надавши `x` значення `x`, та `y` значення `y`.

Вкінці змініть створення екземпляра `v1` та `v2` за допомогою аргументів ключових слів.

# --hints--

Змусьте аргументи ключові слова, додавши `*` як другий параметр в обох методах `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Змініть виклик `super.__init__(x, y)`, щоб використати аргументи ключові слова.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Змініть призначення `v1` на `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Змініть призначення `v2` на `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
