---
id: 65f407ea37ad6e181b90462e
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Ви могли б призначити кожен параметр `i` до `self.i`, як робили раніше. Хоча в цьому випадку у вас небагато рядків для повторення, один зі способів уникнути цього повторення — використати функцію `super()`. `super()` дозволяє вам неявно звертатися до батьківського класу: `super().__init__(x, y)` викликає метод `__init__` батьківського класу.

Додайте виклик `super().__init__(x, y)` до свого методу `__init__`.

# --hints--

У вашому методі `super().__init__(x, y)` має бути виклик `__init__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
