---
id: 65f8709620b2ce1a62608f5a
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Якщо потрібно отримати динамічний доступ до деяких атрибутів, починаючи з рядкового введення, вбудована функція `getattr()` то що вам потрібно. Вона приймає об’єкт як перший аргумент, а рядок, що містить назву атрибута, як другий атрибут.

Почніть виправляти метод `__str__`, замінивши рядок, який повертає `__str__()`, виразом генератора, який ітерує над атрибутами об’єкта та викликає `getattr()` для кожного атрибута `i`.

# --hints--

Поверніть вираз генератора, який ітерує над `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Поверніть вираз генератора, який використовує `i` як змінну ітерації.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Поверніть вираз генератора, який викликає `getattr(self, i)` для кожного елемента `i` в `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
