---
id: 65f9b486989cb90ff3e77ac8
title: Крок 39
challengeType: 20
dashedName: step-39
---

# --description--

Параметр `other` методу `__add__` представляє операнд, розміщений праворуч від оператора `+` в операції додавання.

Ви хочете переконатися, що правий операнд не є об’єктом того самого класу, що й лівий операнд (тобто `self`). Створіть інструкцію `if`, яка верифікує це, перевіряючи тип операндів.

# --hints--

Створіть інструкцію `if`, яка перевіряє, чи відрізняються `type(self)` та `type(other)`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
