---
id: 65f9cb678070ca1668898c70
title: Крок 44
challengeType: 20
dashedName: step-44
---

# --description--

Так само, як під капотом викликається `__add__` при додаванні двох об’єктів, метод `__sub__` викликається неявно у випадку віднімання.

Тепер визначте порожній метод `__sub__` і додайте два параметри: `self` та `other`. Усередині вашого нового методу створіть оператор if, щоб перевірити, чи `self` і `other` не належать до одного класу, і поверніть `NotImplemented`, як ви робили раніше.

# --hints--

Ви повинні визначити метод `__sub__` у класі `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

Ваш метод `__sub__` має мати два параметри, `self` і `other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

Ви повинні створити оператор `if`, який перевіряє, чи `self` і `other` не належать до одного класу.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

Ви повинні повернути `NotImplemented` з тіла `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
