---
id: 65f9dd6e5a08af19c196c2df
title: Крок 45
challengeType: 20
dashedName: step-45
---

# --description--

Вектор, отриманий в результаті віднімання одного вектора з іншого, виходить з обчислення різниці кожного з їх компонентів. Наприклад, віднімання `(2, 4)` від `(7, 3)` створює третій вектор `(5, -1)`, де `5` — це різниця між `7` та `2`, а `-1` — це різниця між `3` та `4`.

Як ви робили раніше в методі `__add__`, оголосіть змінну `kwargs` після інструкції `if`. Призначте розуміння словника, яке ітерує над атрибутами об’єкта та створює пару ключ-значення для кожного ключа `i`, де ключ є тим самим ключем як у поточної ітерації, а його значення є різницею між значеннями атрибутів `i` двох операндів.

# --hints--

Оголосіть змінну `kwargs` та призначте до неї розуміння словника, яке ітерує над атрибутами об’єкта.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

Розуміння словника, призначене до `kwargs`, має використовувати змінну `i`, щоб ітерувати над атрибутами об’єкта.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Розуміння словника, призначене до `kwargs`, має використовувати `i` як ключ.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

Розуміння словника, призначене `kwargs`, має призначати різницю між атрибутом `i` з `self` та атрибута `i` з `other` до ключа `i` для кожного `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
