---
id: 65f9bc10a9fb1612e066e9e8
title: Крок 42
challengeType: 20
dashedName: step-42
---

# --description--

При додаванні двох векторів кожен компонент одного вектора додається до такого самого компонента іншого вектора. Наприклад, додавання `(1, 2)` і `(2, 4)` створює третій вектор `(3, 6)`, де `3` — це сума `1` і `2`, а `6` — сума `2` і `4`.

Словник `kwargs` міститиме пари ключ-значення, необхідні для створення нового вектора того самого класу, що й два додані вектори.

Перетворіть порожній словник на словникове включення, яке ітерує `vars(self)` і для кожного ключа (`i`) створює пару ключ-значення, де ключ — це той самий ключ поточної ітерації, а його значення — сума значень атрибутів `i` двох операндів.

# --hints--

Ви повинні перетворити порожній словник, призначений `kwargs`, на словникове включення, яке ітерує над `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

Словникове включення, призначене `kwargs`, має використовувати змінну `i` для ітерації над `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Словникове включення, призначене `kwargs`, має використовувати `i` як ключ.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

Словникове включення, призначене `kwargs`, має призначати суму атрибутів `i` об’єкта `self` і `i` об’єкта `other` ключу `i` для кожного `i` у `vars(self)`. Для цього використовуйте функцію `getattr()`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
