---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Construir um Gerador de Permutação
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

Neste laboratório, você irá construir um gerador de permutações que receberá uma string e retornará todas as permutações possíveis dos caracteres na string.Por exemplo, as possíveis permutações da string `cat` são `cat`, `cta`, `act`, `atc`, `tac` e `tca`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `permuteString`.
2. A função `permuteString` deve receber três parâmetros - uma string, um valor de prefixo e um array vazio para armazenar e retornar os resultados. O valor do prefixo acumularia caracteres para formar uma permutação.
3. Dentro da função, você deve verificar se o comprimento da string passada é `0`. Se for, adicione o prefixo atual aos resultados e retorne os resultados.
4. Itere sobre cada caractere na string de entrada e, a cada iteração, remova o caractere atual da string e chame a função `permuteString` recursivamente com argumentos atualizados para construir as permutações restantes.
5. Você deve retornar o array de resultados finais.
6. Você deve garantir que as permutações sejam únicas removendo os duplicados.

# --hints--

Você deve ter uma função `permuteString`.

```js
assert.isFunction(permuteString);
```

Você deve usar recursão na sua função `permuteString`.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` deve retornar `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` deve retornar `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` deve retornar `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` deve retornar `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` deve retornar `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` deve retornar os resultados corretos.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
