---
id: 69babfd65c4c4632a0c76f63
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Cuando una puerta no puede procesar a todos los asistentes durante un tick, algunos asistentes permanecen en la cola. Para manejar este desbordamiento, crearás una función que redirija a los asistentes excedentes a otra puerta.

Crea una función vacía llamada `rerouteOverflow` con los siguientes parámetros:

- `gates`: El arreglo completo de objetos de puertas.
- `currentGate`: La puerta donde ocurrió el desbordamiento.
- `tickIndex`: La posición actual del tick en el arreglo `queue`.
- `overflowAmount`: El número de asistentes que no se pudieron procesar.

# --hints--

Debes crear una función llamada `rerouteOverflow`.

```js
assert.isFunction(rerouteOverflow);
```

Tu función `rerouteOverflow` debe tener los parámetros `gates`, `currentGate`, `tickIndex` y `overflowAmount`.

```js
const explorer = await __helpers.Explorer(code);
const { rerouteOverflow } = explorer.allFunctions;
const parameters = rerouteOverflow?.parameters;
assert.lengthOf(parameters, 4);
assert.equal(parameters[0].toString(), "gates");
assert.equal(parameters[1].toString(), "currentGate");
assert.equal(parameters[2].toString(), "tickIndex");
assert.equal(parameters[3].toString(), "overflowAmount");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

--fcc-editable-region--

--fcc-editable-region--
```
