---
id: 69babdf3f896711c11516373
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Durante cada ciclo de simulación:

- Un cierto número de asistentes llega a la puerta (desde su `queue`).
- La puerta procesa a los asistentes según su `capacity`.
- Si llegan más asistentes de los que la puerta puede manejar, algunos quedarán (desbordamiento).

Ahora crearás una función que maneje esta lógica para una puerta en un solo ciclo.

Crea una función vacía llamada `processGateFlow` con los siguientes parámetros:

- `gate`: El objeto gate que se está procesando.
- `tickIndex`: La posición actual del tick en el arreglo `queue`.

# --hints--

Debes crear una función llamada `processGateFlow`.

```js
assert.isFunction(processGateFlow);
```

Tu función `processGateFlow` debe tener los parámetros `gate` y `tickIndex`.

```js
const explorer = await __helpers.Explorer(code);
const { processGateFlow } = explorer.allFunctions;
const parameters = processGateFlow?.parameters;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), "gate");
assert.equal(parameters[1].toString(), "tickIndex");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

--fcc-editable-region--

--fcc-editable-region--
```
