---
id: 69baafd6d683dd9921b75db3
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, usarás JavaScript para simular el flujo de asistentes en un festival de música.

Trabajarás con dos conjuntos de datos que describen las puertas del festival durante los turnos de la mañana y la noche:

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];
```

Cada objeto puerta contiene las siguientes propiedades:

- `id`: Una cadena que identifica la puerta.
- `capacity`: El número de asistentes que la puerta puede procesar por tick.
- `queue`: Un arreglo de números que representa cuántos asistentes llegan a la puerta durante un tick específico.

Para comenzar, crea una función vacía llamada `initializeThroughput` que acepte un parámetro llamado `gates`.

# --hints--

Debes crear una función llamada `initializeThroughput`.

```js
assert.isFunction(initializeThroughput);
```

Tu función `initializeThroughput` debe tener un parámetro `gates`.

```js
const regex = __helpers.functionRegex('initializeThroughput', ['gates']);
assert.match(__helpers.removeJSComments(initializeThroughput.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

--fcc-editable-region--

--fcc-editable-region--
```
