---
id: 69bac067cd15c4685d07906d
title: Paso 15
challengeType: 1
dashedName: step-15
---

# --description--

Ahora que tienes funciones para procesar ticks individuales y manejar el desbordamiento, es momento de crear una función controladora para una sola puerta. Esta función manejará todas las acciones para una puerta durante un solo tick de la simulación.

Crea una función vacía llamada `handleGateAtTick` con los siguientes parámetros:

- `gates`: El arreglo completo de objetos puerta.
- `gate`: La puerta actual que se está procesando.
- `tickIndex`: El índice del tick actual para la simulación.
- `throughputSummary`: Un objeto que rastrea el total de asistentes procesados por puerta.

# --hints--

Debes crear una función llamada `handleGateAtTick`.

```js
assert.isFunction(handleGateAtTick);
```

Tu función `handleGateAtTick` debe tener los parámetros `gates`, `gate`, `tickIndex` y `throughputSummary`.

```js
const regex = __helpers.functionRegex('handleGateAtTick', ['gates', 'gate', 'tickIndex', 'throughputSummary']);
assert.match(__helpers.removeJSComments(handleGateAtTick.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

--fcc-editable-region--

--fcc-editable-region--
```
