---
id: 69bac027df64d4433c708d44
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Al redirigir el exceso de asistentes, debes enviarlos a la siguiente puerta en el arreglo `gates`. Para hacer esto, necesitas encontrar el índice de la siguiente puerta. Normalmente, podrías sumar 1 al índice actual. Sin embargo, si la puerta actual es la última en el arreglo, debes volver a la primera puerta.

Puedes usar el operador módulo (`%`) para manejar esto:

```js
const nextIndex = (currentIndex + 1) % array.length;
```

Esto funciona porque cuando `currentIndex + 1` es igual a `array.length`, el resultado es 0.

Usando este enfoque, crea una variable llamada `nextGateIndex` que almacene el índice de la siguiente puerta en el arreglo `gates`.

# --hints--

Debes usar la técnica del módulo para acceder y asignar el siguiente índice de `currentIndex` en el arreglo `gates` a una variable llamada `nextGateIndex`. Sigue el mismo patrón de una línea en el ejemplo.

```js
const rawFuncStr = __helpers.removeJSComments(rerouteOverflow.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+nextGateIndex/,
  "You must have at least one space between the declaration keyword (let/const) and 'nextGateIndex'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)nextGateIndex=\((currentIndex\+1|1\+currentIndex)\)%gates\.length;?}$/;

assert(
  solution.test(cleaned),
  "Inside your `rerouteOverflow` function, use the modulo technique to create a variable named `nextGateIndex` that stores the index of the next gate."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
--fcc-editable-region--

--fcc-editable-region--
}
```
