---
id: 696920c0c98a1ed58eb86293
title: بناء عرض للمنتجات
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

في هذا المختبر، تم تزويدك بكل ملفات HTML وCSS. ستكمل عرضًا للمنتجات باستخدام TypeScript.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ واجهة باسم `Item` تحتوي على:
   - خاصية `type` لا يمكن أن تكون إلا `"book"` أو `"electronics"` أو `"clothing"`.
   - خاصية `id` من نوع `string`.
   - خاصية `price` من نوع `number`.

1. يجب أن تنشئ واجهة باسم `Book` تمتد من `Item`. يجب أن تكون خاصية `type` فيها محددة بالقيمة الحرفية `"book"`. كما يجب أن تحتوي على خاصيتي `title` و `author` من نوع `string`.

1. يجب أن تنشئ واجهة باسم `Electronics` تمتد من `Item`. يجب أن تكون خاصية `type` فيها محددة بالقيمة الحرفية `"electronics"`. كما يجب أن تحتوي على خاصيتي `item` و`model` من نوع `string`، وخاصية اختيارية `warranty` من نوع `number`.

1. يجب أن تنشئ واجهة باسم `Clothing` تمتد من `Item`. يجب أن تكون خاصية `type` فيها محددة بالقيمة الحرفية `"clothing"`. كما يجب أن تحتوي على خاصيتي `item` و`brand` من نوع `string`، وخاصية اختيارية `size` يمكن أن تكون فقط `"S"` أو `"M"` أو `"L"`.

1. يجب أن تنشئ نوع اتحاد باسم `Product` يجمع بين واجهات `Book` و`Electronics` و`Clothing`.

1. يجب أن تنشئ صنفًا عامًا باسم `Collection` بمعلمة نوع `T`. يجب أن يقبل المُنشئ مصفوفة من العناصر من نوع `T` ويخزنها في خاصية باسم `items`.

1. يجب أن يحتوي صنف `Collection` على طريقة `getAll()` تُرجع جميع العناصر في المجموعة.

1. يجب أن يحتوي صنف `Collection` على طريقة `filter()` تقبل دالة استدعاء راجع وتُرجع فقط العناصر من مصفوفة `items` التي تحقق الشرط المعطى. يجب أن تستقبل الدالة عنصرًا من نوع `T` وتُرجع قيمة منطقية.

1. يجب إنشاء دالة باسم `renderProduct` بمعلمة من نوع `Product`. يجب أن تُرجع سلسلة HTML تمثل منتجًا تحتوي على:
   - عنصر يحمل الصنف `item` ومعرّف بقيمة `id` الخاصة بالمنتج.
   - عنصر يحمل الصنف `price` ونصه هو قيمة `price` الخاصة بالمنتج.
   - عنصر يحتوي على معلومات إضافية عن المنتج.

1. يجب أن تستخدم دالة `renderProduct` تضييق النوع لعرض HTML مختلف حسب نوع المنتج:

- للكتب، يجب أن يكون HTML المعروض بالشكل `Book: [title] by [author]`، حيث يُستبدل `[title]` و`[author]` بعنوان الكتاب والمؤلف على التوالي.
   - للإلكترونيات، يجب أن يكون HTML المعروض بالشكل `Electronics: [item] - [model]`. إذا كان للمنتج ضمان، أضف ` - Warranty: [warranty] year(s)` إلى الصيغة. استبدل `[item]` و`[model]` و`[warranty]` بالقيم المناسبة.
   - للملابس، يجب أن يكون HTML المعروض بالشكل `Clothing: [item] by [brand]`. إذا كان للمنتج حجم، أضف ` - Size [size]` إلى الصيغة. استبدل `[item]` و`[brand]` و`[size]` بالقيم المناسبة.
   - إذا صادفت نوع منتج غير صالح، ارمِ خطأً بالصيغة: `Unknown product type: [product]`، حيث يُستبدل `[product]` بالكائن غير الصالح بعد تحويله إلى سلسلة باستخدام `JSON.stringify`.

1. يجب أن تنشئ مثيلًا من `Collection<Product>` يحتوي على عنصر واحد على الأقل لكل نوع منتج وتُعيّن المثيل إلى متغير باسم `products`.

1. يجب أن تنشئ دالة باسم `showProducts` تعرض المنتجات في العنصر `#output`. يجب أن تقبل دالة اختيارية تمثل فلترًا لنوع المنتج؛ عند حذفها، تعرض كل المنتجات؛ وعند توفيرها، تعرض فقط المنتجات من ذلك النوع.

1. يجب أن تعرض المنتجات بتحويل كل منتج إلى سلسلة HTML باستخدام دالة `renderProduct` ثم حقن النتيجة المجمعة في العنصر `#output`.

1. يجب أن يعرض النقر على أزرار `#all` و`#books` و`#electronics` و`#clothing` مجموعة المنتجات المقابلة في العنصر `#output`.

1. يجب أن تُعرض كل المنتجات افتراضيًا عند تحميل الصفحة.

# --hints--

يجب أن يكون لديك واجهة باسم `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

يجب أن تحتوي واجهة `Item` على خاصية `type` يمكن أن تكون فقط `"book"` أو `"electronics"` أو `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

يجب أن تحتوي واجهة `Item` على خاصية `id` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

يجب أن تحتوي واجهة `Item` على خاصية `price` من النوع `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

يجب أن يكون لديك واجهة باسم `Book`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

يجب أن تمتد واجهة `Book` من `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

يجب أن تحتوي واجهة `Book` على خاصية `type` محددة بالقيمة `"book"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

يجب أن تحتوي واجهة `Book` على خاصية `title` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

يجب أن تحتوي واجهة `Book` على خاصية `author` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

يجب أن يكون لديك واجهة باسم `Electronics`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

يجب أن تمتد واجهة `Electronics` من `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

يجب أن تحتوي واجهة `Electronics` على خاصية `type` محددة بالقيمة `"electronics"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

يجب أن تحتوي واجهة `Electronics` على خاصية `item` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

يجب أن تحتوي واجهة `Electronics` على خاصية `model` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

يجب أن تحتوي واجهة `Electronics` على خاصية اختيارية `warranty` من النوع `number`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

يجب أن يكون لديك واجهة باسم `Clothing`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

يجب أن تمتد واجهة `Clothing` من `Item`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

يجب أن تحتوي واجهة `Clothing` على خاصية `type` محددة بالقيمة `"clothing"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

يجب أن تحتوي واجهة `Clothing` على خاصية `item` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

يجب أن تحتوي واجهة `Clothing` على خاصية `brand` من نوع `string`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

يجب أن تحتوي واجهة `Clothing` على خاصية اختيارية `size` يمكن أن تكون فقط `"S"` أو `"M"` أو `"L"`.

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

يجب أن تنشئ نوع اتحاد باسم `Product` يجمع بين واجهات `Book` و`Electronics` و`Clothing`.

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

يجب أن يكون لديك صنف `Collection` بمعلمة نوع `T`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

يجب أن يقبل المُنشئ في صنف `Collection` مصفوفة من عناصر نوع `T` ويخزنها في خاصية باسم `items`.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

يجب أن يحتوي صنف `Collection` على طريقة `getAll`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

يجب أن تُرجع طريقة `getAll` جميع العناصر في المجموعة.

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

يجب أن يحتوي صنف `Collection` على طريقة `filter`.

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

يجب أن تقبل طريقة `filter` دالة استدعاء راجع وتُرجع مصفوفة تحتوي على العناصر من خاصية `items` التي تحقق الشرط المعطى.

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

يجب أن يكون لديك دالة باسم `renderProduct`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

يجب أن تُرجع `renderProduct(product)` سلسلة تحتوي على عنصر بفئة `item` ومعرف يساوي `product.id`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

يجب أن تُرجع `renderProduct(product)` سلسلة تحتوي على عنصر بفئة `price` ونص يساوي `product.price`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

عند استدعاء `renderProduct` مع منتج `Book` `p`، يجب أن تُرجع سلسلة تحتوي على عنصر بنص بالشكل `Book: [title] by [author]`. استبدل `[title]` و`[author]` بالخصائص المقابلة في `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

عند استدعاء `renderProduct` مع منتج `Electronics` `p`، يجب أن تُرجع سلسلة تحتوي على عنصر بنص بالشكل `Electronics: [item] - [model]` أو `Electronics: [item] - [model] - Warranty: [warranty] year(s)` إذا لم تكن `p.warranty` غير معرفة. استبدل `[item]` و`[model]` و`[warranty]` بالخصائص المقابلة في `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

عند استدعاء `renderProduct` مع منتج `Clothing` `p`، يجب أن تُرجع سلسلة تحتوي على عنصر بنص بالشكل `Clothing: [item] by [brand]` أو `Clothing: [item] by [brand] - Size [size]` إذا لم تكن `p.size` غير معرفة. استبدل `[item]` و`[brand]` و`[size]` بالخصائص المقابلة في `p`.

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

عند استدعاء `renderProduct` مع منتج غير صالح، يجب أن ترمي خطأً بالصيغة `Unknown product type: [p]`. استبدل `[p]` بالمنتج غير الصالح بعد تحويله إلى سلسلة باستخدام `JSON.stringify`.

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

يجب أن تنشئ مثيلًا من `Collection<Product>` وتُعيّنه إلى متغير باسم `products`.

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

يجب أن يحتوي متغير `products` على عنصر واحد على الأقل لكل نوع منتج.

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

يجب أن يكون لديك دالة باسم `showProducts`.

```js
assert.isFunction(showProducts);
```

يجب أن يعرض النقر على زر `#all` كل المنتجات في العنصر `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click", { bubbles: true }));
all.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

يجب أن يعرض النقر على زر `#books` مجموعة المنتجات المقابلة في العنصر `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click", { bubbles: true }));
books.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

يجب أن يعرض النقر على زر `#electronics` مجموعة المنتجات المقابلة في العنصر `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click", { bubbles: true }));
electronics.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

يجب أن يعرض النقر على زر `#clothing` مجموعة المنتجات المقابلة في العنصر `#output`.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click", { bubbles: true }));
clothing.dispatchEvent(new Event("click", { bubbles: true }));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

يجب أن تكون كل المنتجات مرئية افتراضيًا عند تحميل الصفحة. استخدم حدث `DOMContentLoaded` لضمان ذلك.

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
