---
id: 66c057041df6394ca796bf33
title: 创建一个句子生成器
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

在本实验中，你将使用句子模板创建两个不同的故事。 你将使用变量来存储故事的不同部分，然后将故事输出到控制台。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应使用 `let` 声明以下变量：

    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. 你应该将上述变量赋值为你所选择的字符串值。

1. 你应该声明一个 `firstStory` 变量。

1. 你应使用以下故事模板创建第一个故事，并将其分配给 `firstStory` 变量：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`；

1. 你应该将第一个故事输出到控制台，使用的信息是 `"First story: [firstStory]"`。

1. 你应该为 `adjective`、`noun`、`verb`、`place`、`adjective2` 和 `noun2` 变量重新赋值。

1. 你应该声明一个 `secondStory` 变量。

1. 使用相同的模板创建另一个故事，并将其赋值给 `secondStory` 变量。

1. 你应该将第二个故事输出到控制台，使用的信息是 `"Second story: [secondStory]"`。


# --hints--

你应该声明一个 `adjective` 变量。

```js
assert.isNotNull(adjective);
```

你应该声明一个 `noun` 变量。

```js
assert.isNotNull(noun);
```

你应该声明一个 `verb` 变量。

```js
assert.isNotNull(verb);
```

你应该声明一个 `place` 变量。

```js
assert.isNotNull(place);
```

你应该声明一个 `adjective2` 变量。

```js
assert.isNotNull(adjective2);
```

你应该声明一个 `noun2` 变量。

```js
assert.isNotNull(noun2);
```

你应该为 `adjective` 变量赋值一个字符串。

```js
assert.isString(adjective);
```

你应该为 `noun` 变量赋值一个字符串。

```js
assert.isString(noun);
```

你应该为 `verb` 变量赋值一个字符串。

```js
assert.isString(verb);
```

你应该为 `place` 变量赋值一个字符串。

```js
assert.isString(place);
```

你应该为 `adjective2` 变量赋值一个字符串。

```js
assert.isString(adjective2);
```

你应该为 `noun2` 变量赋值一个字符串。

```js
assert.isString(noun2);
```

你应该声明一个 `firstStory` 变量。

```js
assert.isNotNull(firstStory);
```

你应该为第一个故事使用正确的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。 注意空格。

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

You should assemble your first story using the variables you declared in the correct order.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你应该将第一个故事输出到控制台，使用的信息是 `"First story: [firstStory]"`。

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

你应该声明一个 `secondStory` 变量。

```js
assert.isNotNull(secondStory);
```

你应该为第二个故事的 `adjective` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

你应该为第二个故事的 `noun` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

你应该为第二个故事的 `verb` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

你应该为第二个故事的 `place` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

你应该为第二个故事的 `adjective2` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

你应该为第二个故事的 `noun2` 变量重新赋值。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

你应该为第二个故事使用正确的故事格式：`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。 注意空格。

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

You should assemble your second story using the variables you declared in the correct order.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

你应该将第二个故事输出到控制台，使用的格式是 `"Second story: [secondStory]"`。

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` 不能和 `secondStory` 相同。

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

