---
id: 6823f9df49cc206af5471a30
title: Depurar a Página de Perfil do Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot está tentando construir uma página de perfil. Eles pediram a um amigo para revisar o código deles e ele disse que há alguns erros.

Seu trabalho é corrigir todos os erros do Camperbot para que ele possa continuar construindo sua página de perfil. Complete os itens nas histórias de usuário abaixo e clique em "Executar os Testes" para ver se você corrigiu todos os erros.

**Histórias de Usuário:**

1. Camperbot está tentando usar um elemento `heading2`, mas esse elemento não existe. Corrija essas tags para que usem um elemento de título de segundo nível correto.
2. Camperbot está tentando adicionar dois parágrafos com `pp`, mas eles também não existem. Corrija-os para que usem as tags de parágrafo corretas.
3. Camperbot está usando um elemento `h3` para o subtítulo `Background and Interests`, mas ele contém um erro de sintaxe. Identifique o problema e resolva-o.

# --hints--

Você deve ter exatamente um elemento `h1` na página.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Você não deve remover o elemento `<h1>Olá do Camperbot!</h1>` da página.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Você deve ter exatamente um elemento `h2` na página.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Você deve ter um elemento `h2` com o texto `Sobre`. Aqui está um exemplo: `<h2>Texto aqui</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Você deve ter um elemento de parágrafo com o texto `Meu nome é Camperbot e eu adoro aprender coisas novas.` Aqui está um exemplo: `<p>Texto aqui</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Você deve ter um elemento de parágrafo com o texto `Eu gosto de resolver quebra-cabeças.` Aqui está um exemplo: `<p>Texto aqui</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Você deve ter exatamente dois elementos de parágrafo na página.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Você deve ter exatamente um elemento `h3` na página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Você deve ter um elemento `h3` com o texto `Antecedentes e Interesses`. Aqui está um exemplo: `<h3>Texto aqui</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Você deve ter exatamente cinco elementos totais na página.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
