---
id: 60a3e3396c7b40068ad69978
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Usa márgenes para ajustar el espacio exterior de un elemento.

Usando la propiedad `margin`, dale al elemento `.frame` un margen vertical de `20px` y un margen horizontal de `auto`. Esto moverá el marco hacia abajo 20 píxeles y lo centrará horizontalmente en la página.

# --hints--

Debe establecer la propiedad `margin` a `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

El elemento `.frame` debe tener un valor de `margin` de `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
