---
id: 60a3e3396c7b40068ad6996f
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Agrega un elemento `div` en el `body`.

Establece el atributo `class` igual a `canvas`.

Esto actuará como el lienzo para la pintura.

# --hints--

El código debe tener un elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

El elemento `div` debe tener una `class` con el valor `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
