---
id: 60a3e3396c7b40068ad6996e
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El margen es el área exterior de la caja y se puede usar para controlar el espacio entre otras cajas o elementos.

Aquí el elemento inferior tiene un margen superior más grande, empujándolo más abajo en la página.

Ahora que ha revisado rápidamente el modelo de caja CSS, es momento de comenzar con la pintura de Rothko.

Elimine el elemento `img`.

# --hints--

No debe tener un elemento `img` en su código.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
