---
id: 60a3e3396c7b40068ad69970
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Antes de que puedas comenzar a estilizar el `div` que agregaste, necesitas vincular tu CSS a tu HTML.

Agrega un elemento `link` para vincular tu archivo `styles.css`. Establece el `href` a `styles.css`, y recuerda establecer el atributo `rel` a `stylesheet`.

# --hints--

Tu código debe tener un elemento `link`.

```js
assert(/<link/.test(code))
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
