---
id: 60a3e3396c7b40068ad6996b
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comience su proyecto agregando un elemento `img` con un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` dentro del elemento `body`.

# --hints--

Su código debe tener un elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

El elemento `img` debe tener un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
