---
id: 5dc174fcf86c76b9248c6eb2
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, continuarás trabajando con elementos HTML básicos como títulos, párrafos y listas, construyendo una aplicación de fotos de gatos.

Comienza el taller agregando un elemento `h1` con el texto de `CatPhotoApp`.

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Puede que quieras revisar la ortografía.

```js
assert.match(code, /catphotoapp/i);
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen esta sintaxis: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

El texto de tu elemento `h1` debe ser `CatPhotoApp`. Has omitido el texto, cometido un error tipográfico o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
