---
id: 5dc24165f86c76b9248c6ebc
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Todos los elementos `img` deben tener un atributo `alt`. El texto del atributo `alt` se utiliza para lectores de pantalla para mejorar la accesibilidad y se muestra si la imagen no se carga.

Aquí tienes un ejemplo de un elemento `img` con un atributo `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Dentro del elemento `img`, agrega un atributo `alt` con este texto:

`A cute orange cat lying on its back`

# --hints--

Tu código debe tener un elemento `img`. Eliminaste el elemento `img` en un paso anterior.

```js
assert.exists(document.querySelector('img'));
```

Tu elemento `img` no tiene un atributo `alt`. Verifica que haya un espacio después del nombre de la etiqueta de apertura y/o haya espacios antes de todos los nombres de los atributos.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

El valor del atributo `alt` de tu elemento `img` está configurado en algo distinto a `A cute orange cat lying on its back`. Asegúrate de que el valor del atributo `alt` esté rodeado de comillas.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
