---
id: 5dc2385ff86c76b9248c6eb7
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 tiene algunos elementos que identifican diferentes áreas de contenido. Estos elementos hacen que tu HTML sea más fácil de leer y ayudan con la optimización para motores de búsqueda (SEO) y accesibilidad.

El elemento `main` se utiliza para representar el contenido principal del cuerpo de un documento HTML. El contenido dentro del elemento `main` debe ser único para el documento y no debe repetirse en otras partes del documento.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifica la sección principal de esta página añadiendo una etiqueta de apertura `<main>` antes del elemento `h1`, y una etiqueta de cierre `</main>` después del elemento `p`.

# --hints--

Tu elemento `main` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Tu elemento `main` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.match(code, /<\/main\>/);
```

La etiqueta de apertura de tu elemento `main` debe estar debajo de la etiqueta de apertura del elemento `body`. Las tienes en el orden incorrecto.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

La etiqueta de apertura de tu elemento `main` debe estar sobre el elemento `h1`. Las tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

La etiqueta de cierre de tu elemento `main` debe estar debajo del elemento `p`. Las tienes en el orden incorrecto.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
