---
id: 5dc24614f86c76b9248c6ebd
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Puedes enlazar a otra página con el elemento de anclaje (`a`).

Aquí hay un ejemplo que enlaza a `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Añade un elemento de anclaje después del párrafo que enlace a `https://freecatphotoapp.com`. En este punto, el enlace no aparecerá en la vista previa.

# --hints--

Tu elemento de anclaje (`a`) debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Tu elemento de anclaje (`a`) debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.match(code, /<\/a\>/);
```

Tu elemento de anclaje (`a`) debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Tu elemento de anclaje (`a`) no tiene un atributo `href`. Verifica que haya un espacio después del nombre de la etiqueta de apertura y/o que haya espacios antes de todos los nombres de atributos.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

Tu elemento de anclaje (`a`) debe enlazar a `https://freecatphotoapp.com`. O has omitido la URL o tienes un error tipográfico.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Aunque has establecido el atributo `href` del elemento de anclaje (`a`) al enlace correcto, se recomienda rodear siempre el valor de un atributo con comillas.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
