---
id: 5dfa3589eacea3f48c6300ae
title: Paso 20
challengeType: 0
dashedName: step-20
---

# --description--

Dentro del segundo elemento `section`, añade un nuevo elemento `h2` con el texto `Cat Lists`.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Tu elemento `h2` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Tu elemento `h2` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Tu segundo elemento `h2` debe estar justo encima de la etiqueta de cierre del segundo elemento `section`. No está en la posición correcta.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

El segundo elemento `h2` debe tener el texto `Cat Lists`. O has omitido el texto o tienes un error tipográfico.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
