---
id: 5dfa2407b521be39a3de7be1
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Para abrir enlaces en una nueva pestaña, puedes usar el atributo `target` en el elemento de anclaje (`a`).

El atributo `target` especifica dónde abrir el documento vinculado. `target="_blank"` abre el documento vinculado en una nueva pestaña o ventana.

Aquí está la sintaxis básica para un elemento `a` con un atributo `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Agrega un atributo `target` con el valor `_blank` a la etiqueta de apertura del elemento de anclaje (`a`), para que el enlace se abra en una nueva pestaña.

# --hints--

Tu elemento `p` debe tener un elemento de anclaje (`a`) anidado con el texto `cat photos`. Puede que lo hayas eliminado o tengas un error tipográfico.

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

Tu elemento de anclaje (`a`) no tiene un atributo `target`. Verifica que haya un espacio después del nombre de la etiqueta de apertura y/o que haya espacios antes de todos los nombres de atributos.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

El valor del atributo `target` debe ser `_blank`. O has omitido el valor o tienes un error tipográfico. Recuerda que los valores de los atributos deben estar rodeados de comillas.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
