---
id: 5dfb6250eacea3f48c6300b2
title: Paso 24
challengeType: 0
dashedName: step-24
---

# --description--

Después de la lista sin orden, agrega una nueva imagen con un valor de atributo `src` establecido en:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Y su valor de atributo `alt` a:

`A slice of lasagna on a plate.`

# --hints--

Debe haber un elemento `img` justo después de la etiqueta de cierre `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

La nueva imagen no tiene un atributo `alt`. Verifique que haya un espacio después del nombre de la etiqueta de apertura y/o espacios antes de todos los nombres de atributos.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

La nueva imagen debe tener un valor `alt` de `A slice of lasagna on a plate.` Asegúrese de que el valor del atributo `alt` esté rodeado de comillas.

```js
assert.match(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

La nueva imagen no tiene un atributo `src`. Verifique que haya un espacio después del nombre de la etiqueta de apertura y/o espacios antes de todos los nombres de atributos.

```js
assert.isTrue(document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

La nueva imagen debe tener un valor `src` de `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Asegúrese de que el valor del atributo `src` esté rodeado de comillas.

```js
assert.equal(
  document.querySelectorAll('section')[1]?.querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Aunque ha establecido el `src` de la nueva imagen con la URL correcta, se recomienda siempre rodear el valor de un atributo con comillas.

```js
assert.notMatch(code, /\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
