---
id: 5ef9b03c81a63668521804e9
title: Paso 37
challengeType: 0
dashedName: step-37
---

# --description--

Convierte el texto existente `freeCodeCamp.org` en un enlace, envolviéndolo en un elemento de anclaje (`a`). El atributo `href` debe establecerse en `https://www.freecodecamp.org`.

# --hints--

Tu elemento de anclaje (`a`) debe estar anidado dentro del elemento `p` dentro del `footer`. Asegúrate de añadir una etiqueta de apertura y una etiqueta de cierre para el elemento de anclaje (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Tu elemento de anclaje (`a`) debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Tu elemento de anclaje (`a`) debe tener un atributo `href` con el valor `https://www.freecodecamp.org`. Es posible que hayas omitido el atributo/valor o que haya un error tipográfico.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

El texto del enlace debe ser `freeCodeCamp.org`. Has omitido el texto o hay un error tipográfico.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Después de anidar el elemento de anclaje (`a`), el único contenido del elemento `p` visible en el navegador debe ser `No Copyright - freeCodeCamp.org`. Verifica bien el texto, el espaciado o la puntuación tanto del elemento `p` como del elemento de anclaje anidado.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
