---
id: 5ef9b03c81a63668521804ec
title: Paso 40
challengeType: 0
dashedName: step-40
---

# --description--

Note que el contenido completo de la página está anidado dentro de un elemento `html`. El elemento `html` es el elemento raíz de una página HTML y envuelve todo el contenido de la página.

También puedes especificar el idioma de tu página agregando el atributo `lang` al elemento `html`.

Agrega el atributo `lang` con el valor `en` a la etiqueta de apertura `html` para especificar que el idioma de la página es inglés.

# --hints--

Has eliminado el elemento `html` o le falta una etiqueta de apertura o de cierre.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

Tu elemento `html` debe tener un atributo `lang` con el valor `en`. Es posible que hayas omitido el atributo/valor, o que tengas un error tipográfico.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
