---
id: 5ef9b03c81a63668521804ee
title: Paso 41
challengeType: 0
dashedName: step-41
---

# --description--

Todas las páginas deben comenzar con `<!DOCTYPE html>`. Esta cadena especial se conoce como una <dfn>declaración</dfn> y asegura que el navegador intente cumplir con las especificaciones de toda la industria.

`<!DOCTYPE html>` indica a los navegadores que el documento es un documento HTML5, que es la última versión de HTML.

Agrega esta declaración como la primera línea del código.

# --hints--

Tu código debe comenzar con la declaración `<!DOCTYPE html>`. Es posible que hayas omitido la declaración, tengas un error tipográfico o no sea la primera línea de código.

```js
assert.match(code, /\<\s*!DOCTYPE\s+html\s*\>/);
```

Tu `<!DOCTYPE html>` debe estar ubicado en la parte superior del documento.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /^\<\!DOCTYPEhtml\>\<html/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<html lang="en">
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
