---
id: 5ef9b03c81a63668521804e7
title: Paso 35
challengeType: 0
dashedName: step-35
---

# --description--

El elemento `footer` se utiliza para definir un pie de página de un documento o sección. Un pie de página suele contener información sobre el autor del documento, datos de copyright, enlaces a términos de uso, información de contacto y más.

Después del elemento `main`, añade un elemento `footer`.

# --hints--

Has eliminado el elemento `main` o le falta una etiqueta de apertura o cierre.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Tu elemento `footer` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Tu elemento `footer` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.match(code, /<\/footer\>/);
```

Tu elemento `footer` debe estar debajo de la etiqueta de cierre del elemento `main`. Lo has colocado en otro lugar.

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```
