---
id: 66731cd027ef3acb155669f5
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

The first element you will learn about is the `h1` element. The `h1` element is a heading element and is used for the main heading of a webpage.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Tal vez deberías revisar tu ortografía.

```js
assert.match(code, /catphotoapp/i);
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

El texto de tu elemento `h1` debe ser `CatPhotoApp`. O has omitido el texto, tienes un error tipográfico, o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Parece que estás utilizando una extensión del navegador que está modificando la página. Asegúrate de desactivar todas las extensiones del navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
