---
id: 6690bb0b4062b114997f708a
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 has some elements that identify different content areas. These elements make your HTML easier to read and help with Search Engine Optimization (SEO) and accessibility.

El elemento `main` se utiliza para representar el contenido principal del cuerpo de un documento HTML. El contenido dentro del elemento `main` debe ser único para el documento y no debe repetirse en otras partes del documento.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifica la sección principal (main) de esta página añadiendo la etiqueta de apertura `<main>` antes del elemento `h1`, y la de clausura `</main>` después del elemento `p`.

# --hints--

Tu elemento `main` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` element should have a closing tag. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.match(code, /<\/main\>/);
```

Your `main` element's opening tag should be below the `body` element's opening tag. Los tienes en el orden incorrecto.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

La etiqueta de apertura del elemento `main` debe estar por encima del elemento `h1`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Your `main` element's closing tag should be below the `p` element. Los tienes en el orden equivocado.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /see more cat photos in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```
