---
id: 6690be42942f421648b119eb
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

Añade un elemento `img` debajo del elemento `p`. Por ahora no se mostrará ninguna imagen en el navegador.

# --hints--

Tu elemento `img` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Tu elemento `img` no debe tener una etiqueta de cierre `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Solo debes tener un elemento `img`. Elimina los que estén de más.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Tu elemento `img` debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


