---
id: 6690ba16cf76f613b4e36197
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Commenting allows you to leave messages without affecting the browser display. It also allows you to make code inactive. A comment in HTML starts with `<!--`, contains any number of lines of text, and ends with `-->`.

Aquí tienes un ejemplo de un comentario con el `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Añade un comentario sobre el elemento `p` con este texto:

`TODO: Add link to cat photos`

# --hints--

Tu comentario debe iniciar con `<!--`. Faltan uno o más de los caracteres que definen el inicio de un comentario.

```js
assert.match(code, /<!--/);
```

Tu comentario debe terminar con `-->`.  Faltan uno o más de los caracteres que definen el final de un comentario.

```js
assert.match(code, /-->/);
```

Tu código no debe tener caracteres extra de apertura/cierre para crear comentarios. Tienes los caracteres `<!--` o `-->` extra mostrándose en el navegador.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g).length, 2)
assert.isBelow(noSpaces.match(/-->/g).length, 2);
```

Tu comentario debe contener el texto `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

Tu comentario debe estar sobre el elemento `p`. Los tienes en el orden incorrecto.

```js
assert.match(
  code
    .replace(/\s/g, ''),
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```


