---
id: 6690b9626a66f51257ae5a91
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

The `h1` through `h6` heading elements are used to signify the importance of content below them. The lower the number, the higher the importance, so `h2` elements have less importance than `h1` elements.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Only use one `h1` element per page and place lower-importance headings below higher-importance headings.

Debajo del elemento `h1`, añade un elemento `h2` con este texto:

`Cat Photos`

# --hints--

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Solo deberías tener un elemento `h1`. Elimina el que está de más.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

El texto de tu elemento `h1` debe ser 'CatPhotoApp'. Probablemente no has añadido el texto o tienes un error tipográfico.

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

Tu elemento `h2` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Tu elemento `h2` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/h2\>/);
```

Your `h2` element's text should be `Cat Photos`. El texto `Cat Photos` debe estar entre las etiquetas de apertura y cierre del elemento `h2`.

```js
assert.equal(document.querySelector('h2').innerText.toLowerCase(), 'cat photos');
```

Tu elemento `h2` debe estar debajo del elemento `h1`. El elemento `h1` tiene mayor importancia, por lo tanto, debe estar encima del elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```
