---
id: 6698429ba5761b39f8aad814
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

It is time to add a new section. Add a second `section` element below the existing `section` element.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Solo debes añadir una etiqueta de apertura al elemento `section`. Elimina cualquier extra.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Tu elemento `section` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g).length, 2);
```

Solo debes añadir una etiqueta de cierre al elemento `section`. Elimina cualquier extra.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

El segundo elemento `section` no debe estar anidado dentro del primer elemento `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Ambos elementos `section` deben estar entre las etiquetas de apertura y cierre del elemento `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```


