---
id: 66984e0cb4febb50c3794edf
title: Paso 33
challengeType: 0
dashedName: step-33
---

# --description--

The `footer` element is used to define a footer for a document or section. A footer typically contains information about the author of the document, copyright data, links to terms of use, contact information, and more.

Después del elemento `main` añade un elemento `footer`.

# --hints--

Has eliminado el elemento `main` o le falta una etiqueta de apertura o de cierre.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Tu elemento `footer` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Tu elemento `footer` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.match(code, /<\/footer\>/);
```

Tu elemento `footer` debe ir debajo de la etiqueta de cierre del elemento `main`. Lo has puesto en un lugar diferente.

```js
assert.equal(document.querySelector('main').nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
--fcc-editable-region--
    </main>

  </body>
--fcc-editable-region--
</html>
```


