---
id: 671b6e873249bb35c9debfcf
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Añada las palabras `See more` antes del elemento de anclaje y `in our gallery` después del elemento de anclaje.

# --hints--

Debe usar el elemento de anclaje existente, no cree uno nuevo.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

El texto dentro del elemento de anclaje no debe cambiar, debe ser `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

Debe tener `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` en su código.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
