---
id: 69a5f35669099ed52f8563b1
title: Побудуйте розумний автомат для поповнення запасів
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

У цій лабораторній роботі ви створите невелику програму для керування запасами, використовуючи базові концепції JavaScript, такі як масиви, об’єкти, цикли та умовні оператори.

Ви змоделюєте отримання партії товарів для запасів, вирішуватимете, що робити з кожним елементом, і організуєте результати для зберігання.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні реалізувати `parseShipment(rawData)` функцію, яка приймає масив рядків і повертає масив об’єктів із властивостями `{ sku, name, qty, expires, zone }`.

2. Ви повинні реалізувати `planRestock(pantry, shipment)` функцію, яка порівнює поточний запас із вхідною партією та повертає масив дій у формі `{ type: "restock" | "discard" | "donate", item }`.

3. Ви повинні реалізувати `groupByZone(actions)` функцію, яка групує дії за зонами зберігання на основі властивості `zone` кожного елемента.

4. Ви повинні реалізувати `clonePantry(pantry)` функцію, яка повертає глибоку копію запасів, щоб планування змін не впливало на оригінальний список. Глибока копія означає створення нового масиву з новими об’єктами, тому зміни копії не змінять оригінальні запаси.

5. Ви повинні використати всі функції разом, щоб обробити партію товарів і вивести фінальний згрупований об’єкт результатів у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Визначте функцію з іменем `parseShipment`, яка приймає один параметр — масив рядків `rawData`.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Ваша функція `parseShipment` повинна конвертувати рядки партії у масив об’єктів із властивостями: `sku`, `name`, `qty`, `expires` та `zone`.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

Повторювані SKU у партії слід ігнорувати.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Якщо `zone` відсутній у рядку партії, призначте йому значення `general`.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

Значення `qty` слід конвертувати у число.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Визначте функцію з іменем `planRestock`, яка приймає два параметри: `pantry` та `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Ваша функція `planRestock` повинна повертати масив дій із властивостями `type` та `item`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Елементи з кількістю `0` або менше у партії повинні створювати дію з типом `discard`, незалежно від того, чи є цей елемент у запасах.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Якщо елемент партії вже існує у запасах (з тим самим `sku`), дія повинна бути `restock`.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Якщо елемент партії відсутній у запасах, дія повинна бути `"donate"`.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Визначте функцію з іменем `groupByZone`, яка приймає один параметр `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

Ваша функція `groupByZone` повинна повертати дії, згруповані за властивістю `zone` кожного об’єкта.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

Визначте функцію з іменем `clonePantry`, яка приймає один параметр `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` повинна повертати новий масив замість оригінального масиву запасів.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Об’єкти всередині клонованого запасу також повинні бути новими об’єктами (глибока копія).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Функції повинні працювати разом, щоб обробити партію та згрупувати отримані дії.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Виведіть у консоль отримані дії, згруповані за зонами.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
