---
id: 587d78af367417b2b2512b03
title: Erstellen Sie ein Umfrageformular
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten einen Seitentitel in einem `h1`-Element mit einem `id` von `title` haben.
1. Sie sollten eine kurze Erklärung in einem `p`-Element mit einem `id` von `description` haben.
1. Sie sollten ein `form`-Element mit einem `id` von `survey-form` haben.
1. Innerhalb des Formularelements sollten Sie ein erforderliches `input`-Feld haben, um Ihren Namen einzugeben, das ein `id` von `name` und ein `type` von `text` hat.
1. Innerhalb des Formularelements sollten Sie ein erforderliches `input`-Feld haben, um Ihre E-Mail einzugeben, das ein `id` von `email` hat.
1. Wenn Sie eine E-Mail eingeben, die nicht korrekt formatiert ist, sollten Sie einen HTML5-Validierungsfehler sehen.
1. Innerhalb des Formularelements sollten Sie ein `input`-Feld haben, um eine Zahl einzugeben, das ein `id` von `number` hat.
1. Die Zahleneingabe sollte keine Nicht-Zahlen akzeptieren, entweder indem Sie daran gehindert werden, sie einzugeben, oder indem ein HTML5-Validierungsfehler angezeigt wird (abhängig von Ihrem Browser).
1. Wenn Sie Zahlen außerhalb des Bereichs der Zahleneingabe eingeben, die durch die `min`- und `max`-Attribute definiert sind, sollten Sie einen HTML5-Validierungsfehler sehen.
1. Für die Eingabefelder Name, E-Mail und Zahl sollten Sie entsprechende `label`-Elemente im Formular haben, die den Zweck jedes Feldes mit den folgenden IDs beschreiben: `name-label`, `email-label` und `number-label`.
1. Für die Eingabefelder Name, E-Mail und Zahl sollten Sie einen Platzhaltertext haben, der eine Beschreibung oder Anweisungen für jedes Feld gibt.
1. Innerhalb des Formularelements sollten Sie ein `select`-Dropdown-Element mit einem `id` von `dropdown` und mindestens zwei auswählbaren Optionen haben.
1. Innerhalb des Formularelements können Sie eine Option aus einer Gruppe von mindestens zwei Optionsschaltflächen auswählen, die mit dem `name`-Attribut gruppiert sind.
1. Innerhalb des Formularelements können Sie mehrere Felder aus einer Reihe von Kontrollkästchen auswählen, von denen jedes ein `value`-Attribut haben muss.
1. Innerhalb des Formularelements sollten Sie ein `textarea` für zusätzliche Kommentare haben.
1. Innerhalb des Formularelements sollten Sie eine Schaltfläche mit einem `id` von `submit` haben, um alle Eingaben abzusenden.

# --hints--

Sie sollten ein `h1`-Element mit einem `id` von `title` haben.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Ihr `#title` sollte nicht leer sein.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Sie sollten ein `p`-Element mit einem `id` von `description` haben.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Ihr `#description` sollte nicht leer sein.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Sie sollten ein `form`-Element mit einem `id` von `survey-form` haben.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Sie sollten ein `input`-Element mit einem `id` von `name` haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#name` sollte ein `type` von `text` haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Ihr `#name` sollte eine Eingabe erfordern.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ihr `#name` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Sie sollten ein `input`-Element mit einem `id` von `email` haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#email` sollte ein `type` von `email` haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Ihr `#email` sollte eine Eingabe erfordern.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ihr `#email` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Sie sollten ein `input`-Element mit einem `id` von `number` haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#number` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Ihr `#number` sollte ein `type` von `number` haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Ihr `#number` sollte ein `min`-Attribut mit einem numerischen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Ihr `#number` sollte ein `max`-Attribut mit einem numerischen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Sie sollten ein `label`-Element mit einem `id` von `name-label` haben.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Sie sollten ein `label`-Element mit einem `id` von `email-label` haben.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Sie sollten ein `label`-Element mit einem `id` von `number-label` haben.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ihr `#name-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#email-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#number-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#name-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Ihr `#email-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Ihr `#number-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Ihr `#name` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ihr `#email` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ihr `#number` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Sie sollten ein `select`-Feld mit einem `id` von `dropdown` haben.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Ihr `#dropdown` sollte mindestens zwei auswählbare (nicht deaktivierte) `option`-Elemente haben.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Ihr `#dropdown` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Sie sollten mindestens zwei `input`-Elemente mit einem `type` von `radio` (Optionsschaltflächen) haben.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Sie sollten mindestens zwei Optionsschaltflächen haben, die Nachfahren von `#survey-form` sind.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Alle Ihre Optionsschaltflächen sollten ein `value`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Alle Ihre Optionsschaltflächen sollten ein `name`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Jede Optionsschaltflächengruppe sollte mindestens 2 Optionsschaltflächen haben.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

Sie sollten mindestens zwei `input`-Elemente mit einem `type` von `checkbox` (Kontrollkästchen) haben, die Nachfahren von `#survey-form` sind.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Alle Ihre Kontrollkästchen innerhalb von `#survey-form` sollten ein `value`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Sie sollten mindestens ein `textarea`-Element haben, das ein Nachfahre von `#survey-form` ist.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Sie sollten ein `input` oder `button`-Element mit einem `id` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Ihr `#submit` sollte ein `type` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Ihr `#submit` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
