---
id: 662a6bc12cde72c32fb526f0
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

一个 <dfn>interface</dfn> 就像一个类的蓝图。 一个接口包含一个类应该实现的一套方法和属性。  

通过声明一个 `Equation` 类开始这个项目。 你将使用此类来定义一个接口，一个通用方程式的蓝图。

# --hints--

你应该定义一个新的 `Equation` 类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
