---
id: 662bdd364bf2cde1487922a9
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

与其他编程语言不同，Python 在其核心语言中没有实现接口，但是 Python 标准库允许你以简单的方式定义接口。

对于这个项目，你将使用来自 `abc` 模块的实用工具。 因此，将此模块导入你的代码。

# --hints--

你应该导入 `abc` 模块。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
