---
id: 662bd552e1c1d2db1b88ba47
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

现在，定义另一个类名为 `LinearEquation` ，并使它继承 `Equation`。 你将使用此类表示线性方程。

# --hints--

你应该定义一个名为 `LinearEquation` 的类。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

`LinearEquation` 类应该继承 `Equation` 类。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
