---
id: 662bde88dc84f1e249801b1a
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` 代表 *Abstract Base Classes*。 `ABC` 类使你能够将常规类转换为抽象类，抽象类可作为具体类的蓝图。

修改 `import` 语句，只从 `abc` 模块导入 `ABC` 类。 你可以使用以下导入语句 `from y import x` 从模块 `y` 导入特定对象 `x`。

然后，通过继承 `ABC` 类将您的 `Equation` 类变成抽象类。

# --hints--

你应该从 `abc` 模块中导入 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

你的 `Equation` 类应该继承 `ABC`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass

    def analyze(self):
        pass

class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
