---
id: 662bd8260da84bdd5feae419
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

你希望 `LinearEquation` 类实现而不是简单地继承 `Equation` 类内部定义的所有方法，该类应该作为接口。

目前， `Equation` 类只是 `LinearEquation` 的父类。 在接下来的步骤中，你将学习如何将其变成一个正式接口。

现在，创建一个 `Equation` 实例并将其分配给变量 `eq`，创建一个 `LinearEquation` 实例并将其分配给变量 `lin_eq`。

# --hints--

你应该声明一个 `eq` 变量，并为其分配 `Equation` 的实例。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

你应该声明一个 `lin_eq` 变量，并为其分配 `LinearEquation` 的实例。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
