---
id: 662bd456896f16d9bd03f1a6
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在 `Equation` 类中，定义两个新的实例方法名为 `solve` 和 `analyze`。

# --hints--

你应该在 `Equation` 类中定义一个名为 `solve` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

你的 `solve` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

你应该在 `Equation` 类中定义一个名为 `analyze` 的方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

你的 `analyze` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
