---
id: 662fa2e2cf27c09f21f4f5d0
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

在 Python 中，数据类型是在运行期（代码执行时）识别的。因此，你在声明变量时不必指定变量的数据类型。不过，你可以通过 `variable: <data type> = value` 或仅 `variable: <data type>` 来注解变量，以明确它将保存特定的数据类型。请注意，Python 解释器不会强制执行用于注解变量的类型，通常你需要借助外部工具来实现这一点。

在 `Equation` 类中定义一个 `degree` 类属性。 不要设置它的值。 而是使用 `int` 类型注释来表明它将在具体类中存储一个整数。

稍后，你将使用此类属性作为传递的参数验证过程的一部分，以实例化方程对象。

# --hints--

你应该定义名为 `degree` 类属性，并在 `Equation` 类中使用 `int` 来注释。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
