---
id: 662f6d7c92381a3049e4c987
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

接口不必只定义抽象方法，还可以实现可供具体类继承的方法。

在处理 `solve` 和 `analyze` 的实现之前，在 `Equation` 类中，定义一个 `__init__` 方法。 不要使用任何装饰器。

# --hints--

你应该在 `Equation` 类中定义一个 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

你的 `__init__` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
