---
id: 662fbcef5f05e1b84f541a0c
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

每个方程对象都将被实例化，传递与方程系数一样多的参数，从 \\( x \\) 的 n 次方开始一直到零次方，包括可能的系数值 `0`。

例如，`LinearEquation(4, 5)` 将表示方程 \\( 4x + 5 = 0 \\)，其中 `4` 为一阶（此处为最高阶）系数，而 `5` 为零阶系数。

你需要检查传递了正确数量的参数来实例化方程对象。

在 `__init__` 方法中，创建一个 `if` 语句来检查 `args` 的长度是否与方程应具有的系数的数量 (`degree + 1`) 不同。 如果是，则引发 `TypeError` 并使用以下字符串提供自定义消息：`f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`。

然后，通过传递 `2` 和 `3` 来实例化 `lin_eq`，从而修复错误。

# --hints--

你应该创建一个 `if` 语句，检查用于实例化方程的系数数量是否不同于 `degree + 1`。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

你应该在新的 `if` 语句中引发 `TypeError`，并使用提供的字符串返回自定义错误消息。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

你应该传递 `2` 和 `3` 来实例化 `lin_eq`。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
