---
id: 662fc3eba556a6bf800d48c1
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

`isinstance()` 内置函数接收两个参数，并返回一个布尔值，表示作为第一个参数传递的对象是否是作为第二个参数传递的类的实例。

```py
isinstance(7, int) # True
```

你需要检查的另一件事是每个参数都是一个数字。 在第一个 `if` 之后，创建一个 `for` 循环，迭代 `args` 并检查当前迭代的参数是不是 `int` 或 `float` 的实例。 使用 `isinstance()` 函数并传递一个包含 `int` 和 `float` 的元组作为第二个参数。

如果参数不是数字，则引发 `TypeError`，提示 `"Coefficients must be of type 'int' or 'float'"`。

# --hints--

你应该在 `if` 语句后创建一个 `for` 循环来迭代 `args`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_iter().is_equivalent("args")`)) })
```

你应该创建一个 `if` 语句，检查当前系数是不是 `for` 循环中的 `int` 或 `float` 的实例。

```js
({ test: () => assert(runPython(`
var = str(_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_for_vars())
cond1 = f'not isinstance({var}, (int, float))'
cond2 = f'not isinstance({var}, (float, int))'
if_stmt = _Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
if_stmt.is_equivalent(cond1) or if_stmt.is_equivalent(cond2)
`)) })
```

你应该使用提供的字符串在 `if` 语句中引发 `TypeError`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("raise TypeError(\\"Coefficients must be of type 'int' or 'float'\\")")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
```
