---
id: 662f96576ef178927de87975
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

为了被视为抽象的方法，方法应该以 `@abstractmethod` 装饰符装饰。

<dfn>decorator</dfn> 在Python 中被用来修改函数的行为。 这里有一个例子展示怎么使用 `spam` 装饰符：

```py
@spam
def foo():
    pass
```

修改你的导入语句来导入 `abstractmethod` 装饰符，并在 `Equation` 类中装饰 `solve` 和 `analyze` 方法。 这将引起两个异常。

一旦一个类继承 `ABC` 拥有抽象的方法，该类不能再实例化了。 因此，删除 `Equation` 实例以消除错误。

另一个错误是因为 `LinearEquation` 类必须实现接口里定义的所有抽象方法。 请确保在 `LinearEquation` 类中定义它们。 你不能在具体类中使用 `abstractmethod` 装饰器。

# --hints--

你应该从 `abc` 模块中导入 `abstractmethod`。

```js
({ test: () => assert(runPython(`
_Node(_code).has_import("from abc import ABC, abstractmethod") or \\
_Node(_code).has_import("from abc import abstractmethod, ABC") or \\
(_Node(_code).has_import("from abc import abstractmethod") and _Node(_code).has_import("from abc import ABC"))
`)) })
```

你应该使用 `@abstractmethod` 来装饰 `Equation` 类中的 `solve` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_decorators("abstractmethod")`)) })
```

你应该使用 `@abstractmethod` 来装饰 `Equation` 类中的 `analyze` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_decorators("abstractmethod")`)) })
```

你应该在 `LinearEquation` 类中定义一个 `solve` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("solve")`)) })
```

你的 `solve` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_args("self")`)) })
```

你应该在 `LinearEquation` 类中定义一个 `analyze` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("analyze")`)) })
```

你的 `solve` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC


class Equation(ABC):
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass

eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
