---
id: 662fa2e2cf27c09f21f4f5d0
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

在 Python 中，数据类型在运行时（执行代码时）被识别。 因此，当你声明一个变量时，不必指定该变量的数据类型。  Nonetheless, you can annotate a variable to clarify that it will hold a specific data type with `variable: <data type> = value` or just `variable: <data type>`. 请注意，Python 解释器不强制使用注释变量的类型，通常你需要外部工具来执行此操作。

在 `Equation` 类中定义一个 `degree` 类属性。 不要设置它的值。 而是使用 `int` 类型注释来表明它将在具体类中存储一个整数。

稍后，你将使用此类属性作为传递的参数验证过程的一部分，以实例化方程对象。

# --hints--

你应该定义名为 `degree` 类属性，并在 `Equation` 类中使用 `int` 来注释。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
