---
id: 6639fdcc701833a54c364211
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

验证系数后，需要将它们存储在实例属性中。 使用字典推导创建一个字典，其中键是系数的度数，对应的值是系数，并将其分配给名为 `coefficients` 的属性。

例如，用 `2` 和 `4` 实例化的 `LinearEquation` 对象应该具有以下 `coefficients` 属性：`{1: 2, 0: 4}`，因为 `2` 对应于 `x` 的一次，而 `4` 对应于 `x` 的零次。

按照与 `args` 中相同的顺序在新字典中创建键值对。

# --hints--

你应该在 `__init__` 方法中声明一个名为 `coefficients` 的属性。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

你应该使用字典理解来存储你的系数。

```js
({ test: () => runPython(`
import ast
node = _Node(_code).find_class("Equation").find_function("__init__").find_variable("self.coefficients")
assert isinstance(node.tree.value, ast.DictComp)
`) })
```

你的 `coefficients` 属性应该是一个包含度数-系数形式键值对的字典。 记住遵循系数在 `args` 中存储的相同顺序。

```js
({ test: () => runPython(`
actual1 = list(LinearEquation(1, 6).coefficients.items())
expected1 = list({1: 1, 0: 6}.items())
actual2 = list(LinearEquation(-3.5, 0).coefficients.items())
expected2 = list({1: -3.5, 0: 0}.items())
assert actual1 == expected1
assert actual2 == expected2
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass

    
lin_eq = LinearEquation(2, 3)
```
