---
id: 663a32735b317af9812eb0d7
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

在 \\( ax + b = 0 \\) 的线性方程中，斜率就是系数 \\( a \\)，y 截距就是系数 \\( b \\)。

<img alt="线性函数图" src="https://cdn.freecodecamp.org/curriculum/python/linear-equation.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;" />

你将使用 `analyze` 方法提供有关该方程的附加信息。 在 `analyze` 方法内部，将系数解包到变量 `slope` 和 `intercept` 中。

然后，返回一个字典，其中包含键 `'slope'` 和 `'intercept'` 以及斜率和 y 截距的值。 之后，在 `lin_eq` 上调用 `analyze` 并打印结果。


# --hints--

你应该将 `coefficients` 属性中存储的值解包到 `analyze` 方法内的变量 `slope` 和 `intercept` 中。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_stmt("slope, intercept = self.coefficients.values()")`)) })
```

`analyze` 方法应返回一个字典，其中包含键 `'slope'` 和 `'intercept'`，分别包含斜率和 y 截距的值。

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
a = eq.analyze()
assert a['slope'] == 2.2, "Expected different slope"
assert a['intercept'] == 1.5, "Expected different intercept"
`) })
```

你应该调用 `lin_eq` 对象的 `analyze` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq.analyze())")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int

    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        

    @abstractmethod
    def solve(self):
        pass

    @abstractmethod
    def analyze(self):
        pass

class LinearEquation(Equation):
    degree = 1

    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
print(lin_eq.solve())
--fcc-editable-region--
```
