---
id: 663b93aee129b3c4cc07d0db
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

在 `terms` 列表之后，创建一个 `for` 循环并使用 `.items()` 方法遍历存储在 `coefficients` 属性中的键和值。 使用 `n` 和 `coefficient` 作为循环变量。 

在循环内部，创建一个 `if` 语句，检查当前迭代的系数是否具有假值，如果是，则跳过本次循环。 这是因为不能用零值表示系数。

# --hints--

你应该创建一个 `for` 循环来遍历 `coefficients.items()` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_iter().is_equivalent("self.coefficients.items()")`)) })
```

你的 `for` 循环应该使用 `n` 和 `coefficient` 来迭代 `coefficients.items()`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_for_vars().is_equivalent("n, coefficient")`)) })
```

你应该创建一个 `if` 语句来检查 `coefficient` 在 `for` 循环内是否具有假值。

```js
({ test: () => assert(runPython(`
if_cond = _Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_conditions()[0]
conditions = ["not coefficient", "coefficient == 0", "0 == coefficient"]
any(if_cond.is_equivalent(condition) for condition in conditions)
`)) })
```

你应该在新的 `if` 语句中使用 `continue` 关键字。

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_stmt("continue")
`)) })
```

你的 `for` 循环应该放置在声明 `terms` 之后。

```js
({ test: () => assert(runPython(`
loop = str(_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0])
_Node(_code).find_class("Equation").find_function("__str__").is_ordered("terms = []", loop, "equation_string = ' '.join(terms)", "return equation_string")
`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    def __str__(self):
        terms = []
        
--fcc-editable-region--
        equation_string = ' '.join(terms)
        return equation_string        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
