---
id: 663b83a28943e6aa6275a514
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

仍在 `Equation` 类中，定义 `__str__` 方法提供合适的字符串表示你将要创建的方程对象。

现在，在 `__str__` 方法中，声明一个变量 `terms` 并为其分配一个空列表。 你将使用此变量来存储方程的每个项（系数乘以 \\（ x^n \\））。

然后，声明一个变量 `equation_string`，并将 `terms` 列表中元素用空格连接的结果赋给它。 最后，返回 `equation_string`。

# --hints--

你应该在 `Equation` 类中定义一个 `__str__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__str__")`)) })
```

你的 `__str__` 方法应该接收一个 `self` 参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_args("self")`)) })
```

你应该声明一个变量 `terms` 并在 `__str__` 方法中为其分配一个空列表。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("terms = []")`)) })
```

你应该声明一个变量 `equation_string`，并在 `__str__` 方法中将 `terms` 中元素与空格连接的结果分配给它。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_stmt("equation_string = ' '.join(terms)")`)) })
```

你应该从 `__str__` 方法中返回 `equation_string`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").has_return("equation_string")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
    
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
--fcc-editable-region--
    
--fcc-editable-region--    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass


lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
