---
id: 664e4a590b52ba8d2adff19f
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

形式为 \\( ax^2 + bx + c = 0 \\) 的二次方程的判别式（通常用大写希腊字母 delta 表示）等于 \\( Δ = b^2 - 4ac \\)。

在 `QuadraticEquation` 类中，定义一个 `__init__` 方法。 使用 `super()` 从父类调用 `__init__` 方法。 然后，定义一个名为`delta`的新属性，用于存储方程的判别式的值。

# --hints--

你应该在 `QuadraticEquation` 类中定义一个 `__init__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("__init__")`)) })
```

`__init__` 方法应该有两个参数按： `self` 和 `*args` 。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_args("self, *args")`)) })
```

你应该在 `__init__` 方法中调用 `super().__init__(*args)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("__init__").has_call("super().__init__(*args)")`)) })
```

你应该在 `__init__` 方法中声明一个 `delta` 属性，并为其分配方程判别式的值。

```js
({ test: () => runPython(`
eq = QuadraticEquation(2, -3, -4)
assert eq.delta == 41
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    
--fcc-editable-region--    
    def solve(self):
        pass
    
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
