---
id: 6675aaf418b41157f6ccd692
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在是时候回到 `__init__` 方法了。 根据方程类型，你需要在实例化时传递可变数量的参数。

向该方法添加第二个参数 `args` 并使用 `*` 运算符使其接受可变数量的参数。

# --hints--

你的 `__init__` 方法应该采用两个参数，`self` 和 `*args`。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
