---
id: 66759e32b88fb5459b1e0234
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

每当定义它的类被子类化时，就会调用 `__init_subclass__` 方法，并且它能够自定义子类。 该方法采用一个按约定命名的参数 `cls` （代表“class”），代表新的子类。

在你的 `Equation` 类中定义一个 `__init_subclass__` 方法并为其提供 `cls` 参数。

# --hints--

你应该在 `Equation` 类中定义一个带有 `cls` 参数的 `__init_subclass__` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
