---
id: 657efa642593c5746acc5c81
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

У тілі інструкції `if` потрібно перетворити будь-який символ верхнього регістру на нижній регістр і додати підкреслення перед цим символом у нижньому регістрі.

Використовуйте метод рядка `.lower()` для перетворення символів верхнього регістру на символи нижнього регістру. Потім додайте підкреслення перед символом. Призначте результати змінній з ім’ям `converted_character`.

# --hints--

Ви повинні призначити змінену літеру змінній з ім’ям `converted_character`.

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

У вашій інструкції `pass` не повинно бути `if`.

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
