---
id: 657ef2a86d4e545cec9a85fb
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

У camel case і pascal case великі літери позначають початок нових слів. Щоб перетворити вхідний рядок у snake case, потрібно перевірити, чи є символи у вхідному рядку великими.

Ви можете використати метод рядка `.isupper()`, щоб перевірити, чи символ є великим. Цей метод повертає `True`, якщо символ у верхньому регістрі, і `False`, якщо ні.

Всередині циклу `for` додайте інструкцію `if`, щоб перевірити, чи поточний символ у верхньому регістрі. Перемістіть інструкцію `pass` всередину нової інструкції `if`.

# --hints--

Ви повинні написати нову інструкцію if з char.isupper() як умовою. Не забудьте додати двокрапку в кінці та використати ключове слово pass для тіла інструкції if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Ви повинні замінити інструкцію `pass` у тілі циклу на інструкцію `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
