---
id: 657ee28cefc4945568287673
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

З порожнім списком тепер ви можете почати ітерувати по вхідному рядку та перетворювати його у snake case.

Всередині функції, під списком, який ви щойно створили, додайте `for` цикл, щоб ітерувати по `pascal_or_camel_cased_string`. Обов’язково назвіть цільову змінну `char`. Поки що додайте `pass` інструкцію як заповнювач у тілі циклу.

# --hints--

Ви маєте написати новий `for` цикл із цільовою змінною з іменем `char`. Не забудьте двокрапку в кінці та використати `pass` для заповнення тіла циклу.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
