---
id: 657effaa2a5e0277d71f9cbe
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

На цьому етапі змінна `snake_cased_char_list` містить список перетворених символів. Щоб об’єднати ці символи в один рядок, ви можете використати метод `.join()`.

Метод `join` працює шляхом конкатенації кожного елемента списку в рядок, розділених заданим рядком, який називається роздільником.

```py
result_string = ''.join(characters)
```

Приклад вище об’єднує елементи списку `characters` в один рядок, де кожен елемент конкатенується разом, використовуючи порожній рядок як роздільник.

Тепер, одразу після циклу `for`, використайте метод `.join()` для об’єднання елементів у `snake_cased_char_list` з порожнім рядком як роздільником. Призначте результат новій змінній з ім’ям `snake_cased_string`.


# --hints--

Ви повинні об’єднати символи у списку `snake_cased_char_list` у рядок, використовуючи метод `.join()`, і призначити цей рядок змінній з ім’ям `snake_cased_string` після вашого циклу `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
