---
id: 657e928716b77b2277980276
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите програму, яка прийматиме рядок у форматі `camelCase` або `PascalCase` як вхідні дані і перетворюватиме його у рядок у форматі `snake_case` двома способами. Спочатку ви використаєте цикл for, а потім — генератор списку, щоб отримати той самий результат. Ви побачите, як генератор списку може зробити ваш код більш лаконічним.

Почніть із визначення нової функції з назвою `convert_to_snake_case()`, яка прийматиме рядок з ім’ям `pascal_or_camel_cased_string` як вхідні дані. Поки що додайте інструкцію `pass` всередині функції.

# --hints--

Ви повинні визначити нову функцію з назвою `convert_to_snake_case()`. Не забудьте двокрапку в кінці і використайте `pass` для тіла функції.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Ваша функція має приймати параметр з ім’ям `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
