---
id: 657f025ec86c3d7c4177b6be
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Після того, як функція готова, ви тепер можете використовувати її всередині іншої функції.

Створіть нову функцію з назвою `main()`, використовуючи `pass` як тіло функції.

# --hints--

Ви повинні визначити нову функцію з ім’ям `main()`. Не забудьте двокрапку в кінці та використайте `pass` для заповнення тіла функції.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
