---
id: 657f0044be09db790b1eb1c5
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

У pascal case рядки починаються з великої літери. Після перетворення всіх символів у нижній регістр і додавання підкреслення до них, може з’явитися зайве підкреслення на початку рядка.

Найпростіший спосіб це виправити — використати метод рядка `.strip()`, який видаляє з рядка будь-які провідні або кінцеві символи серед набору символів, переданих як аргумент. Наприклад:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

Метод `strip()` застосовується до `original_string`. Це видаляє будь-які провідні та кінцеві підкреслення. Результатом прикладу вище буде рядок `'example_string'`.

Оголосіть нову змінну з ім’ям `clean_snake_cased_string` і призначте їй результат методу `.strip()`, застосованого до `snake_cased_string`, передаючи `'_'` як аргумент методу.

# --hints--

Ви повинні призначити отриманий рядок від `snake_cased_string.strip('_')` змінній з ім’ям `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
