---
id: 657f425dbab54e11993c80f0
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

До цього моменту у цьому проєкті ви використовували цикл `for` для ітерування над вхідним рядком і перетворення його у потрібний вихідний результат. Тепер ви почнете перехід від циклу `for` до спискового включення.

Почніть з того, що перетворіть на коментарі всі рядки коду всередині функції `convert_to_snake_case()`. Не видаляйте їх, оскільки вони будуть корисними, коли ви реалізуєте логіку за допомогою спискового включення.

Не забудьте додати ключове слово `pass` у тіло функції, щоб уникнути помилок під час тестування.

# --hints--

Ви повинні перетворити на коментарі всі 10 рядків коду всередині функції `convert_to_snake_case()` і нічого більше. Додайте `pass`, щоб тимчасово заповнити тіло функції і уникнути помилки.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
