---
id: 657f4add33ea4b1f61ba3dc8
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Нарешті спробуйте цю нову реалізацію, виконавши програму. Змініть вхідний рядок на `'IAmAPascalCasedString'` і перевірте, чи вийде `'i_am_a_pascal_cased_string'`, навіть якщо це неправда.

Якщо ви все зробили правильно, ви повинні побачити, що вхідний рядок перетворився на snake case, як і раніше.

Вітаємо! Тепер ваша функція `convert_to_snake_case()` готова.

# --hints--

Ви повинні змінити вхідний рядок з `'aLongAndComplexString'` на `'IAmAPascalCasedString'` всередині функції `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
