---
id: 657f456223b8c1187b461987
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

У Python <dfn>генератор списку</dfn> — це конструкція, яка дозволяє створити новий список, застосовуючи вираз до кожного елемента існуючого ітерованого об’єкта та за бажанням фільтруючи елементи за умовою. Окрім того, що це коротше, генератори списку часто працюють швидше.

Базовий генератор списку складається з виразу, за яким слідує умова `for`:

```py
spam = [i * 2 for i in iterable]
```

У наведеному прикладі використовується змінна `i` для ітерації над `iterable`. Кожен елемент отриманого списку формується шляхом обчислення виразу `i * 2` на поточній ітерації.

На цьому кроці потрібно заповнити порожній список `snake_cased_char_list` за допомогою синтаксису генератора списку.

Перетворіть свій порожній список на генератор списку, який перетворює кожен символ у `pascal_or_camel_cased_string` на символ нижнього регістру та додає перед ним підкреслення (код, який ви раніше закоментували, може допомогти вам написати вираз). Використовуйте `char` для ітерації над `pascal_or_camel_cased_string`.

# --hints--

Вам слід перетворити `snake_cased_char_list` на генератор списку, який ітерує над `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

Ваш генератор списку має використовувати `char` для ітерації над `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

Ваш генератор списку має обчислювати `'_' + char.lower()` для кожного `char` у `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
