---
id: 657f465f8e718b19c5105ae5
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Спискові включення приймають умовні інструкції, щоб оцінювати наданий вираз лише якщо виконуються певні умови:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Як видно з вихідних даних, список символів, згенерований з `pascal_or_camel_cased_string`, було об’єднано. Оскільки вираз всередині спискового включення оцінюється для кожного символу, результатом є рядок у нижньому регістрі з усіма символами, розділеними підкресленням.

Слідуйте наведеному вище прикладу, щоб додати умову `if` до вашого спискового включення, щоб вираз виконувався лише якщо символ у верхньому регістрі.

# --hints--

Ви повинні додати умову `if` з умовою `char.isupper()` до вашого спискового включення.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
