---
id: 663b16e62fee463b4caf46e9
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Після об’єднання елементів списку `snake_cased_char_list` потрібно видалити будь-які провідні або кінцеві підкреслення з отриманого рядка. Для цього використайте метод `strip` з символом підкреслення `_` як аргументом.

Виклики методів можна ланцюжити, тобто результат одного виклику методу можна використовувати як об’єкт для іншого виклику методу.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

У наведеному вище прикладі метод `.upper()` ланцюжиться до `' '.join(words_list)`, тому `.upper()` викликається на результаті виклику `.join()`.

Змініть інструкцію `return`, додавши ланцюжок виклику методу `''.join(snake_cased_char_list)` до `.strip()`, щоб видалити будь-які провідні або кінцеві підкреслення.

# --hints--

Вам слід змінити оператор return, додавши метод `strip` з символом підкреслення `_` як аргументом.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
