---
id: 6a57c331ae94fa4631b60863
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Il museo ha più di un reperto registrato. Aggiungi un secondo reperto a `collection` con la chiave `102`, seguendo la stessa forma del primo.

Usa questi valori:

```js
{
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
}
```

# --hints--

`collection` dovrebbe avere due elementi.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

Il reperto archiviato sotto `102` dovrebbe corrispondere all'oggetto descritto.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--
  
--fcc-editable-region--
};
```
