---
id: 6a57c331ae94fa4631b60865
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Il personale cerca reperti decine di volte al giorno, quindi il tuo revisore ha bisogno di una ricerca riutilizzabile. Le parentesi quadre possono anche contenere una variabile, quindi `collection[id]` restituisce qualunque reperto a cui si riferisce quell'`id`.

Crea una funzione chiamata `getArtifactTitle` che prende un parametro `id`. Usa `collection[id]` per ottenere il reperto, archivialo in una variabile chiamata `artifact` e restituisci `artifact.title`.

# --hints--

Dovresti avere una funzione chiamata `getArtifactTitle` con un parametro chiamato `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Dovresti assegnare `collection[id]` a una variabile chiamata `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` dovrebbe restituire `artifact.title`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` dovrebbe restituire `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` dovrebbe restituire `"Bronze Tablet"`.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
