---
id: 6a57c331ae94fa4631b6086b
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Sposta l'artefatto con `id` `102` in `"Hall B"` nell'anno `2026`. Poi, chiama `console.log()` con `collection[102].locations` per registrare l'array delle posizioni di quell'artefatto.

# --hints--

Dovresti chiamare `moveArtifact` con `102`, `"Hall B"` e `2026`.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

Dovresti registrare nella console le `locations` del secondo artefatto.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

Dovresti registrare le `locations` dopo aver spostato l'artefatto.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

L'ultima posizione del secondo artefatto dovrebbe essere `{ gallery: "Hall B", year: 2026 }`.

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
