---
id: 6a57c265abbc95dfebc05098
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop costruirai un auditor per la provenienza degli artefatti. Archivierai gli artefatti del museo in un oggetto e scriverai funzioni per cercarli, taggarli, spostarli tra le gallerie e riassumere la loro storia.

Ogni artefatto è archiviato sotto il suo id. Crea un oggetto chiamato `collection` e aggiungigli una voce. Imposta la chiave su `101` e il valore su un oggetto con un `title` di `"Golden Mask"` e una `category` di `"Ceremonial"`.

# --hints--

Dichiara una variabile chiamata `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` dovrebbe essere un oggetto.

```js
assert.isObject(collection);
```

`collection` dovrebbe avere una voce.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

L'artefatto archiviato sotto `101` dovrebbe essere un oggetto.

```js
assert.isObject(collection[101]);
```

L'artefatto archiviato sotto `101` dovrebbe avere un `title` di `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

L'artefatto archiviato sotto `101` dovrebbe avere una `category` di `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
