---
id: 6a57c331ae94fa4631b60866
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

A volte il personale chiederà un reperto che il museo non possiede. Se nessun reperto è archiviato sotto quell'`id`, `collection[id]` è `undefined` e leggere `.title` da esso lancerebbe un errore.

Aggiorna l'istruzione return per usare un condizionale. Restituisci `artifact.title` se viene trovato un reperto, e la stringa `"Artifact not found"` se non lo è.

# --hints--

`getArtifactTitle(101)` dovrebbe restituire ancora `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` dovrebbe restituire `"Artifact not found"` se nessun reperto corrisponde all'`id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` dovrebbe restituire il titolo di qualsiasi reperto trovi, non solo di `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
