---
id: 6a57c331ae94fa4631b60861
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ogni reperto ha una persona responsabile.

Aggiungi una proprietà `curator` all'oggetto archiviato sotto la chiave `101` e imposta il suo valore su un oggetto con un `id` pari a `201` e un `name` pari a `"Earl Sinclair"`.

# --hints--

Il reperto archiviato sotto `101` dovrebbe avere una proprietà `curator`.

```js
assert.property(collection[101], "curator");
```

Il `curator` dovrebbe essere un oggetto.

```js
assert.isObject(collection[101].curator);
```

Il `curator` dovrebbe avere un `id` di `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

Il `curator` dovrebbe avere un `name` pari a `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
