---
id: 6a57c331ae94fa4631b60862
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ogni reperto tiene traccia anche di dove è stato, di come viene descritto e se è esposto.

Aggiungi altre tre proprietà al reperto memorizzato sotto `101`:

- `locations`: un array che contiene due oggetti, `{ gallery: "Hall A", year: 2020 }` e `{ gallery: "Hall C", year: 2024 }`
- `tags`: un array che contiene le stringhe `"gold"` e `"egypt"`
- `onDisplay`: il booleano `true`

# --hints--

La proprietà `locations` del reperto dovrebbe essere un array con i due oggetti di posizione forniti.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

La proprietà `tags` del reperto dovrebbe essere `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

La proprietà `onDisplay` del reperto dovrebbe essere `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
