---
id: 6a57c331ae94fa4631b6086c
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Col tempo, gli artefatti spesso vengono messi in mostra e tolti dalla mostra.

Crea una funzione `toggleDisplayStatus` che prende un parametro `id`. Cerca l'artefatto usando `collection[id]`. Se esiste, imposta la sua proprietà `onDisplay` al valore opposto rispetto a quello attuale.

# --hints--

Dovresti avere una funzione chiamata `toggleDisplayStatus` con un parametro chiamato `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Chiamare `toggleDisplayStatus` dovrebbe invertire il valore `onDisplay` dell'artefatto in entrambe le direzioni.

```js
const before101 = collection[101].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, !before101);
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, before101);
const before102 = collection[102].onDisplay;
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, !before102);
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, before102);
```

`toggleDisplayStatus` non dovrebbe lanciare errori né modificare la collection quando nessun artefatto corrisponde all'`id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => toggleDisplayStatus(999));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

--fcc-editable-region--

--fcc-editable-region--
```
