---
id: 6a57c331ae94fa4631b60868
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

I curatori aggiungono tag agli artefatti per renderli più facili da cercare, e un tag non dovrebbe mai essere aggiunto due volte. Prima di aggiungere un nuovo tag, puoi usare il metodo `includes()` per verificare se esiste già, poiché restituisce `true` quando un array contiene già un valore.

Crea una funzione chiamata `addTag` che prende i parametri `id` e `tag`. Cerca l'artefatto usando `collection[id]`. Se esiste e il suo array `tags` non include già `tag`, aggiungi `tag` a `artifact.tags`.

# --hints--

Dovresti avere una funzione chiamata `addTag` con parametri chiamati `id` e `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Chiamare `addTag` con un nuovo tag dovrebbe aggiungere quel tag all'array `tags` dell'artefatto.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Chiamare `addTag` con un tag che l'artefatto ha già non dovrebbe aggiungere un duplicato.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` dovrebbe aggiungere il tag all'artefatto a cui si riferisce l'`id`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`addTag` non dovrebbe lanciare un errore né modificare la collezione quando nessun artefatto corrisponde all'`id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

Dovresti aggiungere `tag` all'array `tags` esistente dell'artefatto, invece di sostituire l'array.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
