---
id: 612e8279827a28352ce83a72
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

现在，复制这七个 `.key` 元素，并在 `.keys` div 中再粘贴两遍。

# --hints--

应该有 21 个 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

应该有 15 个 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

所有 `.key` 元素都应该在你的 `.keys` 元素中。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

所有 `.black--key` 元素都应该在你的 `.keys` 元素中。

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
