---
id: 612e804c54d5e7308d7ebe56
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `.keys` 元素中，添加七个 `div` 元素，每个都包含类 `key`。

# --hints--

你应该创建七个新的 `div` 元素。

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

你的七个新的 `div` 元素应该在 `.keys` 元素中。

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

你的七个新的 `div` 元素都应该将 `class` 设置为 `key`。

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
