---
id: 612e89d254fe5d3df7d6693d
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

既然已经重置了 html 框模型，就需要对其内部的元素也进行重置。 为此，你需要将所有其他元素的 `box-sizing` 属性设置为继承，以确保目标元素采用与其父元素相同的值。

你还需要定位伪元素，它们是跟随选择器的特殊关键字。 你将使用的两个伪元素是 `::before` 和 `::after` 伪元素。

`::before` 选择器创建一个伪元素，它是被选元素的第一个子元素，而 `::after` 选择器创建一个伪元素，它是最后一个子元素的选定元素。 这些伪元素通常用于创建装饰性内容，你将在本项目后面看到。

现在，创建一个 CSS 选择器以使用 `*` 定位所有元素，包括 `::before` 和 `::after` 伪元素。 将 `box-sizing` 属性设置为 `inherit`。

# --hints--

你应该有一个 `html` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

你的 `html` 选择器应该将 `box-sizing` 属性设置为 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

你应该有一个 `*, ::before, ::after` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

你的 `*, ::before, ::after` 选择器应该将 `box-sizing` 属性设置为 `inherit`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
