---
id: 673b02b03134b04637bf7055
title: створіть багаторазовий нижній колонтитул
challengeType: 25
dashedName: build-a-reusable-footer
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні експортувати компонент `Footer`.
2. Ваш компонент `Footer` має повертати елемент `footer`, який містить усі інші елементи.
3. Ваш елемент `footer` повинен мати:

   - Щонайменше три невпорядковані списки, кожен із яких має щонайменше два елементи списку.
   - Щонайменше один абзац із символом авторського права `©`.
   - Щонайменше три посилання зі значенням `href`, встановленим на `#`, а вміст посилання має бути у вигляді іконки або тексту на ваш вибір.

# --hints--

Ви повинні експортувати компонент `Footer`.

```js
assert.isFunction(window.index.Footer);
```

Ваш компонент `Footer` має повертати елемент `footer`.

```js
const footer = document.getElementById('root').firstElementChild;
assert.equal(footer.tagName, 'FOOTER');
```

Ваш компонент `Footer` повинен містити лише один елемент `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer'), 1);
```

Елемент `footer` не повинен мати братських елементів.

```js
assert.lengthOf(document.getElementById('root').children, 1);
```

Елемент `footer` має містити щонайменше три невпорядковані списки.

```js
assert.isAtLeast(document.querySelectorAll('ul').length, 3);
```

Кожен із ваших невпорядкованих списків має містити щонайменше два елементи списку.

```js
const uls = [...document.querySelectorAll('ul')];
uls.forEach(ul => {
    assert.isAtLeast(ul.querySelectorAll('li').length, 2);
  }
);
```

Ваш `footer` має містити щонайменше один абзац.

```js
assert.isAtLeast(document.querySelectorAll('p').length, 1);
```

У одному з ваших абзаців має бути символ авторського права (`©`).

```js
const texts = [...document.querySelectorAll('p')].map(p => p.textContent);
const hasCopyright = texts.some(text => text.includes('©'));
assert.isTrue(hasCopyright);
```

Ваш `footer` має містити щонайменше три посилання зі значенням `href`, встановленим на `#`.

```js
const links = document.querySelectorAll('a[href="#"]');
assert.isAtLeast(links.length, 3);
```

Жодне з ваших посилань не повинно бути порожнім.

```js
const links = [...document.querySelectorAll('a[href="#"]')];
links.forEach(link => {
    assert.isNotEmpty(link.textContent);
  }
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css

```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css
.footer {
  background-color: #e3e1eb;
  padding: 20px;
  text-align: center;
  border-top: 1px solid #e5e7eb;
  font-family: Arial, sans-serif;
}

.footer-section {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  justify-content: center;
  margin-bottom: 20px;
  align-items: center;
}

.footer ul {
  list-style: none;
  padding: 0;
  margin: 0;
}

.footer li {
  margin-bottom: 10px;
}

.footer a {
  text-decoration: none;
  color: #1f2937;
}

.footer a:hover {
  color: #6366f1;
}

.footer p {
  color: #6b7280;
  margin-top: 20px;
}

.footer-icons {
  margin-top: 20px;
  display: flex;
  justify-content: center;
}

.footer-icons a {
  margin: 0 10px;
  font-size: 24px;
  color: #6b7280;
  text-decoration: none;
}

.footer-icons a:hover {
  color: #6366f1;
}
```

```jsx
export const Footer = () => (
  <footer className='footer'>
    <div className='footer-section'>
      <ul>
        <li>
          <a href='#'>Fitness Dashboard</a>
        </li>
        <li>
          <a href='#'>Services</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Watch Videos</a>
        </li>
        <li>
          <a href='#'>Discord</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Privacy Policy</a>
        </li>
        <li>
          <a href='#'>Terms & Conditions</a>
        </li>
      </ul>
    </div>
    <p>© 2024 Fitness Dashboard. All Rights Reserved.</p>
    <div className='footer-icons'>
      <a href='#'>&#x1F3AE;</a>
      <a href='#'>&#x1F426;</a>
      <a href='#'>&#x1F4BB;</a>
      <a href='#'>&#x1F3C0;</a>
    </div>
  </footer>
);
```
