---
id: 66bb6a9c2dd58b73cd759034
title: بناء تطبيق سجل الدرجات
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `getAverage` تأخذ مصفوفة من درجات الاختبار كمعلمة وتُرجع متوسط الدرجات.
1. يجب أن تملك دالة باسم `getGrade` تأخذ درجة الطالب كمعلمة وتُرجع سلسلة نصية تمثل الدرجة الحرفية بناءً على الدرجة. هذه هي الدرجات والدرجات الحرفية المقابلة لها:

| نطاق الدرجات | الدرجة   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. يجب أن تملك دالة باسم `hasPassingGrade` تأخذ درجة كمعلمة وتُرجع إما `true` أو `false` بناءً على ما إذا كانت الدرجة تمثل درجة ناجحة.
1. يجب أن تستخدم دالة `hasPassingGrade` الدالة `getGrade` للحصول على الدرجة الحرفية، وتستخدمها لتحديد ما إذا كانت الدرجة ناجحة. الدرجة الناجحة هي أي شيء مختلف عن `"F"`.
1. يجب أن تملك دالة باسم `studentMsg` تأخذ مصفوفة من الدرجات ودرجة الطالب كمعلمات. يجب أن تُرجع الدالة سلسلة نصية بالتنسيق التالي:

- `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`، إذا نجح الطالب في المقرر.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`، إذا رسب الطالب في المقرر.

استبدل `average-goes-here` بمتوسط الدرجات الكلي و`grade-goes-here` بدرجة الطالب.
    استخدم `getAverage` للحصول على متوسط الدرجات و`getGrade` للحصول على درجة الطالب.

# --hints--

يجب أن تملك دالة باسم `getAverage`.

```js
assert.isFunction(getAverage);
```

يجب أن تُرجع دالة `getAverage` رقمًا.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

يجب أن تُرجع `getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` القيمة `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

يجب أن تُرجع `getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` القيمة `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

يجب أن تُرجع `getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` القيمة `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

يجب أن تُرجع `getAverage([10, 20, 30, 40, 55, 65, 75, 83])` القيمة `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

يجب أن تُرجع `getAverage([10, 20, 30, 40, 50, 60, 70, 97])` القيمة `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

يجب أن تُرجع دالة `getAverage` متوسط درجات الاختبار.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

يجب أن تملك دالة باسم `getGrade`.

```js
assert.isFunction(getGrade);
```

يجب أن تُرجع دالة `getGrade` القيمة `"A+"` إذا كانت الدرجة `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

يجب أن تُرجع دالة `getGrade` القيمة `"A"` إذا كانت الدرجة بين `90` و`99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

يجب أن تُرجع دالة `getGrade` القيمة `"B"` إذا كانت الدرجة بين `80` و`89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

يجب أن تُرجع دالة `getGrade` القيمة `"C"` إذا كانت الدرجة بين `70` و`79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

يجب أن تُرجع دالة `getGrade` القيمة `"D"` إذا كانت الدرجة بين `60` و`69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

يجب أن تُرجع دالة `getGrade` القيمة `"F"` إذا كانت الدرجة بين `0` و`59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


يجب أن تملك دالة باسم `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

يجب أن تُرجع دالة `hasPassingGrade` قيمة منطقية.

```js
assert.isBoolean(hasPassingGrade(100));
```

يجب أن تُرجع دالة `hasPassingGrade` القيمة `true` إذا كانت الدرجة `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

يجب أن تُرجع دالة `hasPassingGrade` القيمة `false` إذا كانت الدرجة `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

يجب أن تُرجع دالة `hasPassingGrade` القيمة `true` لجميع الدرجات الناجحة.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

يجب أن تملك دالة باسم `studentMsg`.

```js
assert.isFunction(studentMsg);
```

يجب أن تُرجع الدالة `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` الرسالة التالية: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

يجب أن تُرجع الدالة `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` الرسالة التالية: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

يجب أن تُرجع الدالة `studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` الرسالة التالية: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

يجب أن تُرجع الدالة `studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` الرسالة التالية: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

يجب أن تُرجع دالة `studentMsg` الرسالة الصحيحة بناءً على درجة الطالب ومتوسط الصف.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
