---
id: 66cc1a3a39aef47d6473cb2f
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Наступний крок — додати функціональність ділення до вашого калькулятора.

Оголосіть функцію з назвою `calculateQuotient`, яка приймає два параметри, `num1` та `num2`.

Всередині функції поверніть частку двох параметрів.

Потім додайте `console.log`, який викликає функцію `calculateQuotient` з аргументами `7` та `11`.

# --hints--

У вас має бути функція з назвою `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Ваша функція `calculateQuotient` має мати два параметри з іменами `num1` та `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ваша функція `calculateQuotient` має повертати частку двох параметрів.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

У вас має бути `console.log`, який викликає функцію `calculateQuotient` з аргументами `7` та `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
