---
id: 66cc0a9e06e00b75d6782be9
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У попередніх уроках ви дізналися, як викликати функцію. Виклик функції означає виконати код всередині функції.

Ось нагадування, як викликати функцію:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Додайте інструкцію `console.log`, і всередині неї викличте функцію `addTwoAndSeven`.

Тепер ви повинні побачити суму `2` і `7`, виведену в консоль.

# --hints--

У вашому коді має бути `console.log`.

```js
assert.match(code, /console\.log(.*)/);
```

Ви повинні викликати функцію `addTwoAndSeven` всередині `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
