---
id: 66cc0f1ae40802781b2ea972
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

У попередніх уроках ви навчились працювати з параметрами функцій.

Параметр функції — це змінна, яка визначена в оголошенні функції і слугує заповнювачем.

Ось приклад функції, яка має параметр:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Оголосіть функцію з назвою `calculateSum`, яка приймає два параметри, `num1` і `num2`.

# --hints--

Ви повинні оголосити функцію з назвою `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Ваша функція `calculateSum` має приймати два параметри, `num1` і `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
