---
id: 66cc0de2920bec775f610424
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Хоча функції працюють як очікувалося, у вашому коді багато повторень.

Оскільки ви створюєте калькулятор, не хочеться створювати функцію для кожної можливої комбінації чисел, які ви хочете додати.

Краще створити одну багаторазову функцію, яка може додавати будь-які два числа.

Видаліть увесь код, який ви написали досі.

У наступному кроці ви повторите, як працювати з параметрами та аргументами у функціях.

# --hints--

Вам не слід мати функцію з назвою `addTwoAndSeven`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

Вам не слід мати виклик функції `addTwoAndSeven`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

Вам не слід мати функцію з назвою `addThreeAndFour`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

Вам не слід мати виклик функції `addThreeAndFour`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
