---
id: 66cc0ba5881acb7692cfc4de
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер настав час додати ще одну функцію.

Оголосіть функцію з назвою `addThreeAndFour`, яка повертає суму `3` і `4`.

Потім викличте функцію `addThreeAndFour` всередині `console.log`, щоб побачити результат.

# --hints--

Ви повинні оголосити функцію з назвою `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Ваша функція `addThreeAndFour` має повертати суму `3` і `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Ваша функція `addThreeAndFour` має повертати суму, використовуючи оператор `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Ви повинні викликати функцію `addThreeAndFour` всередині `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
