---
id: 66cc14ca14e65e7a9ebecb80
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Тепер, коли ваш калькулятор може додавати два числа, настав час створити функцію, яка відніматиме два числа.

Оголосіть функцію з назвою `calculateDifference`, яка прийматиме два параметри: `num1` і `num2`.

Всередині функції поверніть різницю між двома параметрами.

# --hints--

Ви повинні оголосити функцію з назвою `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Ваша функція `calculateDifference` має мати два параметри з іменами `num1` і `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateDifference } = explorer.allFunctions;
assert.lengthOf(calculateDifference?.parameters, 2);
assert.equal(calculateDifference?.parameters[0].toString(), "num1");
assert.equal(calculateDifference?.parameters[1].toString(), "num2");
```

Ваша функція `calculateDifference` має повертати різницю між двома параметрами.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
