---
id: 66cc172af7159a7c67804544
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер настав час додати функціонал множення до вашого калькулятора.

Оголосіть функцію з назвою `calculateProduct`, яка приймає два параметри, `num1` та `num2`.

Всередині функції поверніть добуток цих двох параметрів.

Потім додайте `console.log`, який викликає функцію `calculateProduct` з аргументами `13` та `5`.

# --hints--

У вас має бути функція з назвою `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Ваша функція `calculateProduct` має мати два параметри з іменами `num1` та `num2`.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ваша функція `calculateProduct` має повертати добуток двох параметрів.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

У вас має бути `console.log`, який викликає функцію `calculateProduct` з аргументами `13` та `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
