---
id: 66cc16abc30b367ba1d1f9ef
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Тепер настав час протестувати вашу функцію `calculateDifference`.

Почніть із додавання `console.log`, який викликає функцію `calculateDifference` з аргументами `22` та `5`.

Потім додайте другий `console.log`, який викликає функцію `calculateDifference` з аргументами `12` та `1`.

Нарешті, додайте третій `console.log`, який викликає функцію `calculateDifference` з аргументами `17` та `9`.

# --hints--

У вас має бути `console.log`, який викликає функцію `calculateDifference` з аргументами `22` та `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

У вас має бути `console.log`, який викликає функцію `calculateDifference` з аргументами `12` та `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

У вас має бути `console.log`, який викликає функцію `calculateDifference` з аргументами `17` та `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
