---
id: 66cc06012e46aa6bc9b8c001
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви повторите, як працювати з функціями, створюючи застосунок калькулятора.

У попередніх уроках ви навчились оголошувати функції ось так:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Почніть зі створення функції з назвою `addTwoAndSeven`.

Ви можете обрати використання звичайного синтаксису функції або стрілкового синтаксису.

# --hints--

Ви маєте оголосити функцію з назвою `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
