---
id: 66cc0a9e06e00b75d6782be9
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У попередніх уроках ви дізналися, як викликати функцію. Виклик функції означає виконати код всередині функції.

Ось нагадування, як викликати функцію:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Додайте інструкцію `console.log`, і всередині неї викличте функцію `addTwoAndSeven`.

Тепер ви повинні побачити суму `2` і `7`, виведену в консоль.

# --hints--

Ви повинні викликати функцію `addTwoAndSeven` всередині `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
