---
id: 66cc172af7159a7c67804544
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер настав час додати функціонал множення до вашого калькулятора.

Оголосіть функцію з назвою `calculateProduct`, яка приймає два параметри, `num1` та `num2`.

Всередині функції поверніть добуток цих двох параметрів.

Потім додайте `console.log`, який викликає функцію `calculateProduct` з аргументами `13` та `5`.

# --hints--

Ви повинні оголосити функцію з назвою `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Ваша функція `calculateProduct` має мати два параметри з іменами `num1` та `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

Ваша функція `calculateProduct` має повертати добуток двох параметрів.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

У вас має бути `console.log`, який викликає функцію `calculateProduct` з аргументами `13` та `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
