---
id: 66cc0de2920bec775f610424
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Хоча функції працюють як очікувалося, у вашому коді багато повторень.

Оскільки ви створюєте калькулятор, не хочеться створювати функцію для кожної можливої комбінації чисел, які ви хочете додати.

Краще створити одну багаторазову функцію, яка може додавати будь-які два числа.

Видаліть увесь код, який ви написали досі.

У наступному кроці ви повторите, як працювати з параметрами та аргументами у функціях.

# --hints--

Вам не слід мати функцію з назвою `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Вам не слід мати виклик функції `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Вам не слід мати функцію з назвою `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Вам не слід мати виклик функції `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
