---
id: 66cc10958d72e57907f898b2
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Усередині вашої `calculateSum` функції потрібно повернути суму двох параметрів.

# --hints--

Ваша `calculateSum` функція має повертати суму двох параметрів.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
