---
id: 66cc095dfe1682753d2ab030
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Одне з понять, яке ви вивчили в попередніх уроках, — це як повертати значення з функції.

Ось нагадування про те, як повернути значення з функції:

```js
function myFunction() {
  return "Hello World";
}
```

Усередині вашої `addTwoAndSeven` функції поверніть суму `2` і `7`.

# --hints--

Ваша `addTwoAndSeven` функція має повертати суму `2` і `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Ваша `addTwoAndSeven` функція має повертати суму, використовуючи оператор `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
const validExpressions = ["2 + 7", "7 + 2", "(2 + 7)", "(7 + 2)"];
assert.isTrue(validExpressions.some(expr => addTwoAndSeven?.hasReturn(expr)));
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
