---
id: 66cc1a3a39aef47d6473cb2f
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Наступний крок — додати функціональність ділення до вашого калькулятора.

Оголосіть функцію з назвою `calculateQuotient`, яка приймає два параметри, `num1` та `num2`.

Всередині функції поверніть частку двох параметрів.

Потім додайте `console.log`, який викликає функцію `calculateQuotient` з аргументами `7` та `11`.

# --hints--

Ви повинні оголосити функцію з назвою `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Ваша функція `calculateQuotient` має мати два параметри з іменами `num1` та `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

Ваша функція `calculateQuotient` має повертати частку двох параметрів.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

У вас має бути `console.log`, який викликає функцію `calculateQuotient` з аргументами `7` та `11`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
