---
id: 66cc12fa504b0479dac479a0
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Перевага функцій у тому, що їх можна викликати з різними аргументами, що дозволяє повторно використовувати один і той самий код з різними значеннями.

Ось кілька прикладів викликів функцій з різними рядковими аргументами:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Додайте `console.log`, який викликає функцію `calculateSum` з аргументами `2` та `5`.

# --hints--

У вас має бути `console.log`, який викликає функцію `calculateSum` з аргументами `2` та `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
