---
id: 66cc095dfe1682753d2ab030
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Однією з концепцій, яку ви вивчали в попередніх уроках, було те, як повертати значення з функції.

Ось нагадування про те, як повернути значення з функції:

```js
function myFunction() {
  return "Hello World";
}
```

Усередині вашої `addTwoAndSeven` функції поверніть суму `2` і `7`.

# --hints--

Ваша `addTwoAndSeven` функція має повертати суму `2` і `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Ваша `addTwoAndSeven` функція має повертати суму, використовуючи оператор `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
