---
id: 66cc2739f687e484d50bb6f1
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

Останньою функціональністю для вашого калькулятора буде функція квадратного кореня.

Квадрат — це число, помножене саме на себе. Отже, квадратний корінь — це число, яке при множенні на себе дає початкове число.

Наприклад, квадратний корінь з `9` дорівнює `3`, тому що `3 * 3 = 9`.

Щоб отримати квадратний корінь числа в JavaScript, ви можете використати метод <dfn>Math.sqrt()</dfn>.

Ось приклад:

```js
// result: 3
Math.sqrt(9);
```

Оголосіть функцію з назвою `calculateSquareRoot`, яка приймає параметр `num` і повертає квадратний корінь з `num`.

# --hints--

У вас має бути функція з назвою `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Ваша функція `calculateSquareRoot` має приймати параметр `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ваша функція `calculateSquareRoot` має повертати квадратний корінь з `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
