---
id: 66cc0ba5881acb7692cfc4de
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер час додати ще одну функцію.

Оголосіть функцію з назвою `addThreeAndFour`, яка повертає суму `3` і `4`.

Потім викличте функцію `addThreeAndFour` всередині `console.log`, щоб побачити результат.

# --hints--

Ви повинні оголосити функцію з назвою `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Ваша функція `addThreeAndFour` має повертати суму `3` і `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

Ваша функція `addThreeAndFour` має повертати суму, використовуючи оператор `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

Ви повинні викликати функцію `addThreeAndFour` всередині `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
