---
id: 66cc0f1ae40802781b2ea972
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

У попередніх уроках ви навчились працювати з параметрами функцій.

Параметр функції — це змінна, яка визначена в оголошенні функції і слугує заповнювачем.

Ось приклад функції, яка має параметр:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Оголосіть функцію з назвою `calculateSum`, яка приймає два параметри, `num1` і `num2`.

# --hints--

Ви повинні оголосити функцію з назвою `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Ваша функція `calculateSum` має приймати два параметри, `num1` і `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
