---
id: 66cc21d23238dc8240a8a182
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Було б чудово, якби ваш калькулятор міг обчислювати квадрат числа. Квадрат числа — це число, помножене саме на себе.

Щоб обчислити квадрат числа в JavaScript, ви можете використати метод <dfn>Math.pow()</dfn>. Метод `Math.pow()` приймає два аргументи: основне число та степінь.

Також можна використати оператор піднесення до степеня (`**`) для обчислення квадрата числа.

Ось приклад:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Оголосіть функцію з назвою `calculateSquare`, яка приймає параметр `num` і повертає квадрат `num`.

# --hints--

У вас має бути функція з назвою `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Ваша функція `calculateSquare` має приймати параметр `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ваша функція `calculateSquare` має повертати квадрат `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
