---
id: 66cdf76685e4cb5a8726e27b
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Щоб далі перевірити вашу функцію `calculateSum`, потрібно викликати її з різними аргументами.

Додайте ще один `console.log`, який викликає функцію `calculateSum` з аргументами `10` та `10`.

Під цим `console.log` додайте ще один `console.log`, який викликає функцію `calculateSum` з аргументами `5` та `5`.

# --hints--

У вас має бути `console.log`, який викликає функцію `calculateSum` з аргументами `10` та `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

У вас має бути `console.log`, який викликає функцію `calculateSum` з аргументами `5` та `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
