---
id: 6557746aad2844a0cd864e12
title: Крок 27
challengeType: 20
dashedName: step-27
---

# --description--

Поки алгоритм досліджує граф, він має відстежувати наразі відому найкоротшу відстань між початковим вузлом і іншими вузлами.

Перед вашим циклом `for` створіть нову змінну з іменем `distances` і призначте їй порожній словник.

# --hints--

У вас має бути змінна з іменем `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Ваша змінна `distances` має бути порожнім словником.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
