---
id: 65577236b056379d5dbc7000
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Тепер ви почнете розробляти алгоритм для обчислення найкоротшого шляху між кожним вузлом у вашому новому графі.

Оголосіть порожню функцію з назвою `shortest_path`. Використайте ключове слово `pass` для заповнення тіла функції.

# --hints--

Ви повинні визначити функцію з назвою `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
