---
id: 65579228c669fcbebffd01d5
title: Крок 52
challengeType: 20
dashedName: step-52
---

# --description--

Інша помилка більш тонка. Коли для сусіднього вузла знаходиться коротша відстань, `paths[current]` призначається шляху сусіднього вузла, `paths[node]`.

Це означає, що обидві змінні вказують на один і той самий список. Оскільки списки змінні, коли ви додаєте сусідній вузол до його шляху, і `paths[node]`, і `paths[current]` змінюються, бо це один і той самий список. Це призводить до неправильних шляхів, хоча відстані правильні.

Ви можете виправити цю помилку, призначивши копію `paths[current]` шляху сусіднього вузла. Для цього можна використати синтаксис зрізу:

```py
my_list[:]
```

Де `my_list` — це список, який ви хочете скопіювати. Змініть існуюче присвоєння `paths[node] = paths[current]` всередині вашого блоку `if`, зробивши зріз `paths[current]`.

# --hints--

Ви повинні призначити `paths[current][:]` шляху сусіднього вузла.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
