---
id: 65578cb031cd93b77a285db2
title: Крок 38
challengeType: 20
dashedName: step-38
---

# --description--

Ваша функція буде досліджувати всі вузли, пов’язані з початковим вузлом. Вона обчислить найкоротші шляхи для всіх них. Потім початковий вузол буде видалений із невідвіданих вузлів.

Далі буде відвіданий найближчий сусідній вузол, і процес повториться, доки всі вузли не будуть відвідані.

Відтепер ви працюватимете над основним циклом, який досліджує вузли в графі. Щоб уникнути проблем із запуском нескінченного циклу під час розробки алгоритму, перетворіть виклик функції на коментар.

# --hints--

Ви повинні перетворити виклик функції на коментар.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
