---
id: 6557712d77ce2d9bd7e63afd
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Додайте ще один вузол, з’єднаний із `B`, до свого графа і назвіть його `C`.

Змініть свій існуючий словник, щоб представити цю структуру: додайте ще один ключ `'C'` до `my_graph` і надайте йому значення рядка `'B'`.

Також змініть значення існуючого ключа `'B'` на список `['A', 'C']`, щоб представити множинні зв’язки вашого вузла `'B'`.

# --hints--

Ваш словник має містити 3 ключі — `'A'`, `'B'` та `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Значення `my_graph['A']` має бути рядком `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` має бути списком.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Значення `my_graph['B']` має бути списком, що містить `'A'` і `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Значення `my_graph['C']` має бути рядком `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
