---
id: 65577791ad8c26a7705e2919
title: Крок 34
challengeType: 20
dashedName: step-34
---

# --description--

За допомогою словникового включення ви можете створити словник, починаючи з існуючого словника:

```py
{key: val for key in dict}
```

У наведеному вище прикладі `val` — це значення, яке матиме `key` у новому словнику, а `dict` — це існуючий словник.

Ви хочете відстежувати шляхи між початковим вузлом і кожним іншим вузлом.

Після змінної `distances` створіть змінну `paths` і призначте їй словник із усіма ключами зі `graph`. Призначте кожному ключу порожній список і використайте словникове включення, щоб побудувати свій словник.

# --hints--

У вас має бути змінна з назвою `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Ваша змінна `paths` має використовувати синтаксис словникового включення, щоб призначити порожній список кожному вузлу в graph.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
