---
id: 655771d889132f9ccd341060
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Граф називають *зваженим*, коли його ребра мають ваги, що представляють відстань, час або інше кількісне значення.

У вашому випадку ці ваги будуть відстанями між кожним вузлом, або точкою в просторі. Щоб представити зважений граф, ви можете змінити свій словник, використовуючи список кортежів для кожного значення.

Перший елемент у кортежі буде з’єднаним вузлом, а другий елемент — цілим числом, що вказує відстань.

Змініть `my_graph['A']` на список кортежів, враховуючи, що відстань `A-B` дорівнює `3`, а відстань `A-D` дорівнює `1`.

# --hints--

`my_graph["A"]` має бути списком, що містить кортежі `('B', 3)` і `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` має містити 4 ключі з іменами `'A'`, `'B'`, `'C'` та `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
