---
id: 6557709b0aee699a6a00528c
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Наприклад, граф можна використати, щоб представити дві точки у просторі, `A` і `B`, з’єднані шляхом. Такий граф складатиметься з двох вузлів, з’єднаних ребром.

Замініть існуючий ключ `'species'` на рядки `'A'`. Потім замініть відповідне значення на рядок `'B'`, щоб представити з’єднання між вузлами `'A'` і `'B'`.

# --hints--

Ваш словник має містити ключ `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ваш ключ `'A'` має мати значення `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Ваш словник має містити два ключі.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
