---
id: 655774955b097ea14897db12
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Відстань від початкового вузла дорівнює нулю, оскільки алгоритм починає оцінку саме з нього.

Після додавання `node` до `unvisited` у вашому циклі створіть інструкцію `if`, яка спрацьовує, якщо вузол дорівнює початковому вузлу. Потім призначте `0` цьому вузлу всередині словника `distances`.

# --hints--

Ви повинні створити інструкцію `if`, яка виконується, коли `node` дорівнює `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

Всередині вашої нової інструкції `if` потрібно призначити `0` вузлу у словнику `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
