---
id: 65576ff7888f9e96f52a4be1
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Словники зберігають дані у вигляді пар *ключ*-*значення*. Ключ відокремлюється від відповідного значення двокрапкою. А кожна пара ключ-значення відокремлюється від наступної пари комою:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Додайте нову пару ключ-значення до свого словника. Використайте рядок `'species'` як ключ і рядок `'guinea pig'` як значення.

# --hints--

У вашому словнику `'species'` має бути ключ `'guinea pig'` зі значенням `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Ваш словник `copper` повинен містити лише одну пару ключ-значення.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
