---
id: 6566195b0a021bb660b2b4b1
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Тепер змініть `my_graph['B']` на список кортежів, де перший елемент кортежу — це пов’язаний вузол, а другий елемент — відстань. Відстань `B-C` дорівнює `4`.

# --hints--

`my_graph['B']` має бути списком, що містить кортежі `('A', 3)` і `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` має містити 4 ключі з іменами `'A'`, `'B'`, `'C'` і `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
