---
id: 6578b57361f2f132a02e2a18
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви можете отримати доступ до даних, збережених у словнику, через його ключі:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Після вашого словника, дотримуйтесь наведеного вище прикладу, щоб отримати доступ до ключа `'species'` у `copper` і вивести результат.

# --hints--

Ви не повинні змінювати присвоєння вашого словника.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Ви повинні використовувати `copper['species']`, щоб отримати значення ключа `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Ви повинні викликати `print()`, передаючи `copper['species']` як аргумент.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
