---
id: 65789506b30453080f77470c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

До цього моменту ви вже зустрічали різні типи даних:

- Незмінні типи даних, такі як цілі числа, рядки, кортежі та булеві.
- Змінні типи даних, такі як списки та словники.

Словник позначається парою фігурних дужок, `{}`.

Почніть зі створення змінної з назвою `copper` і призначте їй порожній словник, використовуючи пару фігурних дужок, так само, як ви створюєте порожній список за допомогою пари квадратних дужок.

# --hints--

У вас має бути змінна з назвою `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Значення вашої змінної `copper` має бути порожнім словником. Для цього використовуйте пару фігурних дужок.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
