---
id: 6578b13757611e2825beb8a5
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Ключі мають бути унікальними в межах словника і можуть бути лише незмінними типами даних. Це означає, що ви не можете використовувати список або інший словник як ключі.

Додайте ще один ключ `'age'` до свого словника і надайте йому ціле число `2` як значення.

# --hints--

У вашому словнику `'age'` має бути ключ `2` зі значенням `'species': 'guinea pig'` одразу після `copper`. Не забудьте кому.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

У вашому словнику `copper` має бути дві пари ключ-значення.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
