---
id: 6579717f0920131304286804
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Щоб додати нову пару ключ-значення після оголошення словника, ви можете вказати ключ так само, як і для доступу до існуючого ключа, і призначити значення нового ключа, використовуючи оператор присвоєння:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Видаліть свій виклик `print()`. Потім, після оголошення `copper`, додайте ключ `'food'` до свого словника і призначте його значення `'hay'`.

# --hints--

У вашому коді не повинно бути `print(copper['age'])`.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Ви повинні додати ключ `'food'` до `copper` після оголошення словника.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Ви повинні призначити `copper['food']` значення `'hay'` після оголошення словника.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
