---
id: 6579ca0923cfa7162089d2f0
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Той самий синтаксис можна використати, щоб змінити значення існуючого ключа.

Безпосередньо перед викликом `print()` отримайте доступ до ключа `'species'` і призначте йому нове значення `'Cavia porcellus'`.

# --hints--

Ви повинні призначити `copper['species']` рядку `'Cavia porcellus'` перед викликом `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
