---
id: 6579dd49fa8a8e1fd06b85a9
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Графи — це структури даних, які представляють відношення між парами елементів.
Ці елементи, які називаються *вузлами*, можуть бути реальними об’єктами, сутностями, точками в просторі або іншими. Зв’язки між вузлами називаються *ребрами*.

Ось візуальне представлення графа:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Перейменуйте словник `copper` у `my_graph`. Це буде представляти граф для тестування вашого алгоритму.

# --hints--

Ви повинні перейменувати свій словник `copper` у `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Змінна `my_graph` має бути словником.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
