---
id: 6579cbab9825b8170974c69a
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Ви можете видалити пару ключ-значення з словника, використовуючи ключове слово `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Безпосередньо перед вашим циклом `for` використайте ключове слово `del`, щоб видалити ключ `'age'` та його значення зі словника `copper`.

# --hints--

Ви повинні використати ключове слово `del`, щоб видалити `copper['age']` перед циклом `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
