---
id: 65affe602a2a828b0382a058
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Замініть ключ `'age'` рядком `'B'` і встановіть його значення рядком `'A'`, щоб представити зв’язок між вузлами в обох напрямках.

# --hints--

Ваш словник має містити ключ `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Ваш ключ `'B'` має мати значення `'A'`.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Ваш словник має містити ключ `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ваш ключ `'A'` має мати значення `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Ваш словник має містити два ключі.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
