---
id: 65b7cd2b7bd9a684ccf1dc16
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Щоб ітерувати над ключами словника, ви можете просто помістити словник у `for` цикл. Код нижче виведе кожен ключ у словнику `dict`:

```py
for i in dict:
   print(i)
```

Замініть виклик `print()` на `for` цикл, який ітерує над `copper` і виводить кожен ключ.

# --hints--

У вашому коді не має бути `print(copper)`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

Ви маєте створити `for` цикл, щоб ітерувати над словником `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

Ви маєте виводити кожен ключ `copper` всередині вашого `for` циклу.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
