---
id: 65b7d25992879698180e6a71
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Нарешті, якщо ви хочете ітерувати над парами ключ-значення, ви можете використати метод `.items()`.

Змініть свій цикл `for`, щоб ітерувати над `copper.items()` замість `copper.values()`.

# --hints--

Вам слід змінити свій цикл `for`, щоб ітерувати над `copper.items()` замість `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
