---
id: 65b7d3319ebec69b983fb91b
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Як видно з вихідних даних, `.items()` створює структуру даних, яка зберігає кожну пару ключ-значення у окремому кортежі. Щоб ітерувати над елементами в цих кортежах, ви можете додати другу змінну циклу:

```py
for i, j in dict.items():
    print(i, j)
```

Змініть свій цикл `for`, щоб він приймав дві змінні циклу і виводив обидві всередині тіла циклу.

# --hints--

Вам слід змінити свій цикл `for`, щоб використовувати дві змінні.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Вам слід виводити обидві змінні циклу всередині вашого циклу `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
