---
id: 587d7dbc367417b2b2512bb1
title: 간단한 JSX 요소 만들기
challengeType: 6
forumTopicId: 301390
dashedName: create-a-simple-jsx-element
---

# --description--

React는 Facebook이 만들고 유지하는 오픈 소스 뷰 라이브러리입니다. 현대 웹 애플리케이션의 사용자 인터페이스(UI)를 렌더링하는 훌륭한 도구입니다.

React는 JavaScript의 구문 확장인 JSX를 사용하여 JavaScript 내에서 직접 HTML을 작성할 수 있게 합니다. 이는 여러 가지 이점이 있습니다. JavaScript의 모든 프로그래밍 기능을 HTML 내에서 사용할 수 있게 하고, 코드를 읽기 쉽게 유지하는 데 도움이 됩니다. 대부분 JSX는 이미 배운 HTML과 비슷하지만, 이 도전 과제들에서 다룰 몇 가지 중요한 차이점이 있습니다.

예를 들어, JSX는 JavaScript의 구문 확장이기 때문에 JSX 내에서 실제로 JavaScript를 직접 작성할 수 있습니다. 이렇게 하려면 JavaScript로 처리되길 원하는 코드를 중괄호 `{ 'this is treated as JavaScript code' }` 안에 포함시키면 됩니다. 이 점을 기억하세요. 앞으로 여러 도전 과제에서 사용됩니다.

하지만 JSX는 유효한 JavaScript가 아니므로 JSX 코드는 JavaScript로 컴파일되어야 합니다. 트랜스파일러 Babel이 이 과정을 위한 인기 있는 도구입니다. 편의를 위해 이 도전 과제들에서는 이미 백그라운드에서 추가되어 있습니다. 만약 구문적으로 유효하지 않은 JSX를 작성하면, 이 도전 과제의 첫 번째 테스트가 실패하는 것을 볼 수 있습니다.

내부적으로 도전 과제들은 `ReactDOM.render(JSX, document.getElementById('root'))`을 호출하고 있다는 점도 주목할 만합니다. 이 함수 호출이 JSX를 React의 자체 경량 DOM 표현에 배치합니다. React는 자체 DOM의 스냅샷을 사용해 실제 DOM의 특정 부분만 최적화하여 업데이트합니다.

# --instructions--

현재 코드는 JSX를 사용해 상수 `div`에 `JSX` 요소를 할당하고 있습니다. `div`를 `h1` 요소로 바꾸고 그 안에 텍스트 `Hello JSX!`를 추가하세요.

# --hints--

상수 `JSX`는 `h1` 요소를 반환해야 합니다.

```js
assert(JSX.type === 'h1');
```

`h1` 태그 안에는 텍스트 `Hello JSX!`가 포함되어야 합니다.

```js
assert(Enzyme.shallow(JSX).contains('Hello JSX!'));
```

# --seed--

## --seed-contents--

```jsx
const JSX = <div></div>;
```

# --solutions--

```jsx
const JSX = <h1>Hello JSX!</h1>;
```
