---
id: 5a24c314108439a4d403618c
title: 배열.filter()를 사용해 배열을 동적으로 필터링하기
challengeType: 6
forumTopicId: 301416
dashedName: use-array-filter-to-dynamically-filter-an-array
---

# --description--

`map` 배열 메서드는 React 작업 시 자주 사용하게 될 강력한 도구입니다. `map`과 관련된 또 다른 메서드는 `filter`로, 조건에 따라 배열의 내용을 필터링한 후 새 배열을 반환합니다. 예를 들어, 모든 사용자 객체가 `online` 속성을 가지고 있고 이 속성은 `true` 또는 `false`로 설정될 수 있다면, 다음과 같이 온라인 상태인 사용자만 필터링할 수 있습니다:

```js
let onlineUsers = users.filter(user => user.online);
```

# --instructions--

코드 편집기에서 `MyComponent`의 `state`은 사용자 배열로 초기화되어 있습니다. 일부 사용자는 온라인이고 일부는 아닙니다. 온라인 상태인 사용자만 보이도록 배열을 필터링하세요. 먼저 `filter`를 사용해 `online` 속성이 `true`인 사용자만 포함하는 새 배열을 반환하세요. 그런 다음 `renderOnline` 변수에서 필터링된 배열을 map으로 순회하며 각 사용자에 대해 `li` 텍스트를 포함하는 `username` 요소를 반환하세요. 마지막 도전 과제처럼 고유한 `key`도 반드시 포함하세요.

# --hints--

`MyComponent`가 존재하고 페이지에 렌더링되어야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('MyComponent').length,
  1
);
```

`MyComponent`의 상태는 6명의 사용자 배열로 초기화되어야 합니다.

```js
assert(
  Array.isArray(
    Enzyme.mount(React.createElement(MyComponent)).state('users')
  ) === true &&
    Enzyme.mount(React.createElement(MyComponent)).state('users').length === 6
);
```

`MyComponent`은 `div`, `h1`, 그리고 온라인 상태가 `li`인 모든 사용자에 대해 `true` 요소를 포함하는 순서 없는 목록을 반환해야 합니다.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const result = () => comp.find('li').length;
  const _1 = result();
  const _2 = () => {
    comp.setState(users(true));
    return result();
  };
  const _3 = () => {
    comp.setState(users(false));
    return result();
  };
  const _4 = () => {
    comp.setState({ users: [] });
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  assert(
    comp.find('div').length === 1 &&
      comp.find('h1').length === 1 &&
      comp.find('ul').length === 1 &&
      _1 === 4 &&
      _2_val === 5 &&
      _3_val === 0 &&
      _4_val === 0
  );
})();
```

`MyComponent`는 각 온라인 사용자의 `li`를 포함하는 `username` 요소를 렌더링해야 합니다.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const ul = () => {
    comp.setState(users(true));
    return comp.find('ul').html();
  };
  const html = ul();
  assert(
    html ===
      '<ul><li>Jeff</li><li>Alan</li><li>Mary</li><li>Jim</li><li>Laura</li></ul>'
  );
})();
```

각 목록 항목 요소는 고유한 `key` 속성을 가져야 합니다.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(MyComponent)).find('ul');
    console.log(ul.debug());
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key()
    ]);
    return keys.size === 4;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = null; // Change this line
    const renderOnline = null; // Change this line
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = this.state.users.filter(user => {
      return user.online;
    });
    const renderOnline = usersOnline.map(user => {
      return <li key={user.username}>{user.username}</li>;
    });
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```
