---
id: 5a24c314108439a4d4036160
title: JSX에서 HTML 클래스 정의하기
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

JSX 작성에 익숙해지면서 HTML과 어떻게 다른지 궁금할 수 있습니다.

지금까지는 HTML과 JSX가 똑같아 보일 수 있습니다.

JSX에서 중요한 차이점 중 하나는 HTML 클래스를 정의할 때 더 이상 `class`를 사용할 수 없다는 점입니다. 이는 `class`가 JavaScript에서 예약어이기 때문입니다. 대신 JSX에서는 `className`를 사용합니다.

사실, JSX에서 모든 HTML 속성과 이벤트 참조의 명명 관례가 카멜 케이스로 바뀝니다. 예를 들어, JSX에서 클릭 이벤트는 `onClick`이며, HTML의 `onclick`와 다릅니다. 마찬가지로 `onchange`는 `onChange`가 됩니다. 이는 미묘한 차이지만 앞으로 기억해야 할 중요한 점입니다.

# --instructions--

JSX 코드에 제공된 `myDiv`에 `div` 클래스를 적용하세요.

# --hints--

상수 `JSX`는 `div` 요소를 반환해야 합니다.

```js
assert.strictEqual(JSX.type, 'div');
```

`div`는 `myDiv` 클래스를 가져야 합니다.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
