---
id: 5a24c314108439a4d403616d
title: 예상하는 프로프스를 정의하기 위해 PropTypes 사용하기
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React는 컴포넌트가 올바른 타입의 프로프스를 받는지 확인하는 유용한 타입 검사 기능을 제공합니다. 예를 들어, 애플리케이션이 배열 형태일 것으로 예상되는 데이터를 가져오는 API 호출을 하고, 이 데이터를 프로프스로 컴포넌트에 전달한다고 가정해 보세요. 컴포넌트에 `propTypes`를 설정하여 데이터가 `array` 타입이어야 한다고 요구할 수 있습니다. 데이터가 다른 타입일 경우 유용한 경고가 발생합니다.

프로프스 타입을 미리 알고 있다면 `propTypes`를 설정하는 것이 모범 사례로 간주됩니다. 컴포넌트에 `propTypes` 속성을 `defaultProps`를 정의한 것과 같은 방식으로 정의할 수 있습니다. 이렇게 하면 특정 키의 프로프스가 주어진 타입으로 존재하는지 검사합니다. 다음은 `function`라는 프로프스에 대해 `handleClick` 타입을 요구하는 예시입니다:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

위 예시에서 `PropTypes.func` 부분은 `handleClick`가 함수인지 확인합니다. `isRequired`를 추가하면 React에 `handleClick`가 해당 컴포넌트에 필수 속성임을 알립니다. 해당 프로프스가 제공되지 않으면 경고가 표시됩니다. 또한 `func`는 `function`를 나타냅니다. 일곱 가지 JavaScript 원시 타입 중 `function`와 `boolean` (`bool`로 표기됨)만 특이한 철자를 사용합니다. 원시 타입 외에도 다른 타입을 검사할 수 있습니다. 예를 들어, 프로프스가 React 엘리먼트인지 확인할 수 있습니다. 모든 옵션은 <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">문서</a>를 참조하세요.

**참고:** React v15.5.0부터 `PropTypes`는 React와 별도로 다음과 같이 가져옵니다: `import PropTypes from 'prop-types';`

# --instructions--

`propTypes` 컴포넌트에 `Items`를 정의하여 `quantity`를 프로프스로 요구하고, 그 타입이 `number`인지 확인하세요.

# --hints--

`ShoppingCart` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`Items` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

`Items` 컴포넌트에 `propTypes` 검사를 포함하여 `quantity`에 값이 필요하며 그 값이 숫자인지 확인하세요.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --before-all--

```js
var PropTypes = {
  number: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
