---
id: 5a24c314108439a4d403617a
title: 자식 컴포넌트에 상태를 프로프스로 전달하기
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

이전 도전 과제에서 자식 JSX 요소와 자식 React 컴포넌트에 프로프스를 전달하는 많은 예제를 보셨을 겁니다. 그 프로프스가 어디서 오는지 궁금할 수 있습니다. 일반적인 패턴은 애플리케이션에 중요한 `state`를 포함하는 상태 유지 컴포넌트를 두고, 그 컴포넌트가 자식 컴포넌트를 렌더링하는 것입니다. 이 컴포넌트들이 `state`의 일부에 접근할 수 있게 하려면, 그 일부를 프로프스로 전달합니다.

예를 들어, `App` 컴포넌트가 `Navbar`를 포함한 여러 컴포넌트를 렌더링한다고 가정해 보세요. `App`에는 많은 사용자 정보를 담고 있는 `state`가 있지만, `Navbar`는 사용자 이름만 필요해서 그것을 표시할 수 있어야 합니다. 이 `state`의 일부를 `Navbar` 컴포넌트에 프로프로 전달합니다.

이 패턴은 React의 중요한 패러다임 몇 가지를 보여줍니다. 첫 번째는 *단방향 데이터 흐름*입니다. 상태는 애플리케이션 컴포넌트 트리에서 상태 유지 부모 컴포넌트에서 자식 컴포넌트로 한 방향으로 흐릅니다. 자식 컴포넌트는 필요한 상태 데이터만 받습니다. 두 번째는 복잡한 상태 유지 애플리케이션이 몇 개 또는 하나의 상태 유지 컴포넌트로 나눌 수 있다는 점입니다. 나머지 컴포넌트는 부모로부터 상태를 프로프로 받아 그 상태로 UI를 렌더링합니다. 이렇게 하면 상태 관리는 코드의 한 부분에서, UI 렌더링은 다른 부분에서 처리하는 분리가 시작됩니다. 상태 로직과 UI 로직을 분리하는 이 원칙은 React의 핵심 원칙 중 하나입니다. 올바르게 사용하면 복잡한 상태 유지 애플리케이션 설계를 훨씬 쉽게 관리할 수 있습니다.

# --instructions--

`MyApp` 컴포넌트는 상태 유지 컴포넌트이며 자식으로 `Navbar` 컴포넌트를 렌더링합니다. `name` 속성을 `state`에서 자식 컴포넌트로 전달한 다음, `name`의 렌더 메서드 일부인 `h1` 태그 안에 `Navbar`을 표시하세요. `name`는 `Hello, my name is:` 텍스트 뒤에 나타나야 합니다.

# --hints--

`MyApp` 컴포넌트는 내부에 `Navbar` 컴포넌트를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

`Navbar` 컴포넌트는 `MyApp` 상태 속성 `name`를 프로프로 받아야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

`h1` 요소는 `Navbar`에서 `name` 프로프를 렌더링해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
