---
id: 5a24c314108439a4d4036165
title: 중첩된 컴포넌트를 렌더링하기 위해 React 사용하기
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

마지막 도전 과제에서는 두 컴포넌트를 간단히 조합하는 방법을 보여주었지만, React로 컴포넌트를 조합하는 방법은 매우 다양합니다.

컴포넌트 조합은 React의 강력한 기능 중 하나입니다. React를 사용할 때는 마지막 도전 과제의 App 예제처럼 사용자 인터페이스를 컴포넌트 단위로 생각하기 시작하는 것이 중요합니다. UI를 기본 빌딩 블록으로 분해하고, 그 조각들이 컴포넌트가 됩니다. 이렇게 하면 UI를 담당하는 코드와 애플리케이션 로직을 처리하는 코드를 분리할 수 있습니다. 복잡한 프로젝트의 개발과 유지 관리를 크게 단순화할 수 있습니다.

# --instructions--

코드 편집기에 `TypesOfFruit`와 `Fruits`이라는 두 개의 함수형 컴포넌트가 정의되어 있습니다. `TypesOfFruit` 컴포넌트를 가져와 `Fruits` 컴포넌트 안에 조합하거나 *중첩*하세요. 그런 다음 `Fruits` 컴포넌트를 `TypesOfFood` 컴포넌트 안에 중첩하세요. 결과는 자식 컴포넌트가 부모 컴포넌트 안에 중첩되고, 그 부모 컴포넌트가 또 다른 부모 컴포넌트 안에 중첩된 형태가 되어야 합니다!

# --hints--

`TypesOfFood` 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

`TypesOfFood` 컴포넌트는 `Fruits` 컴포넌트를 반환해야 합니다.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

`Fruits` 컴포넌트는 `TypesOfFruit` 컴포넌트를 반환해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

`TypesOfFruit` 컴포넌트는 `h2`와 `ul` 요소를 반환해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
