---
id: 5a24c314108439a4d4036163
title: React 컴포넌트 만들기
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

React 컴포넌트를 정의하는 또 다른 방법은 ES6 `class` 구문을 사용하는 것입니다. 다음 예제에서 `Kitten`은 `React.Component`를 확장합니다:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

이 코드는 `Kitten` 클래스를 확장하는 ES6 클래스 `React.Component`를 만듭니다. 따라서 `Kitten` 클래스는 로컬 상태와 생명주기 훅 같은 유용한 React 기능에 접근할 수 있습니다. 아직 이 용어들이 익숙하지 않아도 걱정하지 마세요. 이후 도전 과제에서 더 자세히 다룰 예정입니다. 또한 `Kitten` 클래스 내에 `constructor`가 정의되어 있고 `super()`를 호출하는 것을 확인하세요. 부모 클래스인 `super()`의 생성자를 호출하기 위해 `React.Component`를 사용합니다. 생성자는 `class` 키워드로 생성된 객체를 초기화할 때 사용하는 특별한 메서드입니다. 컴포넌트의 `constructor`를 `super`와 함께 호출하고 `props`를 둘 다에 전달하는 것이 모범 사례입니다. 이렇게 하면 컴포넌트가 올바르게 초기화됩니다. 지금은 이 코드가 포함되는 것이 표준임을 알아두세요. 곧 생성자와 `props`의 다른 용도도 보게 될 것입니다.

# --instructions--

코드 편집기에서 `MyComponent`가 클래스 구문으로 정의되어 있습니다. `render` 메서드를 완성하여 `div` 요소를 반환하게 하세요. 이 요소는 `h1`를 포함하며, 그 안에 `Hello React!`라는 텍스트가 있어야 합니다.

# --hints--

React 컴포넌트는 `div` 요소를 반환해야 합니다.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

반환된 `div`는 그 안에 `h1` 제목 요소를 렌더링해야 합니다.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

`h1` 제목 요소는 `Hello React!` 문자열을 포함해야 합니다.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
