---
id: 5a24c314108439a4d403617b
title: 콜백을 프로프스로 전달하기
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

`state`를 자식 컴포넌트에 프로프스로 전달할 수 있지만, 데이터 전달에만 제한되지 않습니다. 핸들러 함수나 React 컴포넌트에 정의된 메서드도 자식 컴포넌트에 전달할 수 있습니다. 이렇게 하면 자식 컴포넌트가 부모 컴포넌트와 상호작용할 수 있습니다. 메서드를 자식에게 일반 프로프처럼 전달하세요. 이름을 할당하고 자식 컴포넌트에서 `this.props` 아래에 그 메서드 이름에 접근할 수 있습니다.

# --instructions--

코드 편집기에 세 개의 컴포넌트가 나와 있습니다. `MyApp` 컴포넌트가 부모이며 `GetInput`과 `RenderInput` 자식 컴포넌트를 렌더링합니다. `GetInput`의 렌더 메서드에 `MyApp` 컴포넌트를 추가한 다음, `input`의 `inputValue`에서 `MyApp`에 할당된 `state`이라는 프로프를 전달하세요. 또한 `handleChange`이라는 프로프를 만들고 입력 핸들러 `handleChange`를 전달하세요.

다음으로 `RenderInput`의 렌더 메서드에 `MyApp`을 추가한 다음, `input`에서 `inputValue`을 전달하는 `state`라는 프로프를 만드세요. 완료하면 `input` 컴포넌트의 `GetInput` 필드에 입력할 수 있으며, 이 필드는 프로프스를 통해 부모의 핸들러 메서드를 호출합니다. 그러면 부모의 `state`에 입력이 업데이트되고, 이 값이 두 자식에게 프로프스로 전달됩니다. 컴포넌트 간 데이터 흐름과 단일 진실의 출처가 부모 컴포넌트의 `state`임을 관찰하세요. 이 예제는 다소 인위적이지만, React 컴포넌트 간에 데이터와 콜백을 전달하는 방식을 보여줍니다.

# --hints--

`MyApp` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

`GetInput` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

`RenderInput` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

`GetInput` 컴포넌트는 `MyApp` 상태 속성 `inputValue`을 프로프스로 받고, `input` 상태를 수정하는 `MyApp` 요소를 포함해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

`RenderInput` 컴포넌트는 `MyApp` 상태 속성 `inputValue`을 프로프스로 받아야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
