---
id: 5a24c314108439a4d403616a
title: 배열을 프로프스로 전달하기
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

이전 도전 과제에서는 부모 컴포넌트에서 자식 컴포넌트로 `props` 또는 속성으로 정보를 전달하는 방법을 보여주었습니다. 이번 도전 과제에서는 배열을 `props`로 전달하는 방법을 살펴봅니다. 배열을 JSX 요소에 전달하려면 JavaScript로 취급하여 중괄호로 감싸야 합니다.

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

자식 컴포넌트는 배열 속성 `colors`에 접근할 수 있습니다. `join()` 같은 배열 메서드를 속성에 접근할 때 사용할 수 있습니다.

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
```

이 메서드는 모든 `colors` 배열 항목을 쉼표로 구분된 문자열로 합쳐서 `<p>green, blue, red</p>`를 생성합니다. 이후에 React에서 데이터를 배열로 렌더링하는 다른 일반적인 메서드도 배우게 됩니다.

# --instructions--

코드 편집기에는 `List`와 `ToDo` 컴포넌트가 있습니다. `List` 컴포넌트에서 각 `ToDo`를 렌더링할 때, 할 일 배열에 할당된 `tasks` 속성을 전달하세요. 예를 들어 `["walk dog", "workout"]`처럼요. 그런 다음 `tasks` 컴포넌트에서 이 `List` 배열에 접근하여 `p` 요소 안에 값을 표시하세요. `join(", ")`를 사용해 `props.tasks` 배열을 `p` 요소 안에 쉼표로 구분된 목록으로 표시하세요. 오늘 목록에는 최소 2개의 작업이 있어야 하고 내일 목록에는 최소 3개의 작업이 있어야 합니다.

# --hints--

`ToDo` 컴포넌트는 하나의 외부 `div`를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

`ToDo` 컴포넌트의 세 번째 자식은 `List` 컴포넌트의 인스턴스여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

`ToDo` 컴포넌트의 다섯 번째 자식은 `List` 컴포넌트의 인스턴스여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

`List` 컴포넌트 두 인스턴스 모두 `tasks`라는 속성을 가져야 하며 `tasks`는 배열 타입이어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

오늘 작업을 나타내는 첫 번째 `List` 컴포넌트에는 2개 이상의 항목이 있어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

내일 작업을 나타내는 두 번째 `List` 컴포넌트에는 3개 이상의 항목이 있어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

`List` 컴포넌트는 `tasks` 프로프의 값을 `p` 태그 안에 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
