---
id: 5a24c314108439a4d4036182
title: React에서 인라인 스타일 추가하기
challengeType: 6
forumTopicId: 301378
dashedName: add-inline-styles-in-react
---

# --description--

지난 도전 과제에서 `style` 속성이 JavaScript 객체로 설정된 것 외에도 HTML 인라인 스타일과 다른 여러 구문 차이가 있었던 것을 눈치채셨을 겁니다. 우선, 특정 CSS 스타일 속성 이름은 카멜 케이스를 사용합니다. 예를 들어, 지난 도전 과제에서는 `fontSize` 대신 `font-size`로 폰트 크기를 설정했습니다. `font-size`처럼 하이픈으로 연결된 단어는 JavaScript 객체 속성에 유효하지 않은 구문이므로 React는 카멜 케이스를 사용합니다. 일반적으로 하이픈으로 연결된 스타일 속성은 JSX에서 카멜 케이스로 작성합니다.

모든 속성 값 길이 단위(예: `height`, `width`, `fontSize`)는 별도로 명시하지 않는 한 `px` 단위로 간주됩니다. 예를 들어 `em`을 사용하려면 값과 단위를 따옴표로 감싸야 하며, `{fontSize: "4em"}`처럼 작성합니다. `px`이 기본값인 길이 값을 제외한 모든 속성 값은 따옴표로 감싸야 합니다.

# --instructions--

스타일이 많을 경우, 스타일 `object`을 상수에 할당하여 코드를 정리할 수 있습니다. 파일 상단에 전역 변수로 스타일 상수를 선언하세요. `styles` 상수를 초기화하고 세 가지 스타일 속성과 값을 가진 `object`을 할당하세요. `div`에 `purple` 색상, `40` 폰트 크기, `2px solid purple` 테두리를 지정하세요. 그런 다음 `style` 속성을 `styles` 상수와 같게 설정하세요.

# --hints--

`styles` 변수는 세 개의 속성을 가진 `object`이어야 합니다.

```js
assert(Object.keys(styles).length === 3);
```

`styles` 변수는 `color` 속성을 `purple` 값으로 설정해야 합니다.

```js
assert(styles.color === 'purple');
```

`styles` 변수는 `fontSize` 속성을 `40` 값으로 설정해야 합니다.

```js
assert(styles.fontSize == 40);
```

`styles` 변수는 `border` 속성을 `2px solid purple` 값으로 설정해야 합니다.

```js
assert(styles.border === '2px solid purple');
```

컴포넌트는 `div` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return mockedComponent.type() === 'div';
  })()
);
```

`div` 요소는 `styles` 객체로 스타일이 정의되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return (
      mockedComponent.props().style.color === 'purple' &&
      mockedComponent.props().style.fontSize == 40 &&
      mockedComponent.props().style.border === '2px solid purple'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={{color: "yellow", fontSize: 24}}>Style Me!</div>
    );
    // Change code above this line
  }
};
```

# --solutions--

```jsx
const styles = {
  color: "purple",
  fontSize: 40,
  border: "2px solid purple"
};
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={styles}>Style Me!</div>
    );
    // Change code above this line
  }
};
```
