---
id: 5a24c314108439a4d4036189
title: 컴포넌트 상태에 따라 인라인 CSS를 조건부로 변경하기
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

지금까지 조건부 렌더링과 인라인 스타일 사용법을 여러 번 보았습니다. 여기 두 주제를 결합한 또 다른 예제가 있습니다. React 컴포넌트의 상태에 따라 CSS를 조건부로 렌더링할 수도 있습니다. 이를 위해 조건을 확인하고, 그 조건이 충족되면 렌더 메서드에서 JSX 요소에 할당된 스타일 객체를 수정합니다.

이 방식은 DOM 요소를 직접 수정하여 스타일을 적용하는 전통적인 방법(예: jQuery에서 흔히 사용됨)과는 크게 다르기 때문에 이해하는 것이 중요합니다. 전통적인 방법에서는 요소가 변경될 때를 추적하고 실제 조작을 직접 처리해야 합니다. 변경 사항을 추적하기 어려워 UI가 예측 불가능해질 수 있습니다. 조건에 따라 스타일 객체를 설정하면 애플리케이션 상태에 따른 UI 모양을 함수로 설명하는 셈입니다. 정보 흐름이 한 방향으로만 명확하게 이동합니다. React로 애플리케이션을 작성할 때 선호되는 방법입니다.

# --instructions--

코드 편집기에는 스타일이 적용된 테두리가 있는 간단한 제어된 입력 컴포넌트가 있습니다. 사용자가 입력 상자에 15자 이상의 텍스트를 입력하면 이 테두리를 빨간색으로 스타일링하려고 합니다. 이 조건을 확인하는 코드를 추가하고, 조건이 참이면 입력 테두리 스타일을 `3px solid red`로 설정하세요. 입력에 텍스트를 입력해 테스트할 수 있습니다.

# --hints--

`GateKeeper` 컴포넌트는 `div` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`GateKeeper` 컴포넌트는 빈 문자열로 설정된 상태 키 `input`로 초기화해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

`GateKeeper` 컴포넌트는 `h3` 태그와 `input` 태그를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

`input` 태그는 처음에 `1px solid black` 속성에 대해 `border` 스타일을 가져야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

`input` 태그는 상태의 입력 값이 15자보다 길면 `3px solid red` 테두리로 스타일링해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
