---
id: 5a24c314108439a4d4036167
title: DOM에 클래스 컴포넌트 렌더링하기
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

You may remember using the ReactDOM API in an earlier challenge to render JSX elements to the DOM. The process for rendering React components will look very similar. The past few challenges focused on components and composition, so the rendering was done for you behind the scenes. However, none of the React code you write will render to the DOM without making a call to the ReactDOM API.

호출문 코드: `ReactDOM.render(componentToRender, targetNode)`. 첫 번째 인자는 렌더링하고 싶은 리액트 컴포넌트입니다. 두 번째 인자는 해당 컴포넌트를 렌더링하는 DOM 노드 입니다.

리액트 컴포넌트는 JSX 엘리먼트와는 조금 다르게 `ReactDOM.render()`에 전달됩니다. JSX 엘리먼트는, 렌더링할 엘리먼트의 이름을 전달합니다. 반면에 리액트 컴포넌트는 다른 컴포넌트안에 넣는 경우와 같은 문법을 사용합니다. 예를 들면 `ReactDOM.render(<ComponentToRender />, targetNode)`. 이 문법은 ES6 클래스 컴포넌트와 함수형 컴포넌트 모두에 사용합니다.

# --instructions--

`Fruits`와 `Vegetables` 컴포넌트는 안보이는 곳에 정의되어 있습니다. 두 컴포넌트를 `TypesOfFood` 컴포넌트 안에 렌더링하고, `TypesOfFood`를 DOM에 렌더링 해보세요. `id='challenge-node'`를 가진 `div`를 사용해보세요.

# --hints--

`TypesOfFood` 컴포넌트는 하나의 `div` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood` 컴포넌트는 `h1` 엘리먼트 뒤에 있는 `Fruits` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`TypesOfFood` 컴포넌트는 `Fruits` 컴포넌트 뒤에 있는 `Vegetables` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

`TypesOfFood` 컴포넌트는 id가 `challenge-node`인 `div`안의 DOM에 렌더링해야 합니다.

```js
assert(
  (function () {
    const html = document.getElementById('challenge-node').childNodes[0]
      .innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
const Fruits = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <h4>Non-Citrus:</h4>
        <ul>
          <li>Apples</li>
          <li>Blueberries</li>
          <li>Strawberries</li>
          <li>Bananas</li>
        </ul>
      <h4>Citrus:</h4>
        <ul>
          <li>Lemon</li>
          <li>Lime</li>
          <li>Orange</li>
          <li>Grapefruit</li>
        </ul>
    </div>
  );
};
const Vegetables = () => {
  return (
    <div>
      <h2>Vegetables:</h2>
      <ul>
        <li>Brussel Sprouts</li>
        <li>Broccoli</li>
        <li>Squash</li>
      </ul>
    </div>
  );
};
```

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
