---
id: 5a24c314108439a4d403616f
title: 무상태 함수형 컴포넌트에서 프로프스 사용 검토
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

마지막 도전 과제를 제외하고, 무상태 함수형 컴포넌트에 프로프스를 전달해 왔습니다. 이 컴포넌트들은 순수한 함수처럼 작동합니다. 프로프스를 입력으로 받아 같은 프로프스를 전달받으면 항상 같은 뷰를 반환합니다. 상태가 무엇인지 궁금할 수 있는데, 다음 도전 과제에서 더 자세히 다룹니다. 그 전에 컴포넌트 용어를 검토해 보겠습니다.

*무상태 함수형 컴포넌트*는 프로프스를 받아 JSX를 반환하는 함수입니다. 반면에 *무상태 컴포넌트*는 `React.Component`를 확장하는 클래스이지만 내부 상태를 사용하지 않는 클래스입니다(다음 도전 과제에서 다룹니다). 마지막으로, *상태 유지 컴포넌트*는 자체 내부 상태를 유지하는 클래스 컴포넌트입니다. 상태 유지 컴포넌트는 단순히 컴포넌트 또는 React 컴포넌트라고도 부릅니다.

일반적인 패턴은 상태 유지를 최소화하고 가능한 한 무상태 함수형 컴포넌트를 만드는 것입니다. 이렇게 하면 애플리케이션의 특정 영역에 상태 관리를 집중할 수 있습니다. 결과적으로 상태 변경이 동작에 미치는 영향을 더 쉽게 추적할 수 있어 애플리케이션 개발과 유지 관리가 개선됩니다.

# --instructions--

코드 편집기에는 자식으로 `CampSite` 컴포넌트를 렌더링하는 `Camper` 컴포넌트가 있습니다. `Camper` 컴포넌트를 정의하고 기본 프로프스를 `{ name: 'CamperBot' }`로 할당하세요. `Camper` 컴포넌트 내부에서는 원하는 코드를 렌더링하되, `p`로 전달된 `name` 값만 포함하는 `prop` 요소를 반드시 하나 포함하세요. 마지막으로, `propTypes` 컴포넌트에 `Camper`를 정의하여 `name` 프로프스가 반드시 제공되고 `string` 타입인지 확인하세요.

# --hints--

`CampSite` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

`Camper` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

`Camper` 컴포넌트는 `CamperBot` 키에 문자열 `name`를 할당하는 기본 프로프스를 포함해야 합니다.

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

`Camper` 컴포넌트는 `name` 프로프스가 `string` 타입임을 요구하는 프로프 타입을 포함해야 합니다.

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

`Camper` 컴포넌트는 `p` 프로프스의 텍스트만 포함하는 `name` 요소를 포함해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --before-all--

```js
var PropTypes = {
   string: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
