---
id: 5a24c314108439a4d4036172
title: 사용자 인터페이스에서 상태 렌더링하는 또 다른 방법
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

컴포넌트에서 `state`에 접근하는 또 다른 방법이 있습니다. `render()` 메서드에서 `return` 문 이전에 JavaScript를 직접 작성할 수 있습니다. 예를 들어, 함수를 선언하거나 `state` 또는 `props`에서 데이터를 접근하고, 이 데이터에 대해 계산을 수행하는 등의 작업이 가능합니다. 그런 다음, `return` 문에서 접근할 수 있는 변수에 데이터를 할당할 수 있습니다.

# --instructions--

`MyComponent` render 메서드에서 `const`를 `name`라는 이름으로 정의하고, 컴포넌트의 `state`에 있는 name 값을 할당하세요. 이 코드 부분에서는 JavaScript를 직접 작성할 수 있으므로 이 참조를 중괄호로 감쌀 필요가 없습니다.

다음으로, return 문에서 이 값을 `h1` 태그 안에 `name` 변수를 사용해 렌더링하세요. return 문에서는 JSX 구문(자바스크립트용 중괄호)을 사용해야 한다는 점을 기억하세요.

# --hints--

`MyComponent`는 상태에 `name` 키와 `freeCodeCamp` 값을 저장해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent`는 단일 `h1`로 감싼 `div` 제목 요소를 렌더링해야 합니다.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

렌더링된 `h1` 태그는 `{name}`에 대한 참조를 포함해야 합니다.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

렌더링된 `h1` 제목 요소는 컴포넌트 상태에서 렌더링된 텍스트를 포함해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
