---
id: 5a24c314108439a4d4036185
title: 더 간결한 조건문을 위한 && 사용
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

`if/else` 문장은 지난 도전 과제에서 잘 작동했지만, 같은 결과를 더 간결하게 얻는 방법이 있습니다. 컴포넌트에서 여러 조건을 추적하고 각 조건에 따라 다른 요소를 렌더링하고 싶다고 상상해 보세요. 약간씩 다른 UI를 반환하기 위해 많은 `else if` 문장을 작성하면 코드가 반복되어 오류가 발생할 여지가 생깁니다. 대신 `&&` 논리 연산자를 사용해 조건문을 더 간결하게 처리할 수 있습니다. 이는 조건이 `true`인지 확인하고, 그렇다면 마크업을 반환하려는 경우에 가능합니다. 예를 들어 다음과 같습니다:

```jsx
{condition && <p>markup</p>}
```

`condition`가 `true`라면 마크업이 반환됩니다. 조건이 `false`라면, `false`를 평가한 후 즉시 `condition`를 반환하며 아무것도 반환하지 않습니다. 이 문장들을 JSX에 직접 포함하고 각 조건 뒤에 `&&`를 작성해 여러 조건을 연결할 수 있습니다. 이렇게 하면 많은 코드를 반복하지 않고 `render()` 메서드에서 더 복잡한 조건 논리를 처리할 수 있습니다.

# --instructions--

이전 예제를 다시 풀되, `h1`이 `display`가 `true`일 때만 렌더링되도록 하세요. 단, `&&` 문장 대신 `if/else` 논리 연산자를 사용하세요.

# --hints--

`MyComponent`가 존재하고 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

`display`가 `true`로 설정되면 `div`, `button`, `h1`가 렌더링되어야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

`display`가 `false`로 설정되면 `div`와 `button`만 렌더링되어야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

render 메서드는 `&&` 논리 연산자를 사용해 `this.state.display` 조건을 확인해야 합니다.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
