---
id: 5a24c314108439a4d403616e
title: this.props를 사용해 프로프스 접근하기
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

최근 도전 과제들은 자식 컴포넌트에 프로프스를 전달하는 기본 방법을 다뤘습니다. 그런데 전달하려는 자식 컴포넌트가 무상태 함수형 컴포넌트가 아니라 ES6 클래스 컴포넌트라면 어떻게 할까요? ES6 클래스 컴포넌트는 프로프스에 접근하는 관례가 약간 다릅니다.

클래스 컴포넌트 내부에서 자기 자신을 참조할 때는 `this` 키워드를 사용합니다. 클래스 컴포넌트 내에서 프로프스에 접근하려면, 접근하는 코드를 `this`로 시작해야 합니다. 예를 들어, ES6 클래스 컴포넌트에 `data`라는 프로프가 있다면 JSX에서는 `{this.props.data}`라고 작성합니다.

# --instructions--

부모 컴포넌트 `Welcome` 안에 `App` 컴포넌트 인스턴스를 렌더링하세요. 여기서 `Welcome`에 `name` 프로프를 주고 문자열 값을 할당하세요. 자식 컴포넌트 `Welcome` 안에서는 `name` 태그 내에서 `strong` 프로프에 접근하세요.

# --hints--

`App` 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`App`의 자식은 `Welcome` 컴포넌트여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

`Welcome` 컴포넌트는 `name`라는 프로프를 가져야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

`Welcome` 컴포넌트는 `name` 태그 내에서 `strong` 프로프로 전달한 문자열을 표시해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
