---
id: 5a24c314108439a4d4036181
title: 인라인 스타일 소개
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

React 코드에 강력한 기능을 더하는 다른 복잡한 개념들도 있습니다. 하지만 React에서 생성한 JSX 요소에 스타일을 적용하는 더 간단한 문제에 대해 궁금할 수 있습니다. JSX 요소에 클래스를 적용하는 방식이 <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">HTML과 정확히 같지 않다는 점</a>을 아실 겁니다.

스타일시트에서 스타일을 가져오면 크게 다르지 않습니다. JSX 요소에 `className` 속성으로 클래스를 적용하고, 스타일시트에서 해당 클래스에 스타일을 적용합니다. 또 다른 방법은 ReactJS 개발에서 매우 흔한 인라인 스타일을 적용하는 것입니다.

인라인 스타일은 HTML에서 적용하는 방식과 비슷하지만 JSX에서는 몇 가지 차이가 있습니다. HTML에서 인라인 스타일 예시는 다음과 같습니다.

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

JSX 요소는 `style` 속성을 사용하지만, JSX가 변환되는 방식 때문에 값을 `string`로 설정할 수 없습니다. 대신 JavaScript `object`에 할당합니다. 예시는 다음과 같습니다.

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

`fontSize` 속성을 카멜 케이스로 작성하는 점을 주목하세요. React는 스타일 객체에서 케밥 케이스 키를 허용하지 않기 때문입니다. React가 HTML에서 올바른 속성 이름을 적용해 줍니다.

# --instructions--

코드 편집기에서 `style`에 `div` 속성을 추가하여 텍스트 색상을 빨간색으로, 폰트 크기를 `72px`로 설정하세요.

폰트 크기는 선택적으로 숫자로 설정해 단위 `px`를 생략하거나 `"72px"`로 작성할 수 있습니다.
# --hints--

컴포넌트는 `div` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`div` 요소는 색상이 `red`여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

`div` 요소는 폰트 크기가 `72px`여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
