---
id: 5a24c314108439a4d403616c
title: 기본 프로프 재정의하기
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

기본 프로프를 설정하는 기능은 React에서 유용한 기능입니다. 기본 프로프를 재정의하는 방법은 컴포넌트에 대해 프로프 값을 명시적으로 설정하는 것입니다.

# --instructions--

`ShoppingCart` 컴포넌트는 이제 자식 컴포넌트 `Items`를 렌더링합니다. 이 `Items` 컴포넌트는 정수 `quantity`로 설정된 기본 프로프 `0`를 가지고 있습니다. `10`에 대해 `quantity` 값을 전달하여 기본 프로프를 재정의하세요.

**참고:** 컴포넌트에 프로프를 추가하는 구문은 HTML 속성을 추가하는 방식과 비슷하다는 점을 기억하세요. 하지만 `quantity` 값이 정수이기 때문에 따옴표 안에 들어가지 않고 중괄호로 감싸야 합니다. 예를 들어 `{100}`입니다. 이 구문은 JSX가 중괄호 안의 값을 자바스크립트로 직접 해석하도록 지시합니다.

# --hints--

`ShoppingCart` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`Items` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

`Items` 컴포넌트는 `{ quantity: 10 }` 컴포넌트에서 `ShoppingCart` 프로프를 전달받아야 합니다.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
