---
id: 5a24c314108439a4d403617e
title: 이벤트 리스너 추가
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

`componentDidMount()` 메서드는 특정 기능을 위해 추가해야 하는 이벤트 리스너를 붙이기에 가장 좋은 장소이기도 합니다. React는 브라우저에 존재하는 네이티브 이벤트 시스템을 감싸는 합성 이벤트 시스템을 제공합니다. 이는 네이티브 이벤트가 브라우저마다 다르게 동작할 수 있어도, 합성 이벤트 시스템은 사용자의 브라우저와 상관없이 동일하게 동작한다는 뜻입니다.

이미 `onClick()` 같은 합성 이벤트 핸들러를 사용해 보셨을 겁니다. React의 합성 이벤트 시스템은 DOM 요소에서 관리할 대부분의 상호작용에 매우 적합합니다. 하지만 문서(document)나 창(window) 객체에 이벤트 핸들러를 붙이고 싶다면 직접 붙여야 합니다.

# --instructions--

`componentDidMount()` 메서드에서 `keydown` 이벤트에 대한 이벤트 리스너를 붙이고, 이 이벤트가 `handleKeyPress()` 콜백을 실행하도록 하세요. 첫 번째 인수로 이벤트 이름(따옴표 포함), 두 번째 인수로 콜백을 받는 `document.addEventListener()`를 사용할 수 있습니다.

그다음 `componentWillUnmount()`에서 같은 이벤트 리스너를 제거하세요. `document.removeEventListener()`에 같은 인수를 전달하면 됩니다. React 컴포넌트가 언마운트되어 파괴되기 전에 이 생명주기 메서드에서 정리 작업을 하는 것이 좋은 습관입니다. 이벤트 리스너 제거가 그런 정리 작업의 한 예입니다.

# --hints--

`MyComponent`는 `div` 요소를 렌더링해야 하며, 이 요소는 `h1` 태그를 감싸야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

`keydown` 리스너는 `componentDidMount`에서 문서에 붙여야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

`keydown` 리스너는 `componentWillUnmount`에서 문서에서 제거해야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

컴포넌트가 마운트된 후 `enter`를 누르면 상태가 업데이트되고 렌더링된 `h1` 태그가 변경됩니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
