---
id: 5a24c314108439a4d4036176
title: 요소를 토글하기 위해 상태 사용하기
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

상태를 업데이트할 때 이전 상태를 알아야 할 때가 있습니다. 하지만 상태 업데이트는 비동기일 수 있습니다 — 즉, React가 여러 `setState()` 호출을 하나의 업데이트로 묶을 수 있다는 뜻입니다. 따라서 다음 값을 계산할 때 `this.state` 또는 `this.props`의 이전 값에 의존할 수 없습니다. 그래서 다음과 같은 코드를 사용하면 안 됩니다:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

대신, `setState`에 상태와 프로프스에 접근할 수 있는 함수를 전달해야 합니다. `setState`에 함수를 사용하면 상태와 프로프스의 가장 최신 값을 다루고 있다는 보장이 있습니다. 따라서 위 코드는 다음과 같이 다시 작성해야 합니다:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

`props` 없이 `state`만 필요할 때는 다음과 같은 형태도 사용할 수 있습니다:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

객체 리터럴을 소괄호로 감싸야 한다는 점을 주의하세요. 그렇지 않으면 JavaScript가 이를 코드 블록으로 인식합니다.

# --instructions--

`MyComponent`에는 `visibility` 속성이 있으며, 이 속성은 `false`로 초기화되어 있습니다. 렌더링 메서드는 `visibility` 값이 true일 때는 한 뷰를, false일 때는 다른 뷰를 반환합니다.

현재 컴포넌트의 `visibility`에 있는 `state` 속성을 업데이트하는 방법이 없습니다. 이 값은 true와 false 사이를 토글해야 합니다. 버튼에 클릭 핸들러가 있으며, 이 핸들러는 `toggleVisibility()`이라는 클래스 메서드를 실행합니다. `setState`에 함수를 전달하여 이 메서드를 정의하세요. 메서드가 호출될 때 `state`의 `visibility`가 반대 값으로 토글되도록 하세요. `visibility`이 `false`이면 메서드는 이를 `true`로 설정하고, 그 반대도 마찬가지입니다.

마지막으로 버튼을 클릭하여 컴포넌트가 `state`에 따라 조건부로 렌더링되는 것을 확인하세요.

**힌트:** `this`에서 메서드에 `constructor` 키워드를 바인딩하는 것을 잊지 마세요!

# --hints--

`MyComponent`은 `div` 요소를 반환해야 하며, 이 요소는 `button`를 포함해야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

`MyComponent`의 상태는 `visibility` 속성을 `false`로 설정하여 초기화해야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

버튼 요소를 클릭하면 상태의 `visibility` 속성이 `true`과 `false` 사이를 토글하고 `h1` 요소를 조건부로 렌더링해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

익명 함수가 `setState`에 전달되어야 합니다.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this`는 `setState` 내부에서 사용하면 안 됩니다.

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
