---
id: 5a24c314108439a4d4036174
title: "클래스 메서드에 'this' 바인딩하기"
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

`state`를 설정하고 업데이트하는 것 외에도 컴포넌트 클래스에 메서드를 정의할 수 있습니다. 클래스 메서드는 일반적으로 `this` 키워드를 사용해야 메서드 범위 내에서 클래스의 프로퍼티들(`state`, `props` 등)에 접근할 수 있습니다. 클래스 메서드가 `this`에 접근할 수 있도록 하는 방법이 몇 가지 있습니다.

가장 흔한 방법은 생성자에서 `this`를 명시적으로 바인딩하는 것입니다. 이렇게 하면 컴포넌트가 초기화될 때 `this`가 클래스 메서드에 바인딩됩니다. 이전 도전 과제에서 생성자에 `this.handleClick = this.handleClick.bind(this)`를 `handleClick` 메서드에 사용한 것을 기억할 수 있습니다. 그런 다음 클래스 메서드 내에서 `this.setState()` 같은 함수를 호출하면 `this`는 클래스를 가리키고 `undefined`되지 않습니다.

**참고:** `this` 키워드는 JavaScript에서 가장 혼란스러운 부분 중 하나지만 React에서 중요한 역할을 합니다. 여기서의 동작은 완전히 정상적이지만, 이 수업들은 `this`에 대한 심층 검토를 위한 곳이 아니므로 위 내용이 혼란스럽다면 다른 수업을 참조하세요!

# --instructions--

코드 편집기에는 텍스트를 추적하는 `state`가 있는 컴포넌트가 있습니다. 또한 텍스트를 `You clicked!`로 설정할 수 있는 메서드도 있습니다. 하지만 이 메서드는 정의되지 않은 `this` 키워드를 사용하기 때문에 작동하지 않습니다. 컴포넌트 생성자에서 `this`를 `handleClick()` 메서드에 명시적으로 바인딩하여 고치세요.

다음으로, 렌더링 메서드의 `button` 요소에 클릭 핸들러를 추가하세요. 버튼이 클릭 이벤트를 받으면 `handleClick()` 메서드를 실행해야 합니다. `onClick` 핸들러에 전달하는 메서드는 JavaScript로 직접 해석되어야 하므로 중괄호가 필요하다는 점을 기억하세요.

위 단계를 완료하면 버튼을 클릭할 때 `You clicked!`를 볼 수 있습니다.

# --hints--

`MyComponent`는 버튼과 `div` 요소 두 개를 순서대로 감싸는 `h1` 요소를 반환해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

`MyComponent`의 상태는 키-값 쌍 `{ text: "Hello" }`로 초기화되어야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

`button` 요소를 클릭하면 `handleClick` 메서드가 실행되어 상태 `text`이 `You clicked!`로 설정되어야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
