---
id: 5a24c314108439a4d4036170
title: 상태 유지 컴포넌트 만들기
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

React에서 가장 중요한 주제 중 하나는 `state`입니다. 상태는 애플리케이션이 알아야 하며 시간이 지나면서 변할 수 있는 모든 데이터를 포함합니다. 앱이 상태 변경에 반응하고 필요할 때 업데이트된 UI를 보여주길 원합니다. React는 현대 웹 애플리케이션의 상태 관리를 위한 훌륭한 해결책을 제공합니다.

React 컴포넌트에서 상태는 컴포넌트 클래스의 `state`에 `constructor` 속성을 선언하여 만듭니다. 이렇게 하면 컴포넌트가 생성될 때 `state`으로 초기화됩니다. `state` 속성은 JavaScript `object`로 설정해야 합니다. 선언하는 방법은 다음과 같습니다:

```jsx
this.state = {

}
```

컴포넌트가 존재하는 동안 `state` 객체에 접근할 수 있습니다. 이를 업데이트하고 UI에 렌더링하며 자식 컴포넌트에 프로프스로 전달할 수 있습니다. `state` 객체는 필요에 따라 복잡하거나 단순할 수 있습니다. `React.Component`을 이렇게 만들려면 `state`을 확장하여 클래스 컴포넌트를 생성해야 한다는 점을 기억하세요.

# --instructions--

코드 편집기에 `firstName`의 `state` 속성을 렌더링하려는 컴포넌트가 있습니다. 하지만 `state`가 정의되어 있지 않습니다. `state`에서 `constructor`으로 컴포넌트를 초기화하고 `firstName`의 속성에 자신의 이름을 할당하세요.

# --hints--

`StatefulComponent`가 존재하고 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent`는 `div`와 `h1` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

`StatefulComponent`의 상태는 문자열로 설정된 `firstName` 속성으로 초기화되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

`firstName`의 상태에 있는 `StatefulComponent` 속성은 `h1` 요소에 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
