---
id: 5a24c314108439a4d4036169
title: 무상태 함수형 컴포넌트에 프로프스 전달하기
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

이전 도전 과제들은 JSX 요소 생성과 조합, 함수형 컴포넌트, 그리고 React의 ES6 스타일 클래스 컴포넌트에 대해 많이 다뤘습니다. 이 기초를 바탕으로 React에서 매우 흔한 또 다른 기능인 **프로프스**를 살펴볼 시간입니다. React에서는 자식 컴포넌트에 프로프스, 즉 속성을 전달할 수 있습니다. 예를 들어 `App` 컴포넌트가 무상태 함수형 컴포넌트인 `Welcome` 자식 컴포넌트를 렌더링한다고 가정해 보겠습니다. `Welcome`에 `user` 속성을 전달하려면 다음과 같이 작성하면 됩니다:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

사용자가 직접 만든 React가 지원하는 **커스텀 HTML 속성**을 컴포넌트에 전달하는 방식입니다. 이 경우, 생성한 `user` 속성이 `Welcome` 컴포넌트에 전달됩니다. `Welcome`가 무상태 함수형 컴포넌트이므로 다음과 같이 이 값을 접근할 수 있습니다:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

이 값을 `props`라고 부르는 것이 일반적이며, 무상태 함수형 컴포넌트를 다룰 때는 JSX를 반환하는 함수의 인수로 간주합니다. 함수 본문에서 인수 값을 접근할 수 있습니다. 클래스 컴포넌트에서는 이 부분이 조금 다릅니다.

# --instructions--

코드 편집기에는 `Calendar`와 `CurrentDate` 컴포넌트가 있습니다. `CurrentDate` 컴포넌트에서 `Calendar`을 렌더링할 때, JavaScript의 `date` 객체에서 현재 날짜를 할당한 `Date` 속성을 전달하세요. 그런 다음 `prop` 컴포넌트에서 이 `CurrentDate`를 접근하여 `p` 태그 안에 값을 표시하세요. JavaScript로 평가되는 `prop` 값은 중괄호로 감싸야 한다는 점을 기억하세요. 예를 들어 `date={Date()}`처럼 작성합니다.

# --hints--

`Calendar` 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`Calendar` 컴포넌트의 두 번째 자식은 `CurrentDate` 컴포넌트여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

`CurrentDate` 컴포넌트는 `date`라는 프로프를 가져야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

`date` 프로프는 `CurrentDate`에 문자열이어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

`date` 프로프는 `Date()` 호출로 생성되어야 합니다.

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

`CurrentDate` 컴포넌트는 `date` 태그 안에 `p` 프로프의 값을 렌더링해야 합니다.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
