---
id: 5a24c314108439a4d4036162
title: 무상태 함수형 컴포넌트 만들기
challengeType: 6
forumTopicId: 301392
dashedName: create-a-stateless-functional-component
---

# --description--

컴포넌트는 React의 핵심입니다. React의 모든 것은 컴포넌트이며 여기서 컴포넌트를 만드는 방법을 배우게 됩니다.

React 컴포넌트를 만드는 방법은 두 가지가 있습니다. 첫 번째 방법은 JavaScript 함수를 사용하는 것입니다. 이렇게 컴포넌트를 정의하면 *무상태 함수형 컴포넌트*가 만들어집니다. 애플리케이션에서 상태 개념은 이후 도전 과제에서 다룹니다. 지금은 무상태 컴포넌트를 데이터는 받아서 렌더링할 수 있지만, 그 데이터의 변경을 관리하거나 추적하지 않는 컴포넌트라고 생각하세요. (React 컴포넌트를 만드는 두 번째 방법은 다음 도전 과제에서 다룹니다.)

함수로 컴포넌트를 만들려면 JSX나 `null`를 반환하는 JavaScript 함수를 작성하면 됩니다. 중요한 점은 React가 함수 이름이 대문자로 시작해야 한다는 것입니다. 다음은 JSX에서 HTML 클래스를 할당하는 무상태 함수형 컴포넌트 예시입니다:

```jsx
const DemoComponent = function() {
  return (
    <div className='customClass' />
  );
};
```

트랜스파일된 후 `<div>`는 `customClass`라는 CSS 클래스를 갖게 됩니다.

JSX 컴포넌트가 HTML을 나타내므로 여러 컴포넌트를 조합해 더 복잡한 HTML 페이지를 만들 수 있습니다. 이것이 React가 제공하는 컴포넌트 아키텍처의 주요 장점 중 하나입니다. 여러 개의 분리된 독립 컴포넌트로 UI를 구성할 수 있어 복잡한 사용자 인터페이스를 더 쉽게 만들고 유지할 수 있습니다.

# --instructions--

코드 편집기에 `MyComponent`라는 함수가 있습니다. 이 함수가 문자열 텍스트를 포함하는 단일 `div` 요소를 반환하도록 완성하세요.

**참고:** 텍스트는 `div` 요소의 자식으로 간주되므로 자기 닫는 태그를 사용할 수 없습니다.

# --hints--

`MyComponent`는 JSX를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.length === 1;
  })()
);
```

`MyComponent`는 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`div` 요소는 문자열 텍스트를 포함해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').text() !== '';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const MyComponent = function() {
  // Change code below this line



  // Change code above this line
}
```

# --solutions--

```jsx
const MyComponent = function() {
  // Change code below this line
  return (
    <div>
      Demo Solution
    </div>
  );
  // Change code above this line
}
```
