---
id: 5a24c314108439a4d4036188
title: 프로프스에 따라 조건부 렌더링하기
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

지금까지 `if/else`, `&&`, 그리고 삼항 연산자(`condition ? expressionIfTrue : expressionIfFalse`)를 사용해 무엇을 언제 렌더링할지 조건부로 결정하는 방법을 배웠습니다. 하지만 이 개념들을 React의 또 다른 강력한 기능인 프로프스와 결합할 수 있게 해 주는 중요한 주제가 하나 남아 있습니다. 프로프스를 사용해 조건부로 코드를 렌더링하는 것은 React 개발자들 사이에서 매우 흔한 일입니다 — 즉, 특정 프로프스의 값을 사용해 자동으로 렌더링할 내용을 결정합니다.

이번 도전 과제에서는 자식 컴포넌트가 프로프스를 기반으로 렌더링 결정을 내리도록 설정할 것입니다. 삼항 연산자도 사용할 텐데, 지난 도전 과제들에서 다룬 여러 개념이 이 맥락에서도 유용하게 쓰일 수 있음을 알 수 있을 겁니다.

# --instructions--

코드 편집기에는 `GameOfChance`라는 부모 컴포넌트와 `Results`라는 자식 컴포넌트 두 개가 부분적으로 정의되어 있습니다. 이들은 사용자가 버튼을 눌러 이겼는지 졌는지 확인하는 간단한 게임을 만드는데 사용됩니다.

먼저, 실행할 때마다 무작위로 다른 값을 반환하는 간단한 표현식이 필요합니다. `Math.random()`를 사용할 수 있습니다. 이 메서드는 호출될 때마다 `0` (포함)부터 `1` (미포함) 사이의 값을 반환합니다. 따라서 50/50 확률을 위해서는 표현식에 `Math.random() >= .5`을 사용하세요. 통계적으로 이 표현식은 50% 확률로 `true`를, 나머지 50% 확률로 `false`를 반환합니다. render 메서드에서 `null`을 위 표현식으로 바꿔 변수 선언을 완성하세요.

이제 코드에서 무작위 결정을 내릴 수 있는 표현식이 생겼습니다. 다음으로 이를 구현해야 합니다. `Results`의 자식으로 `GameOfChance` 컴포넌트를 렌더링하고, `expression`를 `fiftyFifty`라는 이름의 프로프스로 전달하세요. `Results` 컴포넌트에서는 `h1`에서 전달된 `You Win!` 프로프스를 기반으로 삼항 표현식을 작성해 `You Lose!` 요소를 `fiftyFifty` 또는 `GameOfChance` 텍스트와 함께 렌더링하세요. 마지막으로 `handleClick()` 메서드가 각 턴을 올바르게 세어 사용자가 몇 번 플레이했는지 알 수 있도록 하세요. 이 기능은 사용자가 연속으로 이기거나 져도 컴포넌트가 실제로 업데이트되었음을 알리는 역할도 합니다.

# --hints--

`GameOfChance` 컴포넌트가 존재하며 페이지에 렌더링되어야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance`는 단일 `button` 요소를 반환해야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance`는 `Results`라는 프로프스를 가진 `fiftyFifty` 컴포넌트의 단일 인스턴스를 반환해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

`GameOfChance` 상태는 `counter` 속성이 `1` 값으로 초기화되어야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

`GameOfChance` 컴포넌트가 처음 DOM에 렌더링될 때, 내부 텍스트가 `p`인 `Turn: 1` 요소가 반환되어야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

버튼이 클릭될 때마다 카운터 상태가 1씩 증가하고, 카운터 상태 값인 `p`을 포함하는 텍스트 `Turn: N`가 들어간 단일 `N` 요소가 DOM에 렌더링되어야 합니다.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

`GameOfChance` 컴포넌트가 처음 마운트될 때와 이후 버튼이 클릭될 때마다, `h1` 또는 `You Win!` 중 하나를 무작위로 렌더링하는 단일 `You Lose!` 요소가 반환되어야 합니다. 참고: 이 동작은 무작위 실패할 수 있습니다. 실패하면 다시 시도하세요.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
