---
id: 5a24c314108439a4d403617d
title: 라이프사이클 메서드 componentDidMount 사용하기
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

대부분의 웹 개발자는 어느 시점에서 API 엔드포인트를 호출해 데이터를 가져와야 합니다. React를 사용한다면 이 작업을 어디서 수행해야 하는지 아는 것이 중요합니다.

React의 모범 사례는 API 호출이나 서버에 대한 모든 호출을 라이프사이클 메서드 `componentDidMount()`에 두는 것입니다. 이 메서드는 컴포넌트가 DOM에 마운트된 후 호출됩니다. 여기서 `setState()`를 호출하면 컴포넌트가 다시 렌더링됩니다. 이 메서드에서 API를 호출하고 API가 반환하는 데이터로 상태를 설정하면, 데이터를 받는 즉시 자동으로 업데이트가 발생합니다.

# --instructions--

`componentDidMount()`에 모의 API 호출이 있습니다. 2.5초 후 상태를 설정하여 서버에서 데이터를 가져오는 것을 시뮬레이션합니다. 이 예제는 사이트의 현재 총 활성 사용자 수를 요청합니다. render 메서드에서 `activeUsers` 텍스트 뒤에 `h1` 안에 `Active Users:` 값을 렌더링하세요. 미리보기에서 어떻게 동작하는지 확인하고, 타임아웃 시간을 변경해 다양한 효과를 경험해 보세요.

# --hints--

`MyComponent`는 `div` 요소를 렌더링해야 하며, 이 요소는 `h1` 태그를 감싸야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

컴포넌트 상태는 `componentDidMount`에서 타임아웃 함수로 업데이트해야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

`h1` 태그는 `activeUsers` 상태의 `MyComponent` 값을 렌더링해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
