---
id: 5a24c314108439a4d4036168
title: React 컴포넌트를 처음부터 작성하기
challengeType: 6
forumTopicId: 301424
dashedName: write-a-react-component-from-scratch
---

# --description--

JSX와 React 컴포넌트의 기본을 배웠으니 이제 직접 컴포넌트를 작성할 차례입니다. React 컴포넌트는 React 애플리케이션의 핵심 구성 요소이므로 작성에 익숙해지는 것이 중요합니다. 일반적인 React 컴포넌트는 `class`를 확장하는 ES6 `React.Component`임을 기억하세요. 이 컴포넌트는 JSX에서 생성된 HTML 또는 `null`를 반환하는 render 메서드를 가집니다. 이것이 React 컴포넌트의 기본 형태입니다. 이것을 잘 이해하면 더 복잡한 React 프로젝트를 시작할 준비가 된 것입니다.

# --instructions--

`MyComponent`를 확장하는 `React.Component` 클래스를 정의하세요. 이 클래스의 render 메서드는 `div`를 반환해야 하며, 그 안에 텍스트 `h1`가 들어 있는 `My First React Component!` 태그가 포함되어야 합니다. 이 텍스트는 대소문자와 구두점까지 정확히 일치해야 합니다. 컴포넌트의 생성자도 반드시 호출하세요.

이 컴포넌트를 `ReactDOM.render()`를 사용해 DOM에 렌더링하세요. 사용할 수 있는 `div`와 `id='challenge-node'`가 있습니다.

# --hints--

`MyComponent`라는 React 컴포넌트가 있어야 합니다.

```js
assert(__helpers.removeWhiteSpace(code).includes('classMyComponentextendsReact.Component{'));
```

`MyComponent`는 텍스트 `h1`가 들어 있는 `My First React Component!` 태그를 포함해야 합니다. 대소문자와 구두점이 중요합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('h1').text() === 'My First React Component!';
  })()
);
```

`MyComponent`는 DOM에 렌더링되어야 합니다.

```js
assert(document.getElementById('challenge-node').childNodes.length === 1);
```

`MyComponent`는 `super`를 `props`와 함께 호출하는 생성자를 가져야 합니다.

```js
assert(
  MyComponent.toString().includes('MyComponent(props)') &&
    MyComponent.toString().includes('_super.call(this, props)')
);
```

# --seed--

## --seed-contents--

```jsx
// Change code below this line
```

# --solutions--

```jsx
// Change code below this line
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>My First React Component!</h1>
      </div>
    );
  }
};

ReactDOM.render(<MyComponent />, document.getElementById('challenge-node'));
```
