---
id: 5a24c314108439a4d4036187
title: 조건부 렌더링에 삼항 표현식 사용하기
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

동적 렌더링 기법으로 넘어가기 전에, 내장된 JavaScript 조건문을 사용해 원하는 내용을 렌더링하는 마지막 방법이 있습니다: <dfn>삼항 연산자</dfn>입니다. 삼항 연산자는 JavaScript에서 `if/else` 문을 간단히 처리하는 지름길로 자주 사용됩니다. 전통적인 `if/else` 문만큼 강력하지는 않지만 React 개발자들 사이에서 매우 인기가 많습니다. 그 이유 중 하나는 JSX가 컴파일되는 방식 때문인데, `if/else` 문은 JSX 코드 안에 직접 삽입할 수 없습니다. 몇 가지 도전 과제에서 이미 눈치챘겠지만 — `if/else` 문이 필요할 때는 항상 `return` 문 *외부*에 있었습니다. JSX 내에서 조건부 로직을 구현하고 싶다면 삼항 표현식이 훌륭한 대안이 될 수 있습니다. 삼항 연산자는 세 부분으로 구성되지만 여러 삼항 표현식을 함께 결합할 수도 있습니다. 기본 구문은 다음과 같습니다:

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

코드 편집기에는 `CheckUserAge` 컴포넌트의 `render()` 메서드 내에 세 개의 상수가 정의되어 있습니다. 이들은 `buttonOne`, `buttonTwo`, `buttonThree`라고 불리며, 각각 버튼 요소를 나타내는 간단한 JSX 표현식이 할당되어 있습니다. 먼저 `CheckUserAge`의 상태를 `input`과 `userAge`가 모두 빈 문자열 값으로 초기화하세요.

컴포넌트가 페이지에 정보를 렌더링하면, 사용자가 상호작용할 방법이 있어야 합니다. 컴포넌트의 `return` 문 내에서 다음 로직을 구현하는 삼항 표현식을 설정하세요: 페이지가 처음 로드될 때는 제출 버튼인 `buttonOne`를 페이지에 렌더링합니다. 그다음 사용자가 나이를 입력하고 버튼을 클릭하면, 나이에 따라 다른 버튼을 렌더링합니다. 사용자가 `18`보다 작은 숫자를 입력하면 `buttonThree`을 렌더링하세요. 사용자가 `18` 이상인 숫자를 입력하면 `buttonTwo`을 렌더링하세요.

# --hints--

`CheckUserAge` 컴포넌트는 하나의 `input` 요소와 하나의 `button` 요소를 렌더링해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

`CheckUserAge` 컴포넌트의 상태는 `userAge` 속성과 `input` 속성으로 초기화되어야 하며, 둘 다 빈 문자열 값으로 설정하세요.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

`CheckUserAge` 컴포넌트가 처음 DOM에 렌더링될 때, `button`의 내부 텍스트는 Submit이어야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

`input` 요소에 18보다 작은 숫자가 입력되고 `button`이 클릭되면, `button`의 내부 텍스트는 `You Shall Not Pass`이어야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

`input` 요소에 18 이상인 숫자가 입력되고 `button`가 클릭되면, `button`의 내부 텍스트는 `You May Enter`이어야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

숫자가 제출된 후 `input`의 값이 다시 변경되면, `button`은 다시 `Submit`을 읽어야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

코드에 `if/else` 문이 포함되어 있으면 안 됩니다.

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
