---
id: 5a24c314108439a4d4036173
title: this.setState로 상태 설정하기
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

이전 도전 과제에서는 컴포넌트 `state`와 `constructor`에서 상태를 초기화하는 방법을 다뤘습니다. 컴포넌트의 `state`를 변경하는 방법도 있습니다. React는 컴포넌트 `state`를 업데이트하는 `setState`라는 메서드를 제공합니다. 컴포넌트 클래스 내에서 `setState` 메서드를 다음과 같이 호출하세요: `this.setState()`, 키-값 쌍이 포함된 객체를 전달합니다. 키는 상태 속성이고 값은 업데이트된 상태 데이터입니다. 예를 들어, 상태에 `username`를 저장하고 이를 업데이트하려면 다음과 같이 작성합니다:

```jsx
this.setState({
  username: 'Lewis'
});
```

React는 `state`를 직접 수정하지 않고 상태 변경이 발생할 때 항상 `this.setState()`를 사용하기를 기대합니다. 또한 React는 성능 향상을 위해 여러 상태 업데이트를 일괄 처리할 수 있다는 점을 알아두세요. 이는 `setState` 메서드를 통한 상태 업데이트가 비동기일 수 있음을 의미합니다. 이 문제를 해결하는 대체 구문이 `setState` 메서드에 있습니다. 이 방법은 거의 필요하지 않지만 기억해 두면 좋습니다! 자세한 내용은 <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">React 기사</a>를 참조하세요.

# --instructions--

코드 편집기에는 `button` 요소가 있고 `onClick()` 핸들러가 있습니다. 이 핸들러는 브라우저에서 `button`가 클릭 이벤트를 받을 때 실행되며 `handleClick`에 정의된 `MyComponent` 메서드를 실행합니다. `handleClick` 메서드 내에서 `state`를 사용해 컴포넌트 `this.setState()`를 업데이트하세요. `name`의 `state` 속성을 문자열 `React Rocks!`로 설정하세요.

버튼을 클릭하면 렌더링된 상태가 업데이트되는 것을 확인할 수 있습니다. 클릭 핸들러 코드가 어떻게 작동하는지 완전히 이해하지 못해도 걱정하지 마세요. 다음 도전 과제에서 다룹니다.

# --hints--

`MyComponent`의 상태는 키-값 쌍 `{ name: Initial State }`로 초기화되어야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent`는 `h1` 제목 요소를 렌더링해야 합니다.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

렌더링된 `h1` 제목 요소는 컴포넌트 상태에서 렌더링된 텍스트를 포함해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

`handleClick`에서 `MyComponent` 메서드를 호출하면 상태의 name 속성이 `React Rocks!`와 같아져야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
