---
id: 5a24c314108439a4d4036167
title: 클래스 컴포넌트를 DOM에 렌더링하기
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

이전 도전 과제에서 ReactDOM API를 사용해 JSX 요소를 DOM에 렌더링했던 것을 기억할 수 있습니다. React 컴포넌트를 렌더링하는 과정도 매우 비슷하게 진행됩니다. 지난 몇 개의 도전 과제는 컴포넌트와 조합에 집중했기 때문에 렌더링은 백그라운드에서 자동으로 처리되었습니다. 하지만 React 코드는 ReactDOM API를 호출하지 않으면 DOM에 렌더링되지 않습니다.

구문을 다시 살펴보면: `ReactDOM.render(componentToRender, targetNode)`. 첫 번째 인수는 렌더링하려는 React 컴포넌트입니다. 두 번째 인수는 그 컴포넌트를 렌더링할 DOM 노드입니다.

React 컴포넌트는 JSX 요소와는 약간 다르게 `ReactDOM.render()`에 전달됩니다. JSX 요소는 렌더링하려는 요소 이름을 전달하지만, React 컴포넌트는 중첩 컴포넌트를 렌더링하는 것과 같은 구문을 사용해야 합니다. 예를 들어 `ReactDOM.render(<ComponentToRender />, targetNode)`와 같습니다. 이 구문은 ES6 클래스 컴포넌트와 함수형 컴포넌트 모두에 적용됩니다.

# --instructions--

`Fruits`와 `Vegetables` 컴포넌트는 백그라운드에서 정의되어 있습니다. 두 컴포넌트를 `TypesOfFood` 컴포넌트의 자식으로 렌더링한 다음, `TypesOfFood`를 DOM에 렌더링하세요. `div`와 `id='challenge-node'`가 사용 가능하니 참고하세요.

# --hints--

`TypesOfFood` 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood` 컴포넌트는 `Fruits` 요소 뒤에 `h1` 컴포넌트를 렌더링해야 합니다.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`TypesOfFood` 컴포넌트는 `Vegetables` 뒤에 `Fruits` 컴포넌트를 렌더링해야 합니다.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

`TypesOfFood` 컴포넌트는 id가 `div`인 `challenge-node` 내에 DOM으로 렌더링해야 합니다.

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
