---
id: 5a24c314108439a4d403618a
title: 배열의 map()을 사용해 동적으로 요소 렌더링하기
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

조건부 렌더링은 유용하지만, 컴포넌트가 알 수 없는 개수의 요소를 렌더링해야 할 때가 있습니다. 반응형 프로그래밍에서는 사용자의 상호작용에 따라 애플리케이션 상태가 실행 시점까지 알 수 없는 경우가 많습니다. 프로그래머는 미리 그 알 수 없는 상태를 올바르게 처리하도록 코드를 작성해야 합니다. React에서 `Array.map()`를 사용하는 것이 이 개념을 보여줍니다.

예를 들어, 간단한 "할 일 목록" 애플리케이션을 만든다고 합시다. 프로그래머로서 사용자가 목록에 몇 개의 항목을 추가할지 알 수 없습니다. 누군가가 오늘은 세탁하는 날이라고 결정하기 훨씬 전에 컴포넌트가 올바른 개수의 목록 요소를 동적으로 렌더링하도록 설정해야 합니다.

# --instructions--

코드 편집기에는 `MyToDoList` 컴포넌트 대부분이 설정되어 있습니다. 이 코드는 제어된 폼 도전 과제를 완료했다면 익숙할 것입니다. `textarea`와 `button`가 있고, 상태를 추적하는 몇 가지 메서드도 있지만 아직 페이지에 렌더링되는 것은 없습니다.

`constructor` 내부에 `this.state` 객체를 만들고 두 개의 상태를 정의하세요: `userInput`는 빈 문자열로 초기화하고, `toDoList`는 빈 배열로 초기화하세요. 다음으로 `null` 메서드에서 `render()` 변수 옆에 있는 `items` 값을 삭제하세요. 그 자리에 컴포넌트 내부 상태에 저장된 `toDoList` 배열을 map으로 순회하며 각 항목에 대해 `li`를 동적으로 렌더링하세요. `eat, code, sleep, repeat`에 문자열 `textarea`를 입력한 다음 버튼을 클릭해 결과를 확인하세요.

**참고:** 이처럼 map 연산으로 생성된 모든 형제 자식 요소에는 고유한 `key` 속성을 반드시 부여해야 한다는 점을 알고 있을 것입니다. 걱정하지 마세요, 이 내용은 다음 도전 과제에서 다룹니다.

# --hints--

MyToDoList 컴포넌트가 존재하고 페이지에 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

`MyToDoList`의 첫 번째 자식은 `textarea` 요소여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

`MyToDoList`의 두 번째 자식은 `br` 요소여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

`MyToDoList`의 세 번째 자식은 `button` 요소여야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

`MyToDoList` 상태는 빈 배열 `toDoList`로 초기화되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

`MyToDoList` 상태는 빈 문자열 `userInput`로 초기화되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

`Create List` 버튼을 클릭하면 `MyToDoList` 컴포넌트가 `textarea` 요소에 입력된 쉼표로 구분된 목록의 각 항목에 대해 목록 항목 요소를 포함하는 순서 없는 목록을 동적으로 반환해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
