---
id: 5a24c314108439a4d403617c
title: 라이프사이클 메서드 componentWillMount 사용하기
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React 컴포넌트에는 컴포넌트 라이프사이클의 특정 시점에 작업을 수행할 수 있는 여러 특별한 메서드가 있습니다. 이를 라이프사이클 메서드 또는 라이프사이클 훅이라고 하며, 특정 시점에 컴포넌트를 포착할 수 있게 해줍니다. 이는 컴포넌트가 렌더링되기 전, 업데이트되기 전, 프로프스를 받기 전, 언마운트되기 전 등 다양한 시점일 수 있습니다. 주요 라이프사이클 메서드 몇 가지는 다음과 같습니다: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` 이후 수업에서는 이러한 라이프사이클 메서드의 기본 사용 사례를 다룹니다.

**참고:** `componentWillMount` 라이프사이클 메서드는 16.X의 향후 버전에서 더 이상 지원되지 않으며 17 버전에서 제거될 예정입니다. 자세한 내용은 이 <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">글</a>을 참조하세요.

# --instructions--

컴포넌트가 DOM에 마운트될 때 `componentWillMount()` 메서드가 호출되기 전에 `render()` 메서드가 호출됩니다. `componentWillMount()` 안에서 콘솔에 무언가를 기록하세요 — 출력 결과를 확인하려면 브라우저 콘솔을 열어두는 것이 좋습니다.

# --hints--

`MyComponent`는 `div` 요소를 렌더링해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log`는 `componentWillMount` 안에서 호출되어야 합니다.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
