---
id: 5a24c314108439a4d403618b
title: 형제 요소에 고유한 key 속성 부여하기
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

지난 도전 과제에서는 `map` 메서드를 사용해 사용자 입력에 따라 여러 요소를 동적으로 렌더링하는 방법을 보여주었습니다. 하지만 그 예제에는 중요한 부분이 빠져 있었습니다. 배열로 요소를 만들 때는 각각에 고유한 `key` 속성을 설정해야 합니다. React는 이 key를 사용해 어떤 항목이 추가, 변경, 삭제되었는지 추적합니다. 이 덕분에 목록이 수정될 때 다시 렌더링하는 과정이 더 효율적입니다.

**참고:** key는 형제 요소 사이에서만 고유하면 되며, 애플리케이션 전체에서 전역적으로 고유할 필요는 없습니다.

# --instructions--

코드 편집기에는 몇 가지 프론트엔드 프레임워크가 들어 있는 배열과 `Frameworks()`라는 무상태 함수형 컴포넌트가 있습니다. `Frameworks()`는 지난 도전 과제처럼 배열을 순서 없는 목록으로 매핑해야 합니다. `map` 콜백을 완성해 `li` 배열의 각 프레임워크마다 `frontEndFrameworks` 요소를 반환하세요. 이번에는 각 `li`에 고유한 값으로 설정된 `key` 속성을 반드시 부여하세요. `li` 요소에는 `frontEndFrameworks`의 텍스트도 포함되어야 합니다.

보통 key는 렌더링되는 요소를 고유하게 식별할 수 있는 값으로 만드세요. 최후의 수단으로 배열 인덱스를 사용할 수 있지만, 보통은 고유 식별자를 사용하는 것이 좋습니다.

# --hints--

`Frameworks` 컴포넌트가 존재하며 페이지에 렌더링되어야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks`는 `h1` 요소를 렌더링해야 합니다.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks`는 `ul` 요소를 렌더링해야 합니다.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

`ul` 태그는 6개의 자식 `li` 요소를 렌더링해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

각 목록 항목 요소에는 고유한 `key` 속성이 있어야 합니다.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

각 목록 항목 요소에는 `frontEndFrameworks`의 텍스트가 포함되어야 합니다.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
