---
id: 5a24c314108439a4d403616b
title: 기본 프로프스 사용하기
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

React는 기본 프로프스를 설정하는 옵션도 제공합니다. 컴포넌트 자체의 속성으로 기본 프로프스를 할당할 수 있으며, 필요할 경우 React가 기본 프로프스를 할당합니다. 이를 통해 명시적으로 값이 제공되지 않은 경우 프로프 값이 무엇이어야 하는지 지정할 수 있습니다. 예를 들어 `MyComponent.defaultProps = { location: 'San Francisco' }`를 선언하면, 별도로 지정하지 않는 한 문자열 `San Francisco`로 설정된 location 프로프를 정의한 것입니다. 프로프가 undefined일 때 React가 기본 프로프스를 할당하지만, 프로프 값으로 `null`를 전달하면 그 값은 `null`로 유지됩니다.

# --instructions--

코드 편집기에는 `ShoppingCart` 컴포넌트가 표시됩니다. 이 컴포넌트에 `items` 프로프를 `0` 값으로 지정하는 기본 프로프스를 정의하세요.

# --hints--

`ShoppingCart` 컴포넌트가 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`ShoppingCart` 컴포넌트는 `{ items: 0 }` 기본 프로프스를 가져야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
