---
id: 5a24c314108439a4d4036180
title: shouldComponentUpdate로 재렌더링 최적화하기
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

지금까지는 어떤 컴포넌트가 새로운 `state`나 새로운 `props`를 받으면, 그 컴포넌트와 모든 자식 컴포넌트가 다시 렌더링되었습니다. 보통은 괜찮습니다. 하지만 React는 자식 컴포넌트가 새로운 `state`나 `props`를 받을 때 호출할 수 있는 생명주기 메서드를 제공하며, 컴포넌트를 업데이트할지 여부를 명확히 선언할 수 있습니다. 이 메서드는 `shouldComponentUpdate()`이고, `nextProps`와 `nextState`를 매개변수로 받습니다.

이 메서드는 성능 최적화에 유용합니다. 예를 들어, 기본 동작은 컴포넌트가 새로운 `props`를 받으면 `props`가 변경되지 않았더라도 다시 렌더링한다는 것입니다. `shouldComponentUpdate()`를 사용해 `props`를 비교하여 이를 방지할 수 있습니다. 이 메서드는 컴포넌트를 업데이트할지 여부를 React에 알려주는 `boolean` 값을 반환해야 합니다. 현재 프로프스(`this.props`)와 다음 프로프스(`nextProps`)를 비교해 업데이트가 필요한지 판단하고, 그에 따라 `true` 또는 `false`를 반환할 수 있습니다.

# --instructions--

`shouldComponentUpdate()` 메서드는 `OnlyEvens`라는 컴포넌트에 추가되어 있습니다. 현재 이 메서드는 `true`를 반환해서 `OnlyEvens`가 새로운 `props`를 받을 때마다 다시 렌더링됩니다. 이 메서드를 수정해 `OnlyEvens`가 새로운 프로프스의 `value`가 짝수일 때만 업데이트되도록 하세요. `Add` 버튼을 클릭하고 생명주기 훅이 실행될 때 브라우저 콘솔에 나타나는 이벤트 순서를 확인하세요.

# --hints--

`Controller` 컴포넌트는 `OnlyEvens` 컴포넌트를 자식으로 렌더링해야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

`shouldComponentUpdate` 메서드는 `OnlyEvens` 컴포넌트에 정의되어야 합니다.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

`OnlyEvens` 컴포넌트는 `h1` 값을 렌더링하는 `this.props.value` 태그를 반환해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens`는 `nextProps.value`가 짝수일 때만 다시 렌더링해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
