---
id: 5a24c314108439a4d4036166
title: React 컴포넌트 합성하기
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

React 컴포넌트와 JSX를 사용한 더 복잡한 합성이 계속되면서 중요한 점이 하나 있습니다. ES6 스타일 클래스 컴포넌트를 다른 컴포넌트 안에서 렌더링하는 것은 지난 도전 과제에서 사용한 간단한 컴포넌트를 렌더링하는 것과 다르지 않습니다. JSX 요소, 무상태 함수형 컴포넌트, ES6 클래스 컴포넌트를 모두 다른 컴포넌트 안에서 렌더링할 수 있습니다.

# --instructions--

코드 편집기에서 `TypesOfFood` 컴포넌트는 이미 `Vegetables`라는 컴포넌트를 렌더링하고 있습니다. 또한, 지난 도전 과제에서 사용한 `Fruits` 컴포넌트도 있습니다.

`Fruits` 안에 두 개의 컴포넌트를 중첩하세요 — 먼저 `NonCitrus`, 그다음 `Citrus`입니다. 이 두 컴포넌트는 백그라운드에서 제공됩니다. 다음으로, `Fruits` 클래스 컴포넌트를 `TypesOfFood` 컴포넌트 안에 중첩하세요. `h1` 제목 요소 아래, `Vegetables` 위에 위치해야 합니다. 결과는 두 가지 다른 컴포넌트 유형을 사용하는 일련의 중첩된 컴포넌트가 되어야 합니다.

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

`TypesOfFood` 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood` 컴포넌트는 `Fruits` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`Fruits` 컴포넌트는 `NonCitrus` 컴포넌트와 `Citrus` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

`TypesOfFood` 컴포넌트는 `Vegetables` 컴포넌트 아래에 `Fruits` 컴포넌트를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
