---
id: 5a24c314108439a4d4036179
title: 제어된 폼 만들기
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

마지막 도전 과제에서는 React가 `input`와 `textarea` 같은 특정 요소의 내부 상태를 제어할 수 있음을 보여주었습니다. 이것이 제어된 컴포넌트가 되는 이유입니다. 이 원리는 일반 HTML `form` 요소를 포함한 다른 폼 요소에도 적용됩니다.

# --instructions--

`MyForm` 컴포넌트는 제출 핸들러가 있는 빈 `form`로 설정되어 있습니다. 폼이 제출될 때 제출 핸들러가 호출됩니다.

폼을 제출하는 버튼을 추가했습니다. 이 버튼에는 폼을 제어하는 버튼임을 나타내는 `type`가 `submit`로 설정되어 있습니다. `input` 안에 `form` 요소를 추가하고, 지난 도전 과제처럼 `value`와 `onChange()` 속성을 설정하세요. 그런 다음 `handleSubmit` 메서드를 완성하여 컴포넌트 상태 속성 `submit`를 로컬 `state`에 있는 현재 입력 값으로 설정해야 합니다.

**참고:** 제출 핸들러에서 기본 폼 제출 동작을 막아 웹 페이지가 새로 고침되는 것을 방지하려면 `event.preventDefault()`도 호출해야 합니다. 학습자의 편의를 위해 기본 동작은 여기서 비활성화되어 도전 과제 코드가 새로 고침으로 초기화되지 않도록 했습니다.

마지막으로 `h1` 뒤에 `form` 태그를 만들어 컴포넌트 상태의 `submit`에서 `state` 값을 렌더링하세요. 그러면 폼에 입력하고 버튼을 클릭하거나 엔터를 누르면 입력한 내용이 페이지에 렌더링되는 것을 볼 수 있습니다.

# --hints--

`MyForm`는 `div` 요소를 반환해야 하며, 이 안에는 `form`과 `h1` 태그가 포함되어야 합니다. 폼에는 `input`과 `button`가 포함되어야 합니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

`MyForm`의 상태는 빈 문자열로 설정된 `input`과 `submit` 속성으로 초기화해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

`input` 요소에 입력하면 컴포넌트 상태의 `input` 속성이 업데이트되어야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

폼을 제출하면 `handleSubmit`이 실행되어 상태의 `submit` 속성을 현재 입력 값으로 설정해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit`는 `event.preventDefault`를 호출해야 합니다.

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

`h1` 제목 요소는 컴포넌트 상태의 `submit` 필드 값을 렌더링해야 합니다.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
