---
id: 5a24c314108439a4d403618d
title: renderToString으로 서버에서 React 렌더링하기
challengeType: 6
forumTopicId: 301407
dashedName: render-react-on-the-server-with-rendertostring
---

# --description--

지금까지는 클라이언트에서 React 컴포넌트를 렌더링해왔습니다. 보통은 항상 이렇게 하게 됩니다. 하지만 React 컴포넌트를 서버에서 렌더링하는 것이 의미가 있는 경우도 있습니다. React는 JavaScript 뷰 라이브러리이고 Node로 서버에서 JavaScript를 실행할 수 있기 때문에 가능합니다. 실제로 React는 이 목적을 위해 사용할 수 있는 `renderToString()` 메서드를 제공합니다.

실제 앱에서 서버 렌더링을 사용하는 데는 두 가지 주요 이유가 있습니다. 첫째, 이렇게 하지 않으면 React 앱은 초기 로드 시 비교적 빈 HTML 파일과 큰 JavaScript 번들로 구성됩니다. 이는 페이지 내용을 색인하려는 검색 엔진에 이상적이지 않을 수 있습니다. 서버에서 초기 HTML 마크업을 렌더링하여 클라이언트에 보내면 초기 페이지 로드에 페이지의 모든 마크업이 포함되어 검색 엔진이 크롤링할 수 있습니다. 둘째, 렌더링된 HTML이 전체 앱의 JavaScript 코드보다 작기 때문에 초기 페이지 로드 속도가 빨라집니다. React는 초기 로드 후에도 앱을 인식하고 관리할 수 있습니다.

# --instructions--

`renderToString()` 메서드는 여기서 전역 객체로 사용할 수 있는 `ReactDOMServer`에 제공됩니다. 이 메서드는 React 요소 하나를 인수로 받습니다. 이를 사용해 `App`를 문자열로 렌더링하세요.

# --hints--

`App` 컴포넌트는 `ReactDOMServer.renderToString`를 사용해 문자열로 렌더링해야 합니다.

```js
  assert(
    code.replace(/ /g, '')
      .includes('ReactDOMServer.renderToString(<App/>)') &&
      Enzyme.mount(React.createElement(App)).children().name() === 'div'
  );
```

# --before-all--

```js
var ReactDOMServer = { renderToString(x) { return null; } };
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
```

# --solutions--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
ReactDOMServer.renderToString(<App/>);
```
