---
id: 5a24c314108439a4d4036184
title: if-else 조건으로 렌더링하기
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

JavaScript를 사용해 렌더링된 뷰를 제어하는 또 다른 방법은 렌더링되는 요소를 조건에 연결하는 것입니다. 조건이 참일 때는 한 뷰가 렌더링되고, 거짓일 때는 다른 뷰가 렌더링됩니다. React 컴포넌트의 `if/else` 메서드 안에서 표준 `render()` 문으로 이 작업을 할 수 있습니다.

# --instructions--

MyComponent에는 UI에 어떤 요소를 표시할지 여부를 추적하는 `boolean`가 상태에 있습니다. `button`는 이 값의 상태를 토글합니다. 현재는 매번 같은 UI를 렌더링합니다. `render()` 메서드를 `if/else` 문을 사용해 다시 작성하세요. `display`가 `true`일 경우 현재 마크업을 반환하고, 그렇지 않으면 `h1` 요소가 없는 마크업을 반환하세요.

**참고:** 테스트를 통과하려면 `if/else`를 작성해야 합니다. 여기서는 삼항 연산자 사용 시 통과하지 못합니다.

# --hints--

`MyComponent`가 존재하고 렌더링되어야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

`display`가 `true`로 설정되면 `div`, `button`, `h1`가 렌더링되어야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

`display`가 `false`로 설정되면 `div`와 `button`만 렌더링되어야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

render 메서드는 `if/else` 조건을 확인하기 위해 `this.state.display` 문을 사용해야 합니다.

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
