---
id: 5a24c314108439a4d4036161
title: 자기 닫는 JSX 태그에 대해 배우기
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

지금까지 JSX가 HTML과 어떻게 다른지, HTML 클래스를 정의할 때 `className` 대신 `class`을 사용하는 중요한 차이점을 보았습니다.

JSX가 HTML과 다른 또 다른 중요한 점은 자기 닫는 태그 개념입니다.

HTML에서는 거의 모든 태그가 여는 태그와 닫는 태그를 모두 가집니다: `<div></div>`; 닫는 태그는 항상 닫으려는 태그 이름 앞에 슬래시가 붙습니다. 하지만 HTML에는 <dfn>void elements</dfn>라고 불리는 특별한 경우가 있는데, 이는 다른 요소가 시작되기 전에 여는 태그와 닫는 태그가 모두 필요하지 않은 요소입니다.

예를 들어 줄 바꿈 태그는 `<br>` 또는 `<br />`로 쓸 수 있지만, 내용이 없으므로 `<br></br>`로는 절대 쓰면 안 됩니다.

JSX에서는 규칙이 약간 다릅니다. 모든 JSX 요소는 자기 닫는 태그로 쓸 수 있고, 모든 요소는 반드시 닫혀야 합니다. 예를 들어 줄 바꿈 태그는 유효한 JSX로 변환되려면 항상 `<br />`로 작성해야 합니다. 반면 `<div>`는 `<div />` 또는 `<div></div>`로 쓸 수 있습니다. 첫 번째 구문 버전에서는 `<div />` 안에 아무것도 포함할 수 없다는 차이가 있습니다. 이후 도전 과제에서 이 구문이 React 컴포넌트를 렌더링할 때 유용하다는 것을 알게 될 것입니다.

# --instructions--

코드 편집기에서 오류를 고쳐 유효한 JSX가 되어 성공적으로 변환되도록 하세요. 내용은 변경하지 말고, 태그가 필요한 곳에만 닫는 태그를 추가하세요.

# --hints--

상수 `JSX`은 `div` 요소를 반환해야 합니다.

```js
assert.strictEqual(JSX.type, 'div');
```

`div`은 `br` 태그를 포함해야 합니다.

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

`div`는 `hr` 태그를 포함해야 합니다.

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
