---
id: 5a24c314108439a4d4036171
title: 사용자 인터페이스에 상태 렌더링하기
challengeType: 6
forumTopicId: 301409
dashedName: render-state-in-the-user-interface
---

# --description--

컴포넌트의 초기 상태를 정의하면 렌더링되는 UI에서 그 상태의 어떤 부분이든 표시할 수 있습니다. 컴포넌트가 상태 유지형이라면 항상 `state`의 데이터를 `render()` 메서드에서 접근할 수 있습니다. 이 데이터는 `this.state`로 접근할 수 있습니다.

렌더 메서드의 `return` 내에서 상태 값을 접근하려면 그 값을 중괄호로 감싸야 합니다.

`state`는 React 컴포넌트에서 가장 강력한 기능 중 하나입니다. 앱의 중요한 데이터를 추적하고 이 데이터 변경에 따라 UI를 렌더링할 수 있게 해줍니다. 데이터가 변경되면 UI도 변경됩니다. React는 가상 DOM이라는 것을 사용해 내부에서 변경 사항을 추적합니다. 상태 데이터가 업데이트되면 해당 데이터를 사용하는 컴포넌트들, 자식 컴포넌트가 프로프로 받은 경우도 포함해, 다시 렌더링을 트리거합니다. React는 실제 DOM을 업데이트하지만 필요한 부분만 변경합니다. 즉, DOM 변경에 대해 신경 쓸 필요 없이 UI가 어떻게 보여야 하는지만 선언하면 됩니다.

컴포넌트를 상태 유지형으로 만들면 다른 컴포넌트는 그 `state`를 알지 못합니다. `state`는 완전히 캡슐화되어 해당 컴포넌트에 국한되며, 상태 데이터를 자식 컴포넌트에 `props`로 전달하지 않는 한 외부에 노출되지 않습니다. 이처럼 캡슐화된 `state` 개념은 특정 로직을 작성하고 그 로직을 코드 내 한 곳에 격리해 둘 수 있게 해주기 때문에 매우 중요합니다.

# --instructions--

코드 편집기에서 `MyComponent`는 이미 상태 유지형입니다. 컴포넌트의 render 메서드에 `h1` 태그를 정의해 컴포넌트 상태의 `name` 값을 렌더링하세요.

**참고:** `h1`는 `state`의 값만 렌더링해야 하며 다른 것은 렌더링하지 않아야 합니다. JSX에서 중괄호 `{ }` 안에 작성한 코드는 JavaScript로 처리됩니다. 따라서 `state`의 값을 접근하려면 참조를 중괄호로 감싸면 됩니다.

# --hints--

`MyComponent`는 상태에 저장된 키 `name`와 값 `freeCodeCamp`를 가져야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent`는 단일 `h1`로 감싼 `div` 제목 요소를 렌더링해야 합니다.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

렌더링된 `h1` 제목 요소는 컴포넌트 상태에서 렌더링된 텍스트만 포함해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  const getValue = firstValue.replace(/\s/g, '');
  assert(getValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{this.state.name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
