---
id: 5a24c314108439a4d4036177
title: 간단한 카운터 작성하기
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

지금까지 다룬 개념들을 결합하여 더 복잡한 상태 유지 컴포넌트를 설계할 수 있습니다. 여기에는 `state` 초기화, `state`를 설정하는 메서드 작성, 그리고 이러한 메서드를 실행하는 클릭 핸들러 할당이 포함됩니다.

# --instructions--

`Counter` 컴포넌트는 `count`에 `state` 값을 추적합니다. 두 개의 버튼이 있으며 각각 `increment()`와 `decrement()` 메서드를 호출합니다. 적절한 버튼을 클릭할 때 카운터 값이 1씩 증가하거나 감소하도록 이 메서드들을 작성하세요. 또한, 리셋 버튼을 클릭하면 카운트가 0으로 설정되도록 `reset()` 메서드도 만드세요.

**참고:** 버튼의 `className`를 수정하지 마세요. 그리고 생성자에서 새로 만든 메서드에 필요한 바인딩을 추가하는 것도 잊지 마세요.

# --hints--

`Counter`는 세 개의 버튼을 포함하는 `div` 요소를 반환해야 하며, 버튼의 텍스트 내용은 이 순서대로 `Increment!`, `Decrement!`, `Reset`입니다.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

`Counter`의 상태는 `count`으로 설정된 `0` 속성으로 초기화해야 합니다.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

증가 버튼을 클릭하면 카운트가 `1`만큼 증가해야 합니다.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

감소 버튼을 클릭하면 카운트가 `1`만큼 감소해야 합니다.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

리셋 버튼을 클릭하면 카운트가 `0`로 리셋되어야 합니다.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
