---
id: 5a24c314108439a4d4036164
title: 컴포지션으로 컴포넌트 만들기
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

이제 여러 React 컴포넌트를 함께 컴포지션하는 방법을 살펴보겠습니다. 앱을 만들고 있다고 상상해 보세요. 세 개의 컴포넌트 `Navbar`, `Dashboard`, `Footer`를 만들었습니다.

이 컴포넌트들을 함께 컴포지션하려면, 이 세 컴포넌트를 *자식*으로 렌더링하는 `App` *부모* 컴포넌트를 만들 수 있습니다. React 컴포넌트에서 자식으로 컴포넌트를 렌더링하려면, JSX 안에 사용자 정의 HTML 태그로 작성된 컴포넌트 이름을 포함하면 됩니다. 예를 들어, `render` 메서드 안에 다음과 같이 쓸 수 있습니다:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

React가 다른 컴포넌트를 참조하는 사용자 정의 HTML 태그(이 예시처럼 `< />`로 감싼 컴포넌트 이름)를 만나면, 그 태그 위치에 해당 컴포넌트의 마크업을 렌더링합니다. 이것은 `App` 컴포넌트와 `Navbar`, `Dashboard`, `Footer` 사이의 부모/자식 관계를 보여줍니다.

# --instructions--

코드 편집기에는 `ChildComponent`라는 간단한 함수형 컴포넌트와 `ParentComponent`라는 클래스 컴포넌트가 있습니다. `ChildComponent` 안에 `ParentComponent`를 렌더링하여 두 컴포넌트를 함께 컴포지션하세요. `ChildComponent` 태그를 슬래시로 닫는 것을 잊지 마세요.

**참고:** `ChildComponent`는 ES6 화살표 함수로 정의되어 있는데, React를 사용할 때 매우 흔한 방식입니다.

# --hints--

React 컴포넌트는 단일 `div` 요소를 반환해야 합니다.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

컴포넌트는 두 개의 중첩된 요소를 반환해야 합니다.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

컴포넌트는 두 번째 자식으로 `ChildComponent`를 반환해야 합니다.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
