---
id: 5a24c314108439a4d4036178
title: 제어되는 입력 만들기
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---

# --description--

애플리케이션은 `state`와 렌더링된 UI 사이에 더 복잡한 상호작용이 있을 수 있습니다. 예를 들어, `input`와 `textarea` 같은 텍스트 입력용 폼 컨트롤 요소는 사용자가 입력할 때 DOM에서 자체 상태를 유지합니다. React를 사용하면 이 변경 가능한 상태를 React 컴포넌트의 `state`로 옮길 수 있습니다. 사용자의 입력은 애플리케이션 `state`의 일부가 되어 React가 해당 입력 필드의 값을 제어합니다. 일반적으로 사용자가 입력할 수 있는 입력 필드가 있는 React 컴포넌트가 있다면, 그것은 제어되는 입력 폼일 것입니다.

# --instructions--

코드 편집기에는 제어되는 `ControlledInput` 요소를 만들기 위한 `input`라는 컴포넌트의 뼈대가 있습니다. 컴포넌트의 `state`는 이미 빈 문자열을 가진 `input` 속성으로 초기화되어 있습니다. 이 값은 사용자가 `input` 필드에 입력하는 텍스트를 나타냅니다.

먼저 `handleChange()`라는 메서드를 만들고 `event`라는 매개변수를 받으세요. 이 메서드가 호출되면 `event` 요소에서 온 문자열 텍스트를 포함하는 `input` 객체를 받습니다. 메서드 내에서 `event.target.value`를 사용해 이 문자열에 접근할 수 있습니다. 컴포넌트의 `input`의 `state` 속성을 이 새로운 문자열로 업데이트하세요.

`render` 메서드에서 `input` 태그 위에 `h4` 요소를 만드세요. 컴포넌트의 `value`의 `input` 속성과 같은 값을 갖는 `state` 속성을 추가하세요. 그리고 `onChange` 이벤트 핸들러 메서드로 설정된 `handleChange()` 속성도 추가하세요.

입력 상자에 타이핑하면 그 텍스트가 `handleChange()` 메서드로 처리되고, 로컬 `input`의 `state` 속성으로 설정되며, 페이지의 `input` 상자에 값으로 렌더링됩니다. 컴포넌트 `state`는 입력 데이터에 관한 단일 진실의 출처입니다.

마지막으로 생성자에서 필요한 바인딩을 추가하는 것을 잊지 마세요.

# --hints--

`ControlledInput`는 `div` 요소를 반환해야 하며, 이 요소는 `input`와 `p` 태그를 포함해야 합니다.

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

`ControlledInput`의 상태는 빈 문자열로 설정된 `input` 속성으로 초기화해야 합니다.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

입력 요소에 타이핑하면 상태와 입력 값이 업데이트되어야 하며, `p` 요소가 입력하는 동안 이 상태를 렌더링해야 합니다.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
