---
id: 642ddfdea4200e313f80a4b6
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`document` 객체가 새로운 HTML 요소를 동적으로 생성할 수 있는 `.createElement()` 메서드를 가지고 있다는 점을 기억하세요.

`createLabel` 함수에서 `label` 변수를 선언하고 새로운 `div` 요소를 할당하세요.

# --hints--

`label` 함수에서 `createLabel` 변수를 선언해야 합니다.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

`label` 변수는 `const`로 선언해야 합니다.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

`.createElement()` 객체의 `document` 메서드를 사용하세요.

```js
assert.match(code, /document\.createElement\(/);
```

`"div"` 메서드에 문자열 `.createElement()`을 전달해야 합니다.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

새로운 `div` 요소를 `label`에 할당하세요.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
