---
id: 642df9df4b5216350de7b0d2
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

숫자 범위를 생성하는 함수를 만들어야 합니다.

빈 `range` 함수를 선언하세요. 이 함수는 `start`과 `end` 매개변수를 받습니다. `Array()` 생성자를 사용하고 빈 배열을 암묵적으로 반환하세요.

# --hints--

`range` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

`const`를 사용해 `range` 변수를 선언하세요.

```js
assert.match(code, /const\s+range/);
```

`range` 변수는 함수여야 합니다.

```js
assert.isFunction(range);
```

`range` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

`range` 함수는 먼저 `start` 매개변수를 받아야 합니다.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

`range` 함수는 두 번째로 `end` 매개변수를 받아야 합니다.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

`range` 함수는 암묵적 반환을 사용해야 합니다. 중괄호를 사용하지 않는다는 점을 기억하세요.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

`range` 함수는 `Array()` 생성자를 사용해야 합니다. 원시 생성자는 `new` 키워드가 필요 없습니다.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

`Array()` 생성자에는 아무것도 전달하지 않아야 합니다.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
