---
id: 642df32c0c2db433d8b46d46
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

마지막으로, `.appendChild()` 메서드를 사용해 `label` 요소를 `container` 요소에 추가하세요.

# --hints--

`.appendChild()` 요소의 `container` 메서드에 접근해야 합니다.

```js
assert.match(code, /container\.appendChild\(/);
```

`label` 요소를 `.appendChild()` 메서드에 전달해야 합니다.

```js
assert.match(code, /container\.appendChild\(\s*label\s*\)/);
```

속성을 설정한 후 `label`를 추가해야 합니다.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)\s*;?\s*label\.className\s*=\s*('|"|`)label\2\s*;?\s*label\.textContent\s*=\s*name\s*;?\s*container\.appendChild\(\s*label\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;

  }
}
--fcc-editable-region--
```
