---
id: 642def66e6a60432c9a0371e
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`className` 요소의 `label` 속성을 `"label"`로 설정하고, `textContent`를 `name` 매개변수로 설정하세요.

# --hints--

`className` 요소의 `label` 속성에 접근해야 합니다.

```js
assert.match(code, /label\.className/);
```

`className` 속성을 문자열 `"label"`로 설정해야 합니다.

```js
assert.match(code, /label\.className\s*=\s*('|"|`)label\1/);
```

`textContent` 요소의 `label` 속성에 접근해야 합니다.

```js
assert.match(code, /label\.textContent/);
```

`textContent` 속성을 `name` 값으로 설정해야 합니다.

```js
assert.match(code, /label\.textContent\s*=\s*name/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");

  }
}
--fcc-editable-region--
```
