---
id: 642db8c409d9991d0b3b2f0d
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

프로젝트는 기본 HTML 컨테이너와 일부 CSS로 시작합니다. 첫 번째 과제는 스프레드시트의 셀을 프로그래밍 방식으로 생성하는 것입니다.

전역 `window` 객체는 브라우저 창(또는 탭)을 나타냅니다. 이 객체에는 스타일시트와 스크립트를 포함해 페이지 전체가 로드되었을 때 동작을 정의할 수 있는 `onload` 속성이 있습니다.

먼저 `onload`의 `window` 속성을 매개변수가 없는 화살표 함수로 설정하세요. 함수 안에서 `container` 변수를 선언하고 `id`의 `"container"`로 요소를 가져오는 값으로 할당하세요.

# --hints--

`onload` 객체의 `window` 속성에 접근해야 합니다.

```js
assert.match(code, /window\.onload/);
```

`onload` 속성을 함수로 설정해야 합니다.

```js
assert.isFunction(window.onload);
```

화살표 구문을 사용해야 합니다.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

`onload` 함수는 매개변수를 받지 않아야 합니다.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

`container` 함수 안에 `onload` 변수를 선언해야 합니다.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

`container` 변수는 `const`로 선언해야 합니다.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

`document.getElementById()`를 사용해야 합니다.

```js
assert.match(code, /document\.getElementById\(/);
```

`id`의 `"container"`로 요소를 가져와야 합니다.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

`#container` 요소를 `container`에 할당해야 합니다.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
