---
id: 642dfb07e7fa6736251541c8
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

배열 크기는 범위 크기와 같아야 합니다. `end`와 `start`의 차이를 구한 다음 `1`를 더해서 이 값을 계산할 수 있습니다.

이 계산 결과를 `Array()` 생성자의 인수로 전달하세요.

# --hints--

`start`에서 `end`를 빼야 합니다.

```js
assert.match(code, /end\s*-\s*start/);
```

`1` 계산에 `end - start`를 더해야 합니다.

```js
assert.match(code, /end\s*-\s*start\s*\+\s*1/);
```

계산한 값을 `Array()` 생성자에 전달해야 합니다.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array();
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
