---
id: 642e02be7845f13b014cd2b0
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

이제 `range` 함수를 가지고 있으니, 이를 사용해 문자 범위를 만들 수도 있습니다.

`charRange`와 화살표 구문을 사용해 `const` 함수를 선언하세요. 이 함수는 `start`와 `end` 매개변수를 받아야 합니다. 함수는 `range()`와 `start`를 인수로 하여 `end`를 호출한 결과를 암묵적으로 반환해야 합니다.

# --hints--

`charRange` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

`charRange` 변수는 `const`로 선언해야 합니다.

```js
assert.match(code, /const\s+charRange/);
```

`charRange` 변수는 함수여야 합니다.

```js
assert.isFunction(charRange);
```

`charRange` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

`charRange` 함수는 첫 번째 매개변수로 `start`를 받아야 합니다.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

`charRange` 함수는 두 번째 매개변수로 `end`를 받아야 합니다.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

`charRange` 함수는 암묵적 반환을 사용해야 합니다.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

`charRange` 함수는 `range` 함수를 호출해야 합니다.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

`start`와 `end`를 `range` 호출의 인수로 전달해야 합니다.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
