---
id: 642e0011c45c893845842058
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`Array()` 생성자에는 배열을 특정 값으로 채우는 데 사용할 수 있는 `.fill()` 메서드가 있습니다. 이 메서드를 사용해 배열을 `start` 값으로 채우세요.

`.fill()` 메서드를 `Array()` 생성자에 연결하고 `start` 값을 전달하세요.

# --hints--

`.fill()` 메서드를 사용하세요.

```js
assert.match(code, /\.fill\(/);
```

`.fill()` 생성자에서 `Array()` 메서드를 호출하세요.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(/);
```

`start` 메서드에 `.fill()` 값을 전달하세요.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*end\s*-\s*start\s*\+\s*1\s*\)\.fill\(\s*start\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const range = (start, end) => Array(end - start + 1);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
