---
id: 6434759f78ec812264ff8f34
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

콜백에서 두 개의 함수를 호출해야 합니다. 먼저 `createLabel()`를 호출하고 `number`을 인수로 전달하세요. 스프레드시트에 숫자가 나타나는 것을 볼 수 있습니다.

그다음 `.forEach()` 배열에서 `letters` 메서드를 호출하세요. `letter` 매개변수를 받는 빈 콜백 함수를 전달하세요.

# --hints--

`createLabel()` 함수를 호출해야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/);
```

`number`을 `createLabel()` 호출에 전달해야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*createLabel\(/)
```

`.forEach()` 배열에서 `letters` 메서드를 호출해야 합니다.

```js
assert.lengthOf(code.match(/letters\.forEach\(/g), 2)
```

화살표 구문을 사용한 콜백 함수를 `.forEach()` 메서드에 전달해야 합니다.

```js
assert.match(code, /letters\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>\s*\{/)
```

콜백 함수에 `letter` 매개변수가 있어야 합니다.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{/)
```

콜백 함수는 비어 있어야 합니다.

```js
assert.match(code, /letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}/)
```

`letters.forEach()` 콜백 함수는 `range(1, 99).forEach(number => {}` 콜백 함수 안에 중첩되어야 합니다.

```js
assert.match(code, /range\s*\(\s*1\s*,\s*99\s*\)\s*.forEach\s*\(\s*(\(\s*number\s*\)|number)\s*=>\s*\{\s*[^}]*letters\.forEach\(\s*(\(\s*letter\s*\)|letter)\s*=>\s*\{\s*\}\s*\)\s*;?\s*\}\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--
  range(1, 99).forEach(number => {

  })
--fcc-editable-region--
}
```
