---
id: 6434552bcc0a951a0a99df3b
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`range` 함수는 숫자를 기대하지만, `start`과 `end` 값은 문자열일 것입니다(특히 `A` 같은 단일 문자입니다).

`start` 호출에서 `end`와 `range()` 값을 숫자로 변환하려면, 해당 값들에 `.charCodeAt()` 메서드를 사용하고, 그 메서드에 숫자 `0`을 인수로 전달하세요.

# --hints--

`.charCodeAt()` 메서드를 사용하세요.

```js
assert.match(code, /\.charCodeAt\(/);
```

`.charCodeAt()`에 `start` 메서드를 호출하세요.

```js
assert.match(code, /start\.charCodeAt\(/);
```

`0`의 `.charCodeAt()` 메서드에 `start`를 전달하세요.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

`.charCodeAt()`에 `end` 메서드를 호출하세요.

```js
assert.match(code, /end\.charCodeAt\(/);
```

`0`의 `.charCodeAt()` 메서드에 `end`를 전달하세요.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

`.charCodeAt()` 호출에서 `range` 메서드를 직접 사용하세요.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
