---
id: 643475e13dc727231acd0f72
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

이제 중첩된 `.forEach()` 호출 안에서 `input` 변수를 선언하세요. `.createElement()` 객체의 `document` 메서드를 사용해 `input` 요소를 만드세요. `type` 속성을 `"text"`으로, `id` 속성을 `letter + number`으로 설정하세요.

# --hints--

`input` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

`const`를 사용해 `input` 변수를 선언하세요.

```js
assert.match(code, /const\s+input/)
```

`.createElement()` 객체의 `document` 메서드를 호출해야 합니다.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

`"input"` 메서드에 문자열 `.createElement()`를 전달하세요.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

새로운 `input` 요소를 `input` 변수에 할당하세요.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

`type` 요소의 `input` 속성에 접근하세요.

```js
assert.match(code, /input\.type/);
```

`type` 요소의 `input` 속성을 `"text"`로 설정하세요.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

`id` 요소의 `input` 속성에 접근하세요.

```js
assert.match(code, /input\.id/);
```

`id` 요소의 `input` 속성을 `letter + number`로 설정하세요.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
