---
id: 64345c560591891f64976f7a
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

도우미 함수가 완성되었으니, `onload` 이벤트 핸들러 안에서 `letters` 변수를 선언하세요. `charRange()`와 `"A"` 문자를 인수로 하여 `"J"`를 호출한 결과를 이 변수에 할당하세요.

# --hints--

`letters` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

`const`를 사용해 `letters` 변수를 선언하세요.

```js
assert.match(code, /const\s+letters/);
```

`charRange()` 호출 결과를 `letters` 변수에 할당해야 합니다.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

`"A"`을 `charRange()` 호출의 첫 번째 인수로 전달하세요.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

`"J"`을 `charRange()` 호출의 두 번째 인수로 전달하세요.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
