---
id: 64347464f78cd9209545f35c
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

이제 `.forEach()` 배열의 `letters` 메서드를 호출하고, 콜백으로 `createLabel` 함수 참조를 전달하세요.

스프레드시트 상단에 몇 개의 글자가 나타나는 것을 볼 수 있을 것입니다.

# --hints--

`.forEach()` 배열에서 `letters` 메서드를 호출해야 합니다.

```js
assert.match(code, /letters\.forEach\(/);
```

`createLabel` 메서드에 `.forEach()` 함수 참조를 전달해야 합니다.

```js
assert.match(code, /letters\.forEach\(\s*(?:\(\s*([\S]*)\s*\)|\s*([\S]*))\s*=>\s*createLabel\(\s*\1\2\s*\)\s*\)|letters\.forEach\(\s*createLabel\s*\)/);
```

`createLabel` 함수 호출을 전달하면 안 됩니다.

```js
assert.notMatch(code, /letters\.forEach\(\s*createLabel\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
--fcc-editable-region--

--fcc-editable-region--
}
```
