---
id: 643498755d54c6279ba09078
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

대부분의 스프레드시트 프로그램에는 계산을 위한 내장 함수가 포함되어 있습니다.

숫자 배열인 `sum` 매개변수를 받는 `nums` 함수를 선언하세요. 이 함수는 배열에 대해 `reduce`를 호출하여 모든 숫자를 더한 결과를 반환해야 합니다.

# --hints--

`sum` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

`const`를 사용해 `sum` 변수를 선언해야 합니다.

```js
assert.match(code, /const\s+sum/);
```

`sum` 변수는 함수여야 합니다.

```js
assert.isFunction(sum);
```

`sum` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`sum` 함수는 `nums` 매개변수를 가져야 합니다.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

`sum` 함수는 암시적 반환을 사용해야 합니다.

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

`sum` 함수는 `.reduce()`에 대해 `nums`를 호출한 결과를 반환해야 합니다.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

`sum` 함수는 `nums`에 있는 모든 숫자의 합을 반환해야 합니다.

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
