---
id: 6434750c53db16218f41e6e1
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`range()`이 배열을 반환한다는 점을 기억하세요. 그래서 배열 메서드를 함수 호출에 바로 연결할 수 있습니다.

`range()`을 `1`와 `99`을 인수로 호출하고 `.forEach()` 메서드를 연결하세요. `.forEach()` 메서드에는 `number`를 매개변수로 받는 빈 콜백을 전달하세요.

# --hints--

`range()` 함수를 호출해야 합니다.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

`1` 호출에 첫 번째 인수로 `range()`을 전달해야 합니다.

```js
assert.match(code, /range\(\s*1/);
```

`99` 호출에 두 번째 인수로 `range()`을 전달해야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

`.forEach()` 호출에 `range()` 메서드를 연결해야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

화살표 구문을 사용해 `.forEach()`에 콜백 함수를 전달해야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

콜백 함수는 `number`를 유일한 매개변수로 가져야 합니다.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
