---
id: 6437133052eaf04d7300e622
title: 22단계
challengeType: 0
dashedName: step-22
---

# --description--

숫자 배열을 `average` 매개변수로 받는 `nums` 함수를 선언하세요. 이 함수는 배열에 있는 모든 숫자의 평균을 반환해야 합니다.

평균은 배열에 있는 모든 숫자의 합을 배열 길이로 나누어 계산할 수 있습니다. 사용할 수 있는 `sum` 함수가 있다는 점을 기억하세요.

# --hints--

`average` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

`const`를 사용해 `average` 변수를 선언하세요.

```js
assert.match(code, /const\s+average/);
```

`average` 변수는 함수여야 합니다.

```js
assert.isFunction(average);
```

`average` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`average` 함수는 `nums` 매개변수를 가져야 합니다.

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

`average` 함수는 암시적 반환을 사용해야 합니다.

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

`average` 함수는 `nums` 배열의 평균값을 반환해야 합니다.

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
