---
id: 6437124c4c03dd4c8fb35d56
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

`isEven` 함수를 선언하세요. 이 함수는 `num` 매개변수를 받고, 숫자가 짝수면 `true`를, 그렇지 않으면 `false`를 반환합니다. 숫자가 짝수인지 홀수인지 판단하려면 나머지 연산자 `%`를 사용하세요.

# --hints--

`isEven` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

`const`를 사용해 `isEven` 변수를 선언하세요.

```js
assert.match(code, /const\s+isEven/);
```

`isEven` 변수는 함수여야 합니다.

```js
assert.isFunction(isEven);
```

`isEven` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

`isEven` 함수는 `num` 매개변수를 가져야 합니다.

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

`isEven` 함수는 나머지 연산자 `%`를 사용해야 합니다.

```js
assert.match(isEven.toString(),  /%/);
```

`isEven` 함수는 짝수일 때 `true`를 반환해야 합니다.

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

`isEven` 함수는 홀수일 때 `false`를 반환해야 합니다.

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
