---
id: 643715013330824ecaa70442
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

다음 함수는 숫자 배열의 중앙값을 계산할 것입니다. `median` 화살표 함수를 선언하고 `nums` 매개변수를 받는 것부터 시작하세요.

함수 안에서 `sorted` 변수를 선언하고 `nums` 배열의 복사본을 정렬한 값을 할당하세요.

배열의 얕은 복사를 위해 `slice()` 메서드를 사용해야 합니다.

# --hints--

`median` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

`const`를 사용해 `median` 변수를 선언하세요.

```js
assert.match(code, /const\s+median/);
```

`median` 변수는 함수여야 합니다.

```js
assert.isFunction(median);
```

`median` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

`median` 함수는 `nums` 매개변수를 가져야 합니다.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

`median` 함수는 암시적 반환을 사용하지 않아야 합니다.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

`median` 함수는 `sorted` 변수를 가져야 합니다.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

`const`를 사용해 `sorted` 변수를 선언하세요.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

`.slice()`를 사용해 `nums` 배열의 복사본을 `sorted`에 할당하세요.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

`.sort()` 메서드를 `.slice()` 메서드에 체인으로 연결하세요.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

`sort` 메서드에 콜백 함수를 전달해 숫자를 오름차순으로 정확히 정렬하세요. 명확성을 위해 암시적 반환을 사용하세요.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
