---
id: 6449749d20436c1f1dfadcf2
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

`window.onload` 함수에서 `input` 요소들이 값이 변경될 때 `update` 함수를 호출하도록 알려야 합니다. `onchange` 속성을 직접 설정해서 이 작업을 할 수 있습니다.

`onchange` 속성을 `update` 함수에 대한 참조로 설정하세요.

# --hints--

`window.onload` 함수는 `onchange` 요소의 `input` 속성에 접근해야 합니다.

```js
assert.match(window.onload.toString(), /input\.onchange/);
```

`window.onload` 함수는 `onchange` 속성을 `update`로 설정해야 합니다.

```js
assert.match(window.onload.toString(), /input\.onchange\s*=\s*update/);
```

`window.onload` 함수는 `update` 함수를 호출하지 않아야 합니다.

```js
assert.notMatch(window.onload.toString(), /update\(\s*\)/);
```

`input` 요소들은 모두 `update` 함수를 `onchange` 속성으로 가져야 합니다.

```js
const inputs = document.querySelectorAll('input');
inputs.forEach(input => {
  assert.property(input, 'onchange');
  assert.equal(input.onchange, update);
})
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

const spreadsheetFunctions = {
  sum,
  average,
  median
}

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
--fcc-editable-region--

--fcc-editable-region--
      container.appendChild(input);
    })
  })
}

const update = event => {

}
```
