---
id: 64497e764135bd24b7960dd3
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

스프레드시트 함수를 실행하려면 입력 문자열을 구문 분석하고 평가할 수 있어야 합니다. 이때 다른 함수를 사용하는 것이 좋습니다.

매개변수 `evalFormula`과 `x`를 받는 `cells` 화살표 함수를 선언하세요.

# --hints--

`evalFormula` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:let|const|var)\s+evalFormula/);
```

`const`를 사용해 `evalFormula` 변수를 선언해야 합니다.

```js
assert.match(code, /const\s+evalFormula/);
```

`evalFormula` 변수는 함수여야 합니다.

```js
assert.isFunction(evalFormula);
```

`evalFormula` 함수는 화살표 구문을 사용해야 합니다.

```js
assert.match(code, /const\s+evalFormula\s*=\s*\(/);
```

`evalFormula` 함수의 첫 번째 매개변수는 `x`여야 합니다.

```js
assert.match(code, /const\s+evalFormula\s*=\s*\(\s*x/);
```

`evalFormula` 함수의 두 번째 매개변수는 `cells`여야 합니다.

```js
assert.match(code, /const\s+evalFormula\s*=\s*\(\s*x\s*,\s*cells\s*\)\s*=>/);
```

`evalFormula` 함수는 비어 있어야 합니다.

```js
assert.match(code, /const\s+evalFormula\s*=\s*\(\s*x\s*,\s*cells\s*\)\s*=>\s*{\s*}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

const spreadsheetFunctions = {
  sum,
  average,
  median
}

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      input.onchange = update;
      container.appendChild(input);
    })
  })
}

const update = event => {
  const element = event.target;
  const value = element.value.replace(/\s/g, "");
  if (!value.includes(element.id) && value.startsWith('=')) {

  }
}
```
