---
id: 64496df724dd3716a71fe971
title: 26단계
challengeType: 0
dashedName: step-26
---

# --description--

객체 속성은 키/값 쌍으로 구성됩니다. 객체 리터럴을 선언할 때 단축 속성 이름을 사용할 수 있습니다. 단축 속성 이름 구문을 사용할 때 변수 이름이 속성 키가 되고 그 값이 속성 값이 됩니다.

다음 예제는 `user` 객체를 `userId`, `firstName`, `loggedIn` 속성으로 선언합니다.

```js
const userId = 1;
const firstName = "John";
const loggedIn = true;

const user = {
  userId,
  firstName,
  loggedIn,
};

console.log(user); // { userId: 1, firstName: 'John', loggedIn: true }
```

스프레드시트의 모든 함수를 추적하려면 `spreadsheetFunctions` 객체를 선언하세요. 단축 표기법 구문을 사용하여 `sum`, `average`, `median`을 `spreadsheetFunctions` 객체의 속성으로 설정하세요.

# --hints--

`spreadsheetFunctions` 변수를 선언해야 합니다.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

`const`을 사용해 `spreadsheetFunctions` 변수를 선언해야 합니다.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

`spreadsheetFunctions` 변수는 객체여야 합니다.

```js
assert.isObject(spreadsheetFunctions);
```

`spreadsheetFunctions` 객체는 `sum` 속성을 가져야 합니다.

```js
assert.property(spreadsheetFunctions, "sum");
```

`sum` 속성은 `sum` 함수여야 합니다.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

`spreadsheetFunctions` 객체는 `average` 속성을 가져야 합니다.

```js
assert.property(spreadsheetFunctions, "average");
```

`average` 속성은 `average` 함수여야 합니다.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

`spreadsheetFunctions` 객체는 `median` 속성을 가져야 합니다.

```js
assert.property(spreadsheetFunctions, "median");
```

`median` 속성은 `median` 함수여야 합니다.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

구조 분해 구문을 사용해 속성을 할당해야 합니다.

```js
const objectText = code.replace(/.*const\s+spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
