---
id: 65ae458e23954c3469e0c209
title: 18단계
challengeType: 0
dashedName: step-18
---


# --description--

이전 프로젝트에서 `setAttribute` 메서드에 대해 배웠습니다. JavaScript에서 속성을 업데이트하는 또 다른 방법은 다음 구문을 사용하는 것입니다:

```js
el.attribute = value;
```

하이픈으로 연결된 HTML 속성 값, 예를 들어 `aria-label`의 속성 이름은 카멜 케이스를 따라 `ariaLabel`가 됩니다.

```js
el.ariaLabel = "Aria Label Value";
```

`aria-label` 요소의 `input` 속성을 `id` 속성과 같은 값으로 설정하세요.


# --hints--

`input.ariaLabel`가 있어야 합니다.

```js
assert.match(code, /input\.ariaLabel\s*=/);
```


`input.ariaLabel`에 `input.id`과 같은 값을 주어야 합니다.

```js
assert.match(code, /input\.ariaLabel\s*=\s*((letter\s*\+\s*number)|(input.id))/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
     createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
     const input = document.createElement("input");
     input.type = "text";
     input.id = letter + number;
 
    })   
--fcc-editable-region--
  })
}
```
