---
id: 60a3e3396c7b40068ad6996c
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在之前关于 CSS 盒模型的课程中，你已经了解到每个超文本标记语言元素都被视为具有四个区域的盒子。

想象一下，你从最喜欢的在线零售商那里收到一个盒子——内容就是盒子里的商品，也就是你的标题、段落或图片元素。

将 `<img>` 中的 `src` 属性从 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 更改为 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`。

# --hints--

`img` 元素的 `src` 应该是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

