---
id: 60a3e3396c7b40068ad6998c
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

你的画的颜色和形状太锐利了，不像是罗斯科画。

在 `.canvas` 元素中使用 `filter` 属性，将画作 `blur` `2px`。

下面是一个添加 3px `blur` 规则的示例：

```css
p {
  filter: blur(3px);
}
```

# --hints--

你应该将 `filter` 属性设置为 `blur(2px)`。

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(2px)');
assert(hasFilter);
```

你的 `.canvas` 元素应该有一个 `filter`，值为 `blur(2px)`。

```js
const canvasFilter = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('filter');
assert(canvasFilter === 'blur(2px)');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;

}
--fcc-editable-region--

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
