---
id: 60a3e3396c7b40068ad69975
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

用 `.frame` 类选择器编写新一个规则。

用 `border` 简写声明为 `.frame` 元素赋予一个宽度为 `50px` 的实线黑色边框。

# --hints--

你的代码应有一个 `.frame` 选择器。

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

你应将 `border` 属性设置为 `50px solid black`。

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

你的 `.frame` 元素应有一个 `border` 属性，值为 `50px solid black`。

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
