---
id: 60a3e3396c7b40068ad6996d
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

内容周围有一个称为内边距的空间，类似于气泡膜将物品与包装盒分隔开来。

把边框想象成装运物品的纸箱。

将 `src` 属性更改为 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`。

# --hints--

`img` 元素的 `src` 应该是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
