---
id: 60a3e3396c7b40068ad69974
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

每幅画都需要一个框架。

将 `.canvas` 元素包装在另一个 `div` 中。 为 `div` 赋予 `frame` 类。
# --hints--

你应添加一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 2)
```

你的 `.canvas` 元素应嵌套在新的 `div` 元素中。

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

你的新 `div` 应有一个 `class`，值为 `frame`。

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

你的新 `div` 应在 `body` 元素中。

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
