---
id: 66a37f37ef5823a313de8c26
title: Побудуйте орбіту Місяця
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

У цій лабораторній роботі ви створите просту анімацію обертання Місяця навколо Землі за допомогою HTML і CSS. Земля буде центром системи, а Місяць обертатиметься навколо неї.

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Вирівняйте всі елементи за центром сторінки, встановивши висоту на `100%` відносно вікна перегляду і задавши макет флексбокса в елементі `body`.

2. Ви повинні мати елемент `div` з класом `space`.

3. У межах елемента `.space` має бути ще два елементи `div` з класами `earth` та `orbit` (у такому ж порядку).

4. У межах елемента `.orbit` має бути ще один елемент `div` з класом `moon`.

5. Елемент `div` з `class` зі значенням `space` має бути розташованим по центру сторінки, а також мати ширину і висоту зі значенням `200px`.

6. Елемент `div` з `class` зі значенням `space` має використовувати позиціювання `relative`.

7. Елемент `.earth` має використовувати позиціювання `absolute`. Розташуйте центр елемента `.earth` посередині його батьківського елемента на вертикальній (зверху) і горизонтальній (зліва) осях. Після цього змістіть елемент `.earth` назад на півширини і піввисоти, щоб відцентрувати його в межах батьківського елемента `.space`.

8. Елемент `.earth` повинен мати ширину та висоту зі значенням `100px`.

9.  Елемент `.orbit` повинен мати ширину та висоту зі значенням `200px`.

10.  Елемент `.orbit` має використовувати позиціювання `absolute`. Правий нижній кут має бути в центрі елемента `.space`, тому використайте властивість `transform`, яка зміщує його на `-50%` як по вертикальній, так і по горизонтальній осях.

11. Орбітою Місяця навколо Землі має бути коло.

12. Елемент `.moon` має використовувати позиціювання `absolute`, а також мати ширину і висоту зі значенням `30px`. Елемент `.moon` має розташовуватися у верхній частині елемента `.orbit` і бути горизонтально відцентрованим.

13. Далі потрібно відрегулювати горизонтальне позиціювання елемента `.moon`, перемістивши його вліво на півширини.

14. Елементи `.earth` та `.moon` повинні мати колір фону, а також `border-radius` зі значенням `50%`, щоб виглядати як планети.

15. Визначте анімацію `@keyframes orbit`, що обертатиме елемент `.orbit` на 360 градусів навколо центру. Застосуйте цю анімацію до елемента `.orbit`, задавши тривалість `5` секунд, лінійну часову функцію та нескінченну кількість ітерацій.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Вміст елемента `body` потрібно вирівняти за центром сторінки, встановивши висоту на `100%` відносно висоти вікна перегляду, `display` на `flex`, а також `justify-content` й `align-items` на `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Ви повинні мати основний елемент `div` з класом `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

Елемент `.space` повинен мати ширину та висоту зі значенням `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

Елемент `.space` повинен мати властивість `position` зі значенням `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

У межах `div` з `class` зі значенням `space` має бути `div` з класом `earth`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

Елемент `.earth` має бути першим дочірнім елементом елемента `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

Елемент `.earth` повинен мати ширину та висоту зі значенням `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

Елемент `.earth` має використовувати позиціювання `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

Елемент `.earth` повинен мати позиціювання `top` та `left` зі значенням `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

Елемент `.earth` потрібно перемістити назад і вгору на `50%`, встановивши властивість `transform` на `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

У межах елемента `.space` має бути `div` з класом `orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

Елемент `div` з класом `orbit` має бути після елементу `div` з класом `earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

Елемент `.orbit` повинен мати ширину та висоту зі значенням `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

Елемент `.orbit` має використовувати позиціювання `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

Елемент `.orbit` має бути в центрі елемента `.space`, тому використайте властивість `transform`, яка зміщує його на `-50%` як по вертикальній, так і по горизонтальній осях.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

У межах елемента `.orbit` має бути `div` з класом `moon`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

Елемент `.moon` повинен мати ширину та висоту зі значенням `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

Елемент `.moon` має використовувати позиціювання `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

Верхній край елемента `.moon` має бути вирівняним відносно верхнього краю відповідного блоку.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Перемістіть лівий край елемента `.moon` до центру відповідного блоку (`50%` з лівої сторони).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

Елемент `.moon` потрібно відрегулювати по горизонталі, перемістивши його на `-50%` по осі Х.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.transform, 'translateX(-50%)');
```

Елемент `.earth` повинен мати колір фону.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Елемент `.earth` повинен мати `border-radius` зі значенням `50%`, щоб виглядати круглим.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

Елемент `.moon` повинен мати колір фону.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Елемент `.moon` повинен мати `border-radius` зі значенням `50%`, щоб виглядати круглим.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Ви повинні мати правило `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

Назвою нового правила `@keyframes` має бути `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

Правило `@keyframes orbit` повинне мати селектор `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

Правило `@keyframes orbit` повинне мати селектор `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Селектор `0%` повинен мати властивість `transform` зі значенням `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Селектор `100%` має бути після селектора `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Селектор `100%` повинен мати властивість `transform` зі значенням `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Селектор `.orbit` повинен мати властивість `animation` зі значенням `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

