---
id: 5a858944d96184f06fd60d61
title: Створіть свою першу сітку CSS
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

Перетворіть будь-який HTML-елемент на контейнер сітки, встановивши його властивість `display` у значення `grid`. Це дасть змогу використовувати всі інші властивості, пов’язані з CSS Grid.

**Примітка:** У CSS Grid батьківський елемент називають <dfn>контейнером</dfn>, а його дочірні елементи — <dfn>елементами</dfn>.

# --instructions--

Змініть відображення div з класом `container` на `grid`.

# --hints--

Клас `container` повинен мати властивість `display` зі значенням `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
