---
id: 5a9036ee38fddaf9a66b5d34
title: Використання одиниць CSS Grid для зміни розміру колонок і рядків
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

Ви можете використовувати абсолютні та відносні одиниці, такі як `px` і `em`, у CSS Grid для визначення розміру рядків і колонок. Також можна використовувати:

`fr`: задає колонку або рядок як дріб від доступного простору,

`auto`: автоматично встановлює ширину або висоту колонки чи рядка відповідно до їхнього вмісту,

`%`: регулює колонку або рядок відповідно до відсоткової ширини контейнера.

Ось код, який генерує вихідні дані у прев’ю:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

Цей фрагмент створює п’ять колонок. Перша колонка має ширину, що відповідає її вмісту, друга — 50px, третя — 10% від ширини контейнера, а для останніх двох колонок залишковий простір ділиться на три частини: дві виділяються четвертій колонці, а одна — п’ятій.

# --instructions--

Створіть сітку з трьома колонками, ширини яких такі: 1fr, 100px і 2fr.

# --hints--

Ваш клас `container` має мати властивість `grid-template-columns`, яка задає три колонки з такими ширинами: `1fr`, `100px` і `2fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
