---
id: 5a9036ee38fddaf9a66b5d36
title: Створіть горизонтальний розрив між рядками за допомогою grid-row-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cPbJ2Cv'
forumTopicId: 301125
dashedName: create-a-row-gap-using-grid-row-gap
---

# --description--

Ви можете додати горизонтальний розрив між рядками сітки, використовуючи `grid-row-gap`, так само, як ви додавали розрив між колонками в попередньому завданні.

# --instructions--

Створіть горизонтальний розрив між рядками висотою `5px`.

# --hints--

Клас `container` має містити властивість `grid-row-gap` зі значенням `5px`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-row-gap\s*?:\s*?5px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-row-gap: 5px;
  }
</style>
```
