---
id: 5a90374338fddaf9a66b5d3a
title: Вирівняти елемент по горизонталі за допомогою justify-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

У CSS сітці вміст кожного елемента розташований у блоці, який називають <dfn>клітинкою</dfn>. Ви можете вирівняти положення вмісту всередині клітинки по горизонталі, використовуючи властивість `justify-self` для елемента сітки. За замовчуванням ця властивість має значення `stretch`, що змушує вміст заповнювати всю ширину клітинки. Ця властивість CSS сітки також приймає інші значення:

`start`: вирівнює вміст по лівому краю клітинки,

`center`: вирівнює вміст по центру клітинки,

`end`: вирівнює вміст по правому краю клітинки.

# --instructions--

Використайте властивість `justify-self`, щоб по центру вирівняти елемент з класом `item2`.

# --hints--

Клас `item2` має мати властивість `justify-self` зі значенням `center`.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
