---
id: 5a9036d038fddaf9a66b5d32
title: додати колонки за допомогою grid-template-columns
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

Просто створити елемент сітки недостатньо. Потрібно також визначити структуру сітки. Щоб додати колонки до сітки, використовуйте властивість `grid-template-columns` на контейнері сітки, як показано нижче:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

Це створить сітку з двома колонками, кожна шириною 50px. Кількість параметрів, переданих властивості `grid-template-columns`, визначає кількість колонок у сітці, а значення кожного параметра вказує ширину відповідної колонки.

# --instructions--

Задайте контейнеру сітки три колонки, кожна шириною `100px`.

# --hints--

Клас `container` повинен мати властивість `grid-template-columns` з трьома одиницями `100px`.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
