---
id: 5a9036ee38fddaf9a66b5d35
title: Створіть горизонтальний розрив між колонками за допомогою grid-column-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cVZ8vfD'
forumTopicId: 301124
dashedName: create-a-column-gap-using-grid-column-gap
---

# --description--

До цього часу в сітках, які ви створювали, колонки були щільно прилягаючими одна до одної. Іноді потрібно додати горизонтальний розрив між колонками. Щоб додати розрив між колонками, використовуйте властивість `grid-column-gap` ось так:

```css
grid-column-gap: 10px;
```

Це створює 10px порожнього простору між усіма нашими колонками.

# --instructions--

Додайте колонкам у сітці горизонтальний розрив `20px`.

# --hints--

Клас `container` має містити властивість `grid-column-gap` зі значенням `20px`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-column-gap\s*?:\s*?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

   
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-column-gap: 20px;
  }
</style>
```
