---
id: 5a90376038fddaf9a66b5d3c
title: Вирівняти всі елементи по горизонталі за допомогою justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

Іноді потрібно, щоб усі елементи у вашій CSS сітці мали однакове вирівнювання. Ви можете використати раніше вивчені властивості та вирівнювати їх поодинці, або ж вирівняти всі одразу по горизонталі, використовуючи `justify-items` на вашому контейнері сітки. Ця властивість приймає всі ті ж значення, що ви вивчали в попередніх двох завданнях, але різниця в тому, що вона перемістить **всі** елементи нашої сітки до потрібного вирівнювання.

# --instructions--

Використайте цю властивість, щоб по центру вирівняти всі наші елементи по горизонталі.

# --hints--

`container` клас має мати властивість `justify-items` зі значенням `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
