---
id: 5a94fe6269fb03452672e462
title: Створення гнучких макетів за допомогою auto-fit
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c3dPph8'
forumTopicId: 301127
dashedName: create-flexible-layouts-using-auto-fit
---

# --description--

`auto-fit` працює майже ідентично `auto-fill`. Єдина різниця в тому, що коли розмір контейнера перевищує сумарний розмір усіх елементів, `auto-fill` продовжує додавати порожні рядки або стовпці і відсуває ваші елементи вбік, тоді як `auto-fit` згортає ці порожні рядки або стовпці і розтягує ваші елементи, щоб вони відповідали розміру контейнера.

**Примітка:** Якщо ваш контейнер не може вмістити всі елементи в одному рядку, він перемістить їх на новий рядок.

# --instructions--

У другій сітці використайте `auto-fit` з `repeat`, щоб заповнити сітку стовпцями з мінімальною шириною `60px` і максимальною `1fr`. Потім змініть розмір попереднього перегляду, щоб побачити різницю.

# --hints--

Клас `container2` має мати властивість `grid-template-columns` зі значеннями `repeat` і `auto-fit`, які заповнять сітку стовпцями з мінімальною шириною `60px` і максимальною `1fr`.

```js
assert.match(
  code,
  /.container2\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fit\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
  }
  .container2 {
    grid-template-columns: repeat(auto-fit, minmax(60px, 1fr));
  }
</style>
```
