---
id: 5a94fe3669fb03452672e45f
title: Зменшення повторень за допомогою функції repeat
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

Коли ви використовували `grid-template-columns` і `grid-template-rows` для визначення структури сітки, ви вводили значення для кожного рядка або стовпця, який створювали.

Припустімо, ви хочете сітку з 100 рядками однакової висоти. Не дуже практично вводити 100 значень по одному. На щастя, є кращий спосіб — використати функцію `repeat`, щоб вказати, скільки разів ви хочете повторити свій стовпець або рядок, потім через кому — значення, яке потрібно повторити.

Ось приклад, який створить сітку з 100 рядків, кожен висотою 50px.

```css
grid-template-rows: repeat(100, 50px);
```

Ви також можете повторювати кілька значень за допомогою функції repeat і вставляти цю функцію серед інших значень при визначенні структури сітки. Ось як це виглядає:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Це означає:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Примітка:** `1fr 50px` повторюється двічі, після чого йде 20px.

# --instructions--

Використайте `repeat`, щоб усунути повторення у властивості `grid-template-columns`.

# --hints--

Клас `container` має мати властивість `grid-template-columns`, яка встановлена на повторення 3 стовпців шириною `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
