---
id: 5a94fe1369fb03452672e45d
title: Розміщення елементів у областях сітки за допомогою властивості grid-area
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

Після створення шаблону областей для вашого контейнера сітки, як показано в попередньому завданні, ви можете розмістити елемент у кастомній області, посилаючись на ім’я, яке ви їй дали. Для цього використовуйте властивість `grid-area` на елементі ось так:

```css
.item1 {
  grid-area: header;
}
```

Це дає сітці знати, що ви хочете, щоб клас `item1` розмістився в області з ім’ям `header`. У цьому випадку елемент займе весь верхній ряд, бо весь цей ряд названий областю `header`.

# --instructions--

Розмістіть елемент з класом `item5` в області `footer` за допомогою властивості `grid-area`.

# --hints--

Клас `item5` має мати властивість `grid-area` зі значенням `footer`.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
