---
id: 5a94fe2669fb03452672e45e
title: Використання grid-area без створення шаблону областей
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

Властивість `grid-area`, яку ви вивчили в останньому завданні, можна використовувати інакше. Якщо у вашій сітці немає шаблону областей для референсу, ви можете створити область на льоту, щоб розмістити елемент ось так:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

Тут використані номери ліній, про які ви дізналися раніше, щоб визначити, де буде область для цього елемента. Числа в наведеному прикладі означають такі значення:

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

Отже, елемент у прикладі займатиме рядки між лініями 1 і 2, а також стовпці між лініями 1 і 4.

# --instructions--

Використовуючи властивість `grid-area`, розмістіть елемент з класом `item5` між третьою та четвертою горизонтальними лініями і між першою та четвертою вертикальними лініями.

# --hints--

Клас `item5` повинен мати властивість `grid-area`, щоб заповнити всю область між третьою та четвертою горизонтальними лініями, а також першою та четвертою вертикальними лініями.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
