---
id: 66a8380d911e3f4270d5cadc
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Agora, é hora de adicionar o elemento `main`, que representa o conteúdo principal da página.

# --hints--

Você deve ter uma tag de abertura para `main`.

```js
assert.match(code, /<main>/i);
```

Você deve acrescentar uma tag de fechamento para `main`.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
