---
id: 66a84111965a0c46df6bbd0a
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

In the previous lectures, you learned how to associate a `label` element with an `input` like this:

```html
<label for="name">Name:</label>
<input type="text" id="name" name="name">
```

O atributo `for` no elemento `label` deve corresponder ao atributo `id` no elemento `input`. Isso é conhecido como uma associação explícita.

Below your `legend` element, add a `label` element with the text of `Name (required):`. Set its `for` attribute to the value of `"full-name"`.

Then below your `label` element, add an `input` element with no attributes. Nos próximos passos, você adicionará os atributos necessários.

# --hints--

Você deve ter uma tag de abertura para `label`.

```js
assert.match(code, /<label.*>/);
```

Você deve ter uma tag de fechamento para `label`.

```js
assert.match(code, /<\/label>/);
```

Your `label` element should be placed below your `legend` inside the `fieldset` element.

```js
assert.isNotNull(document.querySelector('fieldset legend + label'));
```

O elemento `label` deve ter um atributo `for` com o valor `"full-name"`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.getAttribute('for'), 'full-name');
```

Your `label` element should have the text of `Name (required):`.

```js
const label = document.querySelector('fieldset legend + label');
assert.strictEqual(label?.innerText, 'Name (required):');
```

You should have an `input` element below your `label` element.

```js
assert.isNotNull(document.querySelector('fieldset label + input'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
--fcc-editable-region--
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>

        </fieldset>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>
```
