---
id: 66a8290a27c2c625e2355042
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá a trabalhar com formulários criando um formulário de feedback de hotéis.

Comece adicionando o `<!DOCTYPE html>` seguido de um elemento `html` com um atributo `lang` de `en`.

Dentro do elemento `html`, adicione um elemento `head`.


# --hints--

Você deve ter a declaração `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter uma tag de abertura para `html` com o idioma definido para inglês.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve acrescentar uma tag de fechamento para `html`.

```js
assert.match(code, /<\/html>/i);
```

O `DOCTYPE` deve vir antes do elemento `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de fechamento para `head`.

```js
assert.match(code, /<\/head>/i);
```

A tag de abertura de `head` deve vir antes da tag de fechamento de `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

O elemento `head` deve estar dentro do elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
