---
id: 66a8290a27c2c625e2355042
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will learn how to work with HTML forms by building a Hotel Feedback Form.

Start by adding the `<!DOCTYPE html>` followed by an `html` element with a `lang` attribute of `en`.

Inside your `html` element, add a `head` element.


# --hints--

You should have the `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

You should have an opening `html` tag with the language set to english.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve acrescentar uma tag de fechamento para `html`.

```js
assert.match(code, /<\/html>/i);
```

Your `DOCTYPE` should come before the `html` element.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de fechamento para `head`.

```js
assert.match(code, /<\/head>/i);
```

Your opening `head` tag should come before the closing `head` tag.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

Your `head` element should be inside the `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
