---
id: 66a962954f4e0d76223b37ed
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Em seguida, adicione outro elemento `label` com o texto `Location` e o atributo `for` definido como `"location"`.

Para o `input` de caixa de seleção, os atributos `id`, `name` e `value` devem ser definidos como `"location"`.

Below that `input` element, add another `label` element with the text of `Reputation` and the `for` attribute set to `"reputation"`.

Para o `input` de caixa de seleção, os atributos `id`, `name` e `value` devem ser definidos como `"reputation"`.

# --hints--

Você deve ter um elemento `label` com o texto `Location`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.textContent, 'Location');
```

O elemento `label` deve ter o atributo `for` definido como `"location"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.getAttribute('for'), 'location');
```

You should have a third checkbox `input` after your `label`.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(3)[type='checkbox']"));
```

O seu terceiro input de ciaxa de seleção deve ter um `id` definido como `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('id'), 'location');
```

A terceira caixa de seleção deve ter um `name` de `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('name'), 'location');
```

A terceira caixa de seleção deve ter um `value` de `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('value'), 'location');
```

You should have a `label` element with the text of `Reputation` below the location checkbox.

```js
assert.isNotNull(document.querySelectorAll('fieldset:nth-of-type(3) input[value="location"] + label')?.textContent, 'Reputation');
```

O elemento `label` deve ter o atributo `for` definido como `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'reputation');
```

You should have a fourth checkbox `input` below your `label`.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(4)[type='checkbox']"));
```

Você deve ter um quarto `input` de caixa de seleção com um `id` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('id'), 'reputation');
```

A quarta caixa de seleção deve ter um `name` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('name'), 'reputation');
```

A quarta caixa de seleção deve ter um `value` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('value'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
