---
id: 66a94ea5df66236ebed260e8
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

A próxima seção no formulário será responsável por perguntar aos usuários se eles já se hospedaram no hotel antes.

Comece adicionando um elemento `fieldset`.

Dentro do elemento `fieldset`, adicione um elemento `legend` com o texto `Was this your first time at our hotel?`.

# --hints--

Você deve ter um elemento `fieldset`.

```js
assert.strictEqual(document.querySelectorAll('fieldset').length, 2);
```

Você deve ter um elemento `legend` dentro do elemento `fieldset`.

```js
assert.strictEqual(document.querySelectorAll('legend').length, 2);
```

O elemento `legend` deve ter o texto de `Was this your first time at our hotel?`.

```js
assert.strictEqual(document.querySelectorAll('legend')[1].textContent, 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--

      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
