---
id: 66a93c95bc58e26a8fe95818
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Next, add an `input` with the type of `"email"` below your email `label`. Este `input` deve ter um atributo `id` definido como o valor de `"email"`. Além disso, defina o atributo `name` com o valor de `"email"`.

Esse `input` também é obrigatório. Então, certifique-se de adicionar o atributo `required`.

Por fim, adicione um atributo `placeholder` definido como `"example@email.com"`.

# --hints--

You should have an `input` element below your email `label`.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

O elemento `input` deve ter um `type` de `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

O elemento `input` deve ter um atributo `id` definido como `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

O elemento `input` deve ter um atributo `name` definido como `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

O elemento `input` deve ter o atributo `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

O elemento `input` deve ter um atributo `placeholder` definido como `"example@email.com"`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address (required):</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
