---
id: 66a9401c9d660d6bb15993e2
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

O formulário de feedback de hotéis deve ter uma opção para os usuários adicionarem a sua idade.

Start by adding a `label` element with the text `Age (optional):` to the form.

O atributo `for` deve estar definido como `"age"`.

# --hints--

Você deve ter um elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

O elemento `label` deve ter o atributo `for` definido como `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

Your `label` element should have the text `Age (optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent, "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--

--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
