---
id: 66a954b2bcddba72076c1857
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Abaixo do primeiro botão de opção, adicione outro elemento `label` com o atributo `for` definido como `"no-option"`. O texto para o `label` deve ser `No`.

Abaixo do segundo elemento `label`, adicione um botão `radio` com o `id` definido como `"no-option"` e o atributo `name` definido como `"hotel-stay"`.

Quando você terminar, poderá testar os botões de opção selecionando uma opção por vez.

# --hints--

You should have a `label` element below your first radio button.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

O elemento `label` deve ter um atributo `for` definido como `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, "no-option");
```

O elemento `label` deve ter o texto `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label[for="no-option"]')?.textContent, 'No');
```

You should have a second `input` below your second `label` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) label + input'), 2);
```

O elemento `input` deve ter um `type` `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.getAttribute('type'), 'radio');
```

O elemento `input` deve ter um `id` de `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.id, 'no-option');
```

O elemento `input` deve ter um `name` de `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
