---
id: 66a969951120be7818d8ee49
title: Passo 26
challengeType: 0
dashedName: step-26
---

# --description--

A próxima seção do formulário fornecerá aos usuários a capacidade de deixar uma avaliação para o hotel.

Comece adicionando um elemento `fieldset` com um elemento `legend` aninhado dentro dele. A `legend` deve ter o texto `Ratings`.

Abaixo do elemento `legend`, adicione um elemento `label` com o texto `How was the service?`. O atributo `for` deve ser configurado como `"service"`.

# --hints--

Você deve ter um elemento `fieldset`.

```js
assert.strictEqual(document.querySelectorAll('fieldset').length, 4);
```

O `fieldset` deve ter um elemento `legend`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend').length, 1);
```

A legenda deve ter o texto `Ratings`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0].textContent, 'Ratings');
```

O `fieldset` deve ter um elemento `label`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) label').length, 1);
```

O `label` deve ter o texto `How was the service?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) label')[0].textContent, 'How was the service?');
```

O `label` deve ter o atributo `for` configurado como `"service"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) label')[0].getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />

          <label for="location">Location</label>
          <input type="checkbox" id="location" name="location" value="location" />

          <label for="reputation">Reputation</label>
          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />

          <label for="price">Price</label>
          <input type="checkbox" id="price" name="price" value="price" />
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
