---
id: 66a95d0eff8168747805f1f3
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Quando você deseja que um usuário selecione múltiplas opções de uma lista, você pode usar caixas de seleção.

Aqui está um exemplo de como trabalhar com caixas de seleção lidando com opções de alimentos:

```html
<fieldset>
  <legend>Food Options</legend>
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="burger">Burger</label>
  <input type="checkbox" id="burger" name="food" value="burger">
</fieldset>
```

O atributo `value` é usado para especificar o valor que o servidor receberá quando o formulário for enviado.

Abaixo do elemento `legend`, adicione um elemento `label` com o texto `Social Media Ads`. O atributo `for` deve estar definido como `"ads"`.

Abaixo do elemento `label`, adicione um input de caixa de seleção com os atributos `id`, `name` e `value` definidos como `"ads"`.

# --hints--

You should have a `label` element below your `legend`.

```js
assert(document.querySelector("fieldset:nth-of-type(3) legend + label"));
```

O elemento `label` deve ter o texto de `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.textContent, "Social Media Ads");
```

O elemento `label` deve ter um atributo `for` definido como `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + label")?.getAttribute("for"), "ads");
```

You should have a checkbox `input` below your `label`.

```js
const inputElement = document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']");
assert.strictEqual(inputElement?.parentElement.tagName, "FIELDSET");
```

A caixa de seleção deve ter um atributo `id` definido como `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.id, "ads");
```

A caixa de seleção deve ter um atributo `name` definido como `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.name, "ads");
```

A caixa de seleção deve ter um atributo `value` definido como `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) label + input[type='checkbox']")?.value, "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
