---
id: 66a9401c9d660d6bb15993e2
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

O formulário de feedback de hotéis deve ter uma opção para os usuários adicionarem a sua idade.

Comece adicionando um elemento `label` com o texto `Age(optional):` ao formulário.

O atributo `for` deve estar definido como `"age"`.

# --hints--

Você deve ter um elemento `label`.

```js
assert.strictEqual(document.querySelectorAll('fieldset label').length, 3);
```

O elemento `label` deve ter o atributo `for` definido como `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

O elemento `label` deve ter o texto `Age(optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]').textContent, "Age(optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
        --fcc-editable-region--

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
