---
id: 66a9521bc70162712caf118d
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

Se você deseja que os usuários selecionem uma opção de uma lista de opções, você pode usar um conjunto de botões de opção (ou radio buttons, em inglês).

Aqui vemos um exemplo de dois botões de seleção:

```html
<input type="radio" id="yes" name="first-time">
<label for="yes">Yes</label>
<input type="radio" id="no" name="first-time">
<label for="no">No</label>
```

Neste exemplo, os botões de seleção são agrupados usando o mesmo valor do atributo `name`. Isso significa que apenas um botão de opção pode ser selecionado por vez.

Abaixo do elemento `legend`, adicione um elemento `label` com o texto `Yes` e um atributo `for` definido como `"yes-option"`.

Abaixo do elemento `label`, adicione um botão `radio` com o `id` definido como `"yes-option"` e com o atributo `name` definido como `"hotel-stay"`.

# --hints--

Você deve ter um elemento `label` dentro do elemento `fieldset`.

```js
assert(document.querySelector('fieldset:nth-of-type(2) label'));
```

O elemento `label` deve ter o atributo `for` definido como `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label').htmlFor, 'yes-option');
```

O elemento `label` deve ter o texto `Yes`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label').innerText, 'Yes');
```

Você deve ter um botão `radio` dentro do elemento `fieldset`.

```js
assert(document.querySelector('fieldset:nth-of-type(2) input[type="radio"]'));
```

O botão `radio` deve ter um atributo `id` definido como `"yes-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"]').id, 'yes-option');
```

O botão `radio` deve ter um atributo `name` definido como `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"]').name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        --fcc-editable-region--
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
