---
id: 66a930b20f589b6664c51cb0
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Quando um usuário fornece seu nome completo, o `input` aceitará texto simples.

Nos vídeos da aula anterior, você aprendeu a trabalhar com o atributo `type` assim:

```html
<input type="text">
```

Para o elemento `input` existente, adicione um atributo `type` definido como `"text"`.

Também dê ao elemento `input` um atributo `id` com o valor de `"full-name"`.

# --hints--

Você não deve remover o elemento `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Você não deve criar um elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

O elemento `input` deve ter um atributo `type` com o valor `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

O elemento `input` deve ter um atributo `id` com o valor de `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
