---
id: 65df3afd233057f6a620a860
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Це показує, як допоміжні функції можна використовувати для аналізу визначень класів учня

# --instructions--

Ми можемо просто знайти визначення класу в Python, повернути його до JavaScript, а потім перевірити, чи містить воно рядок "pass".

Якщо нам потрібне щось більш надійне, ми можемо розширити помічники, щоб спеціально перевіряти по інструкції пропуску. Це було б корисно, скажімо, якщо ми хочемо написати класи, що містять, скажімо, пасажирів.

# --hints--

`TreeNode` не повинен містити жодних інструкцій `pass`.

```js
({ test: () => {
  const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
  assert.notInclude(pyClassStr, "pass")
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass

class Trap:
    pass
--fcc-editable-region--
```

# --solutions--

```py
class TreeNode:
    def __init__(self, thing):
        self.thing = thing

class Trap:
    pass

```
