---
id: 5d822fd413a79914d39e98d4
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Assegna alle parti del tuo edificio `width` e `height` le proprietĂ  con questi valori: `70%` e `10%` a `.bb1a`, `80%` e `10%` a `.bb1b`, `90%` e `10%` a `.bb1c`, e `100%` e `70%` a `.bb1d`. Ricorda che queste percentuali sono relative al genitore e nota che le altezze si sommeranno a 100% - riempiendo verticalmente il contenitore.

# --hints--

Dovresti usare un selettore di classe per definire lo stile di `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Dovresti dare a `.bb1a` una `width` di `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Dovresti dare a `.bb1a` una `height` di `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Dovresti usare un selettore di classe per definire lo stile di `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Dovresti dare a `.bb1b` una `width` di `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Dovresti dare a `.bb1b` una `height` di `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Dovresti usare un selettore di classe per definire lo stile di `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Dovresti dare a `.bb1c` una `width` di `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Dovresti dare a `.bb1c` una `height` di `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Dovresti usare un selettore di classe per definire lo stile di `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Dovresti dare a `.bb1d` una `width` di `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Dovresti dare a `.bb1d` una `height` di `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```
