---
id: 5d822fd413a79914d39e98cc
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno del `head`, annida un elemento `meta` con un `charset` di `UTF-8`, un elemento `title` con un titolo `City Skyline` e un elemento `link` che collega il tuo file `styles.css`.

# --hints--

Dovresti creare un elemento `meta` all'interno dell'elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('meta'));
```

Dovresti dare al tag `meta` un `charset` di `UTF-8`.

```js
const linkElement = document.querySelector('meta');
const charset = linkElement?.getAttribute("charset").toLowerCase();
assert.strictEqual(charset, 'utf-8');
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

L'elemento `title` dovrebbe essere all'interno dell'elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

Il tuo progetto dovrebbe avere un titolo `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Ricorda, la maiuscola e la minuscola e l'ortografia sono importanti per il titolo.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Il tuo codice dovrebbe avere un elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Non dovresti modificare i tuoi tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

Dovresti avere un elemento `link` void.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
