---
id: 5d822fd413a79914d39e98c9
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Benvenuto al progetto CSS Variables Skyline! Inizia aggiungendo la dichiarazione `!DOCTYPE html` in cima al documento così il browser sa che tipo di documento sta leggendo.

# --hints--

Il tuo codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un `>` dopo il tipo.

```js
assert(code.match(/html\s*>/gi));
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
