---
id: 5d822fd413a79914d39e98d1
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Dai al tuo elemento `.background-buildings` una `width` e una `height` di `100%` per farlo occupare tutta la larghezza e l'altezza del suo genitore, il `body`.

# --hints--

Dovresti usare la classe `background-buildings` per selezionare l'elemento corretto.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Il tuo elemento `.background-buildings` dovrebbe avere una `width` di `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Il tuo elemento `.background-buildings` dovrebbe avere una `height` di `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```
