---
id: 5d822fd413a79914d39e98cd
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

In CSS puoi selezionare tutto con un asterisco. Aggiungi un bordo a tutto usando il selettore `*` e impostando la proprietà `border` su `1px solid black`. Questo trucco aiuta a visualizzare dove si trovano gli elementi e le loro dimensioni. Lo rimuoverai più avanti.

# --hints--

Dovresti usare il selettore `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Dovresti usare la proprietà `border` per definire lo stile di tutti gli elementi.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Tutti gli elementi dovrebbero avere un bordo `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
