---
id: 5d822fd413a79914d39e98cb
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Next, add opening and closing `head` and `body` tags within the `html` element.

# --hints--

Dovresti avere un tag di apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` aperto.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

