---
id: 5d822fd413a79914d39e98cf
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Puoi vedere il `body` (Ă¨ il riquadro piĂ¹ interno nella tua pagina); il riquadro intorno Ă¨ l'elemento `html`. Fai in modo che il `body` riempia l'intera viewport impostando una `height` di `100vh`. Rimuovi la `margin` predefinita dal `body` impostando la `margin` su `0`. Infine, imposta la proprietĂ  `overflow` su `hidden` per nascondere eventuali barre di scorrimento che appaiono quando qualcosa si estende oltre la viewport.

# --hints--

Dovresti usare il selettore `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Il tuo `body` dovrebbe avere una `height` di `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Il tuo `body` dovrebbe avere una `margin` di `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Il tuo `body` dovrebbe avere la proprietĂ  `overflow` impostata su `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```
