---
id: 5d822fd413a79914d39e98ca
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi i tag di apertura e chiusura `html` sotto il `DOCTYPE` in modo da avere un punto di partenza per inserire del codice. Assicurati di impostare la lingua su inglese.

# --hints--

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Il tuo elemento `html` dovrebbe avere un tag di apertura e un attributo `lang` con valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

I tuoi tag `html` dovrebbero essere nell'ordine corretto.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Dovresti avere solo un elemento `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
