---
id: 67bf4350777ac6ffdc027745
title: Construir una aplicación selector de color
challengeType: 25
dashedName: build-a-color-picker-app
demoType: onClick
---

# --description--

**Objetivo:** Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debe definir y exportar un componente `ColorPicker`.
1. Debe usar el gancho `useState`.
1. Debe tener un elemento `#color-picker-container` con un fondo blanco.
1. Debe tener un elemento `#color-input` que debe ser un input de `color`.
1. Su `#color-input` debe ser hijo de `#color-picker-container`.
1. Cuando se cambie `#color-input`, `#color-picker-container` debe tener su fondo ajustado a ese nuevo valor.

# --hints--

Debe definir y exportar un componente `ColorPicker`.

```js
async () => {
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);

  assert.property(exports, "ColorPicker");
}
```

Debe usar el gancho `useState`.

```js
async () => {
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.isAtLeast(abuseState.calls.length, 1);
}
```

Debe usar el valor devuelto del gancho `useState`.

```js
assert.match(code, /=\s*(?:React\.)?useState\(/);
```

Debe usar la desestructuración de arreglos para asignar el valor devuelto de `useState` a dos variables.

```js
assert.match(code, /\[.*,.*\]\s*=\s*(?:React\.)?useState\(/);
```

Debe pasar `#ffffff` como el valor inicial a su llamada `useState`.

```js
async () => {
    const abuseState = __helpers.spyOn(React, "useState");
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src === "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const _b = await __helpers.prepTestComponent(exports.ColorPicker);
    assert.match(abuseState.calls[0], /#ffffff/i);
}
```

Debe renderizar un elemento con `color-picker-container` como el `id`.

```js
assert.exists(document.getElementById("color-picker-container"));
```

Su `#color-picker-container` debe tener un fondo blanco inicialmente.

```js
const el = document.getElementById("color-picker-container");
const style = window.getComputedStyle(el);
assert.strictEqual(style.backgroundColor, "rgb(255, 255, 255)");
```

Debe renderizar un elemento `input` con un `id` de `color-input`.

```js
const inputMaybe = document.getElementById("color-input");
assert.exists(inputMaybe);
assert.strictEqual(inputMaybe?.tagName, "INPUT");
```

Su `#color-input` debe ser hijo de `#color-picker-container`.

```js
const inputPossibly = document.getElementById("color-input");
const parent = inputPossibly.parentElement;
assert.exists(parent);
assert.strictEqual(parent.getAttribute("id"), "color-picker-container");
```

Su `#color-input` debe ser un input de `color`.

```js
const inputDefinitely = document.getElementById("color-input");
assert.equal(inputDefinitely?.type, "color");
```

Su `#color-input` debe tener un atributo `onChange`.

```js
const definitelyAnInput = document.getElementById("color-input");
assert.exists(definitelyAnInput);
const key = Object.keys(definitelyAnInput).find(key => key.startsWith("__reactProps"));
const reactInstance = definitelyAnInput[key]
assert.exists(reactInstance);
assert.isFunction(reactInstance.onChange);
```

Su `#color-input` debe tener un valor inicial de `#ffffff`.

```js
const mostCertainlyAnInput = document.getElementById("color-input");
assert.exists(mostCertainlyAnInput);
assert.strictEqual(mostCertainlyAnInput.value, "#ffffff");
```

El valor inicial de su `#color-input` debe ser gestionado por su gancho `useState`.

```js
const declaredState = code.match(/(?:const|let|var)\s*\[\s*(.*)\s*,.*\]\s*=\s*useState/);
const declaredStateVar = declaredState?.[1];
const regex = new RegExp(`value\s*=\s*\{\s*${declaredStateVar}\s*}`);
assert.match(code, regex);
```

Cuando el valor de su `#color-input` cambie, su `#color-picker-container` debe cambiar su color de fondo para reflejar ese valor.

```js
async () => {
  const inputForSure = document.getElementById("color-input");
  assert.exists(inputForSure);
  const nativeValueSetter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "value").set;
  await React.act(async () => {
    nativeValueSetter.call(inputForSure, "#000001");
    inputForSure.dispatchEvent(new Event("input", { bubbles: true}));
  });
  const container = document.getElementById("color-picker-container");
  assert.exists(container);
  const style = window.getComputedStyle(container);
  assert.equal(style?.backgroundColor, "rgb(0, 0, 1)");
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ColorPicker } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<ColorPicker />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

#pick-color-btn {
    padding: 15px 30px;
    font-size: 18px;
    color: white;
    background-color: #14183c;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    margin-top: 20px;
}

#pick-color-btn:hover {
    background-color: #08131f;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}

.stylep{
    color: #14183c;
    font-size: 15px;
    font-weight: bold;
    margin-bottom: 50px;

}
```

```jsx
const { useState } = React;

export const ColorPicker = () => {
    const [bgColor, setBgColor] = useState('#ffffff');

    const handleColorChange = (event) => {
        setBgColor(event.target.value);
    };

    return (
            <div id="color-picker-container" style={{ backgroundColor: bgColor }}>
            <p className="stylep">
                Choose a color using the color input below:
            </p>

            <input
                id="color-input"
                type="color"
                value={bgColor}
                onChange={handleColorChange}
            />
        </div>
    );
}
```
