---
id: 694648acde178bb8202d9515
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

أضف الآن درجة الطالب.

أعلن عن متغير باسم `score` وقم بتعيين القيمة `80.5` له.

مع أن كلًا من `age` و `score` أرقام، قد لا يكونا من نفس النوع. توفر Python دالة تسمى `isinstance()` للتحقق من ذلك.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

استخدم `isinstance()` للتحقق مما إذا كان `score` من نوع `int`، واطبع النتيجة في الطرفية كما في المثال أعلاه.

# --hints--

يجب أن يكون لديك متغير `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

يجب أن يخزن المتغير `score` القيمة `80.5`. لا تضع القيمة بين علامات اقتباس.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

يجب أن تطبع نتيجة `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
