---
id: 694648acde178bb8202d9516
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

الناتج هو `False`، مما يُظهر أن `score` ليس `int`.

نوع آخر شائع من الأعداد في Python هو `float`، الذي يُمثل عددًا يحتوي على كسور عشرية. استبدل `int` بـ `float` في استدعاء `isinstance()` الحالي لتأكيد ذلك.

# --hints--

يجب استبدال `int` بـ `float` في استدعاء `isinstance(score, int)` الحالي.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
