---
id: 694648abde178bb8202d9511
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

يجب أن ترى الآن اسم الطالب مطبوعًا في الطرفية.

يوفر Python دالة باسم `type()` يمكنك استخدامها للتحقق من نوع القيمة.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

في المثال أعلاه، تعني النتيجة `<class 'str'>` أن المتغير الممرر إلى الدالة `type()` هو سلسلة نصية.

استخدم الدالة `type()` مع `name` كمعلمتها واطبع النتيجة مثل المثال. تحقق من النتيجة في الطرفية التي تظهر أن `name` من النوع `str` (سلسلة نصية).

# --hints--

يجب أن تطبع نتيجة استدعاء `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
