---
id: 69461e4f98594012939bc1c6
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستتمرن على أنواع بيانات Python من خلال بناء طابعة بطاقة تقرير بسيطة.

كما تعلمت في الدروس السابقة، تُعيَّن المتغيرات بهذه الطريقة:

```py
greeting = 'Hello World'
```

في المثال أعلاه، المتغير `greeting` يحمل قيمة سلسلة نصية، وهي متتالية من الأحرف محاطة بعلامات اقتباس.

أنشئ متغيرًا اسمه `name` وعيّن له السلسلة النصية `Alice`. تذكر أن تحيط القيمة بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة كما في المثال.

**مهم:** يستخدم Python المسافة البادئة (المسافات في بداية السطر) لتنظيم الشفرة. في هذه الورشة، تأكد من عدم وجود مسافات إضافية في بداية كل سطر من الشفرة. إضافة مسافات زائدة ستسبب `IndentationError` وتمنع تنفيذ الشفرة.

# --hints--

يجب أن يكون لديك متغير `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

يجب أن يخزن المتغير `name` القيمة `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
