---
id: 694648acde178bb8202d9512
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

يجب أن تُظهر بطاقة التقرير أيضًا ما إذا كان الطالب مسجلاً حاليًا. يمكن تمثيل ذلك باستخدام قيمة منطقية.

تمثل القيم المنطقية حالة نعم أو لا، وغالبًا ما تُستخدم لاتخاذ قرارات في الشفرة. هناك قيمتان منطقيتان فقط: `True` و `False`.

أعلن عن متغير باسم `is_student` وقم بتعيين القيمة `True` له.

# --hints--

يجب أن يكون لديك متغير `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

يجب أن يخزن المتغير `is_student` القيمة `True`. لا تحيط القيمة بعلامات اقتباس.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
