---
id: 694648acde178bb8202d9514
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

الآن عليك إضافة عمر الطالب إلى بطاقة التقرير. لهذا ستستخدم عددًا صحيحًا، وهو أحد أنواع البيانات الرقمية في Python.

أعلن عن متغير باسم `age` وقم بتعيين القيمة الصحيحة `20` له.

ثم اطبع القيمة ونوع البيانات لـ `age` معًا مفصولين بفاصلة. تحقق من الإخراج في الطرفية الذي يعرض قيمة `age` ونوعه كـ `int` (عدد صحيح).

# --hints--

يجب أن يكون لديك متغير `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

يجب أن يخزن المتغير `age` القيمة `20`. لا تحيط القيمة بعلامات اقتباس.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

يجب أن تطبع `age` و `type(age)` باستخدام فاصل فاصلة.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
