---
id: 694648acde178bb8202d9513
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

يجب أن يتبع اسم الطالب نفس تنسيق التفاصيل الأخرى.

احذف المخرجات السابقة للمتغير `name`. ثم اطبع `name` و `type(name)` معًا في سطر واحد مفصولين بفاصلة كما في الخطوة السابقة.

# --hints--

يجب ألا يحتوي شفرتك على `print(name)`.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

يجب ألا يحتوي شفرتك على `print(type(name))`.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

يجب أن تطبع `name` و `type(name)` في نفس استدعاء `print()` باستخدام فاصل فاصلة.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
