---
id: 6947a0c2b6de035dd8d3eea1
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

يمكنك طباعة قيمة متغير باستخدام الدالة `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

سوف تتعلم المزيد عن الدوال في الدروس القادمة. الآن، اعرف أن الدالة هي كتلة شفرة قابلة لإعادة الاستخدام يمكن استدعاؤها لتشغيل شفرتها، ويمكن تمرير المعلمات إليها.

في المثال أعلاه، `print(greeting)` هو استدعاء دالة، و`greeting` هي المعلمة الخاصة بالدالة.

ارجع إلى المثال واطبع المتغير `name`. تحقق من الناتج في الطرفية.

# --hints--

يجب أن تستدعي `print()` مع `name` كمعلمتها.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
