---
id: 6947a3e516bb8a4a681274a7
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يمكن للدالة `print()` عرض أكثر من قيمة في نفس الوقت. افصل القيم بفاصلة (`,`) لطباعة القيم على نفس السطر.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

اطبع كلًا من `is_student` و `type(is_student)` على نفس السطر باستخدام فاصلة `,` كما هو موضح في المثال. ثم تحقق من المخرجات في الطرفية التي تعرض قيمة `is_student` ونوعها كـ `bool` (قيمة منطقية).

# --hints--

يجب أن تطبع `is_student` و `type(is_student)` باستخدام فاصل فاصلة.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
