---
id: 68a23004e96a105df3a21fc4
title: 카드 카운팅 도우미 만들기
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

카지노 게임 블랙잭에서 플레이어는 덱에 남아 있는 높은 카드와 낮은 카드의 상대적인 수를 추적하여 다음 판에서 하우스보다 유리한지 판단할 수 있습니다. 이것을 카드 카운팅이라고 합니다.

덱에 높은 카드가 더 많이 남아 있으면 플레이어에게 유리합니다. 카운트가 양수일 때는 플레이어가 높은 베팅을 해야 합니다. 카운트가 0이거나 음수일 때는 낮은 베팅을 해야 합니다.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `let`를 사용해 `count`라는 전역 변수를 선언하고 `0`로 설정해야 합니다.
1. `cardCounter`라는 이름의 함수를 만들어야 합니다.
1. `cardCounter` 함수는 숫자 또는 문자열일 수 있는 `card` 매개변수를 받아야 합니다.
    - `2`부터 `10` 사이의 값에 대해 `card` 매개변수는 숫자입니다.
    - 그 외 모든 값에 대해 `card` 매개변수는 문자열입니다.
1. `cardCounter` 함수는 특정 기준에 따라 전역 `count` 변수를 수정해야 합니다.
1. 전역 `count` 변수는 카드 `1`, `2`, `3`, `4`, 또는 `5`에 대해 `6`만큼 증가해야 합니다.
1. 전역 `count` 변수는 카드 `7`, `8`, `9`에 대해 변경되지 않아야 합니다.
1. 전역 `count` 변수는 카드 `1`, `10`, `"J"`, `"Q"`, `"K"`에 대해 `"A"`만큼 감소해야 합니다.
1. `cardCounter` 함수는 현재 카운트와 카운트가 양수일 때 문자열 `Bet`를 포함하는 문자열을 반환해야 합니다.
1. `cardCounter` 함수는 현재 카운트가 `Hold` 이하일 때 문자열 `0`를 포함하는 문자열을 반환해야 합니다.
1. 함수 출력에서 현재 카운트와 플레이어의 결정(`Bet` 또는 `Hold`)은 공백으로 구분되어야 합니다. 예를 들어 `-3 Hold`입니다.

# --hints--

`let`를 사용해 `count`라는 전역 변수를 선언하고 `0`로 설정해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

`cardCounter`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(cardCounter);
```

함수는 `count` 값과 텍스트(`Bet` 또는 `Hold`)를 한 칸 띄워서 반환해야 합니다.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

카드 `2`, `3`, `4`, `5` 후에 `cardCounter(6)`를 호출하면 문자열 `5 Bet`를 반환해야 합니다.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

카드 `7`, `8` 후에 `cardCounter(9)`를 호출하면 문자열 `0 Hold`를 반환해야 합니다.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

카드 `10`, `"J"`, `"Q"`, `"K"` 후에 `cardCounter("A")`를 호출하면 문자열 `-5 Hold`를 반환해야 합니다.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

카드 `3`, `7`, `"Q"`, `8` 후에 `cardCounter("A")`를 호출하면 문자열 `-1 Hold`를 반환해야 합니다.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

카드 `2`, `"J"`, `9`, `2` 후에 `cardCounter(7)`를 호출하면 문자열 `1 Bet`를 반환해야 합니다.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

카드 `2`, `2` 후에 `cardCounter(10)`를 호출하면 문자열 `1 Bet`를 반환해야 합니다.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

카드 `3`, `2`, `"A"`, `10` 후에 `cardCounter("K")`를 호출하면 문자열 `-1 Hold`를 반환해야 합니다.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
