---
id: 68de7360e1832ea180577a8a
title: N番目のフィボナッチ数計算機を作成する
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `fibonacci` という名前の関数を作成してください。
2. `sequence` 関数内に `fibonacci` というリストを定義し、`[0, 1]` の値で初期化してください。
3. `fibonacci` 関数は、正の整数 `n` を1つのパラメータとして受け取る必要があります。  
4. `fibonacci(n)` を呼び出すと、動的計画法のアプローチを使ってフィボナッチ数列の `n` 番目の数を計算し返す必要があります。フィボナッチ数列の各数は、直前の2つの数の合計です。
5. フィボナッチ数列の `n` の位置で計算された各数は、`sequence` リストの `n - 1` インデックスに保存される必要があります。

# --hints--

`fibonacci` という名前の関数を用意してください。

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

`fibonacci` 関数は1つのパラメータを受け取る必要があります。

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

`sequence` というリストを `fibonacci` 関数内に用意し、`[0, 1]` で初期化してください。

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` は `0` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` は `1` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` は `1` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` は `2` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` は `5` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` は `55` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` は `610` を返す必要があります。

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

コード内で再帰を使わないでください。

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```
