---
id: 6153908a366afb4d57185c8d
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox ina mhimili mkuu na mhimili wa msalaba. Mhimili mkuu unafafanuliwa na sifa ya `flex-direction`, ambayo ina thamani nne zinazowezekana:

- `row` (chaguo-msingi): mhimili wa usawa na vipande vya kubadilika kutoka kushoto kwenda kulia
- `row-reverse`: mhimili wa usawa na vipande vya kubadilika kutoka kulia kwenda kushoto
- `column`: mhimili wima na vipande vya kubadilika kutoka juu kwenda chini
- `column-reverse`: mhimili wima na vipande vya kubadilika kutoka chini kwenda juu

**Kumbuka**: Mhimili na mwelekeo vitakuwa tofauti kulingana na mwelekeo wa maandishi. Thamani zilizoonyeshwa ni kwa mwelekeo wa maandishi kutoka kushoto kwenda kulia.

Jaribu thamani tofauti kuona jinsi zinavyoathiri mpangilio.

Ukipomaliza, weka `flex-direction` wazi ya `row` kwenye kipengele cha `.gallery`.

# --hints--

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `flex-direction` iliyowekwa kuwa `row` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
