---
id: 61537bb9b1a29430ac15ad38
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza mradi huu kwa kuunganisha faili yako ya `styles.css` na ukurasa.

# --hints--

Unapaswa kuwa na kipengele cha `link` ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Kipengele chako cha `link` kinapaswa kuwa na `rel="stylesheet"` na `href="styles.css"`.

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
