---
id: 615380dff67172357fcf0425
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa picha zako ni kubwa sana.

Tengeneza kichaguzi cha `.gallery img` ili kuzilenga. Wape zote `width` ya `100%` na `max-width` ya `350px`.

Pia weka sifa ya `height` kuwa `300px` ili picha zako ziwe na ukubwa sawa.

# --hints--

Unapaswa kuwa na kichaguzi cha `.gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Kichaguzi chako cha `.gallery img` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `100%` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width, '100%');
```

Kichaguzi chako cha `.gallery img` kinapaswa kuwa na sifa ya `max-width` iliyowekwa kuwa `350px` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

Kichaguzi chako cha `.gallery img` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `300px` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
