---
id: 61537c5f81f0cf325b4a854c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza kipengele cha `header` ndani ya kipengele cha `body` na weka darasa la vitu la `header` ndani yake.

Ndani ya `header`, tengeneza `h1` yenye `css flexbox photo gallery` kama maandishi.

# --hints--

Unapaswa kuwa na kipengele cha `header` ndani ya kipengele chako cha `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Kipengele chako cha `header` kinapaswa kuwa na `class` yenye thamani ya `header`.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Kipengele chako cha `header` kinapaswa kuwa na kipengele cha `h1` ndani yake.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `css flexbox photo gallery` ndani yake.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
