---
id: 61537c9eecea6a335db6da79
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Below your `.header` element, create a new `div` element and assign it a `class` of `gallery`. This `div` will act as a container for the gallery images.

Ndani ya kipengele hicho cha `.gallery`, unda vipengele tisa vya `img`.

# --hints--

Unapaswa kuunda kipengele cha `div` katika kipengele chako cha `body`.

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

Kipengele chako kipya cha `div` kinapaswa kuwa na `class` na `gallery` iliyowekwa kama thamani.

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

Kipengele chako kipya cha `div` kinafaa kuja baada ya kipengele chako cha `.header`.

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Kipengele chako cha `.gallery` kinapaswa kuwa na vipengele tisa vya `img`.

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
