---
id: 61539e07e7430b528fbffe21
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Toa sifa ya `.gallery` kwa kichaguzi chako cha `padding` na uiweke kuwa `20px 10px` ili kuunda nafasi karibu na kontena.

Kisha, mpa `max-width` ya `1400px` na ongeza `margin` ya `0 auto` ili kuikatia katikati.

# --hints--

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `padding` iliyowekwa kuwa `20px 10px` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingTop, '20px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingBottom, '20px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingLeft, '10px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.paddingRight, '10px');
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `max-width` iliyowekwa kuwa `1400px` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.maxWidth, '1400px');
```

Kipengele chako cha `.gallery` kinapaswa kuwekewa katikati kwa kutumia `margin` yenye thamani ya `0 auto`.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.gallery').cssText, 'margin: 0px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
