---
id: 6153897c27f6334716ee5abe
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Flexbox ni mpangilio wa CSS wa kipimo kimoja unaoweza kudhibiti jinsi vipande vinavyopangwa na kuunganishwa ndani ya kontena.

Ili kuitumia, mpe kipengele sifa ya `display` yenye thamani ya `flex`. Hii itafanya kipengele kuwa <em>chombo cha kubadilika</em>. Watoto wote wa moja kwa moja wa chombo cha kubadilika huitwa <em>vipande vya kubadilika</em>.

Tengeneza kichaguzi cha `.gallery` na uifanye kuwa chombo cha kubadilika.



# --hints--

Unapaswa kuwa na kichaguzi cha `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `display` iliyowekwa kuwa `flex` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
