---
id: 615392916d83fa4f02f7e2cf
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--


Sifa ya `flex-wrap` inaamua jinsi vipande vyako vya kubadilika vitakavyotenda wakati chombo cha kubadilika ni kidogo sana. Kuweka thamani yake kuwa `wrap` kutaruhusu vipande kuzunguka kwenda safu au safu inayofuata. `nowrap` (thamani ya msingi) itazuia vipande vyako kuzunguka na kuvidhoofisha ikiwa ni lazima.


Fanya vipande vyako vya kubadilika vizunguke kwenda safu inayofuata wakati vinapokosa nafasi.

# --hints--

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `flex-wrap` iliyowekwa kuwa `wrap` kama thamani.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
  flex-direction: row;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
