---
id: 6494da0daf5df5197963671d
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Tazama jinsi mpaka wa picha ya buluu unavyopanuka zaidi ya mpaka mwekundu wa maktaba. Hii ni kwa sababu ya jinsi vivinjari vinavyopima ukubwa wa vipengele vya kontena.

Sifa ya `box-sizing` hutumika kuweka tabia hii. Kwa chaguo-msingi, mfano wa `content-box` hutumika. Kwa mfano huu, wakati kipengele kina upana maalum, upana huo hupimwa kwa kuzingatia maudhui ya kipengele pekee. Thamani za nafasi ya ndani na mpaka huongezwa kwenye jumla ya upana, hivyo kipengele kinakua ili kuendana na thamani hizi.

Jaribu kuweka `box-sizing` kuwa `content-box` waziwazi, kwa kutumia kichaguzi cha `*` cha kimataifa. Hadi sasa, hautaona mabadiliko yoyote, kwa sababu unatumia thamani ya chaguo-msingi.

# --hints--

Unapaswa kuwa na kichaguzi cha `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Kichaguzi chako cha `*` kinapaswa kuwa na sifa ya `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Kichaguzi chako cha `*` kinapaswa kuwa na sifa yake ya `box-sizing` ikiwa imewekwa kuwa `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
