---
id: 587d8259367417b2b2512c86
title: 挿入ソートを実装する
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

このラボでは、挿入ソートのアルゴリズムを実装します。このメソッドは、リストの先頭にソート済みの配列を作り上げることで動作します。最初の要素をソート済み配列の最初の要素として開始します。次に次の要素を調べ、ソート済み配列の中で正しい位置に来るまで後ろに向かって入れ替えます。リストの最後まで繰り返し、新しい要素をソート済み部分に向かって入れ替え続けます。このアルゴリズムは平均および最悪の場合で二次の時間計算量を持ちます。

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー:**

1. `insertionSort` 関数を作成してください。
1. `insertionSort` 関数は整数の配列を受け取り、同じ整数を最小から最大の順にソートした配列を返す必要があります。

# --hints--

`insertionSort` は関数である必要があります。

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` は順序以外は変更されていない配列を返す必要があります。

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` は最小から最大の順にソートされた配列を返す必要があります。

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` は `[2, 4, 5, 8, 33]` を返す必要があります。

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` は組み込みの `.sort()` メソッドを使わないでください。

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
