---
id: 646c59e942f35541923104bf
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

ارجع إلى ملف HTML الخاص بك، وأنشئ عنصر `main`. داخل هذا العنصر `main`، أضف عنصر `div` مع الصنف `cat-head`.

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

يجب أن يكون لديك عنصر `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

يجب أن يحتوي عنصر `div` على الصنف `cat-head`.

```js
assert.equal(document.querySelector('div')?.getAttribute('class'), 'cat-head');
```

يجب أن يكون عنصر `div` داخل وسم `main`.

```js
assert.lengthOf(document.querySelectorAll('main div'), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
