---
id: 646c5d7057c45f432fcdd46c
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

باستخدام محدِّد صنف، أعطِ عنصر `.cat-head` عرضًا مقداره `205px` وارتفاعًا مقداره `180px`. كذلك، أعطه حدودًا بمقدار `1px solid #000` وانحناء زوايا (`border-radius`) بنسبة `46%`.

# --hints--

يجب أن يكون لديك محدِّد `.cat-head`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

يجب أن يحتوي محدِّد `.cat-head` على عرض (`width`) مضبوط إلى `205px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

يجب أن يحتوي محدِّد `.cat-head` على ارتفاع (`height`) مضبوط إلى `180px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

يجب أن يحتوي محدِّد `.cat-head` على حدود (`border`) مضبوطة إلى `1px solid #000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

يجب أن يحتوي محدِّد `.cat-head` على انحناء زوايا (`border-radius`) مضبوط إلى `46%`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
