---
id: 646c48df8674cf2b91020ecc
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

تم تزويدك بقالب جاهز أساسي. ابدأ بإضافة عنصر `link` داخل عنصر `head`. بالنسبة لذلك العنصر `link`، عرّف الخاصية `rel` على أنها `stylesheet` والخاصية `href` على أنها `./styles.css`.

# --hints--

يجب أن يكون لديك عنصر `link`.

```js
assert.exists(document.querySelector("head>link"));
```

يجب أن يحتوي عنصر `link` الخاص بك على خاصيتي `rel` و `href` معرّفتين على `stylesheet` و `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
