---
id: 646c5e727e56e743c9aed4a1
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

تُتيح خاصية تحديد المواقع في CSS تعيين كيفية وضع العنصر في المتصفح. لديها خاصية `position` يمكنك تعيينها إلى `static` أو `absolute` أو `relative` أو `sticky` أو `fixed`.

بمجرد تعيين خاصية `position` للعنصر، يمكنك تحريك العنصر بتعيين قيمة بكسل أو نسبة مئوية لواحدة أو أكثر من الخصائص `top` أو `right` أو `left` أو `bottom`.

تُعد `static` هي وضع التمركز الافتراضي لجميع العناصر. إذا عينتها لعنصر، فلن تتمكن من تحريكه باستخدام `top` أو `right` أو `left` أو `bottom`.

عيّن للعنصر `.cat-head` خاصية `position` بقيمة `static`، ثم عيّن خاصيتي `top` و `left` إلى `100px` لكل منهما.


# --hints--

يجب أن يحتوي محدِّد `.cat-head` على خاصية `position` معينة إلى `static`. تأكد من إضافة فاصلة منقوطة.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `top` معينة إلى `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

يجب أن يحتوي محدِّد `.cat-head` على خاصية `left` معينة إلى `100px`. تأكد من إضافة فاصلة منقوطة.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
