---
id: 587d7fb5367417b2b2512c04
title: Ondoa kifurushi kutoka kwa utegemezi wako
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Sasa umejaribu njia kadhaa za kusimamia utegemezi wa mradi wako kwa kutumia sehemu ya dependencies katika faili ya package.json. Pia umejumuisha vifurushi vya nje kwa kuviweka kwenye faili na hata kuambia npm aina gani za matoleo unazotaka, kwa kutumia herufi maalum kama tilde au ishara ya caret.

Lakini vipi kama unataka kuondoa kifurushi cha nje ambacho hutaki tena? Labda tayari umegundua, ondoa tu jozi ya ufunguo-thamani inayolingana na kifurushi hicho kutoka kwa utegemezi wako.

Njia hii ile ile inatumika pia kuondoa sehemu nyingine katika package.json yako.

# --instructions--

Ondoa kifurushi cha `@freecodecamp/example` kutoka kwa utegemezi wako.

**Kumbuka:** Hakikisha una idadi sahihi ya koma baada ya kuondoa kifurushi hicho.

# --hints--

`"dependencies"` haipaswi kujumuisha `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
