---
id: 587d7fb5367417b2b2512c03
title: Tumia ishara ya caret kutumia toleo dogo la hivi karibuni la utegemezi
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

Kama vile tilde tuliyojifunza katika changamoto iliyopita inavyoruhusu npm kusakinisha PATCH ya hivi karibuni kwa utegemezi, ishara ya caret (`^`) inaruhusu npm pia kusakinisha masasisho ya baadaye. Tofauti ni kwamba ishara ya caret itaruhusu masasisho ya MINOR pamoja na PATCH.

Toleo lako la sasa la `@freecodecamp/example` linapaswa kuwa `~1.2.13` ambalo linaruhusu npm kusakinisha toleo la hivi karibuni la `1.2.x`. Ikiwa ungetumia ishara ya caret (^) kama kiambishi awali cha toleo badala yake, npm itaruhusiwa kusasisha hadi toleo lolote la `1.x.x`.

```json
"package": "^1.3.8"
```

Hii itaruhusu masasisho kwa toleo lolote la `1.x.x` la kifurushi.

# --instructions--

Tumia ishara ya caret (`^`) kama kiambishi awali cha toleo la `@freecodecamp/example` katika utegemezi wako na ruhusu npm kuuisasisha hadi toleo lolote jipya la MINOR.

**Kumbuka:** Nambari za toleo hazipaswi kubadilishwa.

# --hints--

`"dependencies"` inapaswa kujumuisha `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Toleo la `"@freecodecamp/example"` linapaswa kufanana na la `"^1.x.x"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
