---
id: 587d7fb5367417b2b2512c02
title: Tumia kiambishi awali cha tilde ili kila wakati utumie toleo la hivi karibuni la toleo la hitilafu la utegemezi
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

Katika changamoto ya mwisho, uliambia npm kujumuisha toleo maalum la kifurushi tu. Hiyo ni njia muhimu ya kufunga utegemezi wako ikiwa unahitaji kuhakikisha sehemu tofauti za mradi wako zinaendana na kila mmoja. Lakini katika matumizi mengi, hutaki kukosa marekebisho ya hitilafu kwa sababu mara nyingi yanajumuisha vifuniko muhimu vya usalama na (kwa matumaini) hayavunji mambo wakati wa kufanya hivyo.

Ili kuruhusu utegemezi wa npm kusasishwa hadi toleo la hivi karibuni la PATCH, unaweza kuweka kiambishi awali cha toleo la utegemezi na herufi ya tilde (`~`). Hapa kuna mfano wa jinsi ya kuruhusu masasisho kwa toleo lolote la `1.3.x`.

```json
"package": "~1.3.8"
```

# --instructions--

Katika faili la package.json, sheria yako ya sasa kuhusu jinsi npm inavyoweza kusasisha `@freecodecamp/example` ni kutumia toleo maalum (`1.2.13`). Lakini sasa, unataka kuruhusu toleo la hivi karibuni la `1.2.x`.

Tumia kiambishi awali cha tilde (`~`) kuweka toleo la `@freecodecamp/example` katika utegemezi wako, na ruhusu npm kuuisasisha hadi toleo lolote jipya la _patch_.

**Kumbuka:** Nambari za toleo la mwenyewe hazipaswi kubadilishwa.

# --hints--

`"dependencies"` inapaswa kujumuisha `"@freecodecamp/example"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Toleo la `"@freecodecamp/example"` linapaswa kulingana na `"~1.2.13"`.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
