---
id: 587d7fb4367417b2b2512bff
title: Ongeza toleo kwenye package.json yako
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

`version` ni mojawapo ya sehemu zinazohitajika za faili yako ya package.json. Sehemu hii inaelezea toleo la sasa la mradi wako. Hapa kuna mfano:

```json
"version": "1.2.0",
```

# --instructions--

Ongeza `version` kwenye faili ya package.json ya mradi wako.

# --hints--

package.json inapaswa kuwa na ufunguo halali wa "version"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
