---
id: 587d7fb4367417b2b2512bfd
title: Ongeza maneno muhimu katika package.json yako
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

Sehemu ya `keywords` ndiyo unayoweza kuelezea mradi wako kwa kutumia maneno muhimu yanayohusiana. Hapa kuna mfano:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Kama unavyoona, sehemu hii imeundwa kama safu ya data ya mfuatano wa herufi zilizo ndani ya alama za nukuu mbili.

# --instructions--

Ongeza safu ya mfuatano wa herufi inayofaa katika sehemu ya `keywords` katika faili la package.json la mradi wako.

Moja ya maneno muhimu inapaswa kuwa "freecodecamp".

# --hints--

package.json inapaswa kuwa na ufunguo halali wa "keywords"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Sehemu ya "keywords" inapaswa kuwa safu ya data

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" inapaswa kujumuisha "freecodecamp"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
