---
id: 587d7fb4367417b2b2512bfe
title: Ongeza leseni kwenye package.json yako
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

Sehemu ya `license` ndiko unakofahamisha watumizi kile wanachoruhusiwa kufanya na mradi wako.

Leseni maarufu kwa miradi ya chanzo huru ni pamoja na MIT na BSD. Taarifa za leseni hazihitajiki, na sheria za hakimiliki katika nchi nyingi zitakupa umiliki wa kile unachotengeneza kwa chaguo-msingi. Hata hivyo, ni desturi nzuri kila mara kueleza wazi kile watumizi wanaweza na wasipoweza kufanya. Hapa kuna mfano wa sehemu ya leseni:

```json
"license": "MIT",
```

# --instructions--

Jaza sehemu ya `license` katika faili la package.json la mradi wako kama unavyoona inafaa.

# --hints--

package.json inapaswa kuwa na ufunguo halali wa "license"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
