---
id: 587d7fb3367417b2b2512bfc
title: Ongeza maelezo kwenye package.json yako
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

Sehemu inayofuata ya faili nzuri ya package.json ni uwanja wa `description`; ambapo maelezo mafupi, lakini yenye taarifa kuhusu mradi wako huwekwa.

Ikiwa siku moja unapanga kuchapisha kifurushi kwenye npm, huu ndio mfuatano wa herufi unaotakiwa kuuza wazo lako kwa mtumizi anapoweka uamuzi wa kufunga kifurushi chako au la. Hata hivyo, hiyo siyo matumizi pekee ya maelezo, ni njia nzuri ya kufupisha kile mradi unachofanya. Ni muhimu sana katika mradi wowote wa Node.js kusaidia wasanidi programu wengine, watunzaji wa baadaye au hata wewe mwenyewe wa baadaye kuelewa mradi haraka.

Bila kujali unachopanga kwa mradi wako, maelezo ni jambo linalopendekezwa sana. Hapa kuna mfano:

```json
"description": "A project that does something awesome",
```

# --instructions--

Ongeza `description` kwenye faili la package.json la mradi wako.

**Kumbuka:** Kumbuka kutumia alama za nukuu mbili kwa majina ya viwanja (") na koma (,) kutenganisha viwanja.

# --hints--

package.json inapaswa kuwa na ufunguo halali wa "description"

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
