---
id: 6716249b5405164036fd0b0d
title: Побудувати візуалізатор сортування
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---


# --description--

Алгоритм сортування бульбашкою сортує послідовність цілих чисел, порівнюючи пари сусідніх елементів, починаючи з початку послідовності. Якщо перший елемент більший за другий, він їх міняє місцями. Потім переходить до наступної пари. Коли досягається останній елемент послідовності, починається новий цикл з початку послідовності, і процес повторюється, доки елементи не будуть відсортовані. Алгоритм зупиняється після того, як один цикл проходить без жодних обмінів.

Для цієї лабораторної роботи вам надали весь HTML і CSS. Ви використаєте JavaScript, щоб завершити візуалізатор сортування бульбашкою, який показуватиме кожен крок, необхідний алгоритму сортування бульбашкою для сортування масиву з п’яти цілих чисел.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з іменем `generateElement`, яка повертає випадкове ціле число між `1` і `100` включно.
1. У вас має бути функція з іменем `generateArray`, яка використовує функцію `generateElement` для повернення масиву з п’яти випадкових цілих чисел.
1. У вас має бути функція з іменем `generateContainer`, яка створює і повертає порожній елемент `div`.
1. У вас має бути функція з іменем `fillArrContainer`, яка приймає HTML-елемент як перший аргумент і масив як другий.
1. Функція `fillArrContainer` має приймати елемент як перший параметр і масив цілих чисел як другий параметр, а потім заповнювати цей елемент п’ятьма елементами `span`, кожен з яких показує одне з чисел масиву.
1. У вас має бути функція з іменем `isOrdered`, яка приймає два цілі числа і повертає булеве значення, що вказує, чи перше число менше або дорівнює другому.
1. У вас має бути функція з іменем `swapElements`, яка приймає масив цілих чисел і числовий індекс.
1. Функція `swapElements` має змінювати масив на місці, міняючи місцями елемент за переданим індексом і наступний елемент, якщо `isOrdered` повертає `false`.
1. У вас має бути функція з іменем `highlightCurrentEls`, яка приймає HTML-елемент і числовий індекс.
1. Функція `highlightCurrentEls` має встановити `border` у вигляді стилю `dashed`, колір `red` і ширину на ваш вибір для дочірніх елементів заданого елемента за вказаним індексом і наступним за ним.
1. Коли ви натискаєте `#generate-btn`, ви маєте використати функцію `fillArrContainer`, щоб заповнити `#starting-array` п’ятьма елементами `span`, кожен з випадковим числом як текстом. Якщо є інші елементи, їх слід видалити з `#array-container`.
1. Ви маєте реалізувати алгоритм сортування бульбашкою так, щоб після натискання `#sort-btn` `#array-container` містив стільки елементів `div`, скільки кроків потрібно алгоритму для сортування початкового масиву, включно з `div`, що представляє початковий масив, і `div`, що представляє відсортований масив. Функції, які ви створили досі, можуть бути корисними тут.
1. Кожен `div` має містити п’ять елементів `span`, які представляють масив у поточному стані сортування.
1. Після натискання `#sort-btn` `#starting-array` має представляти початковий крок із початковим масивом і першими двома підсвіченими числами.
1. Для кожного кроку сортування ви маєте використовувати `highlightCurrentEls` для підсвічування двох чисел, які порівнюються, і міняти їх місцями на наступному кроці за допомогою `swapElements`.

# --hints--

У вас має бути функція з іменем `generateElement`.

```js
assert.isFunction(generateElement);
```

Ваша функція `generateElement` має повертати випадкове ціле число між `1` і `100` включно.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

У вас має бути функція з іменем `generateArray`.

```js
assert.isFunction(generateArray)
```

Ваша функція `generateArray` має використовувати функцію `generateElement`.

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

Ваша функція `generateArray` має повертати масив, що містить п’ять випадкових цілих чисел між `1` і `100`.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

У вас має бути функція з іменем `generateContainer`.

```js
assert.isFunction(generateContainer);
```

Ваша функція `generateContainer` має повертати порожній елемент `div`.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

У вас має бути функція з іменем `fillArrContainer`.

```js
assert.isFunction(fillArrContainer);
```

Ваша функція `fillArrContainer()` має приймати елемент як перший параметр і масив цілих чисел як другий параметр, а потім заповнювати цей елемент п’ятьма елементами `span`, кожен з яких показує одне з чисел масиву.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

У вас має бути функція з іменем `isOrdered`.

```js
assert.isFunction(isOrdered);
```

Ваша функція `isOrdered` має приймати два цілі числа і повертати булеве значення, що вказує, чи перше число менше або дорівнює другому.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

У вас має бути функція з іменем `swapElements`.

```js
assert.isFunction(swapElements);
```

Ваша функція `swapElements` має приймати масив цілих чисел і числовий індекс як аргументи. Вона має змінювати переданий масив на місці, міняючи місцями елемент за заданим індексом і наступний елемент, якщо перший елемент більший за другий.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

У вас має бути функція з іменем `highlightCurrentEls`.

```js
assert.isFunction(highlightCurrentEls);
```

Ваша функція `highlightCurrentEls` має встановити кордон, який є `dashed`, `red`, і заданої вами ширини, для нащадків вказаного елемента за заданим індексом і наступним індексом.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

Коли ви натискаєте `#generate-btn`, ви маєте заповнити `#starting-array` п’ятьма елементами `span`, кожен з випадковим числом між `1` і `100` як текстом.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

Якщо `#starting-array` вже містить згенерований масив або `#array-container` містить відсортований масив, натискання `#generate-btn` має видалити інші елементи в `#array-container`, залишивши лише `#starting-array` з новозгенерованими числами.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

Після натискання `#sort-btn` `#array-container` має містити стільки елементів `div`, скільки кроків потрібно алгоритму сортування бульбашкою для сортування початкового масиву, включно з `div`, що представляє початковий масив, і `div`, що представляє відсортований масив.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

Після натискання `#sort-btn` кожен `div` у `#array-container` має містити п’ять `span`, кожен з числом як текстом, і бути розташованим так, щоб представляти кроки, необхідні алгоритму сортування бульбашкою для сортування початкового масиву.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

Коли ви натискаєте `#sort-btn`, ви маєте використати функцію `highlightCurrentEls` для підсвічування елементів, які порівнюються на кожному кроці.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

Після натискання `#sort-btn` `#starting-array` має представляти початковий крок із початковим масивом і першими двома підсвіченими числами за допомогою `highlightCurrentEls`.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
