---
id: 6823f9df49cc206af5471a30
title: 调试 Camperbot 的个人资料页面
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
---

# --description--

Camperbot 正在尝试建立一个个人资料页面。 他们请了一位朋友来检查他们的代码，并且表示其中存在一些错误。

你的任务是修复Camperbot的所有错误，以便他们能够继续构建他们的个人资料页面。 完成以下用户需求中的各项内容，然后点击“运行测试”查看是否修复了所有错误。

**用户故事：**

1. Camperbot 尝试使用 `heading2` 元素，但该元素不存在。 修复这些标签，使其使用正确的二级标题元素。
2. Camperbot 尝试添加两个带有 `pp` 的段落，但这些元素也不存在。 修复它们，以便使用正确的段落标签。
3. Camperbot 在对 `Background and Interests` 子标题使用 `h3` 元素，但有一个语法错误。 找到问题并解决它。

# --hints--

你在页面上应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你不应该从页面删除 `<h1>Hello from Camperbot!</h1>` 元素。

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

你在页面上应该只有一个 `h2` 元素。

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

你应该有一个 `h2` 元素，其文本为 `About` 。 这是一个示例：`<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

你应该有一个段落元素，文本为 `My name is Camperbot and I love learning new things.` 。这是一个示例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

你应该有一个段落元素，文本为 `I enjoy solving puzzles.` 。这是一个示例：`<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

你在页面上应该恰好有两个段落元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你在页面上应该只有一个 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

你应该有一个 `h3` 元素，其文本为 `Background and Interests` 。 这是一个示例：`<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

你在页面上应该总共有五个元素。

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
