---
id: a3566b1109230028080c9345
title: Tengeneza algoriti ya jumlisha nambari zote
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

Katika maabara hii, utahitaji kutengeneza algoriti ya jumlisha nambari zote.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `sumAll` kinachokubali safu ya data yenye nambari mbili.
1. `sumAll([n, m])` inapaswa kurudisha jumla ya `n` na `m` pamoja na jumla ya nambari zote zilizo kati yao. Nambari ndogo haitakuwa daima ya kwanza. Kwa mfano, `sumAll([4,1])` inapaswa kurudisha `10` kwa sababu jumla ya nambari zote kati ya `1` na `4` (zote zikiwa ni pamoja) ni `10`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` inapaswa kurudisha nambari.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` inapaswa kurudisha `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` inapaswa kurudisha `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` inapaswa kurudisha `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` inapaswa kurudisha `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
