---
id: 56533eb9ac21ba0edf2244bc
title: Lista de compras
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Crea una lista de compras en la variable `myList`. La lista debe ser un arreglo multidimensional que contenga varios sub-arreglos.

El primer elemento de cada sub-arreglo debe contener una cadena con el nombre del artículo. El segundo elemento debe ser un número que represente la cantidad, por ejemplo.

```js
["Chocolate Bar", 15]
```

Debe haber al menos 5 sub-arreglos en la lista.

# --hints--

`myList` debe ser un arreglo.

```js
assert(Array.isArray(myList));
```

Los primeros elementos de cada sub-arreglo deben ser cadenas.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

Los segundos elementos de cada sub-arreglo deben ser números.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

Debes tener al menos 5 elementos en tu lista.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
