---
id: 56533eb9ac21ba0edf2244b6
title: Escapa secuencias en cadenas
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

Las comillas no son los únicos caracteres que pueden ser <dfn>escapados</dfn> dentro de una cadena. Las secuencias de escape le permiten utilizar caracteres que de otro modo no podrías usar en una cadena.

<table><thead><tr><th>Code</th><th>Output</th></tr></thead><tbody><tr><td><code>\'</code></td><td>single quote</td></tr><tr><td><code>\"</code></td><td>double quote</td></tr><tr><td><code>\\</code></td><td>backslash</td></tr><tr><td><code>\n</code></td><td>newline</td></tr><tr><td><code>\t</code></td><td>tab</td></tr><tr><td><code>\r</code></td><td>carriage return</td></tr><tr><td><code>\b</code></td><td>backspace</td></tr><tr><td><code>\f</code></td><td>form feed</td></tr></tbody></table>

*Ten en cuenta que la barra invertida debe ser escapada para mostrarse como una barra invertida.*

# --instructions--

Asigna las siguientes tres líneas de texto en la variable única `myStr` usando secuencias de escape.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

Necesitarás usar secuencias de escape para insertar correctamente los caracteres especiales. También necesitarás seguir el espaciado tal y como se ve arriba, sin espacios entre secuencias de escape o palabras.

**Nota:** La sangría para la segunda línea (`SecondLine`) se consigue con el carácter de escape de tabulación, no con espacios.

# --hints--

`myStr` no debe contener ningún espacio

```js
assert(!/ /.test(myStr));
```

`myStr` debe contener las cadenas `FirstLine`, `SecondLine` y `ThirdLine` (recuerda la sensibilidad a mayúsculas y minúsculas)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` debe ir seguido del carácter de línea nueva `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` debe contener un carácter de tabulación `\t` seguido de un carácter de línea nueva

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` debe estar precedida por el carácter de barra invertida `\`

```js
assert(/\\SecondLine/.test(myStr));
```

Debe haber un carácter de línea nueva entre `SecondLine` y `ThirdLine`

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` solo debe contener caracteres mostrados en las instrucciones

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
