---
id: 5cfa3679138e7d9595b9d9d4
title: Reemplaza bucles (loops) usando recursión
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

La recursión es el concepto que una función puede expresarse en términos de sí misma. Para ayudar a comprender esto, comienza pensando en la siguiente tarea: multiplica los primeros `n` elementos de un arreglo para crear el producto de esos elementos. Usando un bucle `for`, puedes hacer esto:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

Sin embargo, observa que `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. Eso significa que puedes reescribir `multiply` en función de sí misma sin necesidad de usar un bucle.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

La versión recursiva de `multiply` se descompone así.  In the <dfn>base case</dfn>, where `n <= 0`, it returns 1. Para valores mayores de `n`, se llama a sí misma pero con `n - 1`.  That function call is evaluated in the same way, calling `multiply` again until `n <= 0`. En este momento, todas las funciones terminan de ejecutarse y `multiply` devuelve la respuesta.

**Nota:** Las funciones recursivas deben tener un caso base cuando retornan sin llamar a la función nuevamente (en este ejemplo, cuando `n <= 0`), de lo contrario nunca podrán terminar de ejecutarse.

# --instructions--

Escribe una función recursiva, `sum(arr, n)`, que devuelva la suma de los primeros `n` elementos de un array `arr`.

# --hints--

`sum([1], 0)` debe ser igual a 0.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` debe ser igual a 2.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` debe ser igual a 9.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

Tu código no debe usar ningún tipo de bucle (`for` o `while`, o funciones de orden superior como `forEach`, `map`, `filter` o `reduce`.).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

Debes usar recursión para resolver este problema.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
