---
id: 68551b677a5f6898af364b85
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过使用 flexbox 实用类来构建响应式号召性用语 (CTA) 组件来了解有关 Tailwind CSS 的更多信息。

已经为你提供了基本的 HTML 样板，因此你需要首先构建 HTML 元素。

首先创建一个 `div` 元素。 在 `div` 内部，再创建两个单独的 `div` 元素。

# --hints--

你应该创建一个 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

`div` 元素内应该还有两个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
