---
id: 685523c05fb41d0848c4b52f
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

在第一个 `div` 元素内，创建一个带有文本 `Soundflow` 的 `span` 元素，一个带有文本 `Discover New Music` 的 `h1` 元素，以及一个带有文本 `Stream your favorite tracks and discover new artists.` 的 `p` 元素。

# --hints--

你应该在第一个 `div` 元素内创建一个 `span` 元素，其中包含文本 `Soundflow`。

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

你应该在第一个 `div` 元素内创建一个 `h1` 元素，其中包含文本 `Discover New Music`。

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

你应该在第一个 `div` 元素内创建一个 `p` 元素，其中包含文本 `Stream your favorite tracks and discover new artists.`。

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
