---
id: 6855255d96c1441bd919dc5d
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

移动到带有文本 `Soundflow` 的 `span` 元素，其中的文本需要大写。

下面是一个展示如何使用 `uppercase` 实用类的示例：

```html
<p class="uppercase">Quincy Larson</p>
```

将类 `uppercase` 添加到 `span` 元素。

# --hints--

你的 `span` 元素应该具有实用类 `uppercase`。

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl.classList.contains("uppercase"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        --fcc-editable-region--
        <span>Soundflow</span>
        --fcc-editable-region--
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
