---
id: 68552464ef378e1051a55c28
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在第二个 `div` 内，创建两个锚元素。 第一个应该包含文本 `Learn more`，第二个应该包含文本 `Start listening`。

对于两个锚元素的 `href` 属性，将它们都设置为 `#`。

# --hints--

你应该在第二个 `div` 元素内创建两个 `a` 元素

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

你的第一个 `a` 元素应该包含文本 `Learn more`，并且在第二个 `div` 元素内包含 `#` 的 `href` 属性。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

你的第二个 `a` 元素应该具有文本 `Start listening` 以及第二个 `div` 元素内的 `href` 属性 `#`。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
