---
id: 6858ee73ed8106ccb214942a
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

现在所有内容都粘在左侧，你需要通过在 x 轴上添加自动边距来使整个内容居中。

将与之对应的 Tailwind 类添加到外部 `div` 元素。 要看到该效果，请调整预览选项卡，直到整个内容居中。

# --hints--

你的外层 `div` 元素应该具有实用类 `mx-auto`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
