---
id: 6858e6408a0999603b2a4fee
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

对于 `h1`，添加类 `font-bold` 和 `text-4xl`。 这将使文本更粗、更大，并且行距也更大。

另外，添加类 `my-4` 以将 `h1` 的顶部和底部边距增加 16 像素，或 `1rem`。

# --hints--

你的 `h1` 元素应该具有实用类 `font-bold`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

你的 `h1` 元素应该具有实用类 `text-4xl`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

你的 `h1` 元素应该具有实用类 `my-4`。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
