---
id: 6858e7e23bf9e475e14cb931
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

现在 HTML 结构已经构建完毕，是时候使用 Tailwind 转到适当的样式了。

从包裹整个内容的 `div` 元素开始。 通过赋予 `class` 为 `bg-indigo-600` 来设置其背景颜色。

# --hints--

你的外部 `div` 应该将 `class` 属性设置为 `bg-indigo-600`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
