---
id: 6858f90c54647a48ddb8822d
title: 步骤11
challengeType: 0
dashedName: step-11
---

# --description--

此时，无论屏幕大小如何，内容，即外层 `div` 内的两个 `div` 元素将堆叠在一起。

这是因为包装的 `div` 元素中有一个 `flex-col` 类。 要解决这个问题，请使用 `lg:` 前缀将类 `flex-row` 添加到 `div`。 

这意味着在屏幕更大的设备上，弹性方向将是 `row`。 要再次查看效果，请将浏览器调整到更大的宽度。

# --hints--

你的外层 `div` 元素应该具有实用类 `lg:flex-row`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
