---
id: 6858f3a67aa8580ff37016c7
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

如果你再次将预览选项卡的宽度调整为与移动设备相同的宽度，你会注意到内容始终并排放置。 这并不是你在移动设备上想要的结果。

为了解决这个问题，请将类 `flex-col` 添加到 `div`。 这会将弹性方向改变为列，因此元素会堆叠在一起。

以下是使用 `flex-col` 实用类的方法：

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

你的起始 `div` 元素应该具有实用类 `flex-col`。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
