---
id: 68597114730d7c6b8838d204
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

`h1` 下的段落看起来不错，因此不需要任何其他样式。 这意味着你可以移动到第二个 `div` 元素内的锚元素。

为第一个锚元素赋予类 `bg-white`，`hover:bg-gray-100` 和 `text-indigo-600`。 这样，当你将鼠标悬停在其上时，它将具有白色的背景、靛蓝色的文本颜色和略浅的背景。

对于第二个锚元素，赋予它类 `bg-pink-500`，`hover:bg-pink-600` 和 `text-white`。

# --hints--

你的第一个锚元素应该具有实用类 `bg-white`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

你的第一个锚元素应该具有实用类 `hover:bg-gray-100`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

你的第一个锚元素应该具有实用类 `text-indigo-600`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

你的第二个锚元素应该具有实用类 `bg-pink-500`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

你的第二个锚元素应该具有实用类 `hover:bg-pink-600`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

你的第二个锚元素应该具有实用类 `text-white`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
