---
id: 685973ca82469e940afcc422
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

要完成按钮样式，请将实用类 `font-semibold` 添加到每个 `a` 元素，以便它们的文本可以更粗体。 另外，添加类 `rounded` 来消除按钮上的尖锐边缘。

# --hints--

你的第一个锚元素应该具有实用类 `font-semibold`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("font-semibold"));
```

你的第一个锚元素应该具有实用类 `rounded`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("rounded"));
```

你的第二个锚元素应该具有实用类 `font-semibold`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("font-semibold"));
```

你的第二个锚元素应该具有实用类 `rounded`。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("rounded"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
--fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600 px-4 py-2"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2"
         >
          Start listening
        </a>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
