---
id: 67c562286b29447da020d407
title: crea un generatore di password usa e getta
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

In questo laboratorio, genererai un OTP (One-Time Password) di 6 cifre e lo mostrerai all’utente. L’OTP scadrà dopo 5 secondi e ne verrà generato uno nuovo quando l’utente cliccherà il pulsante `"Generate New OTP"`.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti usare il hook `useEffect` per gestire il timer del conto alla rovescia.
2. Il tuo componente `OTPGenerator` dovrebbe restituire un elemento `div` con il nome della classe `container`.
3. Il `div` con la classe `container` dovrebbe includere i seguenti elementi:

    - Un elemento `h1` con l’ID `otp-title` e il testo `"OTP Generator"`.
    - Un elemento `h2` con l’ID `otp-display` che mostra il messaggio `"Click 'Generate OTP' to get a code"` oppure l’OTP generato se disponibile.
    - Un elemento `p` con l’ID `otp-timer` e l’attributo `aria-live` impostato su un valore valido che:
        - Inizia vuoto.
        - Mostra `"Expires in: X seconds"` dopo che il pulsante è stato cliccato, dove `X` indica il tempo rimanente prima che l’OTP scada.
        - Visualizza il messaggio `"OTP expired. Click the button to generate a new OTP."` una volta che il conto alla rovescia raggiunge `0`.
    - Un elemento `button` con l’ID `generate-otp-button` etichettato `"Generate OTP"`. Quando viene cliccato, deve generare un nuovo OTP e avviare un conto alla rovescia di 5 secondi.
    - Il pulsante `"Generate OTP"` deve essere disabilitato mentre il conto alla rovescia è attivo.

4. Dovresti assicurarti che il timer del conto alla rovescia si fermi automaticamente una volta raggiunti `0` secondi per evitare aggiornamenti inutili.
5. L’OTP generato deve essere lungo 6 cifre.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

Non dovresti rimuovere l’assegnazione `const { useState, useEffect, useRef } = React;` esistente dal codice.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Dovresti renderizzare un `div` con la classe `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

Il `.container` dovrebbe contenere un elemento `h1` con l’ID `otp-title` e il testo `"OTP Generator"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

Il `.container` dovrebbe contenere un elemento `h2` con l’ID `otp-display` che mostra l’OTP quando viene generato.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

Inizialmente, il `h2` dentro `.container` dovrebbe mostrare il messaggio `"Click 'Generate OTP' to get a code"`.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

L’elemento `.container` dovrebbe contenere un elemento `p` con l’ID `otp-timer`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

L’elemento `p` dovrebbe avere un attributo `aria-live` impostato su `assertive` o `polite`.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

Inizialmente, l’elemento `p` dovrebbe essere vuoto.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

Il `.container` dovrebbe contenere un elemento `button` con l’ID `generate-otp-button` e il testo `"Generate OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Quando il pulsante viene cliccato, dovrebbe mostrare un nuovo OTP nell’elemento `h2` con id `otp-display`.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

L’OTP generato deve essere lungo 6 cifre.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Quando il pulsante viene cliccato, l’elemento `p` con id `otp-timer` dovrebbe mostrare un conto alla rovescia di 5 secondi.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Il messaggio nell’elemento `p` con id `otp-timer` dovrebbe aggiornarsi ogni secondo per mostrare il tempo rimanente.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

Il pulsante `"Generate OTP"` dovrebbe essere disabilitato mentre il timer del conto alla rovescia è attivo.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

Il pulsante `"Generate OTP"` dovrebbe essere abilitato una volta che il timer del conto alla rovescia raggiunge 0.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

Quando il timer del conto alla rovescia raggiunge `0`, dovresti mostrare il messaggio `OTP expired. Click the button to generate a new OTP.`.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

Dovresti esportare il componente `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
