---
id: 6a89c69cbe1bb69845992d8b
title: Einen Zahlenbereich-Generator erstellen
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

In diesem Labor bauen Sie eine `range_of_numbers`-Funktion, die Rekursion verwendet, um eine Liste von Zahlen innerhalb eines angegebenen Bereichs zu erzeugen.

Zum Beispiel:

- `range_of_numbers(3, 9)` sollte `[3, 4, 5, 6, 7, 8, 9]` zurückgeben.
- `range_of_numbers(5, 5)` sollte `[5]` zurückgeben.

Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `range_of_numbers` mit zwei Parametern namens `start_num` und `end_num` erstellen.
1. Die Funktion sollte eine Liste aufeinanderfolgender ganzer Zahlen zurückgeben, die mit `start_num` beginnt und mit `end_num` endet, beide Werte eingeschlossen.
1. Der Wert von `start_num` ist immer kleiner oder gleich `end_num`.
1. Ihre Funktion sollte Rekursion verwenden, indem sie sich selbst aufruft.
1. Ihre Funktion sollte keine `for`- oder `while`-Schleifen, keine Comprehensions oder die `range()`-Funktion verwenden.
1. Wenn `start_num` gleich `end_num` ist, sollte die Funktion eine Liste mit `start_num` zurückgeben. Dies ist der Basisfall.
1. Für den rekursiven Fall sollte die Funktion sich selbst mit einem Argument aufrufen, das sich dem Basisfall nähert, und dann die aktuelle Zahl zur zurückgegebenen Liste hinzufügen.
1. Die Funktion sollte keine globalen Variablen verwenden, um die zurückgegebene Liste zu speichern.

# --hints--

Sie sollten eine Funktion namens `range_of_numbers` mit den Parametern `start_num` und `end_num` haben.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` sollte `[1, 2, 3, 4, 5]` zurückgeben.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` sollte `[6, 7, 8, 9]` zurückgeben.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` sollte `[4]` zurückgeben.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` sollte `[10, 11, 12, 13, 14, 15]` zurückgeben.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Ihre `range_of_numbers`-Funktion sollte sich selbst aufrufen.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Ihre `range_of_numbers`-Funktion sollte keine Schleifen, Comprehensions oder die `range()`-Funktion verwenden.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Der Aufruf von `range_of_numbers(1, 3)` sollte einen späteren Aufruf nicht beeinflussen. Danach sollte `range_of_numbers(6, 9)` weiterhin `[6, 7, 8, 9]` zurückgeben.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
