---
id: 66ba762af611169359d9d369
title: Jenga makala ya gazeti
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuweka `font-size` ya mzizi wa hati yako ya HTML kuwa `24px`.
1. Unapaswa kuwa na kipengele chenye darasa la vitu la `newspaper` kinachojumuisha vipengele vyako vyote vingine.
1. Kipengele chako cha `.newspaper` kinapaswa kuwa na sifa ya `font-size` yenye thamani ya `16px` na fonti ya `Open Sans` yenye fonti mbadala ya `sans-serif`.
1. Ndani ya kipengele chako cha `.newspaper`, unapaswa kuwa na vipengele saba zaidi angalau: kimoja kwa jina la gazeti lenye darasa la `name`, kimoja kwa tarehe ya makala lenye darasa la `date`, kimoja kwa kichwa cha habari lenye darasa la `headline`, kimoja kwa kichwa kidogo lenye darasa la `sub-headline`, kimoja kwa mwandishi lenye darasa la `author`, na aya mbili kila moja ikiwa na darasa la `text`. Vipengele vyote hivi vinapaswa kujazwa na taarifa zako za makala.
1. Kipengele chako cha `.name` kinapaswa kuwa na sifa ya `font-size` ambayo ni mara mbili ya `font-size` ya kipengele cha mzizi na kinapaswa kutumia fonti ya `Times New Roman`, ikiwa na fonti mbadala ya `serif`.
1. Vipengele vyako vya `.name` na `.author` vinapaswa kutumia CSS kufanya herufi zao zote kuwa kubwa.
1. Kipengele chako cha `.headline` kinapaswa kuwa na sifa ya `font-size` ambayo ni mara mbili ya `font-size` ya kipengele chake cha mzazi na kinapaswa kuwa na herufi nzito.
1. Kipengele chako cha `.sub-headline` kinapaswa kuwa na sifa ya `font-weight` yenye thamani ya `100`, sifa ya `font-size` ambayo ni mara `1.5` ya `font-size` ya kipengele chake cha mzazi, na kinapaswa kuwa na herufi za mduara.
1. `.author` inapaswa kutumia CSS kuifanya iwe na herufi nzito.
1. Vipengele vya `.text` vinapaswa kuwa na sifa ya `text-indent` yenye thamani ya `20px`.
1. Vipengele vya `.text` vinapaswa kuwa na sifa ya `line-height` ambayo ni mara mbili ya `font-size` ya kipengele chao cha mzazi.
1. Herufi ya kwanza ya vipengele vyako vya `.text` inapaswa kuwa na herufi nzito na ukubwa mara mbili ya `font-size` ya kipengele chao cha mzazi. Tumia kichaguzi cha `::first-letter` kwa hili.

# --hints--

Unapaswa kuwa na kichaguzi cha `html` kinachoweka sifa yake ya `font-size` kuwa `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Unapaswa kuwa na kipengele chenye darasa la `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

Unapaswa kuwa na kichaguzi cha `.newspaper` kinachoweka sifa za vipengele vyake za `font-size` kuwa `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Unapaswa kuwa na kichaguzi cha `.newspaper` kinachoweka sifa za vipengele vyake za `font-family` kuwa `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Unapaswa kuwa na kipengele chenye darasa la `name` ndani ya kipengele chako cha `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Kipengele chako cha `.name` kinapaswa kuwa na jina la gazeti lako.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Unapaswa kuwa na kipengele chenye darasa la `date` ndani ya kipengele chako cha `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Kipengele chako cha `.date` kinapaswa kuwa na tarehe.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Unapaswa kuwa na kipengele chenye darasa la `headline` ndani ya kipengele chako cha `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Kipengele chako cha `.headline` kinapaswa kuwa na kichwa cha habari cha makala yako.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Unapaswa kuwa na kipengele chenye darasa la `sub-headline` ndani ya kipengele chako cha `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Kipengele chako cha `.sub-headline` kinapaswa kuwa na kichwa kidogo cha makala yako.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Unapaswa kuwa na kipengele chenye darasa la `author` ndani ya kipengele chako cha `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Kipengele chako cha `.author` kinapaswa kuwa na jina la mwandishi wa makala yako.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Unapaswa kuwa na vipengele vya aya angalau viwili, kila kikiwa na darasa la `text`, ndani ya kipengele chako cha `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Vipengele vyako vya `.text` vinapaswa kuwa na maandishi ya makala yako.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Unapaswa kuwa na kichaguzi cha `.name` kinachoweka sifa za vipengele vyake za `font-size` kuwa `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Unapaswa kuwa na kichaguzi cha `.name` kinachoweka sifa za vipengele vyake za `font-family` kuwa `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Kipengele cha `.name` kinapaswa kuwa na sifa ya `text-transform` iliyowekwa kuwa `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Kipengele cha `.author` kinapaswa kuwa na sifa ya `text-transform` iliyowekwa kuwa `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Unapaswa kuwa na kichaguzi cha `.headline` kinachoweka sifa za vipengele vyake za `font-size` kuwa `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Unapaswa kuwa na kichaguzi cha `.headline` kinachoweka sifa za vipengele vyake za `font-weight` kuwa `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Unapaswa kuwa na kichaguzi cha `.sub-headline` kinachoweka sifa za vipengele vyake za `font-weight` kuwa `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Unapaswa kuwa na kichaguzi cha `.sub-headline` kinachoweka sifa za vipengele vyake za `font-size` kuwa `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Unapaswa kuwa na kichaguzi cha `.sub-headline` kinachoweka sifa za vipengele vyake za `font-style` kuwa `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Unapaswa kuwa na kichaguzi cha `.author` kinachoweka sifa za vipengele vyake za `font-weight` kuwa `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Unapaswa kuwa na kichaguzi cha `.text` kinachoweka sifa za vipengele vyake za `text-indent` kuwa `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Unapaswa kuwa na kichaguzi cha `.text` kinachoweka sifa za vipengele vyake za `line-height` kuwa `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Unapaswa kuwa na kichaguzi cha `.text::first-letter` kinachoweka sifa za vipengele vyake za `font-weight` kuwa `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Unapaswa kuwa na kichaguzi cha `.text::first-letter` kinachoweka sifa za vipengele vyake za `font-size` kuwa `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
