---
id: 587d78b2367417b2b2512b0e
title: Ongeza vitu kwenye safu ya data kwa kutumia push() na unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

Urefu wa safu ya data, kama vile aina za data inayoweza kuwa nayo, si thabiti. Safu za data zinaweza kufafanuliwa na urefu wa idadi yoyote ya vipengele, na vipengele vinaweza kuongezwa au kuondolewa kwa muda; kwa maneno mengine, safu za data ni <dfn>zinazobadilika</dfn>. Katika changamoto hii, tutaangalia njia mbili ambazo tunaweza kutumia kubadilisha safu ya data kwa mpangilio wa programu: `Array.push()` na `Array.unshift()`.

Njia zote mbili zinachukua kipengele kimoja au zaidi kama vigezo na kuongezea vipengele hivyo kwenye safu ya data ambayo njia hiyo inaitwa; njia ya `push()` huongeza vipengele mwishoni mwa safu ya data, na `unshift()` huongeza vipengele mwanzoni. Fikiria yafuatayo:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals` ingekuwa na thamani ya `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals` ingekuwa na thamani ya `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. Angalia kwamba pia tunaweza kupitisha vigezo, ambavyo hutupa uhuru zaidi katika kubadilisha data ya safu yetu kwa njia ya mabadiliko ya wakati halisi.

# --instructions--

Tumeunda kitendakazi, `mixedNumbers`, ambacho tunapitisha safu ya data kama hoja. Badilisha kitendakazi kwa kutumia `push()` na `unshift()` kuongeza `'I', 2, 'three'` mwanzoni mwa safu ya data na `7, 'VIII', 9` mwishoni ili safu inayorejeshwa iwe na mwakilishi wa nambari 1-9 kwa mpangilio.

# --hints--

`mixedNumbers(["IV", 5, "six"])` inapaswa sasa kurudisha `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

Kitendakazi cha `mixedNumbers` kinapaswa kutumia njia ya `push()`

```js
assert(mixedNumbers.toString().match(/\.push/));
```

Kitendakazi cha `mixedNumbers` kinapaswa kutumia njia ya `unshift()`

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
