---
id: 587d7b7c367417b2b2512b1a
title: Upatikanaji wa majina ya sifa kwa uandishi wa mabano
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

Katika changamoto ya kwanza ya vitu tulitaja matumizi ya uandishi wa mabano kama njia ya kupata thamani za sifa kwa kutumia tathmini ya kigezo. Kwa mfano, fikiria kwamba kitu chetu cha `foods` kinatumika katika programu ya kasha la supermarket. Tuna kitendakazi fulani kinachoweka `selectedFood` na tunataka kukagua kitu chetu cha `foods` kwa uwepo wa chakula hicho. Hii inaweza kuonekana kama ifuatavyo:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

Msimbo huu utathmini thamani iliyohifadhiwa katika kigezo cha `selectedFood` na kurudisha thamani ya ufunguo huo katika kitu cha `foods`, au `undefined` ikiwa haipo. Uandishi wa mabano ni muhimu sana kwa sababu wakati mwingine sifa za vitu hazijulikani kabla ya wakati wa utekelezaji au tunahitaji kuzipata kwa njia ya mabadiliko zaidi.

# --instructions--

Tumeunda kitendakazi, `checkInventory`, ambacho kinapokea kipengee kilichoskanwa kama hoja. Rudisha thamani ya sasa ya ufunguo wa `scannedItem` katika kitu cha `foods`. Unaweza kudhani kwamba vitufe halali tu vitatolewa kama hoja kwa `checkInventory`.

# --hints--

`checkInventory` inapaswa kuwa kitendakazi.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

Kitu cha `foods` kinapaswa kuwa na jozi za ufunguo-thamani zifuatazo tu: `apples: 25`, `oranges: 32`, `plums: 28`, `bananas: 13`, `grapes: 35`, `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")` inapaswa kurudisha `25`.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")` inapaswa kurudisha `13`.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")` inapaswa kurudisha `27`.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
