---
id: 587d7b7a367417b2b2512b12
title: Nakili vitu vya safu ya data kwa kutumia slice()
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

Njia inayofuata tutakayojifunza ni `slice()`. Badala ya kubadilisha safu ya data, `slice()` hunakili au *hutenganisha* idadi fulani ya vipengele na kuviweka katika safu mpya, huku safu ya data inayopigiwa wito ikibaki bila kubadilika. `slice()` huchukua vigezo viwili tu — cha kwanza ni kielezo ambacho kuanza kutenganisha, na cha pili ni kielezo ambacho kuacha kutenganisha (kutenganisha kutafanyika hadi, lakini si pamoja na kipengele kilicho kwenye kielezo hiki). Fikiria hii:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` itakuwa na thamani `['snow', 'sleet']`, wakati `weatherConditions` bado itakuwa na `['rain', 'snow', 'sleet', 'hail', 'clear']`.

Kwa maana hii, tumeunda safu mpya kwa kutenganisha vipengele kutoka safu ya data iliyopo.

# --instructions--

Tumeunda kitendakazi, `forecast`, ambacho kinachukua safu ya data kama hoja. Badilisha kitendakazi hiki kwa kutumia `slice()` kutenganisha taarifa kutoka kwenye safu ya data ya hoja na kurudisha safu mpya inayojumuisha vipengele vya mfuatano wa herufi `warm` na `sunny`.

# --hints--

`forecast` inapaswa kurudisha `["warm", "sunny"]`

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

Kitendakazi cha `forecast` kinapaswa kutumia njia ya `slice()`

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
