---
id: 587d7b7d367417b2b2512b1e
title: Tengeneza safu ya vibonye vyote vya kitu kwa kutumia Object.keys()
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

Tunaweza pia kutengeneza safu ya data inayojumuisha vibonye vyote vilivyohifadhiwa katika kitu kwa kutumia njia ya `Object.keys()`. Njia hii inachukua kitu kama hoja na inarudisha safu ya mfuatano wa herufi zinazowakilisha kila sifa katika kitu hicho. Tena, hakutakuwa na mpangilio maalum wa maingizo katika safu hiyo.

# --instructions--

Malizia kuandika kitendakazi cha `getArrayOfUsers` ili kirudishe safu ya data inayojumuisha sifa zote katika kitu kinachopokelewa kama hoja.

# --hints--

Kitu cha `users` kinapaswa kuwa na vibonye tu vya `Alan`, `Jeff`, `Sarah`, na `Ryan`

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

Kitendakazi cha `getArrayOfUsers` kinapaswa kurudisha safu ya data inayojumuisha vibonye vyote katika kitu cha `users`

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
