---
id: 587d7b7c367417b2b2512b18
title: Ongeza jozi za ufunguo-thamani kwa vitu vya JavaScript
challengeType: 1
forumTopicId: 301153
dashedName: add-key-value-pairs-to-javascript-objects
---

# --description--

Kwa msingi wake, vitu ni mkusanyiko tu wa jozi za <dfn>ufunguo-thamani</dfn>. Kwa maneno mengine, ni vipande vya data (<dfn>thamani</dfn>) vinavyopangwa kwa vitambulisho vya kipekee vinavyoitwa <dfn>sifa</dfn> (<dfn>funguo</dfn>). Angalia mfano huu:

```js
const tekkenCharacter = {
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true
};
```

Msimbo ulio juu unaelezea kitu cha mchezaji wa mchezo wa video wa Tekken kinachoitwa `tekkenCharacter`. Kina sifa tatu, kila moja ikiwa na thamani maalum. Ikiwa unataka kuongeza sifa nyingine, kama vile "asili", inaweza kufanywa kwa kuweka thamani `origin` kwa kitu hicho:

```js
tekkenCharacter.origin = 'South Korea';
```

Hii inatumia uandishi wa nukta. Ukitazama kitu cha `tekkenCharacter`, sasa kitajumuisha sifa ya `origin`. Hwoarang pia alikuwa na nywele za rangi ya chungwa. Unaweza kuongeza sifa hii kwa kutumia uandishi wa mabano kwa kufanya:

```js
tekkenCharacter['hair color'] = 'dyed orange';
```

Uandishi wa mabano unahitajika ikiwa sifa yako ina nafasi ndani yake au ikiwa unataka kutumia kigezo cha jina la sifa. Katika mfano huu, sifa imefungwa kwa nukuu kuonyesha kuwa ni mfuatano wa herufi na itaongezwa kama ilivyoonyeshwa. Bila nukuu, itahesabiwa kama kigezo na jina la sifa litakuwa thamani yoyote ya kigezo hicho. Hapa kuna mfano na kigezo:

```js
const eyes = 'eye color';

tekkenCharacter[eyes] = 'brown';
```

Baada ya kuongeza mifano yote, kitu kitakuwa kama hiki:

```js
{
  player: 'Hwoarang',
  fightingStyle: 'Tae Kwon Doe',
  human: true,
  origin: 'South Korea',
  'hair color': 'dyed orange',
  'eye color': 'brown'
};
```

# --instructions--

Kitu cha `foods` kimeundwa na maingizo matatu. Kwa kutumia sintaksia unayopendelea, ongeza maingizo matatu zaidi: `bananas` yenye thamani ya `13`, `grapes` yenye thamani ya `35`, na `strawberries` yenye thamani ya `27`.

# --hints--

`foods` inapaswa kuwa kitu.

```js
assert(typeof foods === 'object');
```

Kitu cha `foods` kinapaswa kuwa na ufunguo `bananas` wenye thamani ya `13`.

```js
assert(foods.bananas === 13);
```

Kitu cha `foods` kinapaswa kuwa na ufunguo `grapes` wenye thamani ya `35`.

```js
assert(foods.grapes === 35);
```

Kitu cha `foods` kinapaswa kuwa na ufunguo `strawberries` wenye thamani ya `27`.

```js
assert(foods.strawberries === 27);
```

Ufafanuzi wa kitu cha `foods` haupaswi kubadilishwa.

```js
assert(
  __helpers.removeJSComments(code).search(/let foods/) === -1 &&
  __helpers.removeJSComments(code).search(/const\s+foods\s*=\s*{\s*apples:\s*25,\s*oranges:\s*32,\s*plums:\s*28\s*};/
) !== -1
);
```

# --seed--

## --seed-contents--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
const foods = {
  apples: 25,
  oranges: 32,
  plums: 28
};

foods['bananas'] = 13;
foods['grapes']  = 35;
foods['strawberries'] = 27;
```
