---
id: 587d7b7b367417b2b2512b13
title: Nakili safu ya data kwa opereta ya spread
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

Wakati `slice()` inaturuhusu kuchagua ni vipengele gani vya safu ya data tutakavyo nakili, miongoni mwa kazi nyingine nyingi muhimu, opereta mpya ya spread ya ES6 inaturuhusu kwa urahisi kunakili *vipengele vyote* vya safu ya data, kwa mpangilio, kwa sintaksia rahisi na inayosomeka kwa urahisi. Sintaksia ya spread inaonekana hivi tu: `...`

Katika matumizi halisi, tunaweza kutumia opereta ya spread kunakili safu ya data kama ifuatavyo:

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray` ni sawa na `[true, true, undefined, false, null]`. `thisArray` haijabadilika na `thatArray` ina vipengele sawa na `thisArray`.

# --instructions--

Tumeunda kitendakazi, `copyMachine` kinachopokea `arr` (safu ya data) na `num` (nambari) kama hoja. Kitendakazi kinapaswa kurudisha safu mpya iliyoundwa kwa nakala `num` za `arr`. Tumefanya kazi nyingi tayari, lakini bado hakifanyi kazi vizuri. Badilisha kitendakazi kwa kutumia sintaksia ya spread ili kifanye kazi ipasavyo (kidokezo: njia nyingine tuliyoshughulikia tayari inaweza kusaidia hapa!).

# --hints--

`copyMachine([true, false, true], 2)` inapaswa kurudisha `[[true, false, true], [true, false, true]]`

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)` inapaswa kurudisha `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]`

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)` inapaswa kurudisha `[[true, true, null]]`

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)` inapaswa kurudisha `[["it works"], ["it works"], ["it works"]]`

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

Kitendakazi cha `copyMachine` kinapaswa kutumia `spread operator` na safu ya data `arr`

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
