---
id: 587d78b2367417b2b2512b0f
title: Ondoa vitu kutoka kwenye safu ya data kwa kutumia pop() na shift()
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

Zote `push()` na `unshift()` zina njia zinazolingana ambazo karibu ni kinyume cha kazi: `pop()` na `shift()`. Kama unavyoweza kuwa umethibitisha sasa, badala ya kuongeza, `pop()` *huondoa* kipengee kutoka mwisho wa safu ya data, wakati `shift()` huondoa kipengee kutoka mwanzo. Tofauti kuu kati ya `pop()` na `shift()` na ndugu zao `push()` na `unshift()`, ni kwamba hakuna njia yoyote inayochukua vigezo, na kila moja inaruhusu safu ya data kubadilishwa kwa kipengee kimoja tu kwa wakati mmoja.

Hebu tuchunguze:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` ingekuwa na thamani `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` ingekuwa na thamani `['hello']`.

Tunaweza pia kurudisha thamani ya kipengee kilichoondolewa kwa njia yoyote kama hii:

```js
let popped = greetings.pop();
```

`greetings` ingekuwa na thamani `[]`, na `popped` ingekuwa na thamani `hello`.

# --instructions--

Tumeunda kitendakazi, `popShift`, ambacho kinachukua safu ya data kama hoja na kurudisha safu mpya ya data. Badilisha kitendakazi hiki, ukitumia `pop()` na `shift()`, kuondoa vipengee vya kwanza na vya mwisho vya safu ya hoja, na weka vipengee vilivyoondolewa kwenye vigezo vinavyolingana, ili safu inayorudishwa iwe na thamani zao.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` inapaswa kurudisha `["challenge", "complete"]`

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

Kitendakazi cha `popShift` kinapaswa kutumia njia ya `pop()`

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

Kitendakazi cha `popShift` kinapaswa kutumia njia ya `shift()`

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
