---
id: 587d7b7b367417b2b2512b15
title: Rudia kupitia vitu vyote vya safu ya data kwa kutumia mizunguko ya for
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

Wakati mwingine unapotumia safu za data, ni rahisi sana kuwa na uwezo wa kurudia kupitia kila kipengee ili kupata kipengee kimoja au zaidi ambacho tunaweza kuhitaji, au kubadilisha safu ya data kulingana na vipengee vya data vinavyokidhi vigezo fulani. JavaScript hutoa njia kadhaa zilizojengwa ndani ambazo kila moja hurudia juu ya safu za data kwa njia tofauti kidogo ili kufanikisha matokeo tofauti (kama `every()`, `forEach()`, `map()`, n.k.), hata hivyo mbinu inayotoa uhuru zaidi na kudhibiti zaidi ni mzunguko rahisi wa `for`.

Fikiria yafuatayo:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

Kutumia mzunguko wa `for`, kitendakazi hiki hurudia kupitia na kupata upatikanaji wa kila kipengee cha safu ya data, na kukiweka chini ya jaribio rahisi tulilolifanya. Kwa njia hii, tumebaini kwa urahisi na kwa mpangilio ni vipengee gani vya data vinavyozidi `10`, na kurudisha safu mpya, `[12, 14, 80]`, yenye vipengee hivyo.

# --instructions--

Tumetangaza kitendakazi, `filteredArray`, ambacho kinapokea `arr`, safu ya data zilizopangwa ndani, na `elem` kama hoja, na kurudisha safu mpya. `elem` inawakilisha kipengee ambacho kinaweza kuwepo au kuto kuwa kwenye moja au zaidi ya safu zilizopangwa ndani ya `arr`. Badilisha kitendakazi, ukitumia mzunguko wa `for`, ili kurudisha toleo lililosafishwa la safu iliyotumwa kwa njia kwamba safu yoyote iliyopangwa ndani ya `arr` yenye `elem` imeondolewa.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` inapaswa kurudisha `[[10, 8, 3], [14, 6, 23]]`

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` inapaswa kurudisha `[["flutes", 4]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` inapaswa kurudisha `[["amy", "beth", "sam"]]`

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` inapaswa kurudisha `[]`

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

Kitendakazi cha `filteredArray` kinapaswa kutumia mzunguko wa `for`

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
