---
id: 587d7b7b367417b2b2512b17
title: Unganisha safu za data kwa opereta ya spread
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

Faida nyingine kubwa ya opereta ya <dfn>spread</dfn> ni uwezo wa kuunganisha safu za data, au kuingiza vipengele vyote vya safu moja ndani ya nyingine, mahali popote pa kielezo. Kwa sintaksia za kawaida zaidi, tunaweza kuunganisha safu za data, lakini hii inatuwezesha tu kuunganisha safu mwishoni mwa moja, na mwanzoni mwa nyingine. Sintaksia ya spread hufanya operesheni ifuatayo kuwa rahisi sana:

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray` ingekuwa na thamani ya `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']`.

Kwa kutumia sintaksia ya spread, tumefanikisha operesheni ambayo ingekuwa ngumu zaidi na yenye maneno mengi zaidi kama tungetumia njia za kawaida.

# --instructions--

Tumeunda kitendakazi `spreadOut` kinachorudisha variable `sentence`. Badilisha kitendakazi kwa kutumia opereta ya <dfn>spread</dfn> ili kirudishe safu ya data `['learning', 'to', 'code', 'is', 'fun']`.

# --hints--

`spreadOut` inapaswa kurudisha `["learning", "to", "code", "is", "fun"]`

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

Kitendakazi `spreadOut` kinapaswa kutumia sintaksia ya spread

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
