---
id: 587d7b7b367417b2b2512b14
title: Angalia uwepo wa kipengele kwa kutumia indexOf()
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

Kwa kuwa safu za data zinaweza kubadilishwa, au *kubadilishwa hali*, wakati wowote, hakuna uhakika kuhusu mahali kipengele fulani kitakachokuwepo kwenye safu hiyo, au kama kipengele hicho bado kipo. Kwa bahati nzuri, JavaScript hutupatia njia nyingine iliyoijengwa ndani, `indexOf()`, inayotuwezesha kuangalia haraka na kwa urahisi uwepo wa kipengele kwenye safu ya data. `indexOf()` huchukua kipengele kama kigezo, na inapoitwa, hurudisha nafasi, au kielezo, cha kipengele hicho, au `-1` ikiwa kipengele hakipo kwenye safu.

Kwa mfano:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` hurudisha `-1`, `indexOf('oranges')` hurudisha `2`, na `indexOf('pears')` hurudisha `1` (kielezo cha kwanza ambacho kila kipengele kipo).

# --instructions--

`indexOf()` inaweza kuwa na msaada mkubwa kwa kuangalia haraka uwepo wa kipengele kwenye safu ya data. Tumetangaza kitendakazi, `quickCheck`, kinachochukua safu ya data na kipengele kama hoja. Badilisha kitendakazi kwa kutumia `indexOf()` ili kirudishe `true` ikiwa kipengele kilichopitishwa kipo kwenye safu, na `false` ikiwa hakipo.

# --hints--

Kitendakazi cha `quickCheck` kinapaswa kurudisha thamani ya kweli au siyo kweli (`true` au `false`), si mfuatano wa herufi (`"true"` au `"false"`)

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` inapaswa kurudisha `false`

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` inapaswa kurudisha `true`

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` inapaswa kurudisha `true`

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` inapaswa kurudisha `false`

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

Kitendakazi cha `quickCheck` kinapaswa kutumia njia ya `indexOf()`

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
