---
id: 587d7b7b367417b2b2512b16
title: Unda safu za data zilizo na vipimo vingi na zilizopangwa ndani
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

Nzuri sana! Umejifunza mengi kuhusu safu za data! Hii imekuwa muhtasari wa kiwango cha juu, na bado kuna mengi ya kujifunza kuhusu kufanya kazi na safu za data, mengi ambayo utaona katika sehemu zijazo. Lakini kabla ya kuendelea kuangalia <dfn>Vitu</dfn>, hebu tuchunguze tena, na tione jinsi safu za data zinavyoweza kuwa ngumu zaidi kuliko tulivyoshuhudia katika changamoto zilizopita.

Moja ya vipengele vyenye nguvu zaidi unapotafakari safu za data kama miundo ya data, ni kwamba safu za data zinaweza kuwa na, au hata kutengenezwa kabisa kwa safu nyingine za data. Tumeona safu za data zenye safu ndani katika changamoto zilizopita, lakini zilikuwa rahisi. Hata hivyo, safu za data zinaweza kuwa na kina kisicho na kikomo cha safu za data ambazo zinaweza kuwa na safu nyingine, kila moja ikiwa na viwango vyake vya kina, na kadhalika. Kwa njia hii, safu ya data inaweza haraka kuwa muundo wa data mgumu sana, unaojulikana kama <dfn>multi-dimensional</dfn>, au safu zilizopangwa ndani. Fikiria mfano ufuatao:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

Safu ya `deep` imepangwa ndani kwa kina cha ngazi 2. Safu za `deeper` zina kina cha ngazi 3. Safu za `deepest` zina kina cha ngazi 4, na `deepest-est?` ni ngazi 5.

Ingawa mfano huu unaweza kuonekana mgumu, kiwango hiki cha ugumu hakijawa jambo la kushangaza, wala si la kawaida, unapotumia kiasi kikubwa cha data. Hata hivyo, bado tunaweza kwa urahisi kupata upatikanaji wa kina zaidi cha safu hii ngumu kwa kutumia uandishi wa mabano:

```js
console.log(nestedArray[2][1][0][0][0]);
```

Hii inaandika mfuatano wa herufi `deepest-est?`. Na sasa tunapojua mahali data hiyo iko, tunaweza kuibadilisha ikiwa tunahitaji:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

Sasa inaandika `deeper still`.

# --instructions--

Tumeunda kigezo, `myNestedArray`, kilichowekwa sawa na safu ya data. Badilisha `myNestedArray`, ukitumia mchanganyiko wowote wa <dfn>mfuatano wa herufi</dfn>, <dfn>nambari</dfn>, na <dfn>thamani za kweli au siyo kweli</dfn> kama vipengele vya data, ili iwe na kina cha ngazi tano kabisa (kumbuka, safu ya nje kabisa ni ngazi ya 1). Mahali popote kwenye ngazi ya tatu, jumuisha mfuatano wa herufi `deep`, kwenye ngazi ya nne, jumuisha mfuatano wa herufi `deeper`, na kwenye ngazi ya tano, jumuisha mfuatano wa herufi `deepest`.

# --hints--

`myNestedArray` inapaswa kuwa na nambari, thamani za kweli au siyo kweli, na mfuatano wa herufi tu kama vipengele vya data

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` inapaswa kuwa na kina cha ngazi 5 kabisa

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` inapaswa kuwa na tukio moja tu la mfuatano wa herufi `deep` kwenye safu iliyopangwa ndani kwa kina cha ngazi 3

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` inapaswa kuwa na tukio moja tu la mfuatano wa herufi `deeper` kwenye safu iliyopangwa ndani kwa kina cha ngazi 4

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` inapaswa kuwa na tukio moja tu la mfuatano wa herufi `deepest` kwenye safu iliyopangwa ndani kwa kina cha ngazi 5

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
