---
id: 587d7b7c367417b2b2512b1b
title: Tumia neno kuu delete kufuta vigezo vya kitu
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

Sasa unajua ni vitu gani na vipengele vyao vya msingi pamoja na faida zao. Kwa kifupi, ni hifadhi za jozi za ufunguo-thamani zinazotoa njia rahisi na ya kueleweka ya kupanga data, ***na***, hutoa muda wa haraka sana wa kutafuta. Katika changamoto hizi zote, tutaelezea shughuli kadhaa za kawaida unazoweza kufanya kwenye vitu ili uweze kuwa na urahisi wa kutumia miundo hii muhimu ya data katika programu zako.

Katika changamoto zilizopita, tumekuwa tukiongeza na kubadilisha jozi za ufunguo-thamani za kitu. Hapa tutaona jinsi tunavyoweza *kufuta* jozi ya ufunguo-thamani kutoka kwenye kitu.

Turejee mfano wetu wa kitu cha `foods` mara ya mwisho. Ikiwa tungetaka kufuta ufunguo wa `apples`, tunaweza kufuta kwa kutumia neno kuu la `delete` kama hii:

```js
delete foods.apples;
```

# --instructions--

Tumia neno kuu delete kufuta funguo za `oranges`, `plums`, na `strawberries` kutoka kwenye kitu cha `foods`.

# --hints--

Kitu cha `foods` kinapaswa kuwa na funguo tatu tu: `apples`, `grapes`, na `bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

Funguo za `oranges`, `plums`, na `strawberries` zinapaswa kufutwa kwa kutumia `delete`.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
