---
id: 587d78b3367417b2b2512b11
title: Ongeza vitu kwa kutumia splice()
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

Kumbuka katika changamoto iliyopita tulitaja kwamba `splice()` inaweza kuchukua vigezo vitatu? Naam, unaweza kutumia kigezo cha tatu, kilicho na kipengee kimoja au zaidi, kuongeza kwenye safu ya data. Hii inaweza kuwa msaada mkubwa kwa kubadilisha haraka kipengee kimoja, au seti ya vipengee, kwa kingine.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

Kutokea kwa pili kwa `12` kunafutwa, na sisi tunaongeza `13` na `14` katika kielezo hicho hicho. Safu ya data ya `numbers` sasa itakuwa `[ 10, 11, 12, 13, 14, 15 ]`.

Hapa, tunaanza na safu ya nambari. Kisha, tunapitisha yafuatayo kwa `splice()`: Kielezo ambacho kuanza kufuta vipengee (3), idadi ya vipengee kufutwa (1), na hoja zilizobaki (13, 14) zitachapishwa kuanzia kielezo hicho hicho. Kumbuka kwamba kunaweza kuwa na idadi yoyote ya vipengee (vilivyotenganishwa kwa koma) vinavyofuata `amountToDelete`, kila kimoja kikiingizwa.

# --instructions--

Tumeunda kitendakazi, `htmlColorNames`, ambacho kinachukua safu ya rangi za HTML kama hoja. Badilisha kitendakazi kwa kutumia `splice()` kufuta vipengee viwili vya kwanza vya safu ya data na kuongeza `'DarkSalmon'` na `'BlanchedAlmond'` mahali pake.

# --hints--

`htmlColorNames` inapaswa kurudisha `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]`

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

Kitendakazi cha `htmlColorNames` kinapaswa kutumia njia ya `splice()`

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

Haupaswi kutumia `shift()` au `unshift()`.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

Haupaswi kutumia uandishi wa mabano wa safu ya data.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
