---
id: 5a661e0f1068aca922b3ef17
title: Upatikanaji wa Yaliyomo ya Safu ya Data Kutumia Uandishi wa Mabano
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

Kipengele cha msingi cha muundo wowote wa data ni, bila shaka, uwezo wa si tu kuhifadhi data, bali pia kuwa na uwezo wa kupata data hiyo kwa amri. Hivyo, sasa tunapojifunza jinsi ya kuunda safu ya data, tuanze kufikiria jinsi tunavyoweza kupata taarifa za safu hiyo.

Tunapobainisha safu rahisi kama inavyoonekana hapa chini, kuna vitu 3 ndani yake:

```js
let ourArray = ["a", "b", "c"];
```

Katika safu ya data, kila kipengele cha safu kina <dfn>kielezo</dfn>. Kielezo hiki kinafanya kazi pia kama nafasi ya kipengele hicho katika safu, na jinsi unavyorejelea kipengele hicho. Hata hivyo, ni muhimu kutambua kuwa safu za JavaScript ni <dfn>zinaanzia kielezo sifuri</dfn>, ikimaanisha kuwa kipengele cha kwanza katika safu kiko katika nafasi ya ***sifuri***, si ya kwanza. Ili kupata kipengele kutoka katika safu, tunaweza kuweka kielezo ndani ya mabano na kuambatanisha mwisho wa safu, au kwa kawaida, kwenye kigezo kinachorejelea kitu cha safu. Hii inajulikana kama <dfn>uandishi wa mabano</dfn>. Kwa mfano, ikiwa tunataka kupata `a` kutoka `ourArray` na kuwekeza thamani hiyo kwenye kigezo, tunaweza kufanya hivyo kwa msimbo ufuatao:

```js
let ourVariable = ourArray[0];
```

Sasa `ourVariable` ina thamani ya `a`.

Mbali na kupata thamani inayohusiana na kielezo, pia unaweza *weka* thamani kwa kielezo kwa kutumia uandishi huo huo:

```js
ourArray[1] = "not b anymore";
```

Kutumia uandishi wa mabano, sasa tumebadilisha kipengele kilichopo kwenye kielezo 1 kutoka mfuatano wa herufi `b`, kuwa `not b anymore`. Sasa `ourArray` ni `["a", "not b anymore", "c"]`.

# --instructions--

Ili kumaliza changamoto hii, weka thamani yoyote unayotaka kwenye nafasi ya 2 (kielezo `1`) ya `myArray`, isipokuwa herufi `b`.

# --hints--

`myArray[0]` inapaswa kuwa sawa na herufi `a`

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` haipaswi kuwa sawa na herufi `b`

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` inapaswa kuwa sawa na herufi `c`

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` inapaswa kuwa sawa na herufi `d`

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
