---
id: 61437d575fb98f57fa8f7f36
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть зі стандартного шаблонного коду HTML. Додайте `DOCTYPE` оголошення, `html` елемент, який вказує, що ця сторінка англійською, `head` елемент і `body` елемент.

Додайте `<meta>` тег з відповідним `charset` та `<meta>` тег для мобільної адаптивності всередині `head` елемента.

# --hints--

Ваш код має містити `DOCTYPE` референс.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Після `DOCTYPE` референса має бути пробіл.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Закрийте `DOCTYPE` оголошення `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Має бути відкриваючий `<html>` тег, і він має містити `lang` `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш `html` елемент має мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Ваше `DOCTYPE` оголошення має бути на початку вашого HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Має бути відкриваючий `<head>` тег.

```js
assert(code.match(/<head\s*>/i));
```

Має бути кінцевий `</head>` тег.

```js
assert(code.match(/<\/head\s*>/i));
```

Має бути відкриваючий `<body>` тег.

```js
assert(code.match(/<body\s*>/i));
```

Має бути кінцевий `</body>` тег.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` і `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Елемент `head` має бути всередині `html` елемента.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` має бути всередині `html` елемента.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Має бути два `meta` елементи.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Один `meta` елемент має мати множину `name`, встановлену в `viewport`, і `content`, встановлену в `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший `meta` елемент має мати атрибут `charset`, встановлений у `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
