---
id: 6143869bb45bd85e3b1926aa
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

У межах вашого елемента `.heading` створіть елемент `header` з множиною `class`, встановленою на `hero`.

У цьому елементі створіть елемент `img` з множиною `src`, встановленою на `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, множиною `alt`, встановленою на `freecodecamp logo`, і множиною `class`, встановленою на `hero-img`.

Атрибут `loading` на елементі `img` можна встановити в `lazy`, щоб браузер не завантажував ресурс зображення, доки це не буде потрібно (наприклад, коли користувач прокручує сторінку до зображення). Додатковою перевагою є те, що елементи з відкладеним завантаженням не завантажуються, доки не завантажаться елементи без відкладеного завантаження — це означає, що користувачі з повільним інтернет-з’єднанням можуть переглядати вміст вашої сторінки, не чекаючи на завантаження зображень.

Додайте вашому новому елементу `img` атрибут `loading`, встановлений у `lazy`.

# --hints--

Ви повинні створити елемент `header`.

```js
assert.exists(document.querySelector('header'));
```

Ваш елемент `header` має бути всередині вашого елемента `.heading`.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

Ваш елемент `header` має мати множину `class`, встановлену на `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

Ви повинні створити елемент `img`.

```js
assert.exists(document.querySelector('img'));
```

Ваш елемент `img` має бути всередині вашого елемента `header`.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

Ваш елемент `img` має мати множину `src`, встановлену на `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Ваш елемент `img` має мати множину `alt`, встановлену на `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

Ваш елемент `img` має мати атрибут `loading`, встановлений у `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

Ваш елемент `img` має мати множину `class`, встановлену на `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
