---
id: 614387cbefeeba5f3654a291
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Після вашого елемента `header` створіть `div` з множиною `class`, встановленою на `author`.

Всередині цього `div` створіть елемент `p` з множиною `class`, встановленою на `author-name`, і додайте текст `By freeCodeCamp`. Обгорніть частину `freeCodeCamp` у елемент `a` з множиною `href`, встановленою на `https://freecodecamp.org`, та множиною `target`, встановленою на `_blank`.

Нижче додайте другий елемент `p` з класом `publish-date` і текстом `March 7, 2019`.

# --hints--

Ви повинні створити новий елемент `div`.

```js
assert.exists(document.querySelector('div'));
```

Ваш елемент `div` має йти після елемента `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

Ваш елемент `div` має мати множину `class`, встановлену на `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

Ви повинні створити два нові елементи `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

Ваші два нові елементи `p` мають бути всередині елемента `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Перший новий елемент `p` має мати множину `class`, встановлену на `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

Перший новий елемент `p` має містити текст `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Другий новий елемент `p` має мати множину `class`, встановлену на `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

Другий новий елемент `p` має містити текст `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Ви повинні створити новий елемент `a`.

```js
assert.exists(document.querySelector('a'));
```

Ваш елемент `a` має бути всередині першого нового елемента `p`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Ваш елемент `a` має мати множину `href`, встановлену на `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Ваш елемент `a` має мати множину `target`, встановлену на `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Ваш елемент `a` має обгорнути текст `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
