---
id: 6143862a5e54455d262c212e
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У вашому `body` створіть елемент `main`. Потім у цьому елементі створіть `section` з множиною `class`, встановленою на `heading`.

# --hints--

У вас має бути елемент `main`.

```js
assert.exists(document.querySelector('main'));
```

Ваш елемент `main` має бути всередині вашого елемента `body`.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

У вас має бути елемент `section`.

```js
assert.exists(document.querySelector('section'));
```

Ваш елемент `section` має бути всередині вашого елемента `main`.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

Ваш елемент `section` має мати множину `class`, встановлену на `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
