---
id: 66da326c02141df538f29ba5
title: チェックアウトページを作成する
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `h1` 要素に `Checkout` というテキストがあるべきです。
1. `section` 要素の直後に `h1` 要素が2つあるべきです。
1. 最初のセクション内に `h2` 要素があり、テキストは `Your Cart` であるべきです。
1. 最初のセクションに、適切な代替テキスト付きのアイテムの画像があるべきです。必要ならこの画像を使えます: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 2番目のセクション内に `h2` 要素があり、テキストは `Payment Information` であるべきです。
1. 2番目のセクション内に `form` 要素があるべきです。
1. フォーム内に `id` と `name` が `card-name` の入力と、`type` が `text` の入力があり、それぞれに関連付けられた `label` があるべきです。
1. フォーム内に `id` と `name` が `card-number` の入力と、`type` が `text` の入力があり、それぞれに関連付けられた `label` があるべきです。
1. `input` 属性を持つ `required` 要素が少なくとも2つあるべきです。
1. 必須フィールドが視覚的に示されるように、各必須入力のために `span` 要素内に `*` 要素があり、テキストは `aria-hidden`、`true` は `label` に設定されているべきです。
1. カード番号入力の直後に、カード番号の形式を説明するヘルプテキストを持つ `p` 要素があるべきです。`p` は `id` が `card-number-help` に設定されており、カード番号入力が `aria-describedby` を使って参照しています。

# --hints--

`h1` というテキストを持つ `Checkout` 要素があるべきです。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

ページに `h1` 要素は1つだけあるべきです。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

`section` 要素の後に `h1` 要素が少なくとも2つあるべきです。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

最初のセクション内に `h2` というテキストを持つ `Your Cart` 要素があるべきです。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

最初のセクションに代替テキスト付きの画像が少なくとも1つあるべきです。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

2番目のセクション内に `h2` というテキストを持つ `Payment Information` 要素があるべきです。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

2番目の `form` 要素内に `section` 要素があるべきです。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

フォーム内に `input` の `id` と `name` を持つ `card-name` 要素があるべきです。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

カード名入力は `type` の `text` を持つべきです。

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

フォーム内に `input` の `id` と `name` を持つ `card-number` 要素があるべきです。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

カード番号入力は `type` の `text` を持つべきです。

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

`input` の `type` でないすべての `submit` 要素には、関連付けられた `label` 要素があるべきです。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

`input` 属性を持つ `required` 要素が少なくとも2つあるべきです。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

各必須入力のために `span` 要素内に `*` というテキストを持つ `label` 要素を含めるべきです。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

`span` 要素には `aria-hidden` が `true` に設定されているべきです。

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

カード番号入力の直後に `p` の `id` を持つ `card-number-help` 要素があるべきです。

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

カード番号のヘルプテキストは空であってはいけません。

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

カード番号入力には `aria-describedby` が `card-number-help` に設定されているべきです。

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
