---
id: 6a3e73eabddd614f91cc235a
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

さあ、先ほどの謎を解決しましょう！`use_eggs` 関数が残りの卵の数を返したとき、その新しい値をどこにも保存しませんでした。そのまま消えてしまったのです！

実際に在庫を更新するには、その返された値を変数に代入して受け取る必要があります。

`make_fried_egg` 関数の中で、`if/else` 文を使って `has_enough_eggs` 変数をチェックしてください。

- もし `has_enough_eggs` が `True` なら、`available_eggs` 変数に `use_eggs(available_eggs, 1)` の結果を再代入し、メッセージ `Made a fried egg. Yummy!` を出力してください。

- そうでなければ、`else` ブロック内でメッセージ `Could not make a fried egg. Not enough eggs!` を出力してください。

# --hints--

`if` 文を使い、その条件に `has_enough_eggs` を使う必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

`if` ブロックの中で、`available_eggs` に `use_eggs(available_eggs, 1)` の結果を再代入してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

`has_enough_eggs` が `True` のときに、メッセージ `Made a fried egg. Yummy!` を出力してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

`has_enough_eggs` が `False` のときに、メッセージ `Could not make a fried egg. Not enough eggs!` を出力する `else` ブロックを用意してください。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
