---
id: 6a3e70b94f67d9ee3ecd8486
title: ステップ11
challengeType: 20
dashedName: step-11
---

# --description--

今、`use_eggs(available_eggs, 1)` を呼び出して卵を1つ使いました。つまり、キッチンには `0` 個の卵が残っているはずですよね？

その答えがわかります！コードの一番下で、`print()` 関数を使って `available_eggs` 変数の値を出力してください。

コードを実行すると驚くことに気づきます：出力はまだ `1` のままです！なぜ卵の数が変わらなかったのでしょう？この謎は次のいくつかのステップで探っていきます。

# --hints--

コードの最後で `available_eggs` 変数を出力してください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
