---
id: 6a3e6fd08125dfa615cb431b
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`use_eggs` 関数が完成したので、テストしてみましょう。**パラメータ**が必要な関数を呼び出すときは、丸括弧の中に具体的な値を渡します。これらの値は **引数** と呼ばれます。

引数を使って関数を呼び出す例を示します：

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

コードの一番下で、`use_eggs` 関数を呼び出し、グローバル変数 `available_eggs` と数字の `1`（使いたい卵の数を表します）という2つの引数を渡してください。

# --hints--

`use_eggs` 関数を呼び出し、`available_eggs` と `1` を引数として渡してください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
