---
id: 6a3e58891e1ff13d07f4ab48
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

Pythonは`def`ブロックを読み込み、関数を記憶しますが、その中のコードをすぐには実行しません。

コードを実行して出力を確認するには、関数を**呼び出す**（または起動する）必要があります。関数名の後に丸括弧を付けて入力してください。

例えば：

```py
my_function()
```

コードの一番下、関数の外側で`check_kitchen_stock`関数を呼び出してください。

# --hints--

コードの最後に`check_kitchen_stock`関数を呼び出してください。インデントされていないことを確認し、グローバルに実行されるようにしてください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
