---
id: 6a3e5475c35dbc138d77c6f4
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

合計を計算したので、在庫を表示してください。前のレッスンで、変数を出力に含めるために `print()` 関数と **f-strings** を使えることを学びましたね。

`check_kitchen_stock`関数内で、`total_items`変数の下に、次の4つの正確なprint文を追加してください。

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

`print(f'The kitchen has {total_items} total items:')`の文が関数内にある必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'The kitchen has {total_items} total items:')`の文は、関数内で最初の`print()`呼び出しである必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'- {available_eggs} eggs')`の文が関数内にある必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

`print(f'- {available_flour} flour')`の文が関数内にある必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

`print(f'- {available_sugar} sugar')`の文が関数内にある必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
