---
id: 6a3e67fc5ef4045201c28219
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`use_eggs` 関数の中で、調理に十分な卵があるかどうかを確認する必要があります。これは `if` 文を使って行えます。

`pass` キーワードを、`eggs_to_use` が `available_eggs` より大きい（`>`）かどうかをチェックする `if` 文に置き換えてください。`if` ブロック内で、文字列 `The kitchen does not have enough eggs.` を出力してください。

# --hints--

`eggs_to_use > available_eggs` をチェックする `if` 文が必要です。

```js
({
  test: () => runPython(`
  cond = _Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0]
  assert cond.is_equivalent("eggs_to_use > available_eggs") or cond.is_equivalent("available_eggs < eggs_to_use")`)
})
```

`if` 文の中で `The kitchen does not have enough eggs.` を出力してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
