---
id: 6a3ea8a55317ba7a08e21d5b
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

これでようやく先ほどの質問に答えられます：なぜグローバルの卵の数が`1`のままだったのでしょうか？関数が値を返しても、既存の変数が自動的に上書きされるわけではありません。新しい値を明示的に保存する必要があります！

`make_fried_egg`関数内で`available_eggs`変数を使っていますが、それはコードの先頭で定義された変数とは異なります。関数内のものはローカル変数で、関数が実行されている間だけ存在します。関数が終了するとローカル変数は破棄され、グローバルの`available_eggs`は変更されません。

実際のキッチンの在庫を更新するには、`make_fried_egg`関数が返す新しい卵の数でグローバルの`available_eggs`変数を上書きしてください。

コードの一番下で、`make_fried_egg`関数を呼び出し、`available_eggs`を引数として渡し、その結果をグローバルの`available_eggs`変数に代入してください。

# --hints--

`available_eggs`変数を`make_fried_egg(available_eggs)`の結果に再代入してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
