---
id: 6a3e4ae26f2b7e4b5b026cf8
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

キッチンにあるものを確認する関数を作成する時間です。Pythonでは、`def` キーワードの後に関数名、丸括弧 `()`、コロン `:` を続けて関数を定義します。

シンプルな関数を定義する例を示します：

```py
def my_function():
    pass
```

`check_kitchen_stock` という名前の空の関数を定義してください。Pythonが `SyntaxError` を発生させないように、関数の本体には今のところ `pass` キーワードを使ってください。

Pythonはコードブロックを定義するのにインデントを使うことを覚えておいてください。関数本体のインデントを必ず付けてください。

# --hints--

`check_kitchen_stock` という名前の関数を定義してください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
