---
id: 6a3e2b9b7eeca084dd4fb857
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

キッチンへようこそ！シェフがレシピに頼って料理を作るように、ソフトウェアを作るときは関数に頼ります。このワークショップでは、再利用可能なPythonの関数の書き方とスコープをマスターし、コード内でいつでもどの材料（変数）が使えるかを正確に理解できるようになります。

まずはキッチンの在庫の基本データを定義しましょう。`available_eggs`、`available_flour`、`available_sugar`という3つの変数を作成し、それぞれに`1`、`2`、`3`の値を代入してください。

# --hints--

`available_eggs`という名前の変数を定義し、値に`1`を代入してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

`available_flour`という名前の変数を定義し、値に`2`を代入してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

`available_sugar`という名前の変数を定義し、値に`3`を代入してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
