---
id: 6a3e6ad6b9554e31dcc30cd6
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

卵が十分にない場合は、料理を試みるべきではありません！関数の実行をこれ以上続けないようにする方法が必要です。

これには `return` キーワードを使えます。関数が `return` 文に達すると、すぐに実行をストップし、返された値を呼び出された場所に渡します。

例えば：

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

`if` ブロックの中で、print 文の下に、`return` キーワードを使って `available_eggs` の値を返してください。

# --hints--

`if` 文の中で `available_eggs` を返す必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

`if` ブロックの中の `print` 文の後に `available_eggs` を返してください。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
