---
id: 6a3ea771b46a239c67628a07
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`make_fried_egg` 関数は、卵が利用可能かどうかを確認し、使った場合はローカルの卵の数を更新します。次に、その数を返して呼び出し元で使えるようにしてください。

`make_fried_egg` 関数の最後に、`return` キーワードを使って `available_eggs` 変数を返してください。この行は `if` や `else` ブロックの中にインデントされないようにしてください。

# --hints--

`make_fried_egg` 関数の最後で `available_eggs` を返す必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_return("available_eggs")
`)
})
```

`if/else` 文の後に `available_eggs` を返してください。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg")
assert node.is_ordered(str(node.find_ifs()[0]), "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

--fcc-editable-region--
    
--fcc-editable-region--
```
