---
id: 6a3e6c0eb5d050ce07781709
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`if` 条件が `False` と評価される場合（つまり十分な卵があることを意味します）、関数は `if` ブロックをスキップして残りのコードの実行を続けます。

次に、データを処理して新しい計算結果を返す必要があります。

`if` ブロックの外側（`if` キーワードと同じインデントレベル）で、`print()` 呼び出しと f-string を使って `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` を出力してください。

その次の行で、`eggs_to_use` を `available_eggs` から引いた残りの卵の数を `return` してください。

# --hints--

`use_eggs` 関数内に `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` を出力する print 文があるはずです。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

`eggs_to_use` を `available_eggs` から引いた残りの卵の数を返す必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

`print` 文の後に `available_eggs - eggs_to_use` を返してください。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
