---
id: 6a3e71a0a3a44260e675b157
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

卵の数が変わらなかった謎を解き、実際に朝食を調理するために、新しい関数を作成してください。この関数はレシピのようなもので、他の関数を呼び出して作業を完了させます。

`print(available_eggs)` の文の下に、`available_eggs` を唯一のパラメータとして受け取る空の関数 `make_fried_egg` を定義してください。

関数の本体には今のところ `pass` キーワードを使ってください。

# --hints--

`make_fried_egg` という名前の関数を定義してください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

`make_fried_egg` 関数は `available_eggs` をパラメータとして持つ必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
