---
id: 6a3e524fc3394ef52ed5b797
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

関数はそれぞれ独自のローカルスコープを持っています。これは、関数内で作成された変数が、そのキッチンのパントリーに保管された材料のようなもので、その特定の関数内でしかアクセスできないことを意味します。

しかし、関数は完全に孤立しているわけではありません。スコープの外で定義されたグローバル変数を読み取ることができます。それらの変数を使って計算を行ったり、判断をしたり、新しいローカル変数を作成したりできます。

`check_kitchen_stock` 関数内で、`pass` キーワードを新しい変数 `total_items` に置き換え、`available_eggs`、`available_flour`、`available_sugar` の合計を代入してください。

# --hints--

`check_kitchen_stock` 関数内に `total_items` という名前の変数を定義してください。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

`total_items` 変数は `available_eggs`、`available_flour`、`available_sugar` の合計である必要があります。

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
