---
id: 6a3e5ccb99fd6053af9032a1
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

キッチンに何があるか分かったので、それらの材料を使って料理ができます。これを行うには、パラメータを受け取る新しい関数を作成してください。パラメータは、関数を呼び出すときに渡すデータのプレースホルダーと考えてください。

こちらは2つのパラメータ、`a` と `b` を受け取り、その合計をターミナルに出力する関数です。

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

`check_kitchen_stock` 関数の下に、`use_eggs` という名前の空の関数を定義してください。この関数は順番に `available_eggs` と `eggs_to_use` の2つのパラメータを受け取ります。

パラメータにグローバル変数とまったく同じ名前（例えば `available_eggs`）を付けると、その関数はグローバル変数ではなく自分自身のローカルパラメータ `available_eggs` を使います。これは名前のシャドウイングと呼ばれ、ローカル変数がグローバル変数を「隠す」ためです。

関数の本体には今のところ `pass` キーワードを使ってください。

# --hints--

`use_eggs` という名前の関数を定義してください。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

`use_eggs` 関数は `available_eggs` と `eggs_to_use` をパラメータとして持つ必要があります。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
