---
id: 669e81368e52b3a5c35a2dc5
title: Erstellen Sie eine Videokompilationsseite
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `main`-Element als einziges Kind des `body`-Elements haben.
1. Sie sollten ein `h1`-Element mit dem Thema Ihrer Seite haben.
1. Sie sollten einen Absatz haben, der das Thema Ihrer Seite unter Ihrem `h1`-Element einführt.
1. Sie sollten drei `section`-Elemente unter Ihrem ersten Absatz haben.
1. Jeder Abschnitt sollte in dieser Reihenfolge ein `h2`-Element, einen Absatz und ein `iframe`-Element enthalten.
1. Die drei `iframe`-Elemente sollten ein `src`-Attribut mit einem gültigen Video haben.
1. Jedes `iframe`-Element sollte außerdem ein `title`-Attribut zur Beschreibung des eingebetteten Inhalts sowie ein `height`-Attribut und ein `width`-Attribut haben, um das Element auf eine passende Größe einzustellen.

# --hints--

Sie sollten ein `main`-Element innerhalb Ihres `body`-Elements haben.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Ihr `main`-Element sollte das einzige Kind des `body`-Elements sein.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Sie sollten ein `h1`-Element mit dem Thema Ihrer Seite innerhalb des `main`-Elements haben.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Sie sollten einen Absatz haben, der das Thema Ihrer Seite unter Ihrem `h1`-Element einführt.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Sie sollten drei `section`-Elemente unter Ihrem ersten `p`-Element haben.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Jedes `section`-Element sollte mit einem `h2`-Element beginnen, das als Titel für diesen Abschnitt dient.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Jedes `section`-Element sollte ein `p`-Element enthalten, das als zweites Kind den Videoinhalt einführt.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Jedes `section`-Element sollte ein `iframe`-Element als drittes Kind enthalten.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Ihr erstes `iframe`-Element sollte ein `src`-Attribut mit einem gültigen Video haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Ihr erstes `iframe`-Element sollte ein `title`-Attribut zur Beschreibung des eingebetteten Inhalts haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Ihr erstes `iframe`-Element sollte ein `height`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Ihr erstes `iframe`-Element sollte ein `width`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Ihr zweites `iframe`-Element sollte ein `src`-Attribut mit einem gültigen Video haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Ihr zweites `iframe`-Element sollte ein `title`-Attribut zur Beschreibung des eingebetteten Inhalts haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Ihr zweites `iframe`-Element sollte ein `height`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Ihr zweites `iframe`-Element sollte ein `width`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Ihr drittes `iframe`-Element sollte ein `src`-Attribut mit einem gültigen Video haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Ihr drittes `iframe`-Element sollte ein `title`-Attribut zur Beschreibung des eingebetteten Inhalts haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Ihr drittes `iframe`-Element sollte ein `height`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Ihr drittes `iframe`-Element sollte ein `width`-Attribut haben.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
