---
id: a97fd23d9b809dac9921074f
title: 선택적 인수를 더하는 함수 만들기
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

이 실습에서는 최대 두 개의 인수를 받아 더하는 함수를 만들지만, 인수가 하나만 있으면 두 번째 숫자를 받아 더하는 함수를 반환합니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `addTogether` 함수를 만드세요.
2. 함수가 두 개의 인수를 받으면 두 인수의 합을 반환해야 합니다.
3. 인수가 하나만 주어지면 `addTogether` 함수는 함수를 반환해야 합니다.
   - 반환된 함수가 하나의 인수로 호출되면 합을 반환해야 합니다.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. 인수 중 하나라도 유효한 숫자가 아니면 `undefined`를 반환하세요.

# --hints--

`addTogether` 함수를 만들어야 합니다.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)`는 5를 반환해야 합니다.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)`는 53.4를 반환해야 합니다.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)`는 함수를 반환해야 합니다.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)`는 12를 반환해야 합니다.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
