---
id: a202eed8fc186c8434cb6d61
title: Einen String-Inverter bauen
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

In diesem Labor bauen Sie einen einfachen String-Inverter, der die Zeichen eines gegebenen Strings umkehrt.

Zum Beispiel sollte `"hello"` zu `"olleh"` werden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bringen Sie alle Tests zum Bestehen, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `reverseString` erstellen, die einen String als Argument annimmt.
2. Die Funktion sollte den umgekehrten String zurückgeben.

# --hints--

Sie sollten eine Funktion namens `reverseString` haben.

```js
assert.isFunction(reverseString);
```

`reverseString` sollte einen String als Argument annehmen.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` sollte einen String zurückgeben.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` sollte den String `olleh` zurückgeben.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` sollte den String `ydwoH` zurückgeben.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` sollte den String `htraE morf sgniteerG` zurückgeben.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
