---
id: 68e51d8f163d3f8b70f9c4e5
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 `includes()`와 `slice()` 메서드를 다루는 연습을 하세요.

먼저, `fccSentence`라는 변수를 만들고 문자열 `"freeCodeCamp is a great place to learn web development."`을 할당하세요.

# --hints--

`fccSentence`라는 변수를 선언하세요.

```js
assert.exists(fccSentence);
```

`fccSentence` 변수는 문자열이어야 합니다.

```js
assert.isString(fccSentence);
```

`fccSentence` 변수에는 문자열 `"freeCodeCamp is a great place to learn web development."`이 할당되어야 합니다.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
