---
id: 68e6172875008fea742a63c9
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이전 수업에서 기억하듯, `includes()` 메서드는 문자열에 특정 부분 문자열이 포함되어 있는지 확인하고 `true` 또는 `false`를 반환합니다.

다음은 `includes()` 메서드를 사용한 예입니다:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

`hasFreeCodeCamp`라는 변수를 만드세요. 그런 다음 `includes()`에서 `fccSentence`이 포함되어 있는지 확인하기 위해 `"freeCodeCamp"` 메서드를 사용한 결과를 할당하세요.

# --hints--

`hasFreeCodeCamp`라는 변수를 선언해야 합니다.

```js
assert.exists(hasFreeCodeCamp);
```

`fccSentence.includes("freeCodeCamp")`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")`의 결과를 `hasFreeCodeCamp` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")`의 결과는 `true`이어야 합니다.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
