---
id: 68e6172975008fea742a63cb
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 부분 문자열이 발견되지 않았을 때 어떤 일이 일어나는지 확인할 차례입니다.

`hasJavaScript`라는 변수를 만드세요. 그리고 `includes()` 메서드를 사용해 `fccSentence`가 `"JavaScript"`를 포함하는지 확인한 결과를 할당하세요.

그다음 템플릿 리터럴이나 문자열 연결을 사용해 메시지 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`를 콘솔에 출력하세요. `<hasJavaScript>`는 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`hasJavaScript`라는 변수를 선언하세요.

```js
assert.exists(hasJavaScript);
```

`fccSentence.includes("JavaScript")`를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

`fccSentence.includes("JavaScript")`의 결과를 `hasJavaScript` 변수에 할당하세요.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

`fccSentence.includes("JavaScript")`의 결과는 `false`이어야 합니다.

```js
assert.isFalse(hasJavaScript);
```

메시지 `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`를 출력하세요. 여기서 `<hasJavaScript>`는 변수의 실제 값으로 교체해야 합니다.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
