---
id: 68e6172975008fea742a63ca
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 템플릿 리터럴이나 문자열 연결 중 하나를 사용해 메시지 `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`를 콘솔에 출력하세요. `<hasFreeCodeCamp>`은 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

메시지 `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`를 출력해야 하며, 여기서 `<hasFreeCodeCamp>`은 변수의 실제 값으로 교체해야 합니다.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
