---
id: 68e6172975008fea742a63cf
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이제 `slice()` 메서드를 사용한 결과를 기록할 차례입니다.

템플릿 리터럴이나 문자열 연결 중 하나를 사용해 메시지 `The word "<platform>" was sliced from the message.`를 콘솔에 기록하세요. `<platform>`는 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

메시지 `The word "<platform>" was sliced from the message.`를 기록해야 하며, `<platform>`는 변수의 실제 값으로 교체해야 합니다.

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
