---
id: 68e6172975008fea742a63cd
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 문자열의 일부를 추출하는 `slice()` 메서드를 살펴보겠습니다.

`message`이라는 변수를 만들고 문자열 `"Welcome to freeCodeCamp!"`를 할당하세요.

그런 다음 `console.log()`를 사용해 `"Here are some examples of the slice() method:"`를 콘솔에 출력하세요.

# --hints--

`message`라는 변수를 선언하세요.

```js
assert.exists(message);
```

`message` 변수는 문자열이어야 합니다.

```js
assert.isString(message);
```

`message` 변수는 `"Welcome to freeCodeCamp!"`이어야 합니다.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

`"Here are some examples of the slice() method:"`를 콘솔에 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
