---
id: 68e6172975008fea742a63d1
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

`slice()` 메서드는 음수 인덱스를 사용해 문자열 끝에서부터 셀 수 있습니다. `-1`은 마지막 문자를 가리킵니다.

예를 들어 보겠습니다:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

`endPunctuation`라는 변수를 만드세요. 그런 다음 `message.slice(-1)`을 사용해 마지막 문자를 추출한 결과를 할당하세요.

그다음 템플릿 리터럴이나 문자열 연결을 사용해 메시지 `The ending punctuation mark is a "<endPunctuation>"`를 콘솔에 출력하세요. `<endPunctuation>`는 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`endPunctuation`라는 변수를 선언하세요.

```js
assert.exists(endPunctuation);
```

`endPunctuation` 변수는 문자열이어야 합니다.

```js
assert.isString(endPunctuation);
```

`message.slice(-1)`을 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

`message.slice(-1)`의 결과를 `endPunctuation` 변수에 할당하세요.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

`endPunctuation` 변수는 `"!"`와 같아야 합니다.

```js
assert.strictEqual(endPunctuation, "!");
```

메시지 `The ending punctuation mark is a "<endPunctuation>"`를 출력하세요. 여기서 `<endPunctuation>`는 변수의 실제 값으로 교체해야 합니다.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
