---
id: 68e6172975008fea742a63d0
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

문자열의 시작 부분에서 추출하려면 `slice()`을 사용할 수 있습니다.

`greetingWord`이라는 변수를 만드세요. 그런 다음 `message.slice(0, 7)`를 사용해 `"Welcome"`의 첫 단어를 추출한 결과를 할당하세요.

그다음 템플릿 리터럴이나 문자열 연결 중 하나를 사용해 메시지 `The first word is "<greetingWord>".`를 콘솔에 출력하세요. `<greetingWord>`는 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`greetingWord`라는 변수를 선언해야 합니다.

```js
assert.exists(greetingWord);
```

`greetingWord` 변수는 문자열이어야 합니다.

```js
assert.isString(greetingWord);
```

`message.slice(0, 7)`을 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

`message.slice(0, 7)`의 결과를 `greetingWord` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

`greetingWord` 변수는 `"Welcome"`와 같아야 합니다.

```js
assert.strictEqual(greetingWord, "Welcome");
```

메시지 `The first word is "<greetingWord>".`를 출력해야 하며, `<greetingWord>`는 변수의 실제 값으로 교체해야 합니다.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
