---
id: 68e6172975008fea742a63ce
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

`slice()` 메서드는 문자열의 섹션을 추출합니다. 두 개의 인수를 받는데, 시작 인덱스(포함)와 끝 인덱스(미포함)입니다.

예시는 다음과 같습니다:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

문자열 인덱스는 0부터 시작합니다. `message` 문자열에서 인덱스 11의 문자는 `f`입니다("freeCodeCamp"의 시작 부분).

`platform`라는 변수를 선언하세요. 그리고 `message.slice(11, 23)`를 사용해 `"freeCodeCamp"`를 추출한 결과를 할당하세요.

# --hints--

`platform`라는 변수를 선언해야 합니다.

```js
assert.exists(platform);
```

`platform` 변수는 문자열이어야 합니다.

```js
assert.isString(platform);
```

`message.slice(11, 23)`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

`message.slice(11, 23)`의 결과를 `platform` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

`platform` 변수는 `"freeCodeCamp"`와 같아야 합니다.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
