---
id: 68e6172975008fea742a63cc
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`includes()` 메서드는 대소문자를 구분합니다.

`hasLowercaseFCC`라는 변수를 선언하세요. 그런 다음 `includes()` 메서드를 사용해 `fccSentence`이 `"freecodecamp"` (모두 소문자)를 포함하는지 확인한 결과를 할당하세요.

그다음 템플릿 리터럴이나 문자열 연결을 사용해 메시지 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`를 콘솔에 출력하세요. `<hasLowercaseFCC>`는 변수의 실제 값으로 교체하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`hasLowercaseFCC`라는 변수를 선언해야 합니다.

```js
assert.exists(hasLowercaseFCC);
```

`fccSentence.includes("freecodecamp")`를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")`의 결과를 `hasLowercaseFCC` 변수에 할당해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")`의 결과는 `false`여야 합니다.

```js
assert.isFalse(hasLowercaseFCC);
```

변수의 실제 값으로 교체된 `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`를 포함하는 메시지 `<hasLowercaseFCC>`를 콘솔에 출력해야 합니다.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
