---
id: 68ef6e42fa6e8a9a18a401fa
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`console.log()`를 사용해 `"Here are some examples of the includes() method:"`을 콘솔에 출력하세요.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`"Here are some examples of the includes() method:"`를 콘솔에 출력해야 합니다.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
