---
id: 6626b8dcf5057f896f948440
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер, коли вчитель має всю необхідну інформацію, він хоче надіслати студенту повідомлення з результатами.

Завершіть функцію `studentMsg` з параметрами `totalScores` та `studentScore`. Функція має повертати рядок, який представляє повідомлення для студента.

Якщо студент склав курс, рядок має відповідати такому формату:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

Якщо студент не склав курс, рядок має відповідати такому формату:

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

Замініть `average-goes-here` на середній бал за всіма оцінками.
Замініть `grade-goes-here` на оцінку студента.

**Поради**

- Використовуйте функцію `getAverage` для отримання середнього балу класу.
- Використовуйте функцію `getGrade` для отримання оцінки студента.
- Використовуйте конкатенацію рядків (`+`) для створення повідомлення.
- Будьте уважні з пунктуацією та пробілами у повідомленні.

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` має повертати таке повідомлення: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` має повертати таке повідомлення: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

Ваша функція `studentMsg` має повертати правильне повідомлення залежно від балу студента та середнього балу класу.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
