---
id: 662693f82c91a66be46c881b
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

Вчитель закінчив перевіряти тести своїх учнів і потребує вашої допомоги, щоб обчислити середній бал для класу.

Завершіть функцію `getAverage`, яка приймає масив результатів тестів і повертає середній бал.

Середній бал обчислюється шляхом додавання всіх результатів і ділення на загальну кількість результатів.

```md
average = sum of all scores / total number of scores
```

Для вас надано кілька викликів функції, щоб ви могли протестувати свій код.

**Поради**

- Ви можете використати цикл, щоб ітерувати над масивом `scores` і додати всі бали.
- Ви можете використати властивість `length`, щоб отримати загальну кількість результатів.

# --hints--

Ваша функція `getAverage` має повертати число.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` має повертати `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` має повертати `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` має повертати `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Ваша функція `getAverage` має повертати середній бал тестових результатів.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
