---
id: 6626b4c58c027d86478ff5eb
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Вчитель дуже задоволений програмою, яку ви створили до цього моменту. Але тепер він хоче мати простий спосіб перевірити, чи студент має прохідний бал. Прохідний бал — це будь-яке значення, що не є `"F"`.

Завершіть функцію `hasPassingGrade`, яка приймає бал студента як параметр. Ваша функція має повертати `true`, якщо студент має прохідний бал, і `false`, якщо ні.

**Поради**

- Використайте функцію `getGrade`, щоб отримати бал студента. Потім перевірте, чи є цей бал прохідним.

# --hints--

Ваша функція `hasPassingGrade` має повертати булеве значення.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Ваша функція `hasPassingGrade` має повертати `true`, якщо бал є `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Ваша функція `hasPassingGrade` має повертати `false`, якщо бал є `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Ваша функція `hasPassingGrade` має повертати `true` для всіх прохідних балів.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
