---
id: 6626a060c4006f793e10cb33
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер учителю потрібна ваша допомога, щоб перетворити оцінку студента на буквений бал.

Завершіть функцію `getGrade`, яка приймає число `score` як параметр. Ваша функція має повертати рядок, що представляє буквений бал на основі оцінки.

Ось оцінки та їх відповідні буквені бали:

| Діапазон оцінок | Бал   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |


**Поради**

- Пам’ятайте, що ви вивчали умовні інструкції (`if`, `else if`, і `else`).
- Пам’ятайте, що ви вивчали оператори порівняння (`>`, `<`, `>=`, `<=`, `===`).
- Пам’ятайте, що ви вивчали логічний оператор AND (`&&`).

# --hints--

Ваша функція `getGrade` має повертати `"A++"`, якщо оцінка `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Ваша функція `getGrade` має повертати `"A"`, якщо оцінка `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Ваша функція `getGrade` має повертати `"B"`, якщо оцінка знаходиться між `80` та `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Ваша функція `getGrade` має повертати `"C"`, якщо оцінка `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Ваша функція `getGrade` має повертати `"D"`, якщо оцінка знаходиться між `60` та `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Ваша функція `getGrade` має повертати `"F"`, якщо оцінка `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Ваша функція `getGrade` має повертати `"F"`, якщо оцінка знаходиться між `0` та `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
