---
id: 66fbcf750a62784cf11f5633
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Cree una función `getDifficultyLevel` que tome un número indicando el tiempo de cocción como parámetro.

Si el tiempo de cocción es menor o igual a `30`, la función debe devolver `"easy"`. Si es menor o igual a `60`, la función debe devolver `"medium"`. De lo contrario, la función debe devolver `"hard"`.

# --hints--

Debe crear una función `getDifficultyLevel`.

```js
assert.isFunction(getDifficultyLevel);
```

Su función `getDifficultyLevel` debe tener un solo parámetro.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

Su función `getDifficultyLevel` debe devolver `"easy"` cuando el tiempo de cocción es menor o igual a `30`.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

Su función `getDifficultyLevel` debe devolver `"medium"` cuando el tiempo de cocción es mayor que `31` y menor o igual a `60`.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

Su función `getDifficultyLevel` debe devolver `"hard"` cuando el tiempo de cocción es mayor que `60`.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
