---
id: 66fbcf750a62784cf11f5630
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Ahora debes insertar los tres objetos en el array `recipes`. Para hacer esto, puedes usar el método `push()`.

Usa el método `push()` para insertar todos los objetos de la receta en el array `recipes`. Asegúrate de insertar `recipe1`, `recipe2` y `recipe3` en ese orden.

También elimina las variables `recipe1Name`, `recipe2CookingTime` y `recipe3Ingredients`, y las sentencias `console.log` que registran esas variables.

# --hints--

Debes agregar todos los objetos de receta al arreglo `recipes`. Asegúrate de agregarlos en el orden en que están declarados.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

Debes eliminar la variable `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Debes eliminar la variable `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Debes eliminar la variable `recipe3Ingredients`.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

Debes eliminar la instrucción `console.log(recipe3Ingredients);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
