---
id: 66fbcf750a62784cf11f5631
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

En los próximos pasos trabajarás en el cálculo de la calificación promedio, el total de ingredientes y el nivel de dificultad para cada receta del array `recipes`.

Comienza creando una función `getAverageRating` que acepte un solo argumento, que es un array con calificaciones. Dentro de la función, calcula la calificación media del array pasado a la función.

Tu función `getAverageRating` debe devolver un número.

# --hints--

Debes crear una función `getAverageRating`.

```js
assert.isFunction(getAverageRating)
```

Tu función `getAverageRating` debe tener un solo parámetro.

```js
assert.lengthOf(getAverageRating, 1);
```

Tu función `getAverageRating` debe devolver un número.

```js
assert.isNumber(getAverageRating(recipe1.ratings))
```

`getAverageRating(recipe1.ratings)` debe devolver `4.5`.

```js
assert.strictEqual(getAverageRating(recipe1.ratings), 4.5)
```

`getAverageRating(recipe2.ratings)` debe devolver `4.75`.

```js
assert.strictEqual(getAverageRating(recipe2.ratings), 4.75)
```

`getAverageRating(recipe3.ratings)` debe devolver `4`.

```js
assert.strictEqual(getAverageRating(recipe3.ratings), 4)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
