---
id: 66fbcf750a62784cf11f562e
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Crea un objeto `recipe2` con las siguientes propiedades y valores:

| Clave | Valor   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

Debes crear un objeto llamado `recipe2`.

```js
assert.isObject(recipe2);
```

Tu objeto `recipe2` debe tener una propiedad `name`.

```js
assert.property(recipe2, "name");
```

Tu propiedad `name` debe establecerse en `Chicken Curry`.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

Tu objeto `recipe2` debe tener una propiedad `ingredients`.

```js
assert.property(recipe2, "ingredients");
```

La propiedad `ingredients` debe tener un valor de arreglo.

```js
assert.isArray(recipe2.ingredients);
```

El valor de la matriz de tu propiedad `ingredients` debe incluir `chicken breast`, `coconut milk`, `curry powder`, `onion` y `garlic`.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

Tu objeto `recipe2` debe tener una propiedad `cookingTime`.

```js
assert.property(recipe2, "cookingTime");
```

El valor de tu propiedad `cookingTime` debe ser un número.

```js
assert.isNumber(recipe2.cookingTime);
```

Tu propiedad `cookingTime` debe establecerse en `42`.

```js
assert.equal(recipe2.cookingTime, 42);
```

Tu objeto `recipe2` debe tener una propiedad `totalIngredients`.

```js
assert.property(recipe2, "totalIngredients");
```

Tu propiedad `totalIngredients` debe establecerse en `null`.

```js
assert.isNull(recipe2.totalIngredients);
```

Tu objeto `recipe2` debe tener una propiedad `difficultyLevel`.

```js
assert.property(recipe2, "difficultyLevel");
```

Tu propiedad `difficultyLevel` debe establecerse en una cadena vacía.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
