---
id: 66fbcf750a62784cf11f5630
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Se ha completado un tercer objeto `recipe3` para ti. Tiene las mismas propiedades que `recipe1` y `recipe2`.

Ahora debes insertar los tres objetos en el array `recipes`. Para hacer esto, puedes usar el método `push()`.

Usa el método `push()` para insertar todos los objetos de la receta en el array `recipes`. Asegúrate de insertar `recipe1`, `recipe2` y `recipe3` en ese orden.

También elimina las variables `recipe1Name`, `recipe2Name`, `recipe1CookingTime` y `recipe2CookingTime`, y las sentencias `console.log` que registran esas variables.

# --hints--

Debes eliminar la variable `recipe1Name`.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe1Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

Debes eliminar la variable `recipe2Name`.

```js
assert.throws(() => recipe2Name, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe2Name);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2Name\s*\)\s*;?/);
```

Debes eliminar la variable `recipe1CookingTime`.

```js
assert.throws(() => recipe1CookingTime, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe1CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1CookingTime\s*\)\s*;?/);
```

Debes eliminar la variable `recipe2CookingTime`.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

Debes eliminar la sentencia `console.log(recipe2CookingTime);`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

Debes insertar todos los objetos de recetas en el array `recipes`. Asegúrate de insertarlos en el orden en que se declaran.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2Name = recipe2.name;
console.log(recipe2Name);

const recipe1CookingTime = recipe1.cookingTime;
console.log(recipe1CookingTime);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);
--fcc-editable-region--
```
