---
id: 66fbcf750a62784cf11f5632
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Cree una función `getTotalIngredients` que tome un solo argumento, que represente un array con ingredientes, y devuelva el número de ingredientes del array pasado a la función.

# --hints--

Debe crear una función `getTotalIngredients`.

```js
assert.isFunction(getTotalIngredients)
```

La función `getTotalIngredients` debe tener un solo parámetro.

```js
assert.lengthOf(getTotalIngredients, 1);
```

La función `getTotalIngredients` debe devolver un número.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients))
```

`getTotalIngredients(recipe1.ingredients)` debe devolver `4`.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)` debe devolver `5`.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)` debe devolver `3`.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
