---
id: 66fbcf750a62784cf11f5634
title: Paso 10
challengeType: 1
dashedName: step-10
---

# --description--

Es hora de probar cada una de las funciones. Puedes usar cualquiera de las recetas para esto, pero este tutorial comenzará con `recipe1`.

Crea tres nuevas variables: `recipe1AverageRating`, `recipe1TotalIngredients`, y `recipe1DifficultyLevel`. Asígnales los valores llamando a la función correspondiente para cada variable y pasando la propiedad `recipe1` adecuada.

Finalmente, registra cada variable en la consola para ver los resultados.

# --hints--

Debe crear una variable `recipe1AverageRating`.

```js
assert.isNotNull(recipe1AverageRating)
```

Su variable `recipe1AverageRating` debe establecerse en `getAverageRating(recipe1.ratings)`.

```js
assert.equal(recipe1AverageRating, getAverageRating(recipe1.ratings));
```

Debe registrar `recipe1AverageRating` en la consola.

```js
assert.match(code, /console\.log\(recipe1AverageRating\);?/)
```

Debe crear una variable `recipe1TotalIngredients`.

```js
assert.isNotNull(recipe1TotalIngredients)
```

Su variable `recipe1TotalIngredients` debe establecerse en `getTotalIngredients(recipe1.ingredients)`.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

Debe registrar `recipe1TotalIngredients` en la consola.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/)
```

Debe crear una variable `recipe1DifficultyLevel`.

```js
assert.isNotNull(recipe1DifficultyLevel)
```

Su variable `recipe1DifficultyLevel` debe establecerse en `getDifficultyLevel(recipe1.cookingTime)`.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

Debe registrar `recipe1DifficultyLevel` en la consola.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length;
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return 'easy';
  } else if (cookingTime <= 60) {
    return 'medium';
  } else {
    return 'hard';
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
