---
id: 66fbcf750a62784cf11f5635
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

Ahora puedes llenar cada elemento del arreglo `recipes` con valores para las propiedades `totalIngredients` y `difficultyLevel`.

Por ahora, accede a `totalIngredients` y `difficultyLevel` de `recipe1` y asígnalos a los resultados apropiados de llamadas a funciones y argumentos.

# --hints--

Debe acceder a la propiedad `totalIngredients` de `recipe1`.

```js
assert.isNotNull(recipe1.totalIngredients);
```

Debe asignar el resultado de llamar a `getTotalIngredients` con `recipe1.ingredients` a la propiedad `totalIngredients` de `recipe1`.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

Debe acceder a la propiedad `difficultyLevel` de `recipe1`.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

Debe asignar el resultado de llamar a `getDifficultyLevel` con `recipe1.cookingTime` a la propiedad `cookingTime` de `recipe1`.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
