---
id: 66fbcf750a62784cf11f562d
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Crea un objeto llamado `recipe1`. Dentro del objeto `recipe1`, crea una propiedad `name` con el valor `"Spaghetti Carbonara"`.

También dentro del objeto `recipe1`, crea una propiedad `ingredients` con un arreglo como valor. El arreglo debe contener `spaghetti`, `queso parmesano`, `panceta` y `pimienta negra`.

# --hints--

Debe crear un objeto llamado `recipe1`.

```js
assert.isObject(recipe1);
```

El objeto `recipe1` debe tener una propiedad `name`.

```js
assert.property(recipe1, 'name');
```

La propiedad `name` debe establecerse en `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

El objeto `recipe1` debe tener una propiedad `ingredients`.

```js
assert.property(recipe1, 'ingredients');
```

La propiedad `ingredients` debe tener un valor de arreglo.

```js
assert.isArray(recipe1.ingredients);
```

El valor del arreglo de la propiedad `ingredients` debe contener `spaghetti`, `queso parmesano`, `panceta` y `pimienta negra`.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
