---
id: 66ff02792b3042a2b811f475
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Agregue las siguientes propiedades al objeto `recipe1`:

ClaveValor
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

Las propiedades con valores `null` y cadenas vacías se actualizarán más tarde después de que las calcule.

# --hints--

El objeto `recipe1` debe tener una propiedad `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

La propiedad `cookingTime` debe establecerse en `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

El objeto `recipe1` debe tener una propiedad `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

Tu propiedad `totalIngredients` debe establecerse en `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

El objeto `recipe1` debe tener una propiedad `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

Tu propiedad `difficultyLevel` debe establecerse en una cadena vacía.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
