---
id: 66ff02792b3042a2b811f475
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Agregue las siguientes propiedades al objeto `recipe1`:

| Clave              | Valor  |
| ------------------ | ------ |
| `cookingTime`      | `22`   |
| `totalIngredients` | `null` |
| `difficultyLevel`  | `""`   |
| `averageRating`    | `null` |

Las propiedades con valores `null` y cadenas vacías se actualizarán más tarde después de que las calcule.

# --hints--

El objeto `recipe1` debe tener una propiedad `cookingTime`.

```js
assert.property(recipe1, "cookingTime");
```

La propiedad `cookingTime` debe establecerse en `22`.

```js
assert.equal(recipe1.cookingTime, 22);
```

El objeto `recipe1` debe tener una propiedad `totalIngredients`.

```js
assert.property(recipe1, "totalIngredients");
```

La propiedad `totalIngredients` debe establecerse en `null`.

```js
assert.isNull(recipe1.totalIngredients);
```

El objeto `recipe1` debe tener una propiedad `difficultyLevel`.

```js
assert.property(recipe1, "difficultyLevel");
```

La propiedad `difficultyLevel` debe establecerse en una cadena vacía.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

El objeto `recipe1` debe tener una propiedad `averageRating`.

```js
assert.property(recipe1, "averageRating");
```

La propiedad `averageRating` debe establecerse en `null`.

```js
assert.isNull(recipe1.averageRating);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  ratings: [4, 5, 4, 5],
--fcc-editable-region--

--fcc-editable-region--
};
```
