---
id: 670e4f45f7116c0f216a5177
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Antes de continuar, debe practicar cómo acceder a propiedades de un objeto.

Puede utilizar notación de punto (`.`) o corchete (`[]`) para hacer esto. Aquí tienes un ejemplo:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Accede a la propiedad `name` de `recipe1` y asígnala a la variable `recipe1Name`.

Luego, accede a la propiedad `cookingTime` de `recipe2` y asígnala a la variable `recipe2CookingTime`.

Finalmente, accede a la propiedad `ingredients` de `recipe3` y asígnala a la variable `recipe3Ingredients`.

Asegúrese de que todas las variables que creó se registren en la consola.

# --hints--

Debe crear una variable `recipe1Name`.

```js
assert.isDefined(recipe1Name);
```

Debe asignar el valor de la propiedad `name` de `recipe1` a su variable `recipe1Name`.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

Debe registrar `recipe1Name` en la consola.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

Debe crear una variable `recipe2CookingTime`.

```js
assert.isDefined(recipe2CookingTime);
```

Debe asignar el valor de la propiedad `cookingTime` de `recipe2` a su variable `recipe2CookingTime`.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

Debe registrar `recipe2CookingTime` en la consola.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

Debes crear una variable `recipe3Ingredients`.

```js
assert.isDefined(recipe3Ingredients);
```

Debes asignar el valor de la propiedad `ingredients` de `recipe3` a tu variable `recipe3Ingredients`.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

Debes registrar `recipe3Ingredients` en la consola.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
