---
id: 69161f4c0c9c8c48b3c843c5
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Crea un objeto `recipe3` con las siguientes propiedades y valores:

|Clave|Valor|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

Debes crear un objeto llamado `recipe3`.

```js
assert.isObject(recipe3);
```

Tu objeto `recipe3` debe tener una propiedad `name`.

```js
assert.property(recipe3, "name");
```

La propiedad `name` debe estar configurada en `Vegetable Stir Fry`.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

Tu objeto `recipe3` debe tener una propiedad `ingredients`.

```js
assert.property(recipe3, "ingredients");
```

La propiedad `ingredients` debe tener un valor de arreglo.

```js
assert.isArray(recipe3.ingredients);
```

El valor de arreglo de tu propiedad `ingredients` debe contener `broccoli`, `carrot` y `bell pepper`.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

Tu objeto `recipe3` debe tener una propiedad `cookingTime`.

```js
assert.property(recipe3, "cookingTime");
```

El valor de tu propiedad `cookingTime` debe ser un número.

```js
assert.isNumber(recipe3.cookingTime);
```

La propiedad `cookingTime` debe estar configurada en `15`.

```js
assert.equal(recipe3.cookingTime, 15);
```

Tu objeto `recipe3` debe tener una propiedad `totalIngredients`.

```js
assert.property(recipe3, "totalIngredients");
```

Tu propiedad `totalIngredients` debe establecerse en `null`.

```js
assert.isNull(recipe3.totalIngredients);
```

Tu objeto `recipe3` debe tener una propiedad `difficultyLevel`.

```js
assert.property(recipe3, "difficultyLevel");
```

Tu propiedad `difficultyLevel` debe establecerse en una cadena vacía.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
