---
id: 587d8255367417b2b2512c74
title: Unda darasa la foleni ya kipaumbele
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

Katika changamoto hii utaunda foleni ya kipaumbele. Foleni ya kipaumbele ni aina maalum ya foleni ambapo vitu vinaweza kuwa na taarifa za ziada zinazobainisha kipaumbele chao. Hii inaweza kuwakilishwa kwa urahisi kwa nambari kamili. Kipaumbele cha kitu kitaghairi mpangilio wa kuwekwa katika kuamua mfululizo wa vitu vinavyotolewa kutoka kwenye foleni. Ikiwa kitu chenye kipaumbele cha juu kitawekwa baada ya vitu vyenye kipaumbele cha chini, kitu chenye kipaumbele cha juu kitatolewa kwanza kabla ya vingine vyote.

Kwa mfano, fikiria tuna foleni ya kipaumbele yenye vitu vitatu:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Hapa thamani ya pili (nambari kamili) inawakilisha kipaumbele cha kitu. Ikiwa tutaweka `['human', 1]` na kipaumbele cha `1` (tukizingatia kipaumbele cha chini kinapewa kipaumbele cha juu) basi kitakuwa kitu cha kwanza kutolewa kutoka kwenye foleni. Mkusanyiko utakuwa kama huu:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Tumeanza kuandika `PriorityQueue` katika mhariri wa msimbo. Utahitaji kuongeza njia ya `enqueue` kwa kuongeza vitu vyenye kipaumbele, njia ya `dequeue` kwa kuondoa na kurudisha vitu, njia ya `size` kurudisha idadi ya vitu katika foleni, njia ya `front` kurudisha kipengele kilicho mbele ya foleni, na hatimaye njia ya `isEmpty` itakayorudisha `true` ikiwa foleni ni tupu au `false` ikiwa si tupu.

`enqueue` inapaswa kupokea vitu kwa muundo ulioonyeshwa hapo juu (`['human', 1]`) ambapo `1` inawakilisha kipaumbele. `dequeue` na `front` zinapaswa kurudisha jina la kitu tu, si kipaumbele chake.

# --hints--

Darasa lako la `PriorityQueue` linapaswa kuwa na njia ya `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

Darasa lako la `PriorityQueue` linapaswa kuwa na njia ya `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

Darasa lako la `PriorityQueue` linapaswa kuwa na njia ya `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

Darasa lako la `PriorityQueue` linapaswa kuwa na njia ya `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

Darasa lako la `PriorityQueue` linapaswa kuwa na njia ya `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Darasa lako la `PriorityQueue` linapaswa kufuatilia kwa usahihi idadi ya sasa ya vitu kwa kutumia njia ya `size` wakati vitu vinawekwa na kutolewa kwenye foleni.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Njia ya `front` inapaswa kurudisha kitu sahihi kilicho mbele ya foleni wakati vitu vinawekwa na kutolewa.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Njia ya `isEmpty` inapaswa kurudisha `true` wakati foleni ni tupu.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

Foleni ya kipaumbele inapaswa kurudisha vitu vyenye kipaumbele cha juu kabla ya vitu vyenye kipaumbele cha chini na kurudisha vitu kwa mpangilio wa kwanza-kuingia-kwa-kutoka (FIFO) vinginevyo.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
