---
id: 587d8257367417b2b2512c7d
title: Tafuta kiwango cha chini na cha juu cha mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--

Katika changamoto ya mwisho tuliandika hali ambapo mti unaweza kuwa usio na usawa. Ili kuelewa dhana ya usawa, tuchunguze sifa nyingine ya mti: urefu. Urefu katika mti unawakilisha umbali kutoka kwa nodi mzizi hadi nodi yoyote ya majani. Njia tofauti katika muundo wa mti wenye matawi mengi zinaweza kuwa na urefu tofauti, lakini kwa mti fulani kutakuwa na kiwango cha chini na cha juu cha urefu. Ikiwa mti ni sawa, thamani hizi zitakuwa tofauti kwa kiwango cha juu cha moja. Hii ina maana kwamba katika mti ulio sawa, nodi zote za majani zipo katika ngazi moja, au kama hazipo katika ngazi moja basi ziko angalau ngazi moja mbali.

Sifa ya usawa ni muhimu kwa miti kwa sababu ndiyo huamua ufanisi wa operesheni za mti. Kama tulivyoeleza katika changamoto ya mwisho, tunakumbana na ugumu wa wakati mbaya zaidi kwa miti isiyo na usawa kabisa. Miti yenye usawa wa kujitengenezea hutumika sana kushughulikia tatizo hili katika miti yenye seti za data zinazobadilika. Mifano ya kawaida ni miti ya AVL, miti nyekundu-nyeusi, na miti ya B. Miti hii yote ina mantiki ya ndani ya ziada inayorekebisha usawa wa mti wakati kuingiza au kufuta kunasababisha hali ya kutokuwa na usawa.

**Kumbuka:** Sifa inayofanana na urefu ni kina, ambayo inahusu umbali wa nodi fulani kutoka kwa nodi mzizi.

# --instructions--

Andika njia mbili kwa ajili ya mti wetu wa kibainari: `findMinHeight` na `findMaxHeight`. Njia hizi zinapaswa kurudisha thamani ya nambari kamili kwa kiwango cha chini na cha juu cha urefu ndani ya mti wa kibainari uliotolewa, mtawaliwa. Ikiwa nodi ni tupu, naomba weka thamani ya urefu ya `-1` (hiyo ni kesi ya msingi). Mwishowe, ongeza njia ya tatu `isBalanced` ambayo inarudisha `true` au `false` kulingana na kama mti ni sawa au la. Unaweza kutumia njia mbili ulizoandika hapo awali kuamua hili.

# --hints--

Muundo wa data `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `findMinHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `findMaxHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `isBalanced`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

Njia ya `findMinHeight` inapaswa kurudisha kiwango cha chini cha urefu wa mti.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

Njia ya `findMaxHeight` inapaswa kurudisha kiwango cha juu cha urefu wa mti.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

Mti tupu unapaswa kurudisha urefu wa `-1`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

Njia ya `isBalanced` inapaswa kurudisha `false` ikiwa mti ni mti wa utafutaji wa kibainari usio na usawa.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

Njia ya `isBalanced` inapaswa kurudisha `true` ikiwa mti ni mti wa utafutaji wa kibainari ulio sawa.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```
