---
id: 587d8258367417b2b2512c81
title: Futa nodi yenye mtoto mmoja katika mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301638
dashedName: delete-a-node-with-one-child-in-a-binary-search-tree
---

# --description--

Sasa tunapoweza kufuta nodi za majani, twende kwenye kesi ya pili: kufuta nodi yenye mtoto mmoja. Kwa kesi hii, tuseme tuna mti wenye nodi zifuatazo 1 — 2 — 3 ambapo 1 ni mzizi. Ili kufuta 2, tunahitaji tu kufanya rejea ya kulia katika 1 kuelekeze kwenye 3. Kwa ujumla zaidi, kufuta nodi yenye mtoto mmoja, tunafanya nodi ya mzazi wa nodi hiyo kurejea nodi inayofuata katika mti.

# --instructions--

Tumetoa msimbo katika njia yetu ya `remove` inayotekeleza majukumu kutoka kwenye changamoto iliyopita. Tunapata nodi inayolengwa kufutwa na nodi ya mzazi wake na kuamua idadi ya watoto wa nodi inayolengwa. Sasa tutaongeza kesi inayofuata hapa kwa nodi zinazolengwa zenye mtoto mmoja tu. Hapa, tutalazimika kubaini kama mtoto huyo mmoja ni tawi la kushoto au la kulia katika mti kisha kuweka rejea sahihi katika mzazi kuelekeza nodi hii. Zaidi ya hayo, tuchukulie pia kesi ambapo nodi inayolengwa ni mzizi (hii inamaanisha nodi ya mzazi itakuwa `null`). Tafadhali badilisha msimbo wa mwanzo wote kwa wako kama unavyotaka mradi tu unapita vipimo.

# --hints--

Muundo wa data wa `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Kujaribu kuondoa kipengele ambacho hakipo kunapaswa kurudisha `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Kama nodi ya mzizi haina watoto, kufuta inapaswa kuweka mzizi kuwa `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Njia ya `remove` inapaswa kuondoa nodi za majani kutoka kwenye mti.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

Njia ya `remove` inapaswa kuondoa nodi zenye mtoto mmoja.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(1);
    test.add(4);
    test.add(3);
    test.add(2);
    test.add(6);
    test.add(8);
    test.remove(6);
    test.remove(3);
    return test.inorder().join('') == '1248';
  })()
);
```

Kuondoa mzizi katika mti wenye nodi mbili kunapaswa kuweka nodi ya pili kuwa mzizi.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(27);
    test.remove(15);
    return test.inorder().join('') == '27';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  this.remove = function(value) {
    if (this.root === null) {
      return null;
    }
    var target;
    var parent = null;
    // Find the target value and its parent
    (function findValue(node = this.root) {
      if (value == node.value) {
        target = node;
      } else if (value < node.value && node.left !== null) {
        parent = node;
        return findValue(node.left);
      } else if (value < node.value && node.left === null) {
        return null;
      } else if (value > node.value && node.right !== null) {
        parent = node;
        return findValue(node.right);
      } else {
        return null;
      }
    }.bind(this)());
    if (target === null) {
      return null;
    }
    // Count the children of the target to delete
    var children =
      (target.left !== null ? 1 : 0) + (target.right !== null ? 1 : 0);
    // Case 1: Target has no children
    if (children === 0) {
      if (target == this.root) {
        this.root = null;
      } else {
        if (parent.left == target) {
          parent.left = null;
        } else {
          parent.right = null;
        }
      }
    }
    // Case 2: Target has one child
    // Only change code below this line
  };
}
```

# --solutions--

```js
// solution required
```
