---
id: 587d8255367417b2b2512c73
title: Tumia opereta ya spread na maelezo kwa muunganisho wa ES5 Set()
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

Je, unakumbuka opereta ya spread ya ES6 `...`?

`...` inaweza kuchukua vitu vinavyoweza kurudiwa katika ES6 na kuvitengeneza kuwa safu za data.

Tufanye Set, na tuchunguze kitendakazi cha spread.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

Katika zoezi hili tutapitisha kitu cha set kwa kitendakazi cha `checkSet`. Kitapaswa kurudisha safu ya data yenye thamani za Set.

Sasa umejifunza kwa mafanikio jinsi ya kutumia kitu cha `Set()` cha ES6, kazi nzuri!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` inapaswa kurudisha `[1, 2, 3, 4, 5, 6, 7]`.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
