---
id: 587d8250367417b2b2512c60
title: Unda darasa la foleni
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

Kama vile stacki, foleni ni mkusanyiko wa vipengele. Lakini tofauti na stacki, foleni hufuata kanuni ya FIFO (Kwanza Kuingia Kwanza Kutoka). Vipengele vinavyoongezwa kwenye foleni huingizwa kwenye mkia, au mwisho, wa foleni, na kipengele pekee kilicho mbele ya foleni ndicho kinachoruhusiwa kuondolewa.

Tunaweza kutumia safu ya data kuwakilisha foleni, lakini kama vile stacki, tunataka kupunguza kiwango cha udhibiti tunachoweza kuwa nacho juu ya foleni zetu.

Njia kuu mbili za darasa la foleni ni njia ya enqueue na njia ya dequeue. Njia ya enqueue huingiza kipengele kwenye mkia wa foleni, na njia ya dequeue huondoa na kurudisha kipengele kilicho mbele ya foleni. Njia nyingine muhimu ni njia za front, size, na isEmpty.

# --instructions--

Andika njia ya `enqueue` ambayo huingiza kipengele kwenye mkia wa foleni, njia ya `dequeue` ambayo huondoa na kurudisha kipengele cha mbele, njia ya `front` ambayo huturuhusu kuona kipengele cha mbele, njia ya `size` ambayo inaonyesha urefu, na njia ya `isEmpty` ya kuangalia kama foleni ni tupu.

# --hints--

Darasa lako la `Queue` linapaswa kuwa na njia ya `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

Darasa lako la `Queue` linapaswa kuwa na njia ya `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

Darasa lako la `Queue` linapaswa kuwa na njia ya `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

Darasa lako la `Queue` linapaswa kuwa na njia ya `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

Darasa lako la `Queue` linapaswa kuwa na njia ya `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Njia ya `dequeue` inapaswa kuondoa na kurudisha kipengele cha mbele cha foleni

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Njia ya `front` inapaswa kurudisha thamani ya kipengele cha mbele cha foleni

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Njia ya `size` inapaswa kurudisha urefu wa foleni

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Njia ya `isEmpty` inapaswa kurudisha `false` ikiwa kuna vipengele katika foleni

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
