---
id: 587d825a367417b2b2512c87
title: Unda orodha iliyounganishwa mara mbili
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

Orodha zote zilizounganishwa tulizounda hadi sasa ni orodha zilizounganishwa mara moja. Hapa, tutaunda <dfn>orodha iliyounganishwa mara mbili</dfn>. Kama jina linavyosema, nodi katika orodha iliyounganishwa mara mbili zina rejea kwa nodi inayofuata na ile iliyotangulia katika orodha.

Hii inatuwezesha kusafiri katika orodha kwa mwelekeo yote miwili lakini pia inahitaji kumbukumbu zaidi kutumika kwa sababu kila nodi lazima iwe na rejea ya ziada kwa nodi iliyotangulia katika orodha.

# --instructions--

Tumetoa kitu cha `Node` na kuanza `DoublyLinkedList` yetu. Tuweke njia mbili katika orodha yetu iliyounganishwa mara mbili zinazoitwa `add` na `remove`. Njia ya `add` inapaswa kuongeza kipengele kilichotolewa kwenye orodha wakati njia ya `remove` inapaswa kuondoa matukio yote ya kipengele kilichotolewa katika orodha.

Kuwa makini kushughulikia kesi zozote za pembezoni unapoandika njia hizi, kama vile kufuta kipengele cha kwanza au cha mwisho. Pia, kuondoa kipengele chochote kwenye orodha tupu kunapaswa kurudisha `null`.

# --hints--

Muundo wa data wa `DoublyLinkedList` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` inapaswa kuwa na njia inayoitwa `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList` inapaswa kuwa na njia inayoitwa `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Kuondoa kipengele kutoka kwenye orodha tupu kunapaswa kurudisha `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Njia ya `add` inapaswa kuongeza vitu kwenye orodha.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

Kila nodi inapaswa kufuatilia nodi iliyotangulia.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

Kipengele cha kwanza kinapaswa kuondolewa kutoka kwenye orodha.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

Kipengele cha mwisho kinapaswa kuondolewa kutoka kwenye orodha.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
