---
id: 587d8250367417b2b2512c5f
title: Unda darasa la stack
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

Katika sehemu ya mwisho, tulizungumza kuhusu nini stack ni na jinsi tunavyoweza kutumia safu ya data kuwakilisha stack. Katika sehemu hii, tutaunda darasa letu la stack. Ingawa unaweza kutumia safu za data kuunda stack, wakati mwingine ni bora kupunguza kiwango cha udhibiti tunachokuwa nacho na stack zetu. Mbali na njia za `push` na `pop`, stack zina njia nyingine muhimu. Tuweke njia za `peek`, `isEmpty`, na `clear` katika darasa letu la stack.

# --instructions--

Andika njia ya `push` ambayo inasukuma kipengele juu ya stack, njia ya `pop` ambayo inatoa na kurudisha kipengele kilicho juu ya stack, njia ya `peek` ambayo huangalia kipengele cha juu katika stack, njia ya `isEmpty` ambayo inakagua kama stack ni tupu, na njia ya `clear` ambayo huondoa vipengele vyote kutoka kwenye stack. Kawaida stack hazina hii, lakini tumeongeza njia ya msaada ya `print` ambayo huandika kwenye konsoli mkusanyiko.

# --hints--

Darasa lako la `Stack` linapaswa kuwa na njia ya `push`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

Darasa lako la `Stack` linapaswa kuwa na njia ya `pop`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

Darasa lako la `Stack` linapaswa kuwa na njia ya `peek`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

Darasa lako la `Stack` linapaswa kuwa na njia ya `isEmpty`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

Darasa lako la `Stack` linapaswa kuwa na njia ya `clear`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

Njia ya `peek` inapaswa kurudisha kipengele cha juu cha stack

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

Njia ya `pop` inapaswa kuondoa na kurudisha kipengele cha juu cha stack

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

Njia ya `isEmpty` inapaswa kurudisha kweli ikiwa stack haina vipengele vyovyote

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

Njia ya `clear` inapaswa kuondoa vipengele vyote kutoka kwenye stack

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
