---
id: 587d8257367417b2b2512c7b
title: Ongeza kipengele kipya katika mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

Mfululizo huu wa changamoto utaanzisha muundo wa data wa mti. Mitit ni muundo muhimu na wenye matumizi mengi katika sayansi ya kompyuta. Bila shaka, jina lake linatokana na ukweli kwamba unapoangalia picha yake, huonekana kama miti tunayojua katika dunia halisi. Muundo wa data wa mti huanza na nodi moja, kawaida huitwa mzizi, na kutoka hapo hutawanyika hadi nodi nyingine, kila moja ikiwa na nodi za watoto zaidi, na kadhalika. Muundo huu wa data kawaida huonyeshwa ukiwa na nodi ya mzizi juu; unaweza kuufikiria kama mti wa asili ulioangushwa kichwa chini.

Kwanza, tuelezee baadhi ya istilahi za kawaida tutakazokutana nazo kuhusu miti. Nodi ya mzizi ni juu ya mti. Pointi za data katika mti huitwa nodi. Nodi zenye matawi yanayoelekea nodi nyingine huitwa wazazi wa nodi ile tawi linaelekea (nodi ya mtoto). Istilahi nyingine za kifamilia zinatumika kama unavyotarajia. Mti mdogo unahusu mazao yote ya nodi fulani, matawi yanaweza kuitwa edges, na nodi za majani ni nodi za mwisho wa mti ambazo hazina watoto. Mwisho, kumbuka kwamba miti ni miundo ya data yenye muundo wa kurudiwa (recursive). Yaani, watoto wa nodi ni wazazi wa mti mdogo wao, na kadhalika. Asili ya kurudiwa ya miti ni muhimu kuelewa wakati wa kubuni algoriti za shughuli za kawaida za mti.

Kuanza, tutajadili aina maalum ya mti, mti wa kibainari. Kwa kweli, tutajadili mti wa kibainari wa utafutaji. Hebu tuelezee maana yake. Ingawa muundo wa data wa mti unaweza kuwa na matawi yoyote kwa nodi moja, mti wa kibainari unaweza kuwa na matawi mawili tu kwa kila nodi. Zaidi ya hayo, mti wa utafutaji wa kibainari umepangwa kulingana na miti mdogo ya watoto, kwa hivyo thamani ya kila nodi katika mti mdogo wa kushoto ni ndogo au sawa na thamani ya nodi ya mzazi, na thamani ya kila nodi katika mti mdogo wa kulia ni kubwa au sawa na thamani ya nodi ya mzazi. Ni msaada sana kuona uhusiano huu ili kuelewa vyema:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Sasa uhusiano huu uliopangwa ni rahisi sana kuonekana. Kumbuka kwamba kila thamani upande wa kushoto ya 8, nodi ya mzizi, ni ndogo kuliko 8, na kila thamani upande wa kulia ni kubwa kuliko 8. Pia angalia kwamba uhusiano huu unatumika kwa kila mti mdogo pia. Kwa mfano, mtoto wa kwanza wa kushoto ni mti mdogo. 3 ni nodi ya mzazi, na ina nodi za watoto wawili tu — kwa sheria za miti ya utafutaji wa kibainari, tunajua bila hata kuangalia kwamba mtoto wa kushoto wa nodi hii (na watoto wake wote) atakuwa mdogo kuliko 3, na mtoto wa kulia (na watoto wake wote) atakuwa mkubwa kuliko 3 (lakini pia mdogo kuliko thamani ya mzizi wa muundo), na kadhalika.

Mitit ya utafutaji wa kibainari ni miundo ya data inayopatikana sana na yenye manufaa kwa sababu hutoa muda wa logarithmic katika hali ya wastani kwa shughuli kadhaa za kawaida kama utafutaji, kuingiza, na kufuta.

# --instructions--

Tutaanza kwa urahisi. Tumeelezea muundo wa msingi wa mti wa utafutaji wa kibainari hapa pamoja na kitendakazi cha kuunda nodi kwa mti wetu. Angalia kwamba kila nodi inaweza kuwa na thamani ya kushoto na kulia. Hizi zitapewa miti midogo ya watoto ikiwa ipo. Katika mti wetu wa utafutaji wa kibainari, utaunda njia ya kuongeza thamani mpya kwenye mti. Njia hii inapaswa kuitwa `add` na inapaswa kupokea thamani ya nambari kamili kuongeza kwenye mti. Chukua tahadhari kudumisha kanuni ya mti wa utafutaji wa kibainari: thamani katika kila mtoto wa kushoto inapaswa kuwa ndogo au sawa na thamani ya mzazi, na thamani katika kila mtoto wa kulia inapaswa kuwa kubwa au sawa na thamani ya mzazi. Hapa, tufanye mti wetu usiweze kushikilia thamani zinazojirudia. Ikiwa tutajaribu kuongeza thamani ambayo tayari ipo, njia inapaswa kurudisha `null`. Vinginevyo, ikiwa kuongeza itafanikiwa, `undefined` inapaswa kurudishwa.

**Kidokezo:** miti ni miundo ya data yenye muundo wa kurudiwa asili!

# --hints--

Muundo wa data wa `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia inayoitwa `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Njia ya kuongeza inapaswa kuongeza vipengele kulingana na sheria za mti wa utafutaji wa kibainari.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Kuongeza kipengele ambacho tayari kipo kunapaswa kurudisha `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```
