---
id: 587d825b367417b2b2512c8d
title: Unda ramani ya ES6 JavaScript
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

Toleo jipya la JavaScript linatupatia kitu cha ramani (Map) kilichojengwa ndani ambacho kinatoa sehemu kubwa ya utendaji tulioandika kwa mkono katika changamoto iliyopita. Kitu hiki cha ramani, ingawa kinafanana na vitu vya kawaida vya JavaScript, kinatoa utendaji muhimu ambao vitu vya kawaida havina. Kwa mfano, ramani ya ES6 hufuata mpangilio wa kuingizwa kwa vitu vinavyoongezwa. Hapa kuna muhtasari kamili wa njia zake:

- `.has(key)` hurudisha kweli au si kweli kulingana na uwepo wa ufunguo
- `.get(key)` hurudisha thamani inayohusiana na ufunguo
- `.set(key, value)` huweka jozi mpya ya ufunguo, thamani
- `.delete(key)` huondoa jozi ya ufunguo, thamani
- `.clear()` huondoa jozi zote za ufunguo, thamani
- `.keys()` hurudisha kielekezi kipya cha ramani chenye funguo zote kwa mpangilio wa kuingizwa
- `.values()` hurudisha kielekezi kipya cha ramani chenye thamani zote kwa mpangilio wa kuingizwa
- `.entries()` hurudisha kielekezi kipya cha ramani chenye jozi zote za ufunguo, thamani kwa mpangilio wa kuingizwa

# --instructions--

Fafanua kitu cha ramani cha JavaScript na weka thamani yake katika kigezo kinachoitwa myMap. Ongeza jozi ya ufunguo, thamani `freeCodeCamp`, `Awesome!` ndani yake.

# --hints--

Kitu cha `myMap` kinapaswa kuwepo.

```js
assert(typeof myMap === 'object');
```

`myMap` inapaswa kuwa na jozi ya ufunguo na thamani `freeCodeCamp`, `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
