---
id: 587d8255367417b2b2512c75
title: Tengeneza safu ya mzunguko
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

Katika changamoto hii utaunda safu ya mzunguko. Safu ya mzunguko ni safu inayoongeza data mwishoni mwa mkusanyiko kisha huanza kuandika tena data mwanzoni mwa mkusanyiko huo. Aina hii ya muundo wa data ni muhimu katika hali fulani. Kwa mfano, safu ya mzunguko inaweza kutumika kwa uchezaji wa vyombo vya habari vinavyotiririka. Mara safu itakapojaa, data mpya ya vyombo vya habari itabeba data ya zamani.

Njia nzuri ya kuelezea dhana hii ni kwa kutumia safu ya data yenye urefu `5`:

```js
[null, null, null, null, null]
 ^Read @ 0
 ^Write @ 0
```

Hapa kichwa cha kusoma na kuandika vyote viko katika nafasi ya `0`. Sasa safu inapata rekodi mpya 3 `a`, `b`, na `c`. Safu yetu sasa inaonekana kama ifuatavyo:

```js
[a, b, c, null, null]
 ^Read @ 0
          ^Write @ 3
```

Kichwa cha kusoma kinaposoma, kinaweza kuondoa au kuhifadhi thamani:

```js
[null, null, null, null, null]
                   ^Read @ 3
                   ^Write @ 3
```

Sasa tunaandika thamani `d`, `e`, na `f` kwenye safu. Mara kuandika kufikia mwisho wa safu, hujirudia mwanzoni:

```js
[f, null, null, d, e]
                ^Read @ 3
    ^Write @ 1
```

Njia hii inahitaji kiasi cha kumbukumbu kisichobadilika lakini huruhusu faili kubwa zaidi kushughulikiwa.

# --instructions--

Katika changamoto hii tutatekeleza safu ya mzunguko. Safu ya mzunguko inapaswa kutoa njia za `enqueue` na `dequeue` ambazo zitakuwezesha kusoma na kuandika kwenye safu. Darasa la vitu lenyewe linapaswa pia kupokea hoja ya nambari kamili ambayo unaweza kuitumia kubainisha ukubwa wa safu wakati inaundwa. Tumeandika toleo la mwanzo la darasa hili kwako katika mhariri wa msimbo.

Unapoongeza vitu kwenye safu, kidokezo cha kuandika kinapaswa kusonga mbele na kurudi mwanzoni mara kinapofikia mwisho wa safu. Njia ya `enqueue` inapaswa kurudisha kipengee ulichoongeza ikiwa imefanikiwa; vinginevyo itarudisha `null`.

Vivyo hivyo, kidokezo cha kusoma kinapaswa kusonga mbele unapotolewa vitu kwenye safu. Unapotoa kipengee, kipengee hicho kinapaswa kurudishwa. Ikiwa huwezi kutoa kipengee, unapaswa kurudisha `null`.

Kidokezo cha kuandika hakipaswi kuruhusiwa kusonga zaidi ya kidokezo cha kusoma (darasa letu halitaruhusu kuandika juu ya data ambayo hujasoma bado) na kidokezo cha kusoma hakipaswi kusonga zaidi ya data uliyoiandika.

# --hints--

Njia ya `enqueue` inapaswa kuongeza vitu kwenye safu ya mzunguko.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Haupasi kuongeza vitu zaidi ya kidokezo cha kusoma.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Njia ya `dequeue` inapaswa kutoa vitu kutoka kwenye safu.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Baada ya kipengee kutolewa, nafasi yake kwenye safu inapaswa kuwekwa tena kuwa `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Kujaribu kutoa zaidi ya kidokezo cha kuandika kunapaswa kurudisha `null` na hakusogei kidokezo cha kuandika.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```
