---
id: 587d825a367417b2b2512c8a
title: Ingiza kipengele katika rundo la juu zaidi
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

Sasa tutaendelea na muundo mwingine wa data wa mti, rundo la binary. Rundo la binary ni mti wa kibainari ulio na mpangilio wa sehemu ambao unatii sifa ya rundo. Sifa ya rundo inaelezea uhusiano kati ya nodi ya mzazi na nodi za watoto. Unaweza kuwa na rundo la juu zaidi, ambapo nodi zote za mzazi ni kubwa au sawa na nodi zao za watoto, au rundo la chini zaidi, ambapo kinyume chake ndicho kweli. Rundo za binary pia ni miti kamili ya kibainari. Hii ina maana kwamba ngazi zote za mti zimejazwa kikamilifu na ikiwa ngazi ya mwisho imejazwa kwa sehemu, hujazwa kutoka kushoto kwenda kulia.

Ingawa rundo za binary zinaweza kutekelezwa kama miundo ya mti yenye nodi zinazoshikilia rejea za kushoto na kulia, upangaji wa sehemu kulingana na sifa ya rundo huturuhusu kuwakilisha rundo kwa kutumia safu ya data. Uhusiano wa mzazi-watoto ndio tunalolenga na kwa hesabu rahisi tunaweza kuhesabu watoto wa mzazi yeyote na mzazi wa nodi yoyote ya mtoto.

Kwa mfano, fikiria uwakilishi huu wa safu ya data wa rundo la chini zaidi la binary:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Nodi ya mzizi ni kipengele cha kwanza, `6`. Watoto wake ni `22` na `30`. Ikiwa tutaangalia uhusiano kati ya viashiria vya safu ya data vya thamani hizi, kwa kielezo `i` watoto ni `2 * i + 1` na `2 * i + 2`. Vivyo hivyo, kipengele kilicho kwenye kielezo `0` ni mzazi wa watoto hawa wawili walioko kwenye viashiria `1` na `2`. Kwa ujumla, tunaweza kupata mzazi wa nodi yoyote kwa kielezo chochote kwa kutumia ifuatayo: `Math.floor((i - 1) / 2)`. Mifumo hii itakuwa sahihi kadri mti wa binary unavyokua kwa ukubwa wowote. Mwishowe, tunaweza kufanya marekebisho kidogo ili kufanya hesabu hii iwe rahisi zaidi kwa kuruka kipengele cha kwanza katika safu ya data. Kufanya hivi kunaunda uhusiano ufuatao kwa kipengele chochote kilicho kwenye kielezo `i`:

Mfano wa uwakilishi wa safu ya data:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Mtoto wa kushoto wa kipengele: `i * 2`

Mtoto wa kulia wa kipengele: `i * 2 + 1`

Mzazi wa kipengele: `Math.floor(i / 2)`

Mara tu unapopata ufahamu wa hesabu hii, kutumia uwakilishi wa safu ya data ni muhimu sana kwa sababu maeneo ya nodi yanaweza kubainishwa haraka kwa kutumia hesabu hii na matumizi ya kumbukumbu hupunguzwa kwa sababu hauitaji kuhifadhi rejea za nodi za watoto.

# --instructions--

Maelekezo: Hapa tutaunda rundo la juu zaidi. Anza kwa kuunda njia ya `insert` ambayo inaongeza vipengele kwenye rundo letu. Wakati wa kuingiza, ni muhimu kila mara kudumisha sifa ya rundo. Kwa rundo la juu zaidi hii ina maana kipengele cha mzizi kinapaswa kuwa na thamani kubwa zaidi katika mti na nodi zote za mzazi zinapaswa kuwa kubwa kuliko watoto wao. Kwa utekelezaji wa rundo kwa kutumia safu ya data, hii kawaida hufanyika kwa hatua tatu:

<ol>
  <li>Ongeza kipengele kipya mwishoni mwa safu ya data.</li>
  <li>Kama kipengele ni kikubwa kuliko mzazi wake, badilisha nafasi zao.</li>
  <li>Endelea kubadilisha hadi kipengele kipya kiwe kidogo kuliko mzazi wake au ufikie mzizi wa mti.</li>
</ol>

Mwishowe, ongeza njia ya `print` ambayo inarudisha safu ya vitu vyote vilivyoongezwa kwenye rundo.

# --hints--

Muundo wa data wa `MaxHeap` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` inapaswa kuwa na njia inayoitwa `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` inapaswa kuwa na njia inayoitwa `print`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Njia ya `insert` inapaswa kuongeza vipengele kulingana na sifa ya rundo la juu zaidi.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```
