---
id: 587d8254367417b2b2512c71
title: Ondoa vitu kutoka kwenye seti katika ES6
challengeType: 1
forumTopicId: 301713
dashedName: remove-items-from-a-set-in-es6
---

# --description--

Tujaribu kuondoa vitu kutoka kwenye Seti ya ES6 kwa kutumia njia ya `delete`.

Kwanza, tengeneza Seti ya ES6:

```js
var set = new Set([1,2,3]);
```

Sasa ondoa kipengee kutoka kwenye Seti yako kwa kutumia njia ya `delete`.

```js
set.delete(1);
console.log([...set]) // should return [ 2, 3 ]
```

# --instructions--

Sasa, tengeneza seti yenye nambari kamili 1, 2, 3, 4, na 5.

Ondoa thamani za 2 na 5, kisha rudisha seti hiyo.

# --hints--

Seti yako inapaswa kuwa na thamani za 1, 3, na 4

```js
assert(
  (function () {
    var test = checkSet();
    return test.has(1) && test.has(3) && test.has(4) && test.size === 3;
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet(){
  // Only change code below this line
  var set = null;

  // Only change code above this line
  return set;   
}
```

# --solutions--

```js
function checkSet(){
var set = new Set([1,2,3,4,5]);
set.delete(2);
set.delete(5);
return set;}
```
