---
id: 587d8251367417b2b2512c61
title: Fanya kazi na nodi katika orodha iliyounganishwa
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

Muundo mwingine wa data unaokutana nao mara kwa mara katika sayansi ya kompyuta ni <dfn>orodha iliyounganishwa</dfn>. Orodha iliyounganishwa ni mkusanyiko wa mstari wa vipengele vya data, vinavyoitwa 'nodi', kila moja likielekeza kwenye inayofuata. Kila <dfn>nodi</dfn> katika orodha iliyounganishwa ina vipengele viwili muhimu vya taarifa: `element` yenyewe, na rejea kwa `node` inayofuata.

Fikiria uko katika mstari wa conga. Una mikono yako juu ya mtu anayefuata katika mstari, na mtu aliyeko nyuma yako ana mikono yake juu yako. Unaweza kuona mtu aliye mbele yako moja kwa moja, lakini wanazuia mtazamo wa watu wengine walioko mbele zaidi katika mstari. Nodi ni kama mtu katika mstari wa conga: wanajua ni nani wao na wanaweza kuona mtu anayefuata tu, lakini hawajui kuhusu watu wengine walioko mbele au nyuma yao.

# --instructions--

Katika mhariri wetu wa msimbo, tumetengeneza nodi mbili, `Kitten` na `Puppy`, na tumeunganisha nodi ya `Kitten` kwa nodi ya `Puppy` kwa mkono.

Tengeneza nodi ya `Cat` na `Dog` na uziunge kwa mkono kwenye mstari.

# --hints--

Nodi yako ya `Puppy` inapaswa kuwa na rejea kwa nodi ya `Cat`.

```js
assert(Puppy.next.element === 'Cat');
```

Nodi yako ya `Cat` inapaswa kuwa na rejea kwa nodi ya `Dog`.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
