---
id: 587d825b367417b2b2512c8c
title: Tekeleza kupanga kwa rundo la chini zaidi
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

Sasa tunapoweza kuongeza na kuondoa vipengele, tuchunguze baadhi ya programu ambazo rundo zinaweza kutumika. Rundo hutumika sana kutekeleza foleni za kipaumbele kwa sababu daima huhifadhi kipengele chenye thamani kubwa au ndogo zaidi katika nafasi ya kwanza. Zaidi ya hayo, hutumika kutekeleza algoriti ya kupanga inayoitwa heap sort. Tutaona jinsi ya kufanya hivyo hapa. Heap sort hutumia rundo la chini zaidi, ambalo ni kinyume cha rundo la juu zaidi. Rundo la chini zaidi daima huhifadhi kipengele chenye thamani ndogo zaidi katika nafasi ya mzizi.

Heap sort hufanya kazi kwa kuchukua safu ya data isiyopangwa, kuongeza kila kipengele katika rundo la chini zaidi, kisha kutoa kila kipengele kutoka kwenye rundo hilo na kuviweka katika safu mpya. Muundo wa rundo la chini zaidi unahakikisha kwamba safu mpya itakuwa na vipengele vya awali kwa mpangilio kutoka kidogo hadi kikubwa. Hii ni moja ya algoriti bora za kupanga kwa ufanisi, ikiwa na utendaji wa wastani na mbaya wa O(nlog(n)).

# --instructions--

Tutekeleze heap sort kwa kutumia rundo la chini zaidi. Tafadhali badilisha msimbo wako wa rundo la juu zaidi hapa. Tengeneza kitu `MinHeap` chenye njia `insert`, `remove`, na `sort`. Njia ya `sort` inapaswa kurudisha safu ya vipengele vyote katika rundo la chini zaidi vilivyopangwa kutoka kidogo hadi kikubwa.

# --hints--

Muundo wa data `MinHeap` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` inapaswa kuwa na njia iitwayo `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` inapaswa kuwa na njia iitwayo `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` inapaswa kuwa na njia iitwayo `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Njia ya `sort` inapaswa kurudisha safu yenye vipengele vyote vilivyoongezwa kwenye rundo la chini zaidi kwa mpangilio.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
