---
id: 587d8257367417b2b2512c7c
title: Angalia kama Kipengele Kipo katika Mti wa Utafutaji wa Kibainari
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

Sasa tunapokuwa na ufahamu wa jumla kuhusu mti wa utafutaji wa kibainari, hebu tuzungumzie kwa undani zaidi. Miti ya utafutaji wa kibainari hutoa muda wa logarithmic kwa shughuli za kawaida za kutafuta, kuingiza, na kufuta katika hali ya wastani, na muda wa mstari katika hali mbaya zaidi. Kwa nini hivi ni hivyo? Kila moja ya shughuli hizo za msingi inahitaji tufanye ni kutafuta kipengele katika mti (au katika kesi ya kuingiza, kutafuta mahali kinapaswa kuwekwa) na kwa sababu ya muundo wa mti katika kila nodi ya mzazi tunagawanya upande wa kushoto au kulia na kwa ufanisi tunatoa nusu ya ukubwa wa mti uliobaki. Hii inafanya utafutaji kuwa sawa na logarithm ya idadi ya nodi katika mti, ambayo huunda muda wa logarithmic kwa shughuli hizi katika hali ya wastani. Sawa, lakini je kuhusu hali mbaya zaidi? Vizuri, fikiria kujenga mti kutoka kwa thamani zifuatazo, ukiziongeza kutoka kushoto kwenda kulia: `10`, `12`, `17`, `25`. Kufuatia sheria zetu za mti wa utafutaji wa kibainari, tutaongeza `12` upande wa kulia wa `10`, `17` upande wa kulia wa hii, na `25` upande wa kulia wa hii. Sasa mti wetu unaonekana kama orodha iliyounganishwa na kupita kwake kutafuta `25` kutahitaji kupita vitu vyote kwa mpangilio wa mstari. Hivyo basi, muda wa mstari katika hali mbaya zaidi. Tatizo hapa ni kwamba mti hauko sawa. Tutatazama zaidi maana ya hili katika changamoto zinazofuata.

# --instructions--

Katika changamoto hii, tutaunda zana kwa ajili ya mti wetu. Andika njia `isPresent` ambayo inachukua thamani ya nambari kamili kama ingizo na inarudisha thamani ya kweli au siyo kweli kwa kuwepo au kutokuwepo kwa thamani hiyo katika mti wa utafutaji wa kibainari.

# --hints--

Muundo wa data `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia inayoitwa `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Njia `isPresent` inapaswa kuangalia kwa usahihi kuwepo au kutokuwepo kwa vipengele vilivyoongezwa kwenye mti.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` inapaswa kushughulikia kesi ambapo mti ni tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```
