---
id: 587d8255367417b2b2512c72
title: Tumia .has na .size kwenye Set ya ES6
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Tuchunguze njia za .has na .size zinazopatikana kwenye kitu cha Set cha ES6.

Kwanza, tengeneza Set ya ES6

```js
var set = new Set([1,2,3]);
```

Njia ya .has itakagua kama thamani ipo ndani ya seti.

```js
var hasTwo = set.has(2);
```

Njia ya .size itarudisha nambari kamili inayowakilisha ukubwa wa Set

```js
var howBig = set.size;
```

# --instructions--

Katika zoezi hili tutapitia safu ya data na thamani kwa kitendakazi cha checkSet(). Kitendakazi chako kinapaswa kutengeneza seti ya ES6 kutoka kwa hoja ya safu ya data. Tafuta kama seti ina thamani ya hoja. Tafuta ukubwa wa seti. Na rudisha thamani hizo mbili katika safu ya data.

# --hints--

`checkSet([4, 5, 6], 3)` inapaswa kurudisha [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
