---
id: 587d8251367417b2b2512c63
title: Ondoa vipengele kutoka orodha iliyounganishwa
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

Njia muhimu inayofuata ambayo utekelezaji wowote wa orodha iliyounganishwa utahitaji ni njia ya `remove`. Njia hii inapaswa kuchukua kipengele tunachotaka kuondoa kama hoja, kisha itafute orodha ili kupata na kuondoa nodi inayoshikilia kipengele hicho.

Kila tunapoondoa nodi kutoka orodha iliyounganishwa, ni muhimu kuhakikisha hatuachii sehemu nyingine ya orodha bila muunganisho. Kumbuka kuwa sifa ya `next` ya kila nodi inaonyesha nodi inayofuata katika orodha. Ikiwa tunaondoa kipengele cha katikati, kwa mfano, tunataka kuhakikisha kuwa tuna muunganisho kutoka kwa sifa ya `next` ya nodi ya awali ya kipengele hicho hadi sifa ya `next` ya kipengele cha katikati (ambayo ni nodi inayofuata katika orodha!)

Hii inaweza kuonekana kuchanganya sana, kwa hivyo turudi kwenye mfano wa mstari wa conga ili kuwa na mfano mzuri wa dhana. Jiweke katika mstari wa conga, na mtu aliyeko mbele yako anaondoka mstari. Mtu aliyetoa mstari hana tena mikono yake juu ya mtu yeyote katika mstari—na wewe huna tena mikono yako juu ya mtu aliyetoa. Unasogea mbele na kuweka mikono yako juu ya mtu anayefuata uone.

Ikiwa kipengele tunachotaka kuondoa ni kipengele cha `head`, tunarekebisha `head` kuwa nodi ya pili ya orodha iliyounganishwa.

# --instructions--

Andika njia ya `remove` inayochukua kipengele na kuiondoa kutoka orodha iliyounganishwa.

**Kumbuka:** Nambari ya `length` ya orodha inapaswa kupungua kwa moja kila mara kipengele kinaondolewa kutoka orodha iliyounganishwa.

# --hints--

Darasa lako la `LinkedList` linapaswa kuwa na njia ya `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Njia yako ya `remove` inapaswa kuweka upya `head` kuwa nodi ya pili wakati nodi ya kwanza inaondolewa.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Njia yako ya `remove` inapaswa kupunguza `length` ya orodha iliyounganishwa kwa moja kwa kila nodi inayoondolewa.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Njia yako ya `remove` inapaswa kuweka upya rejea ya nodi ya awali ya nodi iliyoondolewa kwa rejea ya `next` ya nodi iliyoondolewa.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Njia yako ya `remove` haipaswi kubadilisha orodha iliyounganishwa ikiwa kipengele hakipo katika orodha iliyounganishwa.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
