---
id: 587d8258367417b2b2512c7f
title: Tumia utafutaji wa upana kwanza katika mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301718
dashedName: use-breadth-first-search-in-a-binary-search-tree
---

# --description--

Hapa tutatambulisha njia nyingine ya kupita mti: utafutaji wa upana kwanza. Tofauti na njia za utafutaji wa kina kwanza kutoka changamoto iliyopita, utafutaji wa upana kwanza hutafuta nodi zote katika ngazi fulani ndani ya mti kabla ya kuendelea kwenye ngazi inayofuata. Kawaida, foleni hutumika kama miundo ya data ya msaada katika muundo wa algoriti za utafutaji wa upana kwanza.

Katika njia hii, tunaanza kwa kuongeza nodi ya mzizi kwenye foleni. Kisha tunaanzisha mzunguko ambapo tunatoa kipengee cha kwanza kwenye foleni, tunaongeza kwenye safu mpya, na kisha kuchunguza miti midogo ya watoto wake wote. Ikiwa watoto wake si null, kila mmoja huongezwa kwenye foleni. Mchakato huu unaendelea hadi foleni iwe tupu.

# --instructions--

Tufanye njia ya utafutaji wa upana kwanza katika mti wetu iitwayo `levelOrder`. Njia hii inapaswa kurudisha safu ya data yenye thamani za nodi zote za mti, zilizotafutwa kwa mtindo wa upana kwanza. Hakikisha unarudisha thamani katika safu, si nodi wenyewe. Ngazi inapaswa kupitiwa kutoka kushoto kwenda kulia. Kisha, andika njia inayofanana iitwayo `reverseLevelOrder` ambayo hufanya utafutaji huo huo lakini kwa mwelekeo wa kinyume (kulia kwenda kushoto) katika kila ngazi.

# --hints--

Muundo wa data `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `levelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.levelOrder == 'function';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia iitwayo `reverseLevelOrder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.reverseLevelOrder == 'function';
  })()
);
```

Njia ya `levelOrder` inapaswa kurudisha safu ya thamani za nodi za mti zilizotafutwa kwa mpangilio wa ngazi.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.levelOrder().join('') == '719038102546';
  })()
);
```

Njia ya `reverseLevelOrder` inapaswa kurudisha safu ya thamani za nodi za mti zilizotafutwa kwa mpangilio wa ngazi wa kinyume.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.reverseLevelOrder().join('') == '791108305264';
  })()
);
```

Njia ya `levelOrder` inapaswa kurudisha `null` kwa mti tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.levelOrder !== 'function') {
      return false;
    }
    return test.levelOrder() == null;
  })()
);
```

Njia ya `reverseLevelOrder` inapaswa kurudisha `null` kwa mti tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.reverseLevelOrder !== 'function') {
      return false;
    }
    return test.reverseLevelOrder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  this.levelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [];
    while(queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.left) queue.push(node.left);
      if(node.right) queue.push(node.right);
    }
    return results;
  }

  this.reverseLevelOrder = (root = this.root) => {
    if(!root) return null;
    let queue = [root];
    let results = [] ;
    while ( queue.length > 0) {
      let node = queue.shift();
      results.push(node.value);
      if(node.right) queue.push(node.right);
      if(node.left ) queue.push(node.left);
    }
    return results;
  }
  // Only change code above this line
}
```
