---
id: 587d8257367417b2b2512c7e
title: Tumia utafutaji wa kina kwanza katika mti wa utafutaji wa kibainari
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

Tunajua jinsi ya kutafuta thamani maalum katika mti wa utafutaji wa kibainari. Lakini vipi kama tunataka tu kuchunguza mti mzima? Au vipi kama hatuna mti uliopangwa na tunahitaji tu kutafuta thamani? Hapa tutatambulisha baadhi ya njia za kutembea kwenye mti ambazo zinaweza kutumika kuchunguza miundo ya data ya miti. Kwanza ni utafutaji wa kina kwanza. Katika utafutaji wa kina kwanza, mti mdogo unaotolewa unachunguzwa kwa kina iwezekanavyo kabla ya utafutaji kuendelea kwenye mti mdogo mwingine. Kuna njia tatu za kufanya hivi: Kwa mpangilio: Anza utafutaji kwenye nodi ya kushoto kabisa na maliza kwenye nodi ya kulia kabisa. Kabla ya mpangilio: Chunguza mizizi yote kabla ya majani. Baada ya mpangilio: Chunguza majani yote kabla ya mizizi. Kama unavyoweza kubashiri, unaweza kuchagua njia tofauti za utafutaji kulingana na aina ya data ambayo mti wako unahifadhi na unachotafuta. Kwa mti wa utafutaji wa kibainari, kutembea kwa mpangilio hurejesha nodi kwa mpangilio uliopangwa.

# --instructions--

Hapa tutaunda njia hizi tatu za utafutaji kwenye mti wetu wa utafutaji wa kibainari. Utafutaji wa kina kwanza ni operesheni ya kurudia yenye asili ambayo inaendelea kuchunguza miti midogo zaidi mradi tu nodi za watoto zipo. Mara tu unapofahamu dhana hii msingi, unaweza kubadilisha tu mpangilio wa jinsi unavyotembelea nodi na miti midogo ili kuzalisha mojawapo ya utafutaji tatu ulio hapo juu. Kwa mfano, katika utafutaji wa baada ya mpangilio tungetaka kurudia hadi nodi ya jani kabla ya kuanza kurudisha nodi yoyote yenyewe, wakati katika utafutaji wa kabla ya mpangilio tungetaka kurudisha nodi kwanza, kisha kuendelea kurudia chini ya mti. Tafsiri njia `inorder`, `preorder`, na `postorder` kwenye mti wetu. Kila njia inapaswa kurudisha safu ya vitu vinavyo wakilisha kutembea kwa mti. Hakikisha kurudisha thamani za nambari kamili katika kila nodi kwenye safu, si nodi yenyewe. Mwisho, rudisha `null` ikiwa mti ni tupu.

# --hints--

Muundo wa data `BinarySearchTree` unapaswa kuwepo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia inayoitwa `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia inayoitwa `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

Mti wa utafutaji wa kibainari unapaswa kuwa na njia inayoitwa `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Njia ya `inorder` inapaswa kurudisha safu ya thamani za nodi zinazotokana na kutembea kwa mpangilio.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Njia ya `preorder` inapaswa kurudisha safu ya thamani za nodi zinazotokana na kutembea kabla ya mpangilio.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Njia ya `postorder` inapaswa kurudisha safu ya thamani za nodi zinazotokana na kutembea baada ya mpangilio.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Njia ya `inorder` inapaswa kurudisha `null` kwa mti tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Njia ya `preorder` inapaswa kurudisha `null` kwa mti tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Njia ya `postorder` inapaswa kurudisha `null` kwa mti tupu.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```
