---
id: 587d8254367417b2b2512c70
title: Unda na Ongeza kwenye Sets katika ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

Sasa baada ya kufanya kazi na ES5, utafanya kitu kinachofanana katika ES6. Hii itakuwa rahisi zaidi sana. ES6 ina muundo wa data uliyojengwa ndani `Set` hivyo shughuli nyingi ulizozitengeneza kwa mkono sasa zipo tayari kwa ajili yako. Tuchunguze:

Kuunda seti mpya tupu:

```js
var set = new Set();
```

Unaweza kuunda seti yenye thamani:

```js
var set = new Set(1);
```

Unaweza kuunda seti yenye safu ya data:

```js
var set = new Set([1, 2, 3]);
```

Mara baada ya kuunda seti, unaweza kuongeza thamani unayotaka kwa kutumia njia ya `add`:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

Kumbuka, seti ni muundo wa data usio na thamani rudufu:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

Kwa zoezi hili, rudisha seti yenye thamani zifuatazo: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

`Set` yako inapaswa kuwa na thamani `1, 2, 3, Taco, Cat, Awesome` pekee.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
