---
id: 587d7db8367417b2b2512ba1
title: 숫자가 아닌 모든 문자 찾기
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

지난 도전 과제에서는 소문자 `\d`와 함께 단축키 `d`를 사용해 숫자를 검색하는 방법을 보여주었습니다. 이번에는 대문자 `D`를 사용해 숫자가 아닌 문자를 검색하는 비슷한 단축키도 있습니다.

숫자가 아닌 문자를 찾는 단축키는 `\D`입니다. 이는 0부터 9 사이의 숫자가 아닌 단일 문자를 찾는 문자 클래스 `[^0-9]`와 같습니다.

# --instructions--

숫자가 아닌 문자를 위한 단축 문자 클래스 `\D`를 사용해 영화 제목에 있는 숫자가 아닌 문자의 개수를 세세요.

# --hints--

정규식은 숫자가 아닌 문자를 매칭하는 단축 문자를 사용해야 합니다.

```js
assert(/\\D/.test(noNumRegex.source));
```

정규식은 전역 플래그를 사용해야 합니다.

```js
assert(noNumRegex.global);
```

정규식은 문자열 `9`에서 숫자가 아닌 문자를 찾지 못해야 합니다.

```js
assert('9'.match(noNumRegex) == null);
```

정규식은 문자열 `Catch 22`에서 숫자가 아닌 문자를 6개 찾아야 합니다.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

정규식은 문자열 `101 Dalmatians`에서 숫자가 아닌 문자를 11개 찾아야 합니다.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

정규식은 문자열 `One, Two, Three`에서 숫자가 아닌 문자를 15개 찾아야 합니다.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

정규식은 문자열 `21 Jump Street`에서 숫자가 아닌 문자를 12개 찾아야 합니다.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

정규식은 문자열 `2001: A Space Odyssey`에서 숫자가 아닌 문자를 17개 찾아야 합니다.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
