---
id: 587d7db4367417b2b2512b90
title: 다양한 가능성을 가진 리터럴 문자열 매칭
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

`/coding/` 같은 정규식을 사용하면 다른 문자열에서 `coding` 패턴을 찾을 수 있습니다.

이 방법은 단일 문자열을 검색하는 데 강력하지만, 하나의 패턴에만 제한됩니다. `alternation` 또는 `OR` 연산자 `|`를 사용하면 여러 패턴을 검색할 수 있습니다.

이 연산자는 앞이나 뒤에 있는 패턴과 일치합니다. 예를 들어, `yes` 또는 `no` 문자열과 일치시키고 싶다면, 원하는 정규식은 `/yes|no/`입니다.

두 개 이상의 패턴도 검색할 수 있습니다. `OR` 연산자를 더 추가하여 패턴을 더 많이 넣으면 됩니다. 예를 들어 `/yes|no|maybe/`처럼요.

# --instructions--

정규식 `petRegex`을 완성하여 애완동물 `dog`, `cat`, `bird`, 또는 `fish`와 일치시키세요.

# --hints--

정규식 `petRegex`는 문자열 `true`에 대해 `John has a pet dog.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

정규식 `petRegex`는 문자열 `false`에 대해 `Emma has a pet rock.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

정규식 `petRegex`는 문자열 `true`에 대해 `Emma has a pet bird.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

정규식 `petRegex`는 문자열 `true`에 대해 `Liz has a pet cat.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

정규식 `petRegex`는 문자열 `false`에 대해 `Kara has a pet dolphin.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

정규식 `petRegex`는 문자열 `true`에 대해 `Alice has a pet fish.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

정규식 `petRegex`는 문자열 `false`에 대해 `Jimmy has a pet computer.`를 반환해야 합니다.

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
