---
id: 587d7db6367417b2b2512b98
title: 지정하지 않은 단일 문자와 일치시키기
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

지금까지 일치시키고 싶은 문자 집합을 만들었지만, 일치시키고 싶지 않은 문자 집합도 만들 수 있습니다. 이러한 문자 집합을 <dfn>부정 문자 집합</dfn>이라고 합니다.

부정 문자 집합을 만들려면, 여는 대괄호 뒤와 일치시키고 싶지 않은 문자 앞에 캐럿 문자(`^`)를 넣으세요.

예를 들어, `/[^aeiou]/gi`은 모음이 아닌 모든 문자와 일치합니다. `.`, `!`, `[`, `@`, `/` 및 공백과 같은 문자도 일치한다는 점에 주의하세요 — 부정 모음 문자 집합은 모음 문자만 제외합니다.

# --instructions--

숫자나 모음이 아닌 모든 문자와 일치하는 단일 정규식을 만드세요. 정규식에 적절한 플래그를 포함하는 것도 잊지 마세요.

# --hints--

정규식 `myRegex`은 9개의 항목과 일치해야 합니다.

```js
assert(result.length == 9);
```

정규식 `myRegex`은 전역 플래그를 사용해야 합니다.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

정규식 `myRegex`은 대소문자 구분 없는 플래그를 사용해야 합니다.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
