---
id: 587d7dbb367417b2b2512bab
title: 캡처 그룹을 사용해 검색 및 교체하기
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

검색은 유용합니다. 하지만 검색한 텍스트를 변경(또는 교체)할 수 있다면 훨씬 더 강력해집니다.

문자열에서 `.replace()`를 사용해 텍스트를 검색하고 교체할 수 있습니다. `.replace()`의 입력값은 먼저 검색할 정규식 패턴이고, 두 번째 매개변수는 일치하는 부분을 교체할 문자열이거나 무언가를 수행하는 함수입니다.

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

`replace` 호출은 문자열 `The sky is blue.`을 반환합니다.

교체 문자열에서 달러 기호(`$`)를 사용해 캡처 그룹에 접근할 수도 있습니다.

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

`replace` 호출은 문자열 `Camp Code`을 반환합니다.

# --instructions--

세 개의 캡처 그룹을 사용하는 정규식 `fixRegex`를 작성해 문자열 `one two three`의 각 단어를 검색하세요. 그런 다음 `replaceText` 변수를 업데이트해 `one two three`을 문자열 `three two one`로 교체하고 결과를 `result` 변수에 할당하세요. 교체 문자열에서 달러 기호(`$`) 구문을 사용해 캡처 그룹을 활용하는지 확인하세요.

# --hints--

검색 및 교체에는 `.replace()`를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

정규식은 문자열 `one two three`를 문자열 `three two one`로 변경해야 합니다.

```js
assert(result === 'three two one');
```

마지막 줄은 변경하지 마세요.

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

`fixRegex`는 최소 세 개의 캡처 그룹을 사용해야 합니다.

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

`replaceText`는 괄호로 묶인 하위 일치 문자열(즉, n번째 괄호로 묶인 하위 일치 문자열 $n이 n번째 캡처 그룹에 해당)을 사용해야 합니다.

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
