---
id: 587d7db7367417b2b2512b9c
title: 사냥에서 한 명 이상의 범인을 찾기
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

새로운 정규식 작성 기술을 시험해 볼 시간입니다. 범인 무리가 감옥에서 탈출해 도망쳤지만 몇 명인지 알 수 없습니다. 하지만 그들은 다른 사람들 주변에 있을 때 서로 가까이 붙어 다닌다는 사실은 알고 있습니다. 한 번에 모든 범인을 찾아야 합니다.

다음은 이를 어떻게 하는지 검토할 예시입니다:

정규식 `/z+/`는 문자 `z`가 연속해서 한 번 이상 나타날 때 일치합니다. 다음 문자열들에서 모두 일치하는 부분을 찾을 수 있습니다:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

하지만 문자 `z`가 없기 때문에 다음 문자열들에서는 일치하지 않습니다:

```js
""
"ABC"
"abcabc"
```

# --instructions--

다른 사람들 무리 속에서 한 명 이상 범인을 찾는 탐욕적 정규식을 작성하세요. 범인은 대문자 `C`로 나타냅니다.

# --hints--

정규식은 문자열 `C`에서 한 명의 범인(`C`)과 일치해야 합니다.

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

정규식은 문자열 `CC`에서 두 명의 범인(`CC`)과 일치해야 합니다.

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

정규식은 문자열 `CCC`에서 세 명의 범인(`P1P5P4CCCcP2P6P3`)과 일치해야 합니다.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

정규식은 문자열 `CCCCC`에서 다섯 명의 범인(`P6P2P7P4P5CCCCCP3P1`)과 일치해야 합니다.

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

정규식은 빈 문자열 `""`에서는 범인과 일치하지 않아야 합니다.

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

정규식은 문자열 `P1P2P3`에서도 범인과 일치하지 않아야 합니다.

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

정규식은 문자열 `CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`에서 50명의 범인(`P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3`)과 일치해야 합니다.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
