---
id: 587d7db9367417b2b2512ba6
title: 매치 개수의 하한만 지정하기
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

중괄호를 사용한 수량 지정자로 패턴의 하한과 상한 개수를 지정할 수 있습니다. 때로는 상한 없이 하한 개수만 지정하고 싶을 때가 있습니다.

패턴의 하한 개수만 지정하려면 첫 번째 숫자 뒤에 쉼표를 붙이세요.

예를 들어, 문자 `hah`가 최소 `a`번 나타나는 문자열 `3`만 매치하려면 정규식은 `/ha{3,}h/`가 됩니다.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

순서대로 세 번의 `test` 호출은 `true`, `false`, `true`를 반환합니다.

# --instructions--

정규식 `haRegex`을 변경하여 문자 `Hazzah`가 네 번 이상 있을 때만 단어 `z`와 매치되게 하세요.

# --hints--

정규식에 중괄호를 사용해야 합니다.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

정규식은 문자열 `Hazzah`과 매치되면 안 됩니다.

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

정규식은 문자열 `Hazzzah`과 매치되면 안 됩니다.

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

정규식은 문자열 `Hazzzzah`과 매치되어야 합니다.

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

정규식은 문자열 `Hazzzzzah`과 매치되어야 합니다.

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

정규식은 문자열 `Hazzzzzzah`과 매치되어야 합니다.

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

정규식은 `Hazzah`가 30개 있는 문자열 `z`과 매치되어야 합니다.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
