---
id: 587d7db5367417b2b2512b95
title: 여러 가능성을 가진 단일 문자 매칭
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

리터럴 패턴(`/literal/`)과 와일드카드 문자(`/./`)를 매칭하는 방법을 배웠습니다. 이 둘은 정규 표현식의 양극단으로, 하나는 정확한 일치를 찾고 다른 하나는 모든 것을 매칭합니다. 이 두 극단 사이에 균형을 이루는 옵션도 있습니다.

<dfn>문자 클래스</dfn>를 사용하면 리터럴 패턴을 약간 유연하게 검색할 수 있습니다. 문자 클래스는 대괄호(`[`와 `]`) 안에 원하는 문자 그룹을 정의하여 매칭할 수 있게 합니다.

예를 들어, `bag`, `big`, `bug`는 매칭하고 `bog`는 매칭하지 않으려면 `/b[aiu]g/` 정규식을 만들 수 있습니다. `[aiu]`는 `a`, `i`, 또는 `u` 문자만 매칭하는 문자 클래스입니다.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

순서대로, 네 개의 `match` 호출은 `["big"]`, `["bag"]`, `["bug"]`, 그리고 `null` 값을 반환할 것입니다.

# --instructions--

정규식 `a`에 모음(`e`, `i`, `o`, `u`, `vowelRegex`)을 포함하는 문자 클래스를 사용하여 문자열 `quoteSample` 안의 모든 모음을 찾으세요.

**참고:** 대문자와 소문자 모음 모두 매칭해야 합니다.

# --hints--

모음 25개를 모두 찾아야 합니다.

```js
assert(result.length == 25);
```

정규식 `vowelRegex`는 문자 클래스를 사용해야 합니다.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

정규식 `vowelRegex`는 전역 플래그를 사용해야 합니다.

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

정규식 `vowelRegex`는 대소문자 구분 안 함 플래그를 사용해야 합니다.

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

정규식은 자음과 매칭하면 안 됩니다.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
