---
id: 587d7db5367417b2b2512b94
title: 와일드카드 마침표로 모든 문자 매칭하기
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

때때로 패턴에 정확한 문자를 알지 못하거나 알 필요가 없을 때가 있습니다. 예를 들어, 오타가 포함된 모든 단어를 생각하는 데는 시간이 오래 걸립니다. 다행히도 와일드카드 문자를 사용하면 시간을 절약할 수 있습니다: `.`

와일드카드 문자 `.`는 임의의 한 문자를 매칭합니다. 와일드카드는 `dot`와 `period`라고도 불립니다. 정규식에서 다른 문자처럼 와일드카드 문자를 사용할 수 있습니다. 예를 들어, `hug`, `huh`, `hut`, `hum`를 모두 매칭하고 싶다면, 정규식 `/hu./`를 사용해 네 단어를 모두 매칭할 수 있습니다.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

이 두 `test` 호출은 모두 `true`를 반환합니다.

# --instructions--

정규식 `unRegex`를 완성하여 `run`, `sun`, `fun`, `pun`, `nun`, `bun` 문자열과 매칭되게 하세요. 정규식에는 와일드카드 문자를 사용해야 합니다.

# --hints--

`.test()` 메서드를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

정규식 `unRegex`에 와일드카드 문자를 사용하세요.

```js
assert(/\./.test(unRegex.source));
```

정규식 `unRegex`는 문자열 `run`에서 `Let us go on a run.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

정규식 `unRegex`는 문자열 `sun`에서 `The sun is out today.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

정규식 `unRegex`는 문자열 `fun`에서 `Coding is a lot of fun.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

정규식 `unRegex`는 문자열 `pun`에서 `Seven days without a pun makes one weak.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

정규식 `unRegex`는 문자열 `nun`에서 `One takes a vow to be a nun.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

정규식 `unRegex`는 문자열 `bun`에서 `She got fired from the hot dog stand for putting her hair in a bun.`와 매칭되어야 합니다.

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

정규식 `unRegex`는 문자열 `There is a bug in my code.`와 매칭되면 안 됩니다.

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

정규식 `unRegex`는 문자열 `Catch me if you can.`와 매칭되면 안 됩니다.

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
