---
id: 587d7db9367417b2b2512ba7
title: 정확한 매치 개수 지정
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

중괄호를 사용한 수량 지정자를 통해 패턴의 최소 및 최대 개수를 지정할 수 있습니다. 때로는 특정 개수의 매치만 원할 때가 있습니다.

특정 개수의 패턴을 지정하려면 중괄호 안에 그 숫자 하나만 넣으면 됩니다.

예를 들어, `hah`라는 단어가 `a` `3`번만 나타나는 경우에 매치하려면 정규식은 `/ha{3}h/`가 됩니다.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

순서대로 세 번의 `test` 호출은 `false`, `true`, `false`를 반환합니다.

# --instructions--

정규식 `timRegex`를 변경하여 `Timber`라는 단어가 정확히 네 글자 `m`를 가질 때만 매치되게 하세요.

# --hints--

정규식에 중괄호를 사용해야 합니다.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

정규식은 문자열 `Timber`와 매치되면 안 됩니다.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

정규식은 문자열 `Timmber`와 매치되면 안 됩니다.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

정규식은 문자열 `Timmmber`와 매치되면 안 됩니다.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

정규식은 문자열 `Timmmmber`와 매치되어야 합니다.

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

정규식은 30개의 `Timber`가 포함된 문자열 `m`와 매치되면 안 됩니다.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
