---
id: 587d7db9367417b2b2512ba4
title: 공백이 아닌 문자 찾기
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

소문자 `\s`가 포함된 `s`를 사용해 공백을 검색하는 방법을 배웠습니다. 공백이 아닌 모든 문자를 검색할 수도 있습니다.

대문자 `\S`인 `s`를 사용해 공백이 아닌 문자를 검색하세요. 이 패턴은 공백, 캐리지 리턴, 탭, 폼 피드, 새 줄 문자를 매치하지 않습니다. `[^ \r\t\f\n\v]` 문자 클래스와 비슷하다고 생각할 수 있습니다.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

`.length` 메서드가 반환하는 값은 `32`입니다.

# --instructions--

정규식 `countNonWhiteSpace`를 변경해 문자열에서 여러 개의 공백이 아닌 문자를 찾으세요.

# --hints--

정규식에 전역 플래그를 사용하세요.

```js
assert(countNonWhiteSpace.global);
```

정규식에 모든 공백이 아닌 문자를 매치하는 축약 문자 `\S`를 사용하세요.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

정규식이 문자열 `Men are from Mars and women are from Venus.`에서 35개의 공백이 아닌 문자를 찾아야 합니다.

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

정규식이 문자열 `Space: the final frontier.`에서 23개의 공백이 아닌 문자를 찾아야 합니다.

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

정규식이 문자열 `MindYourPersonalSpace`에서 21개의 공백이 아닌 문자를 찾아야 합니다.

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
