---
id: 587d7db8367417b2b2512ba3
title: 공백 문자 일치시키기
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

지금까지 도전 과제에서는 알파벳 문자와 숫자를 일치시키는 방법을 다뤘습니다. 문자 사이의 공백이나 띄어쓰기도 일치시킬 수 있습니다.

소문자 `\s`인 `s`을 사용해 공백을 검색할 수 있습니다. 이 패턴은 공백뿐 아니라 캐리지 리턴, 탭, 폼 피드, 새 줄 문자도 일치시킵니다. `[ \r\t\f\n\v]` 문자 클래스와 비슷하다고 생각할 수 있습니다.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

이 `match` 호출은 `[" ", " "]`를 반환합니다.
# --instructions--

정규식 `countWhiteSpace`를 변경해 문자열에서 여러 개의 공백 문자를 찾으세요.

# --hints--

정규식에 전역 플래그를 사용해야 합니다.

```js
assert(countWhiteSpace.global);
```

정규식에 모든 공백 문자를 일치시키는 축약 문자 `\s`를 사용해야 합니다.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

정규식은 문자열 `Men are from Mars and women are from Venus.`에서 8개의 공백을 찾아야 합니다.

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

정규식은 문자열 `Space: the final frontier.`에서 3개의 공백을 찾아야 합니다.

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

정규식은 문자열 `MindYourPersonalSpace`에서 공백을 찾지 못해야 합니다.

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
