---
id: 587d7dba367417b2b2512ba9
title: 긍정 및 부정 전방 탐색
challengeType: 1
forumTopicId: 301360
dashedName: positive-and-negative-lookahead
---

# --description--

<dfn>전방 탐색</dfn>은 JavaScript가 문자열에서 앞으로 나아가 패턴을 확인하도록 지시하는 패턴입니다. 이는 같은 문자열에서 여러 패턴을 검색할 때 유용합니다.

전방 탐색에는 두 가지 종류가 있습니다: <dfn>긍정 전방 탐색</dfn>과 <dfn>부정 전방 탐색</dfn>입니다.

긍정 전방 탐색은 검색 패턴에 해당 요소가 있는지 확인하지만 실제로는 매치하지 않습니다. 긍정 전방 탐색은 `(?=...)` 형태로 사용되며, `...` 부분이 매치되지 않는 필수 요소입니다.

반면에 부정 전방 탐색은 검색 패턴에 해당 요소가 없는지 확인합니다. 부정 전방 탐색은 `(?!...)` 형태로 사용되며, `...` 부분이 포함되면 안 되는 패턴입니다. 부정 전방 탐색 부분이 없으면 나머지 패턴이 반환됩니다.

전방 탐색은 다소 혼란스러울 수 있지만 몇 가지 예제가 도움이 될 것입니다.

```js
let quit = "qu";
let noquit = "qt";
let quRegex= /q(?=u)/;
let qRegex = /q(?!u)/;
quit.match(quRegex);
noquit.match(qRegex);
```

이 두 `match` 호출은 모두 `["q"]`를 반환합니다.

전방 탐색의 보다 실용적인 사용법은 하나의 문자열에서 두 개 이상의 패턴을 확인하는 것입니다. 다음은 3~6자 사이이며 최소 한 개의 숫자가 포함된 (단순한) 비밀번호 검사기입니다:

```js
let password = "abc123";
let checkPass = /(?=\w{3,6})(?=\D*\d)/;
checkPass.test(password);
```

# --instructions--

`pwRegex`에서 전방 탐색을 사용해 5자 이상이고 연속된 두 개의 숫자가 포함된 비밀번호와 매치하세요.

# --hints--

정규식에는 두 개의 긍정 `lookaheads`가 사용되어야 합니다.

```js
assert(pwRegex.source.match(/\(\?=.*?\)\(\?=.*?\)/) !== null);
```

정규식은 `astronaut` 문자열과 매치되지 않아야 합니다.

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('astronaut'));
```

정규식은 `banan1` 문자열과 매치되지 않아야 합니다.

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('banan1'));
```

정규식은 `bana12` 문자열과 매치되어야 합니다.

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('bana12'));
```

정규식은 `abc123` 문자열과 매치되어야 합니다.

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('abc123'));
```

정규식은 `12345` 문자열과 매치되지 않아야 합니다.

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('12345'));
```

정규식은 `8pass99` 문자열과 매치되어야 합니다.

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('8pass99'));
```

정규식은 `1a2bcde` 문자열과 매치되지 않아야 합니다.

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('1a2bcde'));
```

정규식은 `astr1on11aut` 문자열과 매치되어야 합니다.

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('astr1on11aut'));
```

# --seed--

## --seed-contents--

```js
let sampleWord = "astronaut";
let pwRegex = /change/; // Change this line
let result = pwRegex.test(sampleWord);
```

# --solutions--

```js
let pwRegex =  /(?=\w{6})(?=\w*\d{2})/;
```
