---
id: 587d7db7367417b2b2512b9f
title: 모든 문자와 숫자 일치시키기
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

문자 클래스(character classes)를 사용해서 `[a-z]`로 알파벳의 모든 문자를 검색할 수 있었습니다. 이런 종류의 문자 클래스는 꽤 흔해서 이를 위한 단축키도 있습니다. 다만 몇 가지 추가 문자가 포함되어 있습니다.

JavaScript에서 알파벳과 가장 가까운 문자 클래스는 `\w`입니다. 이 단축키는 `[A-Za-z0-9_]`와 같습니다. 이 문자 클래스는 대문자와 소문자, 그리고 숫자를 모두 일치시킵니다. 참고로 이 문자 클래스에는 밑줄 문자(`_`)도 포함되어 있습니다.

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

이 네 가지 `test` 호출은 모두 `true`를 반환합니다.

이 단축 문자 클래스들은 <dfn>단축 문자 클래스(shorthand character classes)</dfn>라고도 불립니다.

# --instructions--

단축 문자 클래스 `\w`를 사용해서 다양한 인용문과 문자열에서 영숫자 문자의 개수를 세세요.

# --hints--

정규식에는 전역 플래그를 사용하세요.

```js
assert(alphabetRegexV2.global);
```

정규식에는 모든 영숫자 문자를 일치시키기 위해 단축 문자 `\w`를 사용하세요.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

정규식은 문자열 `The five boxing wizards jump quickly.`에서 31개의 영숫자 문자를 찾아야 합니다.

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

정규식은 문자열 `Pack my box with five dozen liquor jugs.`에서 32개의 영숫자 문자를 찾아야 합니다.

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

정규식은 문자열 `How vexingly quick daft zebras jump!`에서 30개의 영숫자 문자를 찾아야 합니다.

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

정규식은 문자열 `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`에서 36개의 영숫자 문자를 찾아야 합니다.

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
