---
id: 587d7db9367417b2b2512ba5
title: 일치 횟수의 상한과 하한 지정
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

플러스 기호 `+`는 하나 이상의 문자를 찾을 때, 별표 `*`는 0개 이상의 문자를 찾을 때 사용한다는 점을 기억하세요. 이들은 편리하지만 때로는 특정 범위의 패턴과 일치시키고 싶을 때가 있습니다.

중괄호(`{`, `}`)와 함께 사용하는 <dfn>수량 지정자</dfn>로 패턴의 하한과 상한을 지정할 수 있습니다. 중괄호 안에 두 숫자를 넣어 패턴의 최소와 최대 횟수를 나타냅니다.

예를 들어, 문자열 `a`에서 문자 `3`가 `5`번에서 `ah`번 사이에 나타나는 경우에만 일치시키려면 정규식은 `/a{3,5}h/`가 됩니다.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

첫 번째 `test` 호출은 `true`를 반환하고, 두 번째 호출은 `false`를 반환합니다.

# --instructions--

정규식 `ohRegex`를 변경하여 `Oh no` 전체 구문과 일치시키되, `3`에서 `6` 사이의 `h` 문자가 있을 때만 일치하도록 하세요.

# --hints--

정규식에 중괄호를 사용해야 합니다.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

정규식이 문자열 `Ohh no`와 일치하지 않아야 합니다.

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

정규식이 문자열 `Ohhh no`와 일치해야 합니다.

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

정규식이 문자열 `Ohhhh no`와 일치해야 합니다.

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

정규식이 문자열 `Ohhhhh no`와 일치해야 합니다.

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

정규식이 문자열 `Ohhhhhh no`와 일치해야 합니다.

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

정규식이 문자열 `Ohhhhhhh no`와 일치하지 않아야 합니다.

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
