---
id: 587d7db8367417b2b2512ba2
title: 가능한 사용자 이름 제한
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

사용자 이름은 인터넷 어디에서나 사용됩니다. 사용자에게 좋아하는 사이트에서 고유한 정체성을 부여합니다.

데이터베이스에 있는 모든 사용자 이름을 확인해야 합니다. 사용자 이름을 만들 때 사용자가 따라야 하는 간단한 규칙은 다음과 같습니다.

1) 사용자 이름은 영숫자 문자만 사용할 수 있습니다.

2) 사용자 이름에 있는 숫자는 끝에만 있어야 합니다. 끝에 0개 이상 있을 수 있습니다. 사용자 이름은 숫자로 시작할 수 없습니다.

3) 사용자 이름의 문자는 소문자와 대문자를 모두 사용할 수 있습니다.

4) 사용자 이름은 최소 두 글자여야 합니다. 두 글자 사용자 이름은 문자만 사용할 수 있습니다.

# --instructions--

위에 나열된 제약 조건에 맞게 정규식 `userCheck`를 변경하세요.

# --hints--

정규식은 문자열 `JACK`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

정규식은 문자열 `J`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

정규식은 문자열 `Jo`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

정규식은 문자열 `Oceans11`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

정규식은 문자열 `RegexGuru`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

정규식은 문자열 `007`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

정규식은 문자열 `9`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

정규식은 문자열 `A1`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

정규식은 문자열 `BadUs3rnam3`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

정규식은 문자열 `Z97`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

정규식은 문자열 `c57bT3`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

정규식은 문자열 `AB1`과 일치해야 합니다.

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

정규식은 문자열 `J%4`와 일치하지 않아야 합니다.

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
