---
id: 587d7db4367417b2b2512b92
title: 일치 항목 추출하기
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

지금까지는 문자열 내에 패턴이 존재하는지 여부만 확인했습니다. `.match()` 메서드를 사용해 실제로 찾은 일치 항목을 추출할 수도 있습니다.

`.match()` 메서드를 사용하려면 문자열에 메서드를 적용하고 소괄호 안에 정규식을 전달하세요.

예를 들어 다음과 같습니다:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

여기서 첫 번째 `match`는 `["Hello"]`를 반환하고 두 번째는 `["expressions"]`를 반환합니다.

`.match` 구문은 지금까지 사용해온 `.test` 메서드의 "반대"임을 참고하세요:

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

`.match()` 메서드를 적용해 문자열 `coding`를 추출하세요.

# --hints--

`result`에는 문자열 `coding`가 있어야 합니다.

```js
assert(result.join() === 'coding');
```

정규식 `codingRegex`는 문자열 `coding`를 찾아야 합니다.

```js
assert(codingRegex.source === 'coding');
```

`.match()` 메서드를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
