---
id: 587d7dbb367417b2b2512baa
title: 캡처 그룹을 사용한 패턴 재사용
challengeType: 1
forumTopicId: 301364
dashedName: reuse-patterns-using-capture-groups
---

# --description--

아래처럼 여러 번 반복되는 단어를 찾고 싶다고 가정해보세요.

```js
let repeatStr = "row row row your boat";
```

`/row row row/`를 사용할 수 있지만, 반복되는 특정 단어를 모를 때는 어떻게 할까요? <dfn>캡처 그룹</dfn>을 사용하면 반복되는 부분 문자열을 찾을 수 있습니다.

캡처 그룹은 캡처할 정규식 패턴을 소괄호로 감싸서 만듭니다. 이 경우, 영숫자로 이루어진 단어를 캡처하는 것이 목표이므로 캡처 그룹은 소괄호로 감싼 `\w+`가 됩니다: `/(\w+)/`.

그룹에 매칭된 부분 문자열은 임시 "변수"에 저장되며, 같은 정규식 내에서 역슬래시와 캡처 그룹 번호(예: `\1`)를 사용해 접근할 수 있습니다. 캡처 그룹 번호는 여는 소괄호 위치(왼쪽에서 오른쪽) 순서대로 자동으로 1부터 매겨집니다.

아래 예시는 공백으로 구분된 세 번 반복되는 단어를 매칭합니다:

```js
let repeatRegex = /(\w+) \1 \1/;
repeatRegex.test(repeatStr); // Returns true
repeatStr.match(repeatRegex); // Returns ["row row row", "row"]
```

문자열에서 `.match()` 메서드를 사용하면 매칭된 부분 문자열과 캡처된 그룹이 포함된 배열을 반환합니다.


# --instructions--

`reRegex`에서 캡처 그룹을 사용해, 공백으로 구분된 같은 숫자가 정확히 세 번 반복되는 문자열을 매칭하세요.

# --hints--

정규식에는 숫자에 대한 축약 문자 클래스를 사용하세요.

```js
assert(reRegex.source.match(/\\d/));
```

정규식은 캡처 그룹을 두 번 재사용해야 합니다.

```js
assert(reRegex.source.match(/\\1|\\2/g).length >= 2);
```

정규식은 `42 42 42` 문자열과 매칭되어야 합니다.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('42 42 42'));
```

정규식은 `100 100 100` 문자열과 매칭되어야 합니다.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('100 100 100'));
```

정규식은 `42 42 42 42` 문자열과 매칭되지 않아야 합니다.

```js
assert.equal('42 42 42 42'.match(reRegex.source), null);
```

정규식은 `42 42` 문자열과 매칭되지 않아야 합니다.

```js
assert.equal('42 42'.match(reRegex.source), null);
```

정규식은 `101 102 103` 문자열과 매칭되지 않아야 합니다.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('101 102 103'));
```

정규식은 `1 2 3` 문자열과 매칭되지 않아야 합니다.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('1 2 3'));
```

정규식은 `10 10 10` 문자열과 매칭되어야 합니다.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('10 10 10'));
```

정규식은 `42\t42\t42` 문자열과 매칭되지 않아야 합니다.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42\t42\t42'));
```

정규식은 `42  42  42` 문자열과 매칭되지 않아야 합니다.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42  42  42'));
```

# --seed--

## --seed-contents--

```js
let repeatNum = "42 42 42";
let reRegex = /change/; // Change this line
let result = reRegex.test(repeatNum);
```

# --solutions--

```js
let repeatNum = "42 42 42";
let reRegex = /^(\d+) \1 \1$/;
let result = reRegex.test(repeatNum);
```
