---
id: 587d7db6367417b2b2512b99
title: 한 번 이상 반복되는 문자 매칭하기
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

때때로, 한 번 이상 연속해서 나타나는 문자(또는 문자 그룹)를 매칭해야 할 때가 있습니다. 이는 최소 한 번 이상 나타나며 반복될 수도 있다는 뜻입니다.

이럴 때 `+` 문자를 사용해 그런 경우인지 확인할 수 있습니다. 기억하세요, 문자나 패턴은 연속해서 나타나야 합니다. 즉, 문자가 하나씩 이어서 반복되어야 합니다.

예를 들어, `/a+/g`는 `abc`에서 한 번 매칭되어 `["a"]`를 반환합니다. `+` 때문에 `aabc`에서도 한 번 매칭되어 `["aa"]`를 반환합니다.

만약 `abab` 문자열을 검사한다면, `["a", "a"]` 문자가 연속하지 않고 `a`가 그 사이에 있어서 두 번 매칭되어 `b`를 반환합니다. 마지막으로 `a` 문자열에 `bcd`가 없으므로 매칭을 찾지 못합니다.

# --instructions--

문자 `s`가 `Mississippi`에서 한 번 이상 나타날 때 매칭을 찾으세요. `+` 기호를 사용하는 정규식을 작성하세요.

# --hints--

정규식 `myRegex`는 한 번 이상 `+` 문자를 매칭하기 위해 `s` 기호를 사용해야 합니다.

```js
assert(/\+/.test(myRegex.source));
```

정규식 `myRegex`는 2개의 항목과 매칭되어야 합니다.

```js
assert(result.length == 2);
```

`result` 변수는 `ss` 두 개의 매칭을 가진 배열이어야 합니다.

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
