---
id: 587d7db5367417b2b2512b96
title: 알파벳 문자 일치시키기
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

<dfn>문자 집합</dfn>을 사용해 일치시킬 문자 그룹을 지정하는 방법을 배웠지만, 알파벳의 모든 문자처럼 넓은 범위의 문자를 일치시켜야 할 때는 입력해야 할 내용이 많아집니다. 다행히도 이 과정을 간단하고 짧게 만들어 주는 내장 기능이 있습니다.

문자 집합 안에서는 하이픈 문자 `-`를 사용해 일치시킬 문자 범위를 정의할 수 있습니다.

예를 들어, 소문자 `a`부터 `e`까지 일치시키려면 `[a-e]`를 사용하면 됩니다.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

순서대로 세 개의 `match` 호출은 `["cat"]`, `["bat"]`, `null` 값을 반환합니다.

# --instructions--

문자열 `quoteSample`에 있는 모든 문자를 일치시키세요.

**참고**: 대문자와 소문자 모두 일치시키는지 확인하세요.

# --hints--

정규식 `alphabetRegex`는 35개의 항목과 일치해야 합니다.

```js
assert(result.length == 35);
```

정규식 `alphabetRegex`는 전역 플래그를 사용해야 합니다.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

정규식 `alphabetRegex`는 대소문자 구분 안 함 플래그를 사용해야 합니다.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
