---
id: 587d7dbb367417b2b2512bac
title: 시작과 끝의 공백 제거
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

문자열 주변에 원하지 않는 공백 문자가 있을 때가 있습니다. 문자열을 처리할 때 일반적으로 시작과 끝의 공백을 제거합니다.

# --instructions--

정규식을 작성하고 적절한 문자열 메서드를 사용해 문자열의 시작과 끝에 있는 공백을 제거하세요.

**참고:** `String.prototype.trim()` 메서드도 사용할 수 있지만, 이 도전 과제는 정규식을 사용해 완료해야 합니다.

# --hints--

`result`은 문자열 `Hello, World!`와 같아야 합니다.

```js
assert(result === 'Hello, World!');
```

해결책에 `String.prototype.trim()` 메서드를 사용하면 안 됩니다.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

`result` 변수는 문자열로 직접 설정하면 안 됩니다.

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

`hello` 변수의 값은 변경하면 안 됩니다.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
