---
id: 587d7db7367417b2b2512b9d
title: 문자열 시작 패턴 일치시키기
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

이전 도전 과제에서는 정규 표현식을 사용해 여러 개의 일치 항목을 찾을 수 있음을 보여주었습니다. 정규 표현식은 문자열 내 특정 위치의 패턴을 찾는 데에도 사용됩니다.

이전 도전 과제에서 캐럿 문자(`^`)를 문자 집합 안에 넣어 `[^thingsThatWillNotBeMatched]` 형태의 부정 문자 집합을 만들었습니다. 문자 집합 밖에서 캐럿은 문자열의 시작 부분에서 패턴을 찾는 데 사용됩니다.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

첫 번째 `test` 호출은 `true`를 반환하는 반면, 두 번째 호출은 `false`를 반환합니다.

# --instructions--

정규식에서 캐럿 문자를 사용해 `Cal` 문자열의 시작 부분에서만 `rickyAndCal`를 찾으세요.

# --hints--

정규식은 대문자로 시작하는 `Cal` 문자열을 찾아야 합니다.

```js
assert(calRegex.source == '^Cal');
```

정규식에 플래그를 사용하지 마세요.

```js
assert(calRegex.flags == '');
```

정규식은 문자열의 시작 부분에서 `Cal` 문자열과 일치해야 합니다.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

정규식은 문자열 중간에 있는 `Cal` 문자열과 일치하지 않아야 합니다.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
