---
id: 587d7db6367417b2b2512b9a
title: 0회 이상 반복되는 문자와 일치시키기
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

이전 도전 과제에서는 더하기 `+` 기호를 사용해 1회 이상 반복되는 문자를 찾았습니다. 0회 이상 반복되는 문자와 일치시키는 옵션도 있습니다.

이 기능을 하는 문자는 별표 또는 아스터리스크입니다: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

순서대로 세 개의 `match` 호출은 `["goooooooo"]`, `["g"]`, `null` 값을 반환합니다.

# --instructions--

대문자 `chewieRegex` 문자 바로 뒤에 0회 이상 소문자 `*` 문자가 오는 `A`에서 일치시키기 위해 `a` 문자를 사용하는 정규식 `chewieQuote`를 만드세요. 정규식에 플래그나 문자 클래스를 사용할 필요 없으며, 다른 인용구와는 일치하지 않아야 합니다.

# --hints--

정규식 `chewieRegex`는 0회 이상 `*` 문자를 일치시키기 위해 `a` 문자를 사용해야 합니다.

```js
assert(/\*/.test(chewieRegex.source));
```

정규식은 `A`에서 문자열 `chewieQuote`와 일치해야 합니다.

```js
assert(result[0][0] === 'A');
```

정규식은 `Aaaaaaaaaaaaaaaa`에서 문자열 `chewieQuote`와 일치해야 합니다.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

정규식 `chewieRegex`는 `chewieQuote`에서 16개의 문자와 일치해야 합니다.

```js
assert(result[0].length === 16);
```

정규식은 문자열 `He made a fair move. Screaming about it can't help you.`의 어떤 문자와도 일치하지 않아야 합니다.

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

정규식은 문자열 `Let him have it. It's not wise to upset a Wookiee.`의 어떤 문자와도 일치하지 않아야 합니다.

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
