---
id: 587d7dba367417b2b2512ba8
title: 모두 또는 없음 확인하기
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

가끔 찾으려는 패턴에 존재할 수도 있고 없을 수도 있는 부분이 있을 수 있습니다. 하지만 그 부분을 확인하는 것이 중요할 때가 있습니다.

물음표 `?`를 사용해 요소가 있을 수도 있고 없을 수도 있음을 지정할 수 있습니다. 이는 바로 앞 요소가 0개 또는 1개임을 검사합니다. 이 기호는 이전 요소가 선택 사항임을 나타낸다고 생각하세요.

예를 들어, 미국 영어와 영국 영어에는 약간의 차이가 있고, 물음표를 사용해 두 가지 철자 모두에 맞출 수 있습니다.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

`test` 메서드를 두 가지 모두 사용하면 `true`가 반환됩니다.

# --instructions--

정규식 `favRegex`를 변경해 미국 영어(`favorite`)와 영국 영어(`favourite`) 버전의 단어 모두에 맞추세요.

# --hints--

정규식에 선택적 기호 `?`를 사용해야 합니다.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

정규식은 문자열 `favorite`에 맞아야 합니다.

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

정규식은 문자열 `favourite`에 맞아야 합니다.

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

정규식은 문자열 `fav`에는 맞지 않아야 합니다.

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
