---
id: 587d7db4367417b2b2512b93
title: 첫 번째 일치 이상 찾기
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

지금까지는 패턴을 한 번만 추출하거나 검색할 수 있었습니다.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

여기서 `match`는 `["Repeat"]`를 반환합니다.

패턴을 한 번 이상 검색하거나 추출하려면 전역 검색 플래그인 `g`를 사용할 수 있습니다.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

그리고 여기서 `match`는 값 `["Repeat", "Repeat", "Repeat"]`를 반환합니다.

# --instructions--

정규식 `starRegex`를 사용하여 문자열 `Twinkle`에서 `twinkleStar` 단어 두 개를 찾아 추출하세요.

**참고**  
정규식에 `/search/gi`처럼 여러 플래그를 사용할 수 있습니다.

# --hints--

정규식 `starRegex`는 전역 플래그 `g`를 사용해야 합니다.

```js
assert(starRegex.flags.match(/g/).length == 1);
```

정규식 `starRegex`는 대소문자 구분 안 함 플래그 `i`를 사용해야 합니다.

```js
assert(starRegex.flags.match(/i/).length == 1);
```

일치는 단어 `Twinkle`의 두 번 모두에 일치해야 합니다.

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

일치 `result`에는 두 개의 요소가 있어야 합니다.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
