---
id: 587d7db3367417b2b2512b8e
title: 테스트 메서드 사용하기
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

정규 표현식은 프로그래밍 언어에서 문자열의 일부를 찾을 때 사용합니다. 일치시킬 패턴을 만드세요.

문자열 `the`에서 단어 `The dog chased the cat`를 찾고 싶다면, 다음 정규 표현식 `/the/`를 사용할 수 있습니다. 정규 표현식 안에는 따옴표가 필요하지 않다는 점을 기억하세요.

JavaScript에는 정규식을 사용하는 여러 방법이 있습니다. 정규식을 테스트하는 한 가지 방법은 `.test()` 메서드를 사용하는 것입니다. `.test()` 메서드는 정규식을 받아 소괄호 안에 넣은 문자열에 적용하고, 패턴이 일치하면 `true`를, 그렇지 않으면 `false`를 반환합니다.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

여기서 `test` 메서드는 `true`를 반환합니다.

# --instructions--

정규식 `myRegex`를 문자열 `myString`에 `.test()` 메서드를 사용해 적용하세요.

# --hints--

정규식을 테스트하려면 `.test()`를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

결과는 `true`를 반환해야 합니다.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
