---
id: 587d7db4367417b2b2512b91
title: 대소문자 구분 없이 일치시키기
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

지금까지는 문자열을 문자 그대로 일치시키는 정규식을 살펴봤습니다. 하지만 때로는 대소문자 차이도 일치시키고 싶을 때가 있습니다.

대소문자(또는 문자 대소)는 대문자와 소문자의 차이를 말합니다. 대문자의 예로는 `A`, `B`, `C`가 있습니다. 소문자의 예로는 `a`, `b`, `c`가 있습니다.

대소문자 모두를 일치시키려면 플래그라고 하는 것을 사용할 수 있습니다. 다른 플래그도 있지만 여기서는 대소문자를 무시하는 `i` 플래그에 집중합니다. 이 플래그는 정규식 뒤에 붙여서 사용할 수 있습니다. 이 플래그를 사용한 예는 `/ignorecase/i`입니다. 이 정규식은 `ignorecase`, `igNoreCase`, `IgnoreCase` 문자열과 일치할 수 있습니다.

# --instructions--

대소문자에 상관없이 `fccRegex`와 일치하는 정규식 `freeCodeCamp`을 작성하세요. 정규식은 공백이 포함된 약어나 변형과는 일치하지 않아야 합니다.

# --hints--

정규식은 `freeCodeCamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

정규식은 `FreeCodeCamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

정규식은 `FreecodeCamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

정규식은 `FreeCodecamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

정규식은 `Free Code Camp` 문자열과 일치하지 않아야 합니다.

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

정규식은 `FreeCOdeCamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

정규식은 `FCC` 문자열과 일치하지 않아야 합니다.

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

정규식은 `FrEeCoDeCamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

정규식은 `FrEeCodECamp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

정규식은 `FReeCodeCAmp` 문자열과 일치해야 합니다.

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
