---
id: 587d7db6367417b2b2512b9b
title: 게으른 매칭으로 문자 찾기
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

정규 표현식에서 <dfn>탐욕적</dfn> 매칭은 정규식 패턴에 맞는 문자열의 가장 긴 부분을 찾아서 매칭 결과로 반환합니다. 반대로 <dfn>게으른</dfn> 매칭은 정규식 패턴을 만족하는 문자열의 가장 짧은 부분을 찾습니다.

정규식 `/t[a-z]*i/`를 문자열 `"titanic"`에 적용할 수 있습니다. 이 정규식은 기본적으로 `t`로 시작하고 `i`로 끝나며 그 사이에 몇 개의 문자가 있는 패턴입니다.

정규 표현식은 기본적으로 탐욕적이므로, 매칭 결과는 `["titani"]`가 됩니다. 패턴에 맞는 가장 큰 부분 문자열을 찾는 것입니다.

하지만 `?` 문자를 사용하면 게으른 매칭으로 바꿀 수 있습니다. `"titanic"`를 조정된 정규식 `/t[a-z]*?i/`에 적용하면 `["ti"]`가 반환됩니다.

**참고:** 정규 표현식으로 HTML을 파싱하는 것은 피해야 하지만, HTML 문자열에서 패턴 매칭을 하는 것은 완전히 괜찮습니다.

# --instructions--

정규식 `/<.*>/`를 수정하여 텍스트 `<h1>`가 아니라 HTML 태그 `"<h1>Winter is coming</h1>"`를 반환하도록 하세요. 정규식에서 와일드카드 `.`는 어떤 문자든 매칭한다는 점을 기억하세요.

# --hints--

`result` 변수는 `<h1>`를 포함하는 배열이어야 합니다

```js
assert(result[0] == '<h1>');
```

`myRegex`는 게으른 매칭을 사용해야 합니다

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

`myRegex`는 문자열 `h1`를 포함하면 안 됩니다

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
