---
id: 587d7db5367417b2b2512b97
title: 숫자와 알파벳 문자 일치시키기
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

하이픈(`-`)을 사용해 문자 범위를 일치시키는 것은 문자에만 국한되지 않습니다. 숫자 범위를 일치시키는 데도 사용할 수 있습니다.

예를 들어, `/[0-5]/`은 `0`부터 `5` 사이의 모든 숫자와 `0`, `5`를 포함해 일치시킵니다.

또한, 문자 범위와 숫자 범위를 하나의 문자 집합에 결합하는 것도 가능합니다.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

`h`부터 `s` 사이의 문자 범위와 `2`부터 `6` 사이의 숫자 범위를 일치시키는 단일 정규식을 만드세요. 정규식에 적절한 플래그를 포함하는 것도 잊지 마세요.

# --hints--

정규식 `myRegex`은 17개의 항목과 일치해야 합니다.

```js
assert(result.length == 17);
```

정규식 `myRegex`은 전역 플래그를 사용해야 합니다.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

정규식 `myRegex`은 대소문자 구분 없는 플래그를 사용해야 합니다.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
