---
id: 587d7db3367417b2b2512b8f
title: 리터럴 문자열 일치
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

지난 도전 과제에서는 정규 표현식 `Hello`을 사용해 단어 `/Hello/`를 검색했습니다. 그 정규식은 문자열 `Hello`와 리터럴 일치를 검색했습니다. 여기 문자열 `Kevin`와 리터럴 일치를 검색하는 또 다른 예가 있습니다:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

이 `test` 호출은 `true`를 반환합니다.

다른 형태의 `Kevin`는 일치하지 않습니다. 예를 들어, 정규식 `/Kevin/`는 `kevin`나 `KEVIN`와 일치하지 않습니다.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

이 `test` 호출은 `false`를 반환합니다.

앞으로의 도전 과제에서는 다른 형태들도 일치시키는 방법을 보여줄 것입니다.

# --instructions--

정규식 `waldoRegex`를 완성해 문자열 `"Waldo"`에서 `waldoIsHiding`를 리터럴로 찾으세요.

# --hints--

정규식 `waldoRegex`는 문자열 `Waldo`를 찾아야 합니다.

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

정규식 `waldoRegex`는 다른 것은 검색하지 않아야 합니다.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

정규식으로 리터럴 문자열 일치를 수행해야 합니다.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
