---
id: 587d7db8367417b2b2512ba0
title: 영숫자를 제외한 모든 문자 매칭하기
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

영숫자를 `[A-Za-z0-9_]`로 매칭할 때 `\w`라는 단축키를 사용할 수 있다는 것을 배웠습니다. 여러분이 찾고 싶어 할 자연스러운 패턴 중 하나는 영숫자의 반대입니다.

`\w`를 사용해 `\W`의 반대 패턴을 검색할 수 있습니다. 반대 패턴은 대문자를 사용한다는 점에 주의하세요. 이 단축키는 `[^A-Za-z0-9_]`와 동일합니다.

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

첫 번째 `match` 호출은 `["%"]` 값을 반환하고 두 번째 호출은 `["!"]` 값을 반환합니다.

# --instructions--

단축 문자 클래스 `\W`를 사용해 다양한 인용문과 문자열에서 영숫자가 아닌 문자의 개수를 세세요.

# --hints--

정규식에 전역 플래그를 사용하세요.

```js
assert(nonAlphabetRegex.global);
```

정규식은 문자열 `The five boxing wizards jump quickly.`에서 6개의 영숫자가 아닌 문자를 찾아야 합니다.

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

정규식은 영숫자가 아닌 문자를 매칭하기 위해 단축 문자를 사용해야 합니다.

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

정규식은 문자열 `Pack my box with five dozen liquor jugs.`에서 8개의 영숫자가 아닌 문자를 찾아야 합니다.

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

정규식은 문자열 `How vexingly quick daft zebras jump!`에서 6개의 영숫자가 아닌 문자를 찾아야 합니다.

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

정규식은 문자열 `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`에서 12개의 영숫자가 아닌 문자를 찾아야 합니다.

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
