---
id: 587d7db7367417b2b2512b9e
title: 문자열 끝 패턴 일치시키기
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

지난 도전 과제에서는 캐럿 문자로 문자열 시작 부분의 패턴을 검색하는 방법을 배웠습니다. 문자열 끝 부분의 패턴을 검색하는 방법도 있습니다.

정규식 끝에 달러 기호 `$`를 사용하여 문자열 끝을 검색할 수 있습니다.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

첫 번째 `test` 호출은 `true`를 반환하고, 두 번째 호출은 `false`를 반환합니다.

# --instructions--

앵커 문자(`$`)를 사용하여 문자열 `caboose`의 끝에 있는 문자열 `caboose`와 일치시키세요.

# --hints--

정규식에서 달러 기호 `caboose` 앵커로 `$`를 검색해야 합니다.

```js
assert(lastRegex.source == 'caboose$');
```

정규식에 플래그를 사용하지 않아야 합니다.

```js
assert(lastRegex.flags == '');
```

문자열 `caboose`의 끝에 `The last car on a train is the caboose`와 일치시켜야 합니다.

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
