---
id: 5c3dda8b4d8df89bea71600f
title: 문자 그룹 혼합 검사
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

가끔 정규 표현식을 사용해 문자 그룹을 검사하고 싶을 때가 있습니다. 이를 위해 소괄호 `()`를 사용합니다.

문자열에서 `Penguin` 또는 `Pumpkin`를 찾고 싶다면, 다음 정규 표현식을 사용할 수 있습니다: `/P(engu|umpk)in/g`

그런 다음 `test()` 메서드를 사용해 원하는 문자열 그룹이 테스트 문자열에 있는지 확인하세요.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

여기서 `test` 메서드는 `true`를 반환합니다.

# --instructions--

정규식을 수정해 `Franklin Roosevelt` 또는 `Eleanor Roosevelt`의 이름을 대소문자 구분하여 검사하고, 중간 이름에 대해서는 허용하도록 하세요.

그런 다음 작성한 정규식을 `myString`에 대해 검사하도록 코드를 수정하고, 정규식이 일치하는지 여부에 따라 `true` 또는 `false`를 반환하게 하세요.

# --hints--

정규식 `myRegex`는 문자열 `true`에 대해 `Franklin D. Roosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

정규식 `myRegex`는 문자열 `true`에 대해 `Eleanor Roosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

정규식 `myRegex`는 문자열 `false`에 대해 `Franklin Rosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

정규식 `myRegex`는 문자열 `false`에 대해 `Frank Roosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

정규식 `myRegex`는 문자열 `false`에 대해 `FranklinRoosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

정규식 `myRegex`는 문자열 `false`에 대해 `EleanorRoosevelt`를 반환해야 합니다.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

정규식을 테스트하려면 `.test()`를 사용하세요.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

결과는 `true`를 반환해야 합니다.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
