---
id: 587d7fb0367417b2b2512bef
title: تقديم ملف HTML
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

يمكنك الرد على الطلبات بملف باستخدام الطريقة `res.sendFile(path)`. يمكنك وضعها داخل معالِج المسار `app.get('/', ...)`. خلف الكواليس، تضبط هذه الطريقة الرؤوس المناسبة لتوجيه المتصفح حول كيفية التعامل مع الملف الذي تريد إرساله، حسب نوعه. ثم تقرأ الملف وترسله. تحتاج هذه الطريقة إلى مسار ملف مطلق. نوصي باستخدام المتغير العام في Node `__dirname` لحساب المسار بهذه الطريقة:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

أرسل الملف `/views/index.html` كاستجابة لطلبات GET على المسار `/`. إذا شاهدت تطبيقك المباشر، يجب أن ترى عنوان HTML كبير (ونموذجًا سنستخدمه لاحقًا...)، بدون تطبيق أي تنسيق.

**ملاحظة:** يمكنك تعديل حل التحدي السابق أو إنشاء حل جديد. إذا أنشأت حلًا جديدًا، تذكر أن Express يقيم المسارات من الأعلى إلى الأسفل، وينفذ المعالِج لأول تطابق. عليك إضافة تعليق للكود للحل السابق، وإلا سيستمر الخادم في الرد بسلسلة نصية.

# --hints--

يجب أن يخدم تطبيقك الملف views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```

