---
id: 587d7fb0367417b2b2512bee
title: بدء خادم Express يعمل
challengeType: 2
forumTopicId: 301519
dashedName: start-a-working-express-server
---

# --description--

في السطرين الأولين من الملف `myApp.js`، ترى مدى سهولة إنشاء كائن تطبيق Express. يحتوي هذا الكائن على عدة طرق، وستتعلم العديد منها في هذه التحديات. إحدى الطرق الأساسية هي `app.listen(port)`. تخبر هذه الطريقة الخادم بالاستماع على منفذ معين، مما يضعه في حالة تشغيل. لأسباب الاختبار، نحتاج أن يكون التطبيق قيد التشغيل في الخلفية، لذلك أضفنا هذه الطريقة في ملف `server.js` لك.

لنقدم أول سلسلة نصية لنا! في Express، تتبع المسارات البنية التالية: `app.METHOD(PATH, HANDLER)`. METHOD هي طريقة http بأحرف صغيرة. PATH هو مسار نسبي على الخادم (يمكن أن يكون سلسلة نصية أو حتى تعبيرًا نمطيًا). HANDLER هي دالة يستدعيها Express عند مطابقة المسار. يأخذ المعالِجون الشكل `function(req, res) {...}`، حيث req هو كائن الطلب، وres هو كائن الاستجابة. على سبيل المثال، المعالِج

```js
function(req, res) {
  res.send('Response String');
}
```

سيقدم السلسلة النصية 'Response String'.

# --instructions--

استخدم الطريقة `app.get()` لتقديم السلسلة النصية "Hello Express" لطلبات GET التي تطابق المسار `/` (الجذر). تأكد من أن شفرتك تعمل من خلال مراجعة السجلات.

**ملاحظة:** يجب إضافة كل الشفرة لهذه الدروس بين الأسطر القليلة من الشفرة التي بدأناها لك.

# --hints--

يجب أن يقدم تطبيقك السلسلة النصية 'Hello Express'

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.equal(
    data,
    'Hello Express',
    'Your app does not serve the text "Hello Express"'
  );
```

