---
id: 587d7fb0367417b2b2512bed
title: تعرّف على وحدة التحكم في Node
challengeType: 2
forumTopicId: 301515
dashedName: meet-the-node-console
---

# --description--

سيتطلب منك العمل على هذه التحديات كتابة شفرتك باستخدام إحدى الطرق التالية:

- استنسخ <a href="https://github.com/freeCodeCamp/boilerplate-express/" target="_blank" rel="noopener noreferrer nofollow">مستودع GitHub هذا</a> وأكمل هذه التحديات محليًا.
- استخدم منشئ مواقع من اختيارك لإكمال المشروع. تأكد من دمج جميع الملفات من مستودع GitHub الخاص بنا.

خلال عملية التطوير، من المهم أن تتمكن من التحقق مما يحدث في شفرتك.

Node هو مجرد بيئة عمل لجافاسكريبت. مثل جافاسكريبت على جانب العميل، يمكنك استخدام وحدة التحكم لعرض معلومات معالجة الأخطاء المفيدة. على جهازك المحلي، سترى مخرجات وحدة التحكم في الطرفية.

ننصحك بالحفاظ على الطرفية مفتوحة أثناء العمل على هذه التحديات. بقراءة المخرجات في الطرفية، يمكنك رؤية أي أخطاء قد تحدث.

يجب إعادة تشغيل الخادم بعد إجراء تغييرات على ملفاته.

يمكنك إيقاف الخادم من الطرفية باستخدام `Ctrl + C` وتشغيله مباشرة باستخدام Node (`node mainEntryFile.js`) أو باستخدام سكربت تشغيل في ملف `package.json` عبر `npm run`.

على سبيل المثال، السكربت `"start": "node server.js"` يُشغّل من الطرفية باستخدام الأمر `npm run start`.

لتنفيذ إعادة تشغيل الخادم تلقائيًا عند حفظ الملف، يوفر Node العلامة `--watch` التي يمكنك إضافتها إلى سكربت البدء `"start": "node --watch server.js"` أو يمكنك تثبيت حزمة npm مثل `nodemon`. سنترك هذا لك كتحدٍّ.

# --instructions--

عدّل ملف `myApp.js` لتسجيل "Hello World" في وحدة التحكم.

# --hints--

`"Hello World"` يجب أن تظهر في وحدة التحكم

```js
  const response = await fetch(code + '/_api/hello-console');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.passed, '"Hello World" is not in the server console');
```

