---
id: 587d7fb2367417b2b2512bf6
title: الحصول على مُدخَل معلمة الاستعلام من العميل
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

طريقة شائعة أخرى للحصول على مُدخَل من العميل هي ترميز البيانات بعد مسار المسار، باستخدام سلسلة استعلام. تُحدد سلسلة الاستعلام بعلامة استفهام (؟)، وتشمل أزواج الحقل=القيمة. يُفصل كل زوج بعلامة وع (&). يمكن لـ Express تحليل البيانات من سلسلة الاستعلام، وملء كائن `req.query`. بعض الأحرف، مثل النسبة المئوية (%)، لا يمكن أن تكون في عناوين URL ويجب ترميزها بصيغة مختلفة قبل الإرسال. إذا استخدمت API من JavaScript، يمكنك استخدام طرق محددة لترميز/فك ترميز هذه الأحرف.

<blockquote>مسار_المسار: '/library'<br>عنوان_الطلب_الفعلي: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

أنشئ نقطة نهاية API، مركبة على `GET /name`. رد بوثيقة JSON، تأخذ البنية `{ name: 'firstname lastname'}`. يجب ترميز معلمات الاسم الأول والأخير في سلسلة استعلام مثل `?first=firstname&last=lastname`.

**ملاحظة:** في التمرين التالي ستتلقى بيانات من طلب POST، على نفس مسار `/name`. إذا أردت، يمكنك استخدام الطريقة `app.route(path).get(handler).post(handler)`. تتيح هذه الصيغة ربط معالجات أفعال مختلفة على نفس مسار المسار. يمكنك توفير بعض الكتابة، والحصول على شفرة أنظف.

# --hints--

اختبار 1 : يجب أن ترد نقطة نهاية API الخاصة بك بـ `{ "name": "Mick Jagger" }` عندما يُستدعى مسار `/name` مع `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

اختبار 2 : يجب أن ترد نقطة نهاية API الخاصة بك بـ `{ "name": "Keith Richards" }` عندما يُستدعى مسار `/name` مع `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```

