---
id: 587d7fb0367417b2b2512bf0
title: تقديم الأصول الثابتة
challengeType: 2
forumTopicId: 301518
dashedName: serve-static-assets
---

# --description--

عادةً ما يحتوي خادم HTML على مجلد واحد أو أكثر يمكن للمستخدم الوصول إليها. يمكنك وضع الأصول الثابتة التي يحتاجها التطبيق هناك (أنماط CSS، سكريبتات، صور).

في Express، يمكنك تفعيل هذه الوظيفة باستخدام برمجيات وسيطة `express.static(path)`، حيث تكون معلمة `path` هي المسار المطلق للمجلد الذي يحتوي على الأصول.

إذا لم تكن تعرف ما هي البرمجيات الوسيطة... لا تقلق، سنشرحها بالتفصيل لاحقًا. باختصار، البرمجيات الوسيطة هي دوال تعترض معالجات المسارات، مضيفة نوعًا من المعلومات. يجب تركيب البرمجيات الوسيطة باستخدام الطريقة `app.use(path, middlewareFunction)`. المعلمة الأولى `path` اختيارية. إذا لم تمررها، ستُنفذ البرمجيات الوسيطة لجميع الطلبات.

# --instructions--

ركب البرمجيات الوسيطة `express.static()` على المسار `/public` باستخدام `app.use()`. المسار المطلق لمجلد الأصول هو `__dirname + /public`.

الآن يجب أن يكون تطبيقك قادرًا على تقديم ملف أنماط CSS. لاحظ أن ملف `/public/style.css` مذكور في `/views/index.html` ضمن قالب المشروع. يجب أن تبدو الصفحة الرئيسية أفضل قليلًا الآن!

# --hints--

يجب أن يخدم تطبيقك ملفات الأصول من مجلد `/public` إلى مسار `/public`

```js
  const response = await fetch(code + '/public/style.css');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /body\s*\{[^\}]*\}/,
    'Your app does not serve static assets'
  );
```

يجب ألا يخدم تطبيقك ملفات من أي مجلدات أخرى سوى مجلد `/public`

```js
  const response = await fetch(code + '/server.js');
  if (response.ok) {
    const data = await response.text();
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  } else {
    assert.equal(
      response.status + '',
      404 + '',
      'Your app must serve files only from "public" directory'
    );
  }
```

