---
id: 587d7fb2367417b2b2512bf5
title: الحصول على مُدخَل معلمة المسار من العميل
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

عند بناء واجهة برمجة التطبيقات (API)، يجب السماح للمستخدمين بإبلاغنا بما يريدون الحصول عليه من خدمتنا. على سبيل المثال، إذا كان العميل يطلب معلومات عن مستخدم مخزن في قاعدة البيانات، يحتاج إلى طريقة ليخبرنا أي مستخدم يهتم به. إحدى الطرق الممكنة لتحقيق ذلك هي استخدام معلمات المسار. معلمات المسار هي أجزاء مسماة من عنوان URL، تفصلها شرطات مائلة (/). كل جزء يلتقط قيمة الجزء من عنوان URL الذي يتطابق مع موقعه. يمكن العثور على القيم الملتقطة في كائن `req.params`.

<blockquote>مسار_المسار: '/user/:userId/book/:bookId'<br>عنوان_الطلب_الفعلي: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

أنشئ خادم صدى، مركب على المسار `GET /:word/echo`. رد بكائن JSON، يأخذ البنية `{echo: word}`. يمكنك العثور على الكلمة التي يجب تكرارها في `req.params.word`. يمكنك اختبار مسارك من شريط عنوان المتصفح، بزيارة بعض المسارات المطابقة، مثل `your-app-rootpath/freecodecamp/echo`.

# --hints--

اختبار 1 : يجب أن يكرر خادم الصدى الكلمات بشكل صحيح

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

اختبار 2 : يجب أن يكرر خادم الصدى الكلمات بشكل صحيح

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```

