---
id: 587d7fb1367417b2b2512bf4
title: ربط برمجيات وسيطة لإنشاء خادم وقت
challengeType: 2
forumTopicId: 301510
dashedName: chain-middleware-to-create-a-time-server
---

# --description--

يمكن تركيب البرمجيات الوسيطة على مسار محدد باستخدام `app.METHOD(path, middlewareFunction)`. كما يمكن ربط البرمجيات الوسيطة داخل تعريف المسار.

انظر إلى المثال التالي:

```js
app.get('/user', function(req, res, next) {
  req.user = getTheUserSync();  // Hypothetical synchronous operation
  next();
}, function(req, res) {
  res.send(req.user);
});
```

تُفيد هذه الطريقة في تقسيم عمليات الخادم إلى وحدات أصغر. يؤدي ذلك إلى هيكلية أفضل للتطبيق، وإمكانية إعادة استخدام الشفرة في أماكن مختلفة. يمكن أيضًا استخدام هذه الطريقة لأداء بعض التحقق على البيانات. في كل نقطة من مكدس البرمجيات الوسيطة يمكنك حجب تنفيذ السلسلة الحالية وتمرير التحكم إلى دوال مصممة خصيصًا لمعالجة الأخطاء. أو يمكنك تمرير التحكم إلى المسار المطابق التالي، لمعالجة الحالات الخاصة. سنرى كيف في قسم Express المتقدم.

# --instructions--

في المسار `app.get('/now', ...)` اربط دالة برمجيات وسيطة والمعالِج النهائي. في دالة البرمجيات الوسيطة يجب إضافة الوقت الحالي إلى كائن الطلب في المفتاح `req.time`. يمكنك استخدام `new Date().toString()`. في المعالِج، استجب بكائن JSON يأخذ البنية `{time: req.time}`.

**ملاحظة:** لن ينجح الاختبار إذا لم تربط البرمجيات الوسيطة. إذا ركبت الدالة في مكان آخر، سيفشل الاختبار حتى لو كانت النتيجة صحيحة.

# --hints--

يجب أن يحتوي نقطة النهاية /now على برمجيات وسيطة مركبة

```js
  const response = await fetch(code + '/_api/chain-middleware-time');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.stackLength,
    2,
    '"/now" route has no mounted middleware'
  );
```

يجب أن تُرجع نقطة النهاية `/now` الوقت الحالي.

```js
  const response = await fetch(code + '/_api/chain-middleware-time');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  var now = new Date();
  assert.isAtMost(
    Math.abs(new Date(data.time) - now),
    20000,
    'the returned time is not between +- 20 secs from now'
  );
```

