---
id: 587d7fb2367417b2b2512bf8
title: الحصول على بيانات من طلبات POST
challengeType: 2
forumTopicId: 301511
dashedName: get-data-from-post-requests
---

# --description--

ثبت معالِج POST على المسار `/name`. إنه نفس المسار السابق. أعددنا نموذجًا في الصفحة الرئيسية html. سيرسل نفس بيانات التمرين 10 (سلسلة الاستعلام). إذا تم تكوين محلل محتوى الطلبات بشكل صحيح، يجب أن تجد المعلمات في الكائن `req.body`. ألق نظرة على المثال المعتاد لمكتبة:

<blockquote>مسار: POST '/library'<br>urlencoded_body: userId=546&#x26;bookId=6754 <br>req.body: {userId: '546', bookId: '6754'}</blockquote>

استجب بنفس كائن JSON السابق: `{name: 'firstname lastname'}`. اختبر ما إذا كانت نقطة النهاية تعمل باستخدام نموذج html الذي وفرناه في الصفحة الرئيسية للتطبيق.

نصيحة: هناك عدة طرق http أخرى غير GET و POST. ووفقًا للاصطلاح، هناك تطابق بين فعل http والعملية التي ستنفذها على الخادم. التوافق الاصطلاحي هو:

POST (وأحيانًا PUT) - إنشاء مورد جديد باستخدام المعلومات المرسلة مع الطلب،

GET - قراءة مورد موجود دون تعديله،

PUT أو PATCH (وأحيانًا POST) - تحديث مورد باستخدام البيانات المرسلة،

DELETE - حذف مورد.

هناك أيضًا بعض الطرق الأخرى التي تُستخدم للتفاوض على اتصال مع الخادم. باستثناء GET، يمكن لجميع الطرق الأخرى المذكورة أعلاه أن تحتوي على حمولة (أي البيانات في جسم الطلب). تعمل برمجيات وسيطة محلل محتوى الطلبات مع هذه الطرق أيضًا.

# --hints--

اختبار 1 : يجب أن تستجيب نقطة النهاية في API بالاسم الصحيح

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({ first: 'Mick', last: 'Jagger' }).toString()
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "POST /name" route does not behave as expected'
  );
```

اختبار 2 : يجب أن تستجيب نقطة النهاية في API بالاسم الصحيح

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({
      first: 'Keith',
      last: 'Richards'
    }).toString()
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "POST /name" route does not behave as expected'
  );
```

