---
id: 68e97fe79367ad7b5dd6c9cd
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, construirás una página de librería creando tarjetas de libros que muestran información sobre diferentes libros. Practicarás organizar contenido usando elementos `div`, clases e IDs.

Comienza agregando un elemento `h1` con el texto `XYZ Bookstore`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

El texto de tu elemento `h1` debe ser `XYZ Bookstore`. Verifica dos veces errores de ortografía o mayúsculas.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
