---
id: 68e97fe79367ad7b5dd6c9cd
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

En este paso, agrega un elemento `h1` con el texto `XYZ Bookstore`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

El texto de tu elemento `h1` debe ser `XYZ Bookstore`. Verifica dos veces errores de ortografía o mayúsculas.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Tu elemento `h1` debe estar dentro del elemento `body`.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
</html>
```
