---
id: 68ebdbacdd3fa474132cc975
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del elemento `h1`, agrega un elemento `p` con este texto: `Browse our collection of amazing books!`.

# --hints--

Debes tener un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

El texto de tu elemento `p` debe ser: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Tu elemento `p` debe estar debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>

--fcc-editable-region--
```
