---
id: 68ebdbacdd3fa474132cc975
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Debajo del elemento `h1`, agrega un elemento `p` con este texto: `Browse our collection of amazing books!`.

# --hints--

Debes tener un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

El texto de tu elemento `p` debe ser: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Tu elemento `p` debe estar debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  --fcc-editable-region--

  --fcc-editable-region--
</body>

</html>
```
