---
id: 68eca6364c5616d393389a30
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Debajo del elemento con la clase `card-container`, agrega un nuevo elemento `p` con este texto:

```md
Review your selections and continue to checkout.
```

Debajo del elemento `p`, crea un elemento `div` con el atributo `class` establecido en `btn-container`. Este contenedor agrupará tus elementos de botón de navegación.

# --hints--

El nuevo elemento `p` debe estar debajo del elemento con una clase `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

El texto de tu nuevo elemento `p` debe ser `Review your selections and continue to checkout.`.

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

Debes tener un elemento `div` con una `class` de `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Tu elemento `div` con una `class` de `btn-container` debe ir después del elemento `p` recién agregado.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>

--fcc-editable-region--
</body>

</html>
```
