---
id: 68eca5412b5dd9d06b3d6404
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Debajo del elemento `h2` en la segunda tarjeta, agrega un elemento `p` con este texto:

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

Debes tener un elemento `p` anidado dentro del segundo elemento que tiene una clase `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

El texto del elemento `p` de tu segunda tarjeta debe ser `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

El elemento `p` de tu segunda tarjeta debe estar debajo de tu elemento `h2`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>

  </div>
  --fcc-editable-region--
</div>
```
