---
id: 68eca2a795b333ca5fee30a8
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Ahora crea una segunda tarjeta de libro. Añade otro elemento `div` con el atributo `class` establecido en `card`. Observa cómo puedes reutilizar el mismo nombre de clase para múltiples elementos y aplicar un estilo consistente.

# --hints--

Debes tener dos elementos `div` con un atributo `class` que tenga un valor de `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

El segundo elemento con una clase `card` debe estar dentro del elemento que tiene una clase `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

Los dos elementos con una clase de `card` deben ser hermanos.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>

  --fcc-editable-region--
  </div>
</body>

</html>
```
