---
id: 68eca6e5a6759ed4ea0034dc
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Dentro del elemento con una clase `btn-container`, agrega dos elementos `button`:

Primer botón:

- Id: `view-cart-btn`
- Clase: `btn`
- Texto: `View Cart`

Segundo botón:

- Id: `checkout-btn`
- Clase: `btn`
- Texto: `Checkout`


¡Felicidades! Has construido con éxito la estructura de una página de librería usando divs, clases e ids para organizar tu contenido.

# --hints--

Debes tener dos elementos `button` anidados dentro del elemento `.btn-container`.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Tu primer botón dentro del elemento con la clase `btn-container` debe tener un atributo `id` con el valor `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Tu primer botón dentro del elemento que tiene una clase `btn-container` debe tener un atributo `class` con el valor `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

El texto de tu primer botón debe ser `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Tu segundo botón dentro del elemento con una clase de `btn-container` debe tener un atributo `id` con el valor de `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Tu segundo botón dentro del elemento con `class="btn-container"` debe tener un atributo `class` con el valor de `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

El texto de tu segundo botón debe ser `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
--fcc-editable-region--
  <p>Review your selections and continue to checkout.</p>
  <div class="btn-container">

  </div>
--fcc-editable-region--
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
