---
id: 68ec6cd5b7e8f5a8f7319e32
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

El elemento `div` se usa como contenedor para agrupar otros elementos HTML. Principalmente usarás el elemento `div` cuando quieras agrupar elementos HTML que compartirán un conjunto de estilos CSS.

Debajo del elemento `p`, agrega un elemento `div`. Este `div` será un contenedor para tus tarjetas de libros.

# --hints--

Debe tener un elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Tu elemento `div` debe estar debajo de tu elemento `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
--fcc-editable-region--

--fcc-editable-region--
</body>

</html>
```
