---
id: 68ec6d9a315221aa31e54816
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El atributo `class` se usa para identificar uno o más elementos para el estilo. A diferencia del atributo `id`, los nombres de clase no necesitan ser únicos: varios elementos pueden compartir la misma clase.

Aquí hay un ejemplo:

```html
<p class="example">example paragraph</p>
```

Agrega un atributo `class` a tu elemento `div` y establece su valor en `card-container`.

# --hints--

Tu elemento `div` debe tener un atributo `class`.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Tu elemento `div` debe tener un atributo `class` con el valor `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div>
    
</div>
--fcc-editable-region--
```
