---
id: 68eca4c4b0f952cf09fabe09
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Dentro del segundo elemento que tiene una clase `card`, agrega un elemento `h2` con el texto `Dave's Cooking Adventure`.

# --hints--

Debes tener un elemento `h2` anidado dentro del segundo elemento que tiene una clase `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

El elemento `h2` dentro del segundo elemento que tiene una clase `card` debe tener un texto de `Dave's Cooking Adventure`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
