---
id: 68ec6cd5b7e8f5a8f7319e32
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

El elemento `div` se usa como contenedor para agrupar otros elementos HTML. Principalmente usarás el elemento `div` cuando quieras agrupar elementos HTML que compartirán un conjunto de estilos CSS.

Debajo del elemento `p`, agrega un elemento `div`. Este `div` será un contenedor para tus tarjetas de libros.

# --hints--

Debe tener un elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Tu elemento `div` debe estar debajo de tu elemento `p`.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>

--fcc-editable-region--
```
