---
id: 68ec99e478211dc578699944
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Debajo del elemento `h2` en el primer elemento que tenga una clase `card`, agrega un elemento `p` con el siguiente texto:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

Debes tener un elemento `p` anidado dentro del elemento que tiene una clase `card`.

```js
assert.exists(document.querySelector('.card p'));
```

El texto de tu elemento `p` debe ser `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`.

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Su elemento `p` debe estar debajo de su elemento `h2`.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
