---
id: 68eca11e99e3c5c894ca9d69
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

El elemento `button` se usa para crear botones clicables en una página web. Los botones son elementos interactivos que los usuarios pueden clicar para realizar acciones.

Agrega un elemento `button` dentro del elemento que tiene un `class` de `card`, dale al botón un atributo `class` establecido en `btn` y el texto `Buy Now`.

# --hints--

Debes tener un elemento `button` anidado dentro del elemento que tiene una clase `card`.

```js
assert.exists(document.querySelector('.card button'));
```

Tu elemento `button` debe tener un atributo `class`.

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

Tu elemento `button` debe tener un `class` que tenga un valor de `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

El texto de tu elemento `button` debe ser `Buy Now`. 

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
