---
id: 68ec9332a9b5b2b32487bd00
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

El atributo `id` añade un identificador único a un elemento HTML. Cada `id` debe ser único dentro de una página y solo debe usarse una vez.

Los valores de `id` no pueden contener espacios y solo deben contener letras, dígitos, guiones bajos y guiones.

Aquí hay un ejemplo:

```html
<p id="para">example paragraph</p>
```

Agrega un atributo `id` a tu elemento que tenga una clase `card` y establece su valor en `sally-adventure-book`.

# --hints--

Tu elemento con una clase de `card` debe tener un atributo `id`.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Tu elemento que tenga una clase `card` debe tener un `id` con el valor `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
