---
id: 68eca5c0065c82d256d29ca3
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Dentro de la segunda tarjeta, agrega un elemento `button` con el atributo `class` establecido en `btn` y el texto `Buy Now`.

Ambos elementos `button` ahora comparten la misma `class`, lo que significa que pueden ser estilizados de manera consistente juntos.

# --hints--

Debes tener un elemento `button` anidado dentro del segundo elemento que tiene una clase `card`.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

El elemento `button` de tu segunda tarjeta debe tener un atributo `class`.

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

El elemento `button` de tu segunda tarjeta debe tener una `class` con el valor `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

El texto del `button` del segundo `card` debe ser `Buy Now`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  --fcc-editable-region--
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
  </div>
  --fcc-editable-region--
</div>
```
