---
id: 69134b6cb43aa5f254950a10
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, construirás una página de librería creando tarjetas de libros que muestran información sobre diferentes libros. Practicarás organizar contenido usando elementos `div`, clases e IDs.

Comienza tu página de librería creando la estructura básica de HTML.

Agrega la declaración `<!DOCTYPE html>` y los elementos `html` y `head`.

Agrega un atributo `lang` al elemento `html` y configúralo en `"en"`.

# --hints--

Tu código debe incluir una declaración `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Tu código debe incluir tanto las etiquetas de apertura como de cierre `<html>`.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Tu elemento `html` debe tener un atributo `lang` establecido en `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Tu código debe incluir tanto las etiquetas de apertura como de cierre `<head>`.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Tu elemento `head` debe estar anidado dentro del elemento `html`.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
