---
id: 69134d14150d51f8066a3246
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Ahora, mejora la estructura de tu documento HTML para asegurar que tu página esté codificada correctamente y se vea bien en todos los tamaños de pantalla.

Agrega un atributo de idioma al elemento `html` y configúralo en `en`. Luego, dentro del elemento `head` agrega:

- una etiqueta `<meta charset="UTF-8">`
- una etiqueta `<meta name="viewport" content="width=device-width, initial-scale=1.0">`

Ahora agrega el elemento `body` debajo de la sección head. Aquí es donde irá todo el contenido visible de tu página.

# --hints--

Tu elemento `html` debe tener un atributo `lang` establecido en `"en"`.

```js
assert.match(code, /<html[^>]*\blang=["']en["'][^>]*>/i);
```

Tu código debe incluir un elemento `<meta charset="UTF-8">`.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Tu código debe incluir un elemento `meta` de viewport.

```js
assert.exists(document.querySelector('head > meta[name="viewport"][content="width=device-width, initial-scale=1.0"]'));
```

Tu código debe incluir las etiquetas de apertura y cierre `<body>`.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

El elemento `body` debe estar debajo del elemento `head`.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
<head>
  <title>XYZ Bookstore Page</title>
</head>

</html>
--fcc-editable-region--
```
