---
id: 691f85085787a8c4ca589fbb
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Agrega el elemento `title` dentro del elemento `head`. 

Establece el título de la página en `XYZ Bookstore Page`.

# --hints--

Tu código debe incluir tanto las etiquetas de apertura como de cierre `<title>`. 

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Tu título debe ser `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Tu elemento `title` debe estar anidado dentro del elemento `head`.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
  --fcc-editable-region--
<!DOCTYPE html>
<html>
<head>

</head>
  --fcc-editable-region-- 
</html>
```
