---
id: 6690e10ebe2181212abc9652
title: تصميم بطاقة عمل
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `div` له خاصية `class` بقيمة `business-card` يحتوي على جميع العناصر الأخرى.
1. داخل عنصر `.business-card`، ينبغي أن يكون هناك عنصر `img` له خاصية `class` بقيمة `profile-image`. يمكنك ضبط مصدر الصورة على `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` إذا أردت. وينبغي أن يكون هناك `alt` بوصف ذي معنى.
1. داخل عنصر `.business-card`، بعد عنصر `img`، يجب أن يكون لديك ثلاثة عناصر `p` لها خاصية `class` بقيمة `full-name` و`designation` و`company` على الترتيب.
1. يجب أن يحتوي عنصر `p` الأول على اسمك.
1. يجب أن يحتوي عنصر `p` الثاني على مسماك الوظيفي.
1. يجب أن يحتوي عنصر `p` الثالث على اسم شركتك.
1. ينبغي أن يكون هناك عنصر `hr` أسفل عنصر `p` الثالث.
1. بعد عنصر `hr`، ينبغي أن يكون هناك عنصران `p`. في عنصر `p` الأول، يجب أن يكون لديك عنوان بريد إلكتروني كنص. وفي عنصر `p` الثاني، ينبغي أن يكون النص رقم هاتف.
1. بعد عنصري `p`، ينبغي أن يكون هناك عنصر `a` له صنف بقيمة `portfolio-link` ونصّه `Portfolio`، وينبغي أن يرتبط بعنوان URL صالح.
1. ينبغي أن يكون هناك عنصر `hr` بعد عنصر `a` الذي يحتوي على رابط معرض الأعمال.
1. يجب أن يكون لديك عنصر `div` آخر له خاصية `class` بقيمة `social-media`. وداخل هذا العنصر، ينبغي أن يكون هناك عنصر `h2` نصّه `Connect with me`.
1. داخل عناصر `.social-media`، ينبغي أن يكون هناك ثلاثة عناصر `a`. وينبغي أن تحمل عناصر `a` النصوص `Twitter` و`LinkedIn` و`GitHub` على الترتيب مع روابط إلى مواقع صالحة.
1. يجب أن تربط ورقة الأنماط `styles.css` باستخدام وسم `link` في قسم `head` من ملف HTML.
1. يجب أن تكون خلفية صفحتك `rosybrown`. وينبغي أن يكون الخط العام `Arial`، مع خط بديل `sans-serif`.
1. يجب أن يحتوي محدِّد `.business-card` على خصائص تجعل عرض العنصر `300px` ولون خلفية من اختيارك. كما ينبغي أن تضبط `padding` بمقدار `20px` من جميع الجهات و`margin` علويًا بمقدار `100px`. وينبغي أن يكون النص محاذيًا للوسط وحجم الخط `16px`. وينبغي ضبط الهامشين الأيمن والأيسر على `auto`.
1. يجب أن يحتوي محدِّد `.profile-image` على خاصية `max-width` بقيمة `100%`.
1. يجب أن يكون لجميع عناصر `p` هامش `margin` علوي وسفلي بمقدار `5px`.
1. يجب ألا تكون جميع الروابط في الصفحة مُسطَّرة بخط سفلي.

# --hints--

يجب أن يكون لديك عنصر `div` رئيسي واحد مع خاصية `class` بقيمة `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

داخل عنصر `.business-card`، يجب أن يكون هناك عنصر `img` مع خاصية `class` بقيمة `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

يجب تعيين مصدر الصورة إلى صورة صالحة.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

يجب تعيين خاصية `alt` للصورة بوصف ذي معنى.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

داخل عنصر `.business-card`، بعد عنصر `img`، يجب أن يكون هناك عنصر `p` مع خاصية `class` بقيمة `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

بعد عنصر `.full-name`، يجب أن يكون هناك عنصر `p` مع خاصية `class` بقيمة `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

بعد عنصر `.designation`، يجب أن يكون هناك عنصر `p` مع خاصية `class` بقيمة `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

يجب أن يحتوي أول عنصر `p`، الذي يحمل صنف `full-name`، على اسمك.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

يجب أن يحتوي ثاني عنصر `p`، الذي يحمل خاصية `class` بقيمة `designation`، على مسماك الوظيفي.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

يجب أن يحتوي ثالث عنصر `p`، الذي يحمل خاصية `class` بقيمة `company`، على اسم شركتك.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

بعد ثالث عنصر `p`، الذي يحمل خاصية `class` بقيمة `company`، يجب أن يكون هناك عنصر `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

بعد أول عنصر `hr`، يجب أن يكون هناك عنصر `p` يحتوي على عنوان بريد إلكتروني كنص.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

بعد عنصر `p` الخاص بالبريد الإلكتروني، يجب أن يكون هناك عنصر `p` آخر يحتوي على رقم هاتف كنص.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

بعد عنصر `p` الخاص برقم الهاتف، يجب أن يكون هناك عنصر `a` مع صنف `portfolio-link`.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

يجب أن يحتوي عنصر `.portfolio-link` على نص `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

يجب أن يحتوي عنصر `.portfolio-link` على خاصية `href` صالحة.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

بعد عنصر `a` الذي يحتوي على معرض الأعمال، يجب أن يكون هناك عنصر `hr` آخر كفاصل.

```js
assert.exists(document.querySelector('a + hr'));
```

بعد ثاني عنصر `hr`، يجب أن يكون هناك عنصر `div` مع خاصية `class` بقيمة `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

داخل عنصر `.social-media`، يجب أن يكون هناك عنصر `h2` بالنص `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

داخل عنصر `.social-media` يجب أن يكون هناك ثلاثة عناصر `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

يجب أن يحتوي أول عنصر `a` على النص `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

يجب أن تشير خاصية `href` لأول عنصر `a` إلى رابط صالح.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

يجب أن يحتوي ثاني عنصر `a` على النص `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

يجب أن تشير خاصية `href` لثاني عنصر `a` إلى رابط صالح.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

يجب أن يحتوي ثالث عنصر `a` على النص `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

يجب أن تشير خاصية `href` لثالث عنصر `a` إلى رابط صالح.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

يجب أن يكون عنصر `link` داخل عنصر `head` الخاص بك.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

يجب تعيين لون خلفية الصفحة في محدد `body` إلى `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

يجب تعيين خط الصفحة إلى `Arial` في عنصر `body` مع نظام بديل `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

يجب أن يحتوي محدد `.business-card` على خاصية `width` بقيمة `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

يجب أن يحتوي محدد `.business-card` على خاصية `background-color` بقيمة لون صالحة.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

يجب أن يحتوي محدد `.business-card` على خاصية `padding` بقيمة `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

يجب أن يحتوي المحدِّد `.business-card` على هامش علوي بقيمة `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

يجب أن يكون النص داخل عنصر `.business-card` محاذيًا للوسط.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

يجب أن يكون حجم الخط للنص داخل عنصر `.business-card` بقيمة `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

يجب تعيين الهوامش اليمنى واليسرى لعنصر `.business-card` إلى `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

يجب أن يحتوي محدد `.profile-image` على خاصية `max-width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

يجب أن يكون لعناصر `p` لديك هامش `margin` علوي وسفلي بمقدار `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

يجب ألا تحتوي روابط الصفحة على خط تحت النص.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
