---
id: 683ec95f2f8781355556ff83
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Ви можете побачити, що перший виклик повертає `[3]`. Це тому, що `3` є середнім значенням у `[1, 2, 3, 4, 5]`.

З іншого боку, другий виклик повертає порожній список. Це відбувається тому, що після перевірки середнього значення цикл наразі не має логіки для звуження діапазону пошуку шляхом оновлення змінних `low` або `high`. Він лише виконує одноразову перевірку середини і завершується без подальшого пошуку.

Щоб дозволити бінарному пошуку продовжувати звужувати свій пошук, додайте блок `elif`, який перевіряє, чи `value` більше за `value_at_middle`. Тимчасово додайте ключове слово `pass` всередину `elif`.

# --hints--

Ви повинні додати блок `elif` для випадку, коли `value` більше за `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
