---
id: 683ec95f2f8781355556ff7b
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У бінарному пошуку існує найнижчий можливий індекс і найвищий можливий індекс.

Найнижчий можливий індекс представляє ліву межу в поточній частині значень, у якій відбувається пошук, у цьому випадку `search_list`. Найвищий можливий індекс, навпаки, представляє праву межу в значеннях, у яких відбувається пошук.

Щоб врахувати ці дві межі, потрібно розглядати весь список, у якому відбувається пошук. Тож всередині функції `binary_search` визначте змінну `low` зі значенням `0` і змінну `high` зі значенням, яке відповідає останньому індексу в списку, у якому відбувається пошук.

# --hints--

У вашій функції `low` має бути змінна з ім’ям `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Значення вашої змінної `low` має бути встановлене в `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

У вашій функції `high` має бути змінна з ім’ям `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Значення вашої змінної `high` має бути встановлене в `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
