---
id: 683ec95f2f8781355556ff80
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Наступним кроком є створення умови, яка перевірятиме, чи знаходиться ціль `value` посередині.

Всередині циклу `while` створіть інструкцію `if`, яка перевіряє, чи дорівнює ціль `value` `value_at_middle`. Якщо так, поверніть змінну `path_to_target`.

# --hints--

Ви повинні створити інструкцію `if`, яка перевіряє, чи дорівнює `value` `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Ви повинні повернути змінну `path_to_target` всередині інструкції `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
