---
id: 683ec95f2f8781355556ff7a
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Алгоритм бінарного пошуку, який ви створюєте, повертатиме не лише `True` або `False`, щоб показати, чи було знайдено значення. Він також повертатиме шлях, яким він знаходить це значення. Для цього у вас буде змінна `path_to_target`, щоб відстежувати цей шлях.

Усередині вашої функції `binary_search` замініть ключове слово `pass` на змінну `path_to_target` і ініціалізуйте її як порожній список на початку.

# --hints--

Ви повинні створити змінну `path_to_target`, яка є множиною порожнього списку, всередині вашої функції `binary_search`. Не забудьте видалити ключове слово `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
