---
id: 683ec95f2f8781355556ff81
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Вам потрібно протестувати деякі речі, щоб зрозуміти потік алгоритму на цьому ініціальному етапі.

Для цього спочатку потрібно вийти з циклу. Це тому, що поточна реалізація дозволить лише одну ітерацію, тож якщо умова не виконується, виникне нескінченний цикл.

Одразу після `if`інструкції використайте ключове слово `break`, щоб вийти з циклу `while`. Потім, після циклу `while`, поверніть порожній список, щоб позначити, що значення не було знайдено.

# --hints--

Ви повинні використати ключове слово `break`, щоб вийти з циклу `while` після інструкції `if`.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

Після вашого циклу `while` слід повернути порожній список.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
