---
id: 683ec95f2f8781355556ff7c
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер, коли ви визначили межі свого пошуку, настав час створити цикл, який виконуватиме алгоритм бінарного пошуку.

Бінарний пошук працює шляхом багаторазового звуження області пошуку. Цей процес триває доти, доки існує дійсний діапазон елементів для перевірки.

Ви можете виразити це за допомогою `while` циклу, який триватиме, поки ваш `low` вказівник менший або дорівнює `high` вказівнику.

Якщо `low` коли-небудь стане більшим за `high`, це означає, що область пошуку стала порожньою, і `value` немає у списку.

Всередині `binary_search` функції створіть `while` цикл з умовою, яка перевіряє, чи `low` менший або дорівнює high. Усередині циклу додайте ключове слово `pass` як заповнювач поки що.

# --hints--

Ви повинні створити `while` цикл, який працюватиме, поки `low` менший або дорівнює `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
