---
id: 683ec39565dbc7e3375263b1
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У попередньому уроці ви дізналися про бінарний пошук і чим він відрізняється від лінійного пошуку.

У цьому практичному занятті ви застосуєте це на практиці, створивши алгоритм бінарного пошуку.

Для початку створіть `binary_search` функцію з параметрами `search_list` та `value`.

`search_list` — це список елементів, у якому функція буде шукати, а `value` — це ціль, яку функція намагатиметься знайти в `search_list`.

Поки що додайте ключове слово `pass` всередину функції, щоб тест пройшов.

# --hints--

Ви повинні створити `binary_search` функцію.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Ваша `binary_search` функція має мати параметри `search_list` та `value`.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
