---
id: 683ec95f2f8781355556ff7d
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Щоб почати цикл, потрібно знайти середній індекс поточного простору пошуку.

Ви можете знайти середній індекс, обчисливши середнє значення змінних `low` і `high`. Потрібно буде використати цілочисельне ділення після додавання двох значень, щоб отримати середнє, яке завжди буде цілим числом, округленим вниз, оскільки індекси мають бути цілими числами.

Видаліть ключове слово `pass` з вашого циклу `while` і створіть змінну `mid`, встановлену на середнє значення `low` і `high`.

# --hints--

Ви повинні видалити ключове слово `pass` з циклу while і створити змінну `mid`, встановлену на `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
