---
id: 683ec95f2f8781355556ff82
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Пам’ятайте, що функція приймає список під назвою `search_list` і `value`, який ви шукаєте.

Тепер викличте функцію з `binary_search([1, 2, 3, 4, 5], 3)` і `binary_search([1, 2, 3, 4, 5, 9], 4)` та одразу виведіть виклики.

# --hints--

Ви повинні викликати функцію з `binary_search([1, 2, 3, 4, 5], 3)` і одразу її вивести.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

Ви повинні викликати функцію з `binary_search([1, 2, 3, 4, 5, 9], 4)` і одразу її вивести.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
