---
id: 683ec95f2f8781355556ff7e
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви обчислили індекс `mid`, потрібно фактично отримати значення з `search_list` за цим індексом. Це значення ви порівнюватимете з вашим параметром `value`, який є ціллю, яку ви шукаєте.

Створіть змінну `value_at_middle` і призначте їй елемент з `search_list`, що знаходиться за індексом `mid`.

# --hints--

Ви маєте створити змінну `value_at_middle`, призначивши їй `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
