---
id: 684169b00a073c9b9f0a703c
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Ви помітили, що другий виклик функції тепер має значення `3, 5, 4` у списку? Це означає, що бінарний пошук працює так, як задумано.

Ось як це сталося: алгоритм спочатку перевірив `3` як середину початкового списку. Оскільки `4` більше за `3`, пошук змістився до правої половини. Потім він розглянув `5` як нову середину. Оскільки `4` менше за `5`, пошук перемістився вліво, врешті-решт визначивши `4` як середину кінцевого діапазону.

Щоб знову протестувати функцію, викличте її з `[1, 3, 5, 9, 14, 22], 10` і одразу виведіть виклик. Це ситуація, коли значення не буде знайдено.

# --hints--

Ви повинні викликати функцію з `binary_search([1, 3, 5, 9, 14, 22], 10)` і одразу вивести результат.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
