---
id: 66a37f37ef5823a313de8c26
title: "Creare un'orbita lunare"
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

In questo laboratorio, creerai una semplice animazione dell'orbita della Luna intorno alla Terra usando HTML e CSS. La Terra sarà al centro del sistema e la Luna orbiterà intorno ad essa.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti allineare tutti gli elementi al centro della pagina impostando l'altezza a `100%` del viewport e impostando un layout flexbox nell'elemento `body`.
   
2. Dovresti avere un `div` con la classe `space`.

3. All'interno dell'elemento `.space` dovrebbero esserci altri due elementi `div` con le classi `earth` e `orbit`, rispettivamente, in ordine.

4. All'interno dell'elemento `.orbit` dovrebbe esserci un altro `div` con la classe `moon`.
   
5. L'elemento `div` con una `class` di `space` dovrebbe essere centrato nella pagina e avere una larghezza e un'altezza di `200px`.

6. L'elemento `div` con una `class` di `space` dovrebbe usare il posizionamento `relative`.

7. L'elemento `.earth` dovrebbe usare il posizionamento `absolute`. Posiziona il centro dell'elemento `.earth` a metà del suo genitore sia sull'asse verticale (top) sia su quello orizzontale (left). Dopodiché, sposta l'elemento `.earth` indietro di metà della sua larghezza e altezza per centrarlo all'interno del suo genitore, l'elemento `.space`.

8. L'elemento `.earth` dovrebbe avere una larghezza e un'altezza di `100px`.

9. L'elemento `.orbit` dovrebbe avere una larghezza e un'altezza di `200px`.

10. L'elemento `.orbit` dovrebbe essere posizionato usando il posizionamento `absolute`. Il suo angolo in basso a destra dovrebbe trovarsi al centro dell'elemento `.space` usando una proprietà `transform` che lo sposta di `-50%` sia sull'asse verticale sia su quello orizzontale.

11. Il percorso dell'orbita della luna intorno alla Terra dovrebbe essere un cerchio.

12. L'elemento `.moon` dovrebbe essere posizionato usando il posizionamento `absolute` e avere una larghezza e un'altezza di `30px`. L'elemento `.moon` dovrebbe posizionarsi in cima all'elemento `.orbit` ed essere centrato orizzontalmente.

13. Dovresti regolare ulteriormente il posizionamento orizzontale dell'elemento `.moon` spostandolo a sinistra di metà della sua larghezza.

14. Il selettore `.earth` dovrebbe avere un colore di sfondo e una `border-radius` di `50%`.

15. Il selettore `.moon` dovrebbe avere un colore di sfondo e una `border-radius` di `50%`.

16. Dovresti definire un'animazione `@keyframes orbit` che ruota l'elemento `.orbit` di 360 gradi intorno al suo centro. Dovresti applicare questa animazione all'elemento `.orbit` con una durata di `5` secondi, una funzione di temporizzazione lineare e iterazioni infinite.

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e di applicare il CSS.

# --hints--

Il contenuto dell'elemento `body` dovrebbe essere centrato nella pagina impostando l'altezza a `100%` dell'altezza del viewport, `display` a `flex`, `justify-content` e `align-items` a `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Dovresti avere un elemento principale `div` con la classe `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

L'elemento `.space` dovrebbe avere una larghezza e un'altezza di `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

L'elemento `.space` dovrebbe avere una proprietà `position` di `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Dovresti avere un `div` con una classe `earth` all'interno del `div` con una `class` di `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

L'elemento `.earth` dovrebbe essere il primo figlio dell'elemento `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

L'elemento `.earth` dovrebbe avere una larghezza e un'altezza di `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

L'elemento `.earth` dovrebbe usare il posizionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

L'elemento `.earth` dovrebbe avere una posizione `top` e `left` di `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

L'elemento `.earth` dovrebbe essere spostato indietro e verso l'alto di `50%` impostando la proprietà `transform` a `translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Dovresti avere un `div` con la classe `orbit` all'interno dell'elemento `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

L'elemento `.orbit` dovrebbe venire dopo l'elemento `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

L'elemento `.orbit` dovrebbe avere una larghezza e un'altezza di `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

L'elemento `.orbit` dovrebbe usare il posizionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

L'elemento `.orbit` dovrebbe essere posizionato al centro dell'elemento `.space` usando una proprietà `transform` che lo sposta di `-50%` sia sull'asse verticale sia su quello orizzontale.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Dovresti avere un `div` con la classe `moon` all'interno dell'elemento `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

L'elemento `.moon` dovrebbe avere una larghezza e un'altezza di `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

L'elemento `.moon` dovrebbe essere posizionato usando il posizionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

Il bordo superiore dell'elemento `.moon` dovrebbe essere allineato con il bordo superiore del suo blocco contenitore.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Dovresti spostare il bordo sinistro dell'elemento `.moon` al centro del blocco contenitore (`50%` dal lato sinistro).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

L'elemento `.moon` dovrebbe essere regolato orizzontalmente traducendolo di `-50%` lungo l'asse x.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

Il selettore `.earth` dovrebbe avere un colore di sfondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Il selettore `.earth` dovrebbe avere una `border-radius` di `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

Il selettore `.moon` dovrebbe avere un colore di sfondo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Il selettore `.moon` dovrebbe avere una `border-radius` di `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Dovresti avere una regola `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

La tua nuova regola `@keyframes` dovrebbe chiamarsi `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

La tua regola `@keyframes orbit` dovrebbe avere un selettore `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

La tua regola `@keyframes orbit` dovrebbe avere un selettore `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Il selettore `0%` dovrebbe avere una proprietà `transform` impostata a `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Il selettore `100%` dovrebbe venire dopo il selettore `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Il selettore `100%` dovrebbe avere una proprietà `transform` impostata a `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Il selettore `.orbit` dovrebbe avere una proprietà `animation` impostata a `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```
