---
id: 687d8c6f066877768f7ee441
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro de tu clase `LinkedList`, reemplaza la palabra clave `pass` con una estructura interna llamada `Node`. Esto representará nodos individuales que almacenan datos en la lista.

# --hints--

Debes crear una clase `Node` dentro de la clase `LinkedList`. Recuerda agregar la palabra clave `pass` dentro de la clase `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
