---
id: 687d8e05a529a4838f1c961c
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

A continuación, elimina la palabra clave `pass` dentro de la clase `Node`, luego dale a la clase `Node` un método `__init__` con dos parámetros, `self` y `element`.

Dentro del método, asigna `element` a `self.element`. Esto almacena el valor pasado al crear un nodo.

# --hints--

Debes definir un método `__init__` dentro de tu clase `Node`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Tu método `__init__` debe tener dos parámetros, `self` y `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Debes asignar `element` a `self.element` dentro del método `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
