---
id: 687d7deea6a2a22fa8816a73
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En lecciones anteriores, aprendiste sobre algunas estructuras de datos comunes. En este taller, implementarás una lista enlazada.

Primero, crea una clase llamada `LinkedList` para representar la estructura que contendrá los nodos de tu lista. Agrega la palabra clave `pass` dentro de la clase para que la prueba pueda pasar.

# --hints--

Debes crear una clase llamada `LinkedList`. Recuerda agregar la palabra clave `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
