---
id: 687e344cc1993f7316af030c
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro de la clase `LinkedList`, agrega un método `__init__` que establezca `self.length` en `0`. Esto inicializa un contador para rastrear cuántos nodos (elementos) hay actualmente en la lista.

# --hints--

Debes establecer `self.length` en `0` dentro del método `__init__` de la clase `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
