---
id: 687e366b3004a9831e4cbd41
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Agrega un método llamado `is_empty` que retorne `self.length == 0`.

Usarás esto para llevar un registro de cuántos nodos hay en la lista. Entonces, con `self.length == 0`, devuelve `True`, lo que significa que la lista no tiene nodos. 

Si `self.length` es mayor que `0`, devuelve `False`, lo que significa que la lista tiene al menos un nodo.

# --hints--

Debes crear un método `is_empty` que devuelva `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
