---
id: 68800a64b06e80d9a31c131b
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Todavía en el método `__init__` de la clase `LinkedList`, establece `self.head` en `None`. Esto actúa como una referencia al primer nodo de la lista. La lista está vacía por el momento, así que tienes que establecerlo en `None`.


# --hints--

Debes asignar `self.head` a `None` dentro del método `__init__` de la clase `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
