---
id: 688003f589b875af60273dd9
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Todavía dentro del método `__init__`, asigna `None` a `self.next`. Esto establece el puntero `next` en `None`, lo que significa que el nuevo nodo aún no apunta a ningún otro nodo en la lista.

Después, `next` puede actualizarse para apuntar a otro nodo y formar la cadena de una lista enlazada.

# --hints--

Dentro del método `__init__`, debes inicializar el atributo `next` a `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
