---
id: 68858a1b4e867acd87f6f7e8
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Si la lista está vacía, entonces el nuevo nodo se convierte en el primer elemento (el `head`). Dentro del método `add`, crea una estructura `if` que verifique si la lista está vacía usando `self.is_empty()` y asigna `node` a `self.head`.

# --hints--

Debes verificar si la lista está vacía usando el método `is_empty()`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

Debes asignar `node` a `self.head` dentro de la sentencia `if`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
