---
id: 68857e0a49baf4751f737a2a
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Ahora, después del método `is_empty`, crea un método llamado `add` dentro de la clase `LinkedList`. El nuevo método debe tener dos parámetros, `self` y `element`.

Dentro de tu nuevo método, crea una variable llamada `node` y asígnale una nueva instancia de `Node` usando `self.Node(element)`. Esto creará un nuevo nodo que se añadirá a la lista enlazada.

# --hints--

Debes crear un método `add` dentro de la clase `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Tu método `add` debe tener dos parámetros, `self` y `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Debes asignar `self.Node(element)` a `node` en el método `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
