---
id: 688649c8673f842f49228a5e
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Si la lista no está vacía, el método debe encontrar el último nodo y enlazar el nuevo nodo a su puntero `next`.

Para hacer esto, crea un bloque `else` después del bloque `if`. Dentro del bloque `else`, crea una variable `current_node` y asígnale `self.head`.

# --hints--

Debes asignar `self.head` a `current_node` dentro del bloque `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
