---
id: 6888fe6c745b6852ed240229
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

A continuación, recorrerás la lista enlazada para encontrar el nodo que contiene el elemento especificado para eliminar.

Crea un ciclo `while` que verifique cuando `current_node` no sea `None` y cuando `current_node.element` no sea igual a `element`. Dentro del ciclo `while`, asigna `current_node` a `previous_node`.

# --hints--

Debes tener un ciclo `while` dentro del método `remove`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

Tu ciclo `while` debe verificar cuándo `current_node` no es `None` y cuándo `current_node.element` no es igual a `element`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

Debes asignar `current_node` a `previous_node` dentro del ciclo `while`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
