---
id: 688acb2053ac7978a06345f1
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

A continuación, crea una estructura `elif` que verifique que `previous_node` no sea `None`. Asigna `current_node.next` a `previous_node.next` dentro de ella.

Con eso, si se encuentra el elemento a eliminar y no está en el nodo cabeza, el puntero `next` del nodo anterior se actualizará para saltar el nodo actual (el elemento a eliminar), removiéndolo de la lista enlazada.

# --hints--

Debes crear una estructura `elif` que verifique que `previous_node` no sea `None`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

Debes asignar `current_node.next` a `previous_node.next` dentro de la sentencia `elif`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
