---
id: 68b84df2648b4c96b25022f3
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

¡Probemos tu lista enlazada hasta ahora! Al final de tu código, crea una instancia de `LinkedList` llamada `my_list`. Luego, usando el método `is_empty` que creaste en el paso anterior, imprime `my_list.is_empty()`.

# --hints--

Debes crear una instancia de `LinkedList` llamada `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Debes imprimir `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
