---
id: 68b9e951e4892721632acc26
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

¡Genial! Ahora vamos a probar agregando elementos a tu lista enlazada. Al final de tu código, después de crear una instancia de `LinkedList` llamada `my_list` e imprimir `my_list.is_empty()`, agrega `1` y `2` a la lista usando el método `add`.

Luego, imprime `my_list.is_empty()` nuevamente para confirmar que la lista ya no está vacía. Finalmente, imprime `my_list.length` para ver la longitud actual de la lista enlazada.

# --hints--

Debes agregar `1` a la lista usando el método `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Debes agregar `2` a la lista usando el método `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Debes imprimir `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Debes imprimir `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
