---
id: 69de4308e652aaeb5ccdb03e
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستتمرن على العمل مع طرق السلاسل النصية في Python.

ابدأ بإنشاء متغير باسم `email` وقم بتعيين السلسلة النصية `alice.johnson@company.com` له.

# --hints--

يجب إنشاء متغير باسم `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

يجب تعيين السلسلة النصية `alice.johnson@company.com` إلى `email`. تذكر أن تحيط النص إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
