---
id: 69e20876baed33de636b989a
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

أحيانًا، قد تريد إيجاد موقع حرف معين داخل سلسلة نصية. يوفر Python الطريقة المدمجة `.find()` لهذا الغرض. تُرجع عددًا صحيحًا يمثل الفهرس (الموقع) لأول ظهور للسلسلة الفرعية المعطاة.

إليك مثالاً:

```py
text = 'hello'
print(text.find('e'))  # 1
```

الآن، استدعِ الطريقة `.find()` على `email` ومرر لها `'@'` لإيجاد موقع علامة @ في سلسلة `email`. عيّن النتيجة في متغير اسمه `at_position`.

ثم اطبع `at_position` بحيث يكون الناتج: `Position of @: 13`.

# --hints--

يجب إنشاء متغير اسمه `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

يجب تعيين `at_position` القيمة `email.find('@')`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

يجب طباعة `at_position` بحيث يكون الناتج `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
