---
id: 69e20a0451c7a3ef189514d3
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

بما أنك تعرف موقع `@` في البريد الإلكتروني، يمكنك استخراج اسم المستخدم الموضوع قبل علامة `@`.

في ورشة العمل السابقة، استخرجت جزءًا من سلسلة نصية باستخدام التقطيع. للتذكير، الصيغة النحوية هي:

```python
text = 'hello world'
print(text[:5])  # hello
```

استخدم التقطيع على `email` لاستخراج كل شيء قبل `at_position`، وخزن النتيجة في متغير اسمه `username`.

ثم اطبع `username` بحيث يظهر الإخراج كالتالي: `Username: alice.johnson`.

# --hints--

يجب إنشاء متغير اسمه `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

يجب استخدام التقطيع مع `at_position` على `email` لتعيين قيمة لـ `username`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

يجب طباعة `username` بحيث يظهر الإخراج كالتالي: `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
