---
id: 69e6284f3b7659c5a05c2a79
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يمكنك التحقق مما إذا كانت سلسلة نصية فرعية موجودة داخل سلسلة نصية باستخدام الكلمة المفتاحية `in`. تُرجع `True` إذا كانت السلسلة الفرعية موجودة و`False` إذا لم تكن كذلك. إليك مثالًا:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

أنشئ متغيرًا باسم `has_at` يتحقق مما إذا كان الحرف `@` موجودًا في `email`.

ثم اطبع `has_at` بحيث تكون النتيجة: `Has @ symbol: True`.

# --hints--

يجب أن تنشئ متغيرًا باسم `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

يجب أن يحتوي متغير `has_at` على القيمة `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

يجب أن تطبع `has_at` بحيث تكون النتيجة `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
