---
id: 69e62ce96df237884da25282
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

تُحوّل الطريقة `.lower()` كل الأحرف في سلسلة نصية إلى أحرف صغيرة. هذا مفيد لتوحيد عناوين البريد الإلكتروني:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

يمكنك استدعاء كل من `.strip()` و `.lower()` على سلسلة نصية بهذه الطريقة:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

أنشئ متغيرًا باسم `messy_email` وعيّن له السلسلة النصية `SARAH.DAVIS@COMPANY.COM` مع وجود مسافتين في البداية ومسافتين في النهاية.

أيضًا، استدعِ `.strip().lower()` على `messy_email` وعيّن النتيجة إلى متغير `clean_email`.

ثم اطبع كلا المتغيرين بحيث يكون الإخراج كما يلي:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

يجب أن تنشئ متغيرًا باسم `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

يجب أن تعيّن <code>'  SARAH.DAVIS@COMPANY.COM  '</code> إلى `messy_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

يجب أن تنشئ متغيرًا باسم `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

يجب أن تستدعي `.strip().lower()` على `messy_email` وتعيّن النتيجة إلى متغير `clean_email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

يجب أن تطبع `messy_email` بحيث يكون الإخراج: `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

يجب أن تطبع `clean_email` بحيث يكون الإخراج: `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
