---
id: 69e6299710319e7cb8ebdc37
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

غالبًا ما تكون البيانات الواقعية غير مرتبة مع مسافات إضافية وأحرف بحالة غير متناسقة.

أنشئ متغيرًا باسم `messy_name` وقم بتعيين السلسلة النصية `sARaH dAVis` له مع مسافتين في البداية ومسافتين في النهاية. ثم اطبع المتغير بحيث يكون الناتج:

```md
Original name: '  sARaH dAVis  '
```

ستقوم بتنظيف هذا في الخطوات القادمة.

# --hints--

يجب أن تنشئ متغيرًا باسم `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

يجب أن تعيّن <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> إلى `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

يجب أن تطبع `messy_name` بحيث يكون الناتج `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
