---
id: 69e736c39dfa403496e8071e
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

تحول الطريقة `.upper()` جميع الأحرف إلى أحرف كبيرة. هذا مفيد لشارات الأسماء أو العناوين:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

أنشئ متغيرًا باسم `name_badge` وقم بتعيين نتيجة استدعاء الطريقة `.upper()` على `clean_name` إليه.

ثم اطبع `name_badge` بحيث يكون الناتج: `Name badge: SARAH DAVIS`.

# --hints--

يجب أن تنشئ متغيرًا باسم `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

يجب أن تستدعي `.upper()` على `clean_name` وتعيّن النتيجة إلى المتغير `name_badge`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

يجب أن تطبع `name_badge` بحيث يكون الناتج `Name badge: SARAH DAVIS`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
