---
id: 69e73a141addc6f442924b0d
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

لقد استخدمت بالفعل `.find()` والتقطيع على المتغير `email`. الآن ستستخدم نفس التقنية على `clean_email`.

ابحث عن موقع الحرف `@` في `clean_email` وقم بتخزينه في `at_pos`. ثم استخدم التقطيع لاستخراج كل شيء قبل `@` وخزنه في `email_user`.

ثم اطبع `clean_email` و `email_user` بحيث يكون الناتج كما يلي:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

يجب إنشاء متغير باسم `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

يجب استدعاء `.find('@')` على `clean_email` وتعيين النتيجة إلى المتغير `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

يجب إنشاء متغير باسم `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

يجب استخدام التقطيع مع `at_pos` على `clean_email` وتعيين النتيجة إلى المتغير `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

يجب طباعة `clean_email` بحيث يكون الناتج `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

يجب طباعة `email_user` بحيث يكون الناتج `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
