---
id: 69e774b82fe0bfdb684f95b4
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

أحيانًا تحتاج إلى التحقق من كيفية بدء سلسلة نصية. يوفر Python طريقة `.startswith()` لهذا الغرض. تُرجع `True` إذا بدأت السلسلة النصية بـبادئة معينة؛ وإلا تُرجع `False`.

إليك مثالاً:

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

أنشئ متغيرًا باسم `employee_code` وعيّن له السلسلة النصية `EMP-2024-SD`. ثم أنشئ متغيرًا `starts_with_emp` يتحقق مما إذا كانت `employee_code` تبدأ بـ `EMP`.

# --hints--

يجب إنشاء متغير باسم `employee_code`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('employee_code')
`) })
```

يجب تعيين `EMP-2024-SD` إلى `employee_code`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('employee_code').is_equivalent("employee_code = 'EMP-2024-SD'")
`) })
```

يجب إنشاء متغير باسم `starts_with_emp`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('starts_with_emp')
`) })
```

يجب استخدام `.startswith('EMP')` على `employee_code` وتعيين النتيجة إلى المتغير `starts_with_emp`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('starts_with_emp').is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

--fcc-editable-region--

--fcc-editable-region--
```
