---
id: 69e76d1997661be7f0040e2d
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

طريقة `.split(separator)` تقسم سلسلة نصية إلى قائمة من السلاسل الفرعية عند كل ظهور لـ `separator`. على سبيل المثال:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

أنشئ متغيرًا باسم `full_address` وعيّن له السلسلة النصية `123 Main Street, Springfield, IL`. ثم أنشئ `address_parts` بتقسيم السلسلة عند الفاصلة متبوعة بمسافة واحدة.

ثم اطبع `address_parts` بحيث يكون الناتج: `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

# --hints--

يجب أن تنشئ متغيرًا باسم `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

يجب أن تعيّن السلسلة النصية `123 Main Street, Springfield, IL` إلى `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

يجب أن تنشئ متغيرًا باسم `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

يجب أن تستخدم `.split(', ')` على `full_address` وتعيّن النتيجة إلى المتغير `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

يجب أن تطبع `address_parts` بحيث يكون الناتج: `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
