---
id: 69e73f4f5f986657647bbd63
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

طريقة `.replace(old, new)` تستبدل كل ظهور للسلسلة الفرعية `old` بأخرى، وهي `new`. يمكنك أيضًا ربطها مع طرق أخرى. على سبيل المثال:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

استبدل `.` بمسافة واحدة في `email_user`، ثم استدعِ `.title()` على النتيجة. عيّن النتيجة في متغير اسمه `display_name`.

ثم اطبع `display_name` بحيث يكون الإخراج: `Display name: Sarah Davis`.

# --hints--

يجب إنشاء متغير اسمه `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

يجب استخدام `.replace('.', ' ')` و`.title()` على `email_user` وتعيين النتيجة في المتغير `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

يجب طباعة `display_name` بحيث يكون الإخراج `Display name: Sarah Davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
