---
id: 6a71b3e49668c3be461257b5
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

تُحوّل الطريقة `.title()` الحرف الأول من كل كلمة إلى حرف كبير. إليك مثالاً:

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

أنشئ متغيرًا باسم `clean_name` وقم بتعيين نتيجة استدعاء `.title()` على `messy_name` له.

# --hints--

يجب أن تنشئ متغيرًا باسم `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

يجب أن تستخدم `.title()` على `messy_name` لتعيين قيمة لـ `clean_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
