---
id: 6a86b9e64da746838e67fb13
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

تُزيل الطريقة `.strip()` الفراغات من كلا طرفي السلسلة النصية.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

يمكنك ربط الطريقتين `.strip()` و `.title()` لإزالة الفراغات من كلا طرفي السلسلة النصية وتكبير الحرف الأول من كل كلمة:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

حدّث قيمة `clean_name` لاستدعاء `.strip()` على `messy_name` قبل `.title()`.

بعد ذلك، اطبع `clean_name` بحيث تكون النتيجة `Cleaned name: 'Sarah Davis'`.

# --hints--

يجب تحديث `clean_name` لاستدعاء `.strip()` على `messy_name` قبل `.title()`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

يجب طباعة `clean_name` بحيث تكون النتيجة `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
