---
id: 6a870cc60c23f0332a2fe788
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

تذكر من ورشة العمل السابقة أنه يمكنك استخدام f-string لتضمين متغير داخل سلسلة نصية، بكتابة `f` قبل علامة الاقتباس الافتتاحية وتغليف المتغير بـ `{}`.

إليك تذكير بكيفية عملها:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

الآن، اطبع `email` بحيث تكون النتيجة: `Email: alice.johnson@company.com`.

# --hints--

يجب أن تطبع `email` بحيث تكون النتيجة `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
