---
id: 66a37f37ef5823a313de8c26
title: 创建一个月球轨道
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

在本实验室中，你将使用 HTML 和 CSS 制作一个简单的月球绕地球运行的动画。 地球将位于系统的中心，月球将围绕地球运行。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该通过将 `body` 元素高度设置为视口高度的 `100%`，并将其设置为弹性布局，使所有元素在页面中居中对齐。
   
2. 你应该有一个带有 `space` 类的 `div` 元素。

3. 在 `.space` 元素中，应该有两个 `div` 元素，依次分别具有类 `earth` 和 `orbit`。

4. 在 `.orbit` 元素中，应该有另一个 `div`，类为 `moon`。
   
5. `class` 为 `space` 的 `div` 元素应该在页面上居中，其宽度和高度为 `200px`。

6. `class` 为 `space` 的 `div` 元素应该使用 `relative` 定位。

7. `.earth` 元素应该使用 `absolute` 定位。 将 `.earth` 元素的中心定位在父元素垂直轴（顶部）和水平轴（左侧）的中间位置。 然后，将 `.earth` 元素向后移动其宽度和高度的一半，使其在父元素 `.space` 中居中。

8. `.earth` 元素的宽度和高度应该为 `100px`。

9.  `.orbit` 元素的宽度和高度应该为 `200px`。

10.  `.orbit` 元素应该使用 `absolute` 定位。 通过使用 `transform` 属性将其垂直轴和水平轴均偏移 `-50%`，使其右下角位于 `.space` 元素的中心。

11. 月球环绕地球的轨道路径应该是一个圆。

12. The `.moon` 元素应该使用 `absolute` 定位，其宽度和高度为 `30px`。 `.moon` 元素应位于 `.orbit` 元素的顶部，并水平居中。

13. 你应进一步调整 `.moon` 元素的水平定位，将该元素向左移动其宽度的一半。

14. `.earth` 选择器应有背景色和 `50%` 的 `border-radius`。

15. `.moon` 选择器应有背景色和 `50%` 的 `border-radius`。

16. 你应定义一个 `@keyframes orbit` 动画，使 `.orbit` 元素围绕其中心旋转 360 度。 你应将此动画应用于 `.orbit` 元素，持续时间为 `5` 秒，采用线性缓动函数，并可无限循环。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

`body` 元素的内容应通过设置高度为视口高度的 `100%`、`display` 为 `flex`、`justify-content` 和 `align-items` 为 `center`，在页面中居中。

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

你应该有一个类为 `space` 的主 `div`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

`.space` 元素的宽度和高度应该为 `200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

`.space` 元素应该有一个 `position` 属性，值为 `relative`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

你应该在 `class` 为 `space` 的 `div` 中有一个类为 `earth` 的 `div`。

```js
assert.exists(document.querySelector('div.space div.earth'));
```

`.earth` 元素应是 `.space` 元素的第一个子元素。

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

`.earth` 元素的宽度和高度应该为 `100px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

`.earth` 元素应该使用 `absolute` 定位。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

`.earth` 元素的 `top` 和 `left` 位置应该为 `50%`。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

应该通过将 `transform` 属性设置为 `translate(-50%, -50%)`，将 `.earth` 元素向后和向上移动 `50%`。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

你应该在 `.space` 元素内部，有一个带有 `orbit` 类的 `div` 元素。

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

`.orbit` 元素应该放在 `.earth` 元素之后。

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

`.orbit` 元素的宽度和高度应该为 `200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

`.orbit` 元素应该使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

应该通过 `transform` 属性将 `.orbit` 元素沿垂直轴和水平轴均偏移 `-50%`，使其位于 `.space` 元素的中心。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

你应该在 `.orbit` 元素内部，有一个带有 `moon` 类的 `div` 元素。

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

`.moon` 元素的宽度和高度应该为 `30px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

`.moon` 元素应该使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

`.moon` 元素的顶部边缘应与其容器的顶部边缘对齐。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

你应该将 `.moon` 元素的左边缘移动到容器的中心位置（左侧 `50%`）。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

应该水平调整 `.moon` 元素，将其在 x 轴上平移 `-50%`。

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

`.earth` 选择器应有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

`.earth` 选择器应有 `50%` 的 `border-radius`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

`.moon` 选择器应有背景颜色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

`.moon` 选择器应有 `50%` 的 `border-radius`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

你应该有一个 `@keyframes` 规则。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

你的新的 `@keyframes` 规则应命名为 `orbit`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

你的 `@keyframes orbit` 规则应该有一个 `0%` 选择器。

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

你的 `@keyframes orbit` 规则应该有一个 `100%` 选择器。

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

你的 `0%` 选择器应该有一个 `transform` 属性，设置为 `rotate(0deg) translate(-50%, -50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

你的 `100%` 选择器应该在 `0%` 选择器之后。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

你的 `100%` 选择器应该有一个 `transform` 属性，设置为 `rotate(360deg) translate(-50%, -50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

你的 `.orbit` 选择器应该有一个 `animation` 属性，设置为 `orbit 5s linear infinite`。

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

