---
id: 669e2f60e83c011754f711f9
title: Crea una página de agencia de viajes
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una declaración `DOCTYPE`.
1. Debes tener un elemento `html` con `lang` configurado en `en`.
1. Debes tener un elemento `head` que contenga un elemento void `meta` con `charset` configurado en `utf-8` y un `title` con el texto `Travel Agency Page`.
1. Debes tener una etiqueta `meta` en tu elemento `head` que contenga una breve descripción de tu sitio web para SEO.
1. Debes tener un elemento `h1` para presentar tus destinos de viaje. Debe ser el único elemento `h1` en tu página.
1. Debes tener un párrafo debajo del elemento `h1` que introduzca las oportunidades de viaje.
1. Debes tener un elemento `h2` con el texto `Packages`.
1. Debes tener un elemento `p` que introduzca brevemente los distintos paquetes.
1. Debes tener un elemento de lista desordenada con dos elementos de lista. Los dos elementos de lista deben tener el texto `Group Travels` y `Private Tours`, respectivamente. El texto de cada elemento de lista debe estar encerrado por un elemento ancla.
1. Debes tener un elemento `h2` con el texto `Top Itineraries`.
1. Debes tener al menos tres elementos `figure`, cada uno conteniendo un elemento ancla y un elemento `figcaption` que describa la imagen de la figura.
1. Los tres elementos ancla deben tener un elemento `img` con un atributo `alt` apropiado y un atributo `src` configurado a una imagen válida como su contenido. Puedes usar `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` y `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` si quieres.
1. Todos tus cinco elementos ancla deben tener un atributo `href` con el valor `https://www.freecodecamp.org/learn` y un atributo `target` con el valor `_blank`.

# --hints--

Su página de agencia de viajes debería tener una declaración `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Debes tener un elemento `html` con el atributo `lang` configurado a `en`.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/html\s*>/i);
```

Debes tener un elemento `head` dentro del elemento `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

Tu elemento `head` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/head\s*>/i);
```

Debes tener un elemento `title` dentro de tu elemento `head`.

```js
assert.isNotNull(document.querySelector('head title'));
```

Tu elemento `title` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/title\s*>/i);
```

Debería tener dos elementos `meta` dentro de su elemento `head`.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

Un elemento `meta` debería tener un atributo `name` con el valor de `description` y un atributo `content` no vacío.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

Un elemento `meta` debe tener su atributo `charset` configurado en `utf-8`.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Tu elemento `title` debe tener el nombre de tu agencia de viajes.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Debes tener un elemento `body` dentro de tu elemento `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

Tu elemento `body` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/body\s*>/i);
```

Debes tener un elemento `h1` para presentar tus destinos de viaje.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Debes tener solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debes tener un elemento `p` debajo del elemento `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Tu primer elemento `p` debe presentar las oportunidades de viaje.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Tu primer elemento `h2` debe tener el texto `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Debes tener un elemento `p` debajo de tu primer elemento `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Tu segundo elemento `p` debe presentar brevemente los distintos paquetes.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Debes tener un elemento de lista desordenada debajo de tu segundo elemento `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Debes tener dos elementos en tu lista desordenada.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Ambos elementos de la lista deben contener un elemento ancla.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

El elemento ancla de tu primer elemento de lista debe envolver el texto `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

El elemento ancla de tu segundo elemento de lista debe envolver el texto `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Debes tener un elemento `h2` después de tu lista desordenada.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Tu segundo elemento `h2` debe tener el texto `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Debes tener al menos tres elementos `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Cada elemento `figure` debe contener un elemento ancla como su primer elemento anidado.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Cada elemento `figure` debe contener un elemento `figcaption` como su segundo elemento anidado.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Cada `figcaption` debe contener algo de texto.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Cada uno de los elementos `a` anidados en tus elementos `figure` debe contener una imagen.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Cada elemento `img` debe tener un atributo `src` válido.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Cada elemento `img` debe tener un atributo `alt` con un valor apropiado.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Cada elemento `a` debe tener un atributo `href` con el valor `https://www.freecodecamp.org/learn`. No olvides los enlaces en los elementos de la lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Cada elemento `a` debe tener un atributo `target` con el valor `_blank`. No olvides los enlaces en los elementos de la lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
