---
id: 680900675ae3d54ee19590c9
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Dentro de su `select` elemento, necesita crear opciones para que el usuario elija como su fuente de energía. El `select` debe tener siete opciones listadas a continuación:

| Fuente de Energía              |
| ------------------------------ |
| `Select one`                   |
| `Bitten by a strange creature` |
| `Radioactive exposure`         |
| `Science experiment`           |
| `Alien heritage`               |
| `Ancient artifact discovery`   |
| `Other`                        |

La primera opción (`Select one`) debe tener una cadena vacía como su atributo `value` (es decir, `value=''`). Esto actúa como un marcador de posición en la interfaz de usuario.

# --hints--

Debe crear el primer `option` elemento con el texto `Select one` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

Debe dar al primer `option` elemento un atributo `value` con una cadena vacía.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.getAttribute('value'), '');
```

Debe crear un segundo `option` elemento con el texto `Bitten by a strange creature` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

Debe crear un tercer `option` elemento con el texto `Radioactive exposure` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

Debe crear un cuarto `option` elemento con el texto `Science experiment` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

Debe crear un quinto `option` elemento con el texto `Alien heritage` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

Debe crear un sexto `option` elemento con el texto `Ancient artifact discovery` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

Debe crear un séptimo `option` elemento con el texto `Other` dentro de su `select` elemento.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
        </div>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>

            </select>
            --fcc-editable-region--
          </label>
      </form>
    </div>
  )
};
```
