---
id: 641ceed81533263283835c3d
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

除了使用 `.match()` 方法之外，你還可以使用正則表達式的 `.test()` 方法來測試字符串是否與模式匹配。 與 `.match()` 不同，`.test()` 返回一個布爾值，指示字符串是否與模式匹配。

更新你的 `isSpam()` 函數以在 `helpRegex` 上使用 `.test()` 方法來測試 `msg` 是否匹配。

然後，嘗試在你的頁面上輸入一些消息並查看結果。

# --hints--

你的 `isSpam()` 函數應該使用 `.test()` 方法。

```js
const testArr = [];
const temp = RegExp.prototype.test;
RegExp.prototype.test = (msg) => {testArr.push("test")};
try {
  isSpam("");
  assert.strictEqual(testArr[0], "test");
} finally {
  RegExp.prototype.test = temp;
}
```

你的 `isSpam()` 函數應該在 `helpRegex` 上調用 `.test()` 方法，並以 `msg` 作爲參數。

```js
const testArr = [];
const temp = helpRegex.test;
helpRegex.test = (msg) => {testArr.push(msg)};
try {
  isSpam("spam");
  assert.strictEqual(testArr[0], "spam");
  assert.lengthOf(testArr, 1);
} finally {
  helpRegex.test = temp;
}
```

你的 `isSpam()` 函數應該返回 `helpRegex.test(msg)` 的結果。

```js
const testArr = [];
const temp = helpRegex.test;
helpRegex.test = (msg) => testArr.push(msg);
const spam = "pleASe HeLp";
try {
  isSpam(spam);
  assert.strictEqual(testArr[0], spam);
} finally {
  helpRegex.test = temp;
  assert.isTrue(isSpam(spam));
  assert.isFalse(isSpam("not spam"));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help/i;

--fcc-editable-region--
const isSpam = (msg) => msg.match(helpRegex);
--fcc-editable-region--

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
