---
id: 642208bc4d44701c6fd6f65e
title: 步驟14
challengeType: 0
dashedName: step-14
---

# --description--

美元價值可能不止一位數字。 爲了匹配這一點，可以使用 `+` 量詞 - 它匹配一個或多個連續的出現。 例如，正則表達式 `/a+/` 與一個或多個連續的 `a` 字符匹配。

更新正則表達式以匹配一個或多個連續的數字。

# --hints--

你的 `dollarRegex` 應該使用 `+` 量詞。

```js
assert.match(dollarRegex.source, /\+/);
```

你的 `dollarRegex` 應該在 `[0-9]` 字符類上使用 `+` 量詞。

```js
assert.match(dollarRegex.source, /\[0-9\]\+/);
```

你的 `dollarRegex` 應該仍然與 `"100 dollars"` 匹配。

```js
const hundredDollars = '100 dollars'
assert.deepEqual(hundredDollars.match(dollarRegex), hundredDollars.match(/[0-9]+ dollars/i));
```

你的 `dollarRegex` 應該仍然與 `"3 dollars"` 匹配。

```js
const threeDollars = '3 dollars'
assert.deepEqual(threeDollars.match(dollarRegex), threeDollars.match(/[0-9]+ dollars/i));
```

你的 `dollarRegex` 應該不區分大小寫。

```js
assert.include(dollarRegex.flags, 'i');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /[0-9] dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
