---
id: 642206e054eef81b5e3092ed
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你需要在文字 `dollars` 前面匹配一個數字。雖然你可以寫出 `0|1|2` 等等，但正規表達式有一個讓這更簡單的特性。

<dfn>字元類別</dfn>由方括號定義，並且會符合括號內的任何字元。<dfn>例如</dfn>，`[aeiou]` 會符合列表 `aeiou` 中的任何字元。你也可以使用連字號定義要符合的字元範圍。<dfn>例如</dfn>，`[a-z]` 會符合從 `a` 到 `z` 的任何字元。

為你的 `dollarRegex` 表達式添加一個字元類別以匹配數字 `0` 到 `9` —— 請記住數字必須出現在字 `dollars` 之前，且數字與字之間應該有一個空格。

# --hints--

你的 `dollarRegex` 應該使用字元類別。

```js
assert.match(dollarRegex.source, /\[.*\]/);
```

你的字元類別應該是 `0-9`。

```js
assert.match(dollarRegex.source, /\[0-9\]/);
```

你的 `dollarRegex` 應該仍然能匹配 `"1 dollars"`。

```js
const oneDollar = '1 dollars'
assert.deepEqual(oneDollar.match(dollarRegex), oneDollar.match(/[0-9] dollars/i));
```

你的 `dollarRegex` 應該是不區分大小寫的。

```js
assert.include(dollarRegex.flags, 'i');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result-message"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result-message");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
