---
id: 680b3ef395479b0e449ecb6e
title: 선택 정렬 알고리즘 구현하기
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

선택 정렬은 대부분의 컴퓨터 과학 강의에서 가르치는 또 다른 인기 있는 정렬 알고리즘입니다.

이 알고리즘은 리스트의 정렬되지 않은 부분에서 가장 작은 요소를 반복해서 찾아 첫 번째 정렬되지 않은 요소와 교환하는 방식으로 작동합니다. 전체 리스트에서 최솟값을 선택해 첫 번째 요소와 교환하는 것으로 시작합니다. 그런 다음 두 번째 위치로 이동해 남은 정렬되지 않은 요소 중 가장 작은 값을 찾아 두 번째 요소와 교환합니다. 이 과정을 리스트를 한 요소씩 순회하며 전체 리스트가 정렬될 때까지 계속합니다.

선택 정렬은 최선, 평균, 최악의 경우 모두에서 이차 시간 복잡도를 가집니다. 공간 복잡도는 리스트 크기와 상관없이 일정한 메모리를 사용하며 제자리에서 정렬하기 때문에 `O(1)`입니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `selection_sort`라는 이름의 함수를 정의해야 합니다.
1. `selection_sort` 함수는 항목 리스트를 나타내는 매개변수 하나를 가져야 합니다.
1. `selection_sort` 함수는 리스트를 받아 가장 작은 값부터 가장 큰 값 순서로 제자리에서 정렬해야 합니다.
1. `selection_sort` 함수는 입력 리스트를 제자리에서 수정하고 정렬된 리스트를 반환해야 합니다.
1. `selection_sort` 함수는 선택 정렬 알고리즘을 따라 리스트의 정렬되지 않은 부분에서 가장 작은 요소를 첫 번째 정렬되지 않은 요소와 교환해야 합니다.
1. `selection_sort` 함수는 가장 작은 요소가 이미 올바른 위치에 있을 때 불필요한 교환을 수행하지 않아야 합니다.
1. `selection_sort` 함수는 내장 `sort()` 메서드나 `sorted()` 함수를 사용하지 않아야 합니다.

# --hints--

`selection_sort`라는 이름의 함수가 있어야 합니다.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

`selection_sort` 함수는 매개변수 하나를 가져야 합니다.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

코드에서 어떤 모듈도 가져오거나 내장 정렬 메서드를 사용하지 않아야 합니다.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

`selection_sort`는 입력 리스트와 같은 리스트를 반환해야 합니다.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

`selection_sort`는 입력 리스트를 제자리에서 수정해야 합니다. 리스트에 항목을 추가하거나 제거하는 메서드를 사용하지 않아야 합니다.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

`selection_sort` 함수는 선택 정렬 알고리즘을 따라 리스트의 정렬되지 않은 부분에서 최솟값을 첫 번째 정렬되지 않은 요소와 교환해야 합니다. 최솟값이 이미 올바른 위치에 있을 때는 불필요한 교환을 피하세요.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])`는 `[1, 2, 33, 67, 89, 245]`를 반환해야 합니다.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])`는 `[3, 5, 12, 15, 16, 23, 72, 99, 567]`를 반환해야 합니다.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])`는 `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]`를 반환해야 합니다.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

`selection_sort` 함수는 어떤 숫자 리스트도 올바르게 정렬해야 합니다.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
