---
id: 68ed661e9f463c155bb677a1
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

在本次研讨会中，你将通过设计一组彩色盒子来练习使用 CSS flexbox。

首先，在 `body` 元素内添加一个 `header` 元素。在 `header` 元素内嵌套一个带有文本 `Colored Boxes Layout` 的 `h1` 元素。

# --hints--

你应该在 `body` 元素内有一个 `header` 元素。

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

`header` 元素内部应该有 `h1` 元素。

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

你的 `h1` 元素中应包含 `Colored Boxes Layout` 文本。

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
