---
id: 68ef0b3237fef991c92d9c19
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

现在在你的 `header` 下方添加一个带有 `flex-container` 类的 `div` 元素。然后在你刚创建的 `div` 内部创建六个带有 `box` 类的 `div` 元素。

请记住，类属性通常用于指向样式表中的类名。在这种情况下，所有的 `div` 元素将根据 `.flex-container` 和 `.box` 样式定义被同等地设置样式。

# --hints--

你应该在 `header` 元素下方有一个 `div` 元素。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

你在 `header` 元素下的 `div` 元素应具有 `flex-container` 类。

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

你应该在 `.flex-container` 元素内嵌套六个 `div` 元素。

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

你 `.flex-container` 元素内的所有六个 `div` 元素都应该有 `box` 类。

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
