---
id: 68f471410dbccba9aac8234f
title: 步骤 41
challengeType: 0
dashedName: step-41
---

# --description--

现在创建一个 `.box5` 选择器，并为它添加值为 `#457b9d` 的 `background` 属性。

然后为你的新选择器添加 `order` 属性，值为 `4`。你会看到这个盒子移动到带有 `box4` 类的盒子之后。

# --hints--

你应该有一个 `.box5` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box5'));
```

你的 `.box5` 选择器应具有属性 `background`，其值为 `#457b9d`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box5')?.getPropVal('background'), 'rgb(69, 123, 157)');
```

你的 `.box5` 选择器应具有属性 `order`，其值为 `4`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box5')?.getPropVal('order'), '4');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box box5">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3;
}

--fcc-editable-region--

--fcc-editable-region--
```
