---
id: 68f46b758eb8ea7043abd73c
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

现在你将组织你的 flexbox 容器内的盒子。首先，创建一个 `.box1` 选择器，并为它添加 `background` 属性，值为 `#f16e79`。

然后，添加属性 `order`，值为 `1`。你会看到第一个盒子移动到末尾。之所以发生这种情况，是因为 `box` 类已经有一个 `order` 值为 `0`，而具有较高 order 值的项会出现在后面。

还要添加值为 `1` 的 `flex-grow` 属性。这将使盒子在其所在行中增长以填充额外空间。

# --hints--

你应该有一个 `.box1` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

你的 `.box1` 选择器应具有属性 `background`，其值为 `#f16e79`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

你的 `.box1` 选择器应具有属性 `order`，其值为 `1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

你的 `.box1` 选择器应具有属性 `flex-grow`，其值为 `1`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
