---
id: 68f91df5776e693049fb8942
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

现在是时候为 `styles.css` 文件添加一些 CSS 规则了。首先为 `h1` 元素创建一个选择器。 

通过将 `h1` 元素的 `text-align` 属性设置为 `center` 来居中内容。还创建一个值为 `10px` 的 `margin-bottom` 属性，以在 `h1` 元素底部和任何位于其下方的超文本标记语言元素之间添加外边距。

# --hints--

应该有一个 `h1` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 `h1` 选择器应该将 `text-align` 属性设置为 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

你的 `h1` 选择器应将 `margin-bottom` 属性设置为 `10px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

