---
id: 692f22a2ad44ae887a24c8f9
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在每个 `.box` 元素内，添加一个带有文本 `Box`，后跟一个空格和从 `1` 开始的循序数字的 `h2` 元素。

# --hints--

你应该在第一个带有 `box` 类的 `div` 内有一个 `h2` 元素，里面包含 `Box 1` 文本。

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

你应该在第二个带有 `box` 类的 `div` 内有一个 `h2` 元素，里面的文本是 `Box 2`。

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

你应该在第三个带有 `box` 类的 `div` 内有一个 `h2` 元素，里面的文本是 `Box 3`。

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

你应该在第四个带有 `box` 类的 `div` 内有一个 `h2` 元素，里面的文本是 `Box 4`。

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

你应该在第五个具有 `box` 类的 `div` 内有一个 `h2` 元素，里面的文本是 `Box 5`。

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

你应该在最后一个带有 `box` 类的 `div` 内有一个 `h2` 元素，里面的文本是 `Box 6`。

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
