---
id: 692f3001c9fc019f006f6b4e
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

`align-items` 属性沿交叉轴（垂直于主轴）对 flex 项进行对齐。

现在为 `.box` 选择器添加 `align-items` 属性，值为 `center`。由于 `box` 类具有 `flex-direction: column`，交叉轴是水平的，因此 `center` 会使项在盒子内水平居中。

# --hints--

你的 `.box` 选择器应具有属性 `align-items`，其值为 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('align-items'), 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
