---
id: 692f22933ba1cb86cd75a70c
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

现在是为这些盒子创建颜色名称的时候了。为此，在每个 `h2` 元素下方添加一个 `p` 元素，按顺序包含这些颜色名称：`Red`、`Orange`、`Yellow`、`Green`、`Blue` 和 `Indigo`。每个盒子对应一种颜色。

# --hints--

你应该在第一个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Red` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

你应该在第二个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Orange` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

你应该在第三个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Yellow` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

你应该在第四个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Green` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

你应该在第五个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Blue` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

你应该在第六个 `.box` 元素内的 `h2` 元素下方添加一个带有文本 `Indigo` 的 `p` 元素。

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
