---
id: 680a7e200400089967b8b1f4
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

El método `str.maketrans()` toma dos cadenas de igual longitud y devuelve una tabla de traducción que asigna cada carácter de la primera cadena con el carácter correspondiente de la segunda cadena. Cada carácter en la tabla de traducción se almacena como un ordinal Unicode, un número que identifica de manera única al carácter.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Aprenderás más sobre el tipo de estructura que devuelve `str.maketrans()` más adelante en el plan de estudios. Por ahora, crea una tabla de traducción que asigne los caracteres en `alphabet` a los caracteres en `shifted_alphabet` y asígnala a una variable llamada `translation_table`.

# --hints--

Debes declarar una variable llamada `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Debes llamar a `str.maketrans()` pasando `alphabet` y `shifted_alphabet` como los argumentos, y asignar el resultado a `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
