---
id: 680a836235461fce1ff53b3f
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

El método `translate()` toma como argumento la tabla de traducción generada por `str.maketrans()`. Se llama en una cadena y devuelve una copia de la cadena original donde los caracteres han sido reemplazados según la tabla de traducción:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Llama al método `translate()` en `text` pasando como argumento `translation_table` y asigna el resultado a una variable llamada `encrypted_text`.

# --hints--

Debe haber una variable llamada `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Debes asignar `text.translate(translation_table)` a `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
