---
id: 680a549f60f5a3fedfc0edd9
title: Pasos 4
challengeType: 20
dashedName: step-4
---

# --description--

Como puedes ver en la salida, el alfabeto desplazado empieza en la letra `f` porque `shift` tiene el valor `5`. Pero ahora las primeras cinco letras del alfabeto – `a`, `b`, `c`, `d` y `e` – faltan en el alfabeto desplazado, por lo que necesitarás agregarlas al final del alfabeto desplazado.

El operador `+` se utiliza para combinar dos o más cadenas juntas en un proceso llamado concatenación de esta manera:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Usa la sintaxis de corte para extraer la primera parte faltante de `alphabet` y concatenarla a `alphabet[shift:]`.

Como recordatorio, `sentence[start:stop]` devuelve los caracteres de `sentence` desde la posición `start` (incluida) hasta `stop` (excluida).

# --hints--

Debes concatenar la primera parte faltante de `alphabet` a `alphabet[shift:]` usando la sintaxis de corte. Usa `shift` para especificar dónde detener el corte.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:]
print(shifted_alphabet)
--fcc-editable-region--
```
