---
id: 680a41ef5bfdc059cfd9d4b5
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Declara otra variable llamada `alphabet` y asigna la cadena `abcdefghijklmnopqrstuvwxyz` a esta variable.

# --hints--

Debes declarar una variable llamada `alphabet`. Presta atención para colocar el nombre de la variable al principio de la línea.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Debes asignarle la cadena `abcdefghijklmnopqrstuvwxyz` a tu variable `alphabet`. Recuerde utilizar comillas simples o dobles para encerrar la cadena.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
