---
id: 680a549f60f5a3fedfc0edd9
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Como puedes ver en la salida, el alfabeto desplazado empieza en la letra `f` porque `shift` tiene el valor `5`. Pero ahora las primeras cinco letras del alfabeto – `a`, `b`, `c`, `d` y `e` – faltan en el alfabeto desplazado, por lo que necesitarás agregarlas al final del alfabeto desplazado.

El operador `+` se utiliza para combinar dos o más cadenas juntas en un proceso llamado concatenación de esta manera:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Modifica la asignación existente de la variable `shifted_alphabet`: usa la sintaxis de slicing para extraer la primera porción faltante de `alphabet` y concatenarla a `alphabet[shift:]`.

Como recordatorio, `sentence[start:stop]` devuelve los caracteres de `sentence` desde la posición `start` (incluida) hasta `stop` (excluida).

# --hints--

No debes modificar directamente la variable `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Debes asignar la concatenación de `alphabet[shift:]` y la primera porción faltante de `alphabet` a la variable `shifted_alphabet` en la misma línea en que fue declarada. Usa `shift` para especificar dónde detener el corte.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:]
print(shifted_alphabet)
--fcc-editable-region--
```
