---
id: 680a431692a5f863f726c438
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En este taller, vas a construir un *cifrado de César*. Ésta es una de las técnicas más simples para encriptar texto, que consiste en sustituir cada letra del texto plano con la letra que se encuentra a un número fijo de posiciones hacia abajo en el alfabeto. Por ejemplo, con un desplazamiento de 5, `a` se reemplazaría por `f`, `b` por `g` y así sucesivamente.

Para implementar este cifrado, necesitas crear una nueva versión del alfabeto que comience en la posición indicada por el desplazamiento. As you learned in a previous lesson, you can extract part of a string using string slicing:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Crea una variable llamada `shifted_alphabet` y usa la sintaxis de segmentación para asignarle la porción de `alphabet` que comienza en el índice de `shift`. Luego, llama a la función incorporada `print()` para imprimir `shifted_alphabet` en la terminal y observa el resultado.

# --hints--

Debes declarar una nueva variable llamada `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Debes asignar `alphabet[shift:]` a `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Debes llamar a la función `print()` pasando `shifted_alphabet` como argumento.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5

--fcc-editable-region--
```
