---
id: 680a3d4c065557260de2ff33
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, crearás un programa llamado cifrado César, que es un método de encriptación que desplaza letras en el alfabeto para codificar mensajes.

Comienza creando una variable llamada `shift` y asigna el valor `5` a tu nueva variable.

# --hints--

Debes declarar una variable llamada `shift`. Presta atención para colocar el nombre de la variable al principio de la línea.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Debes asignar el valor `5` a tu variable `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
