---
id: 680a82806c4c6bc3567f2708
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Declare una nueva variable llamada `text` y asígnale la cadena `hello world`. Este será el mensaje para encriptar.

# --hints--

Debes declarar una variable llamada `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Debes asignar la cadena `hello world` a `text`. Recuerda encerrar la cadena entre comillas simples o dobles y prestar atención a las mayúsculas y minúsculas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
