---
id: 6818b80519e6c6c47a8bb6e8
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Como aprendiste en una lección anterior, las funciones pueden tener parámetros, que son variables que pueden ser referenciadas dentro de la función. Aquí tienes una función `sum` con dos parámetros, `num1` y `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

El mensaje para cifrar y el desplazamiento siguen estando codificados de manera fija en tu función. Dale a tu función dos parámetros llamados `text` y `shift`, y elimina las declaraciones de ambas variables, `text` y `shift`, del cuerpo de la función `caesar`.

# --hints--

La función `caesar` debe tener dos parámetros, `text` y `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

No debes tener una variable `text` declarada dentro del cuerpo de la función `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

No debes tener una variable `shift` declarada dentro del cuerpo de la función `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--
```
