---
id: 6819ca464db219ab86f37dc4
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Como ya has verificado, el desplazamiento para cifrar el texto debe ser positivo. Sin embargo, no puede exceder <codigo>25</codigo> (el último índice de <codigo>alfabeto</codigo>).

Agrega una segunda condición a la declaración <codigo>if</codigo> que verifique que <codigo>shift</codigo> es mayor que <codigo>25</codigo>. Recuerda que la operación OR lógico en Python se implementa a través del operador <codigo>or</codigo>.

Además, actualiza el mensaje devuelto a <codigo>Shift must be an integer between 1 and 25.</codigo>.

# --hints--

Debes usar el operador `or` para agregar una segunda condición después de `shift < 1` que verifique si `shift` es mayor que `25`. No debes modificar la condición existente.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Tu segunda declaración <codigo>if</codigo> debería devolver <codigo>Shift must be an integer between 1 and 25.</codigo>.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
