---
id: 6819c5c9faf2548e556163d5
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Como recuerdas de lecciones anteriores, el operador lógico unario `not` niega una expresión:

```py
print(not True) # False
print(not False) # True
```

Utiliza el operador `not` para corregir la condición de la instrucción `if` para que la función devuelva `Shift must be an integer value.` cuando `shift` *no* es un entero.

# --hints--

Debes usar el operador `not` para negar `isinstance(shift, int)` en tu instrucción `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    if isinstance(shift, int):
        return 'Shift must be an integer value.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
