---
id: 6819cf38880d8cc9b9f6af7a
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Python te permite especificar un valor por defecto para los parámetros en una función, creando una función que puede ser llamada con menos o ningún argumento. Aquí hay cómo crear una función con un parámetro de `nombre` que tenga un valor por defecto:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Agrega un tercer parámetro llamado `encrypt` a la función y asígnale un valor por defecto de `True`. 

# --hints--

Debes agregar un tercer parámetro llamado `encrypt` con un valor por defecto de `True` a tu función. No modifiques los parámetros existentes.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
