---
id: 6819bc0637b80256a33adccc
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Ahora que has implementado las funcionalidades básicas del cifrado, es momento de agregar algo de validación. Para eso, necesitarás una declaración `if`. Aquí hay un recordatorio de la sintaxis para una declaración `if`:

```py
if condition:
    # code to run when condition is true
```

Al comienzo del cuerpo de tu función, crea una estructura `if`. Por ahora, usa `True` como condición, y dentro del cuerpo de la estructura `if` devuelve la cadena `Shift must be an integer value.`

# --hints--

Deberías tener una declaración `if` dentro de tu función `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Tu declaración `if` debería usar `True` como su condición.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Deberías devolver la cadena `Shift must be an integer value.` desde tu declaración `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
