---
id: 681a2a6cf95e82ca30866364
title: Paso 23
challengeType: 20
dashedName: step-23
---

# --description--

Declara dos funciones llamadas `encrypt` y `decrypt`, ambas con parámetros `text` y `shift`.

Usarás `encrypt` para la encriptación y `decrypt` para la desencriptación, dando a cada acción un nombre descriptivo.

Desde ambas funciones nuevas, devuelve una llamada a `caesar` que pase `text` y `shift`. En `decrypt`, también pasa `False` como tercer argumento.

# --hints--

Debes tener una función llamada `encrypt` con dos parámetros, `text` y `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Tu función `encrypt` debe devolver `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Debes tener una función llamada `decrypt` con dos parámetros, `text` y `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Tu función `decrypt` debe devolver `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
