---
id: 681a1d60e037127e7c8e0230
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Vas a utilizar el parámetro agregado en el paso anterior para determinar si la función debe cifrar el texto que se le pasa (comportamiento predeterminado, `encrypt=True`), o si debe descifrar un mensaje cifrado.

Crea una declaración `if` que verifique si `encrypt` no es verdadero. Dentro de la nueva declaración `if`, asigna `shift` a `- shift`. Esto es necesario para que el desplazamiento se realice en la dirección opuesta con respecto al proceso de cifrado.

# --hints--

Debes asignar `shift` a `- shift` cuando `encrypt` no es verdadero.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
