---
id: 681a2e35c406ade2f0a5f9b2
title: Pasos 24
challengeType: 20
dashedName: step-24
---

# --description--

Es hora de probar la función `encrypt`. Usando los mismos argumentos, reemplaza tu llamada a `caesar` con una llamada a `encrypt`. Verás el mismo resultado en la terminal.

# --hints--

Tu variable `encrypted_text` debe tener el valor de `encrypt('freeCodeCamp', 3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = encrypt('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift

    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)

def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)
--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```
