---
id: 587d78b1367417b2b2512b0a
title: استخدام صورة شبكية العين لشاشات دقة الوضوح العالية
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

مع زيادة الأجهزة المتصلة بالإنترنت، تختلف أحجامها ومواصفاتها، وقد تختلف الشاشات التي تستخدمها داخليًا وخارجيًا. كثافة البكسل هي جانب قد يختلف من جهاز لآخر، وتُعرف هذه الكثافة باسم بكسل لكل إنش (PPI) أو نقاط لكل إنش (DPI). أشهر هذه الشاشات هي تلك المعروفة باسم "شاشة شبكية العين" في أحدث دفاتر MacBook Pro من Apple، ومؤخرًا في حواسيب iMac. بسبب الفرق في كثافة البكسل بين شاشات "شبكية العين" و"غير شبكية العين"، قد تبدو بعض الصور التي لم تُصمم لشاشات دقة وضوح عالية "مبكسلة" عند عرضها على شاشة دقة وضوح عالية.

أسهل طريقة لجعل صورك تظهر بشكل صحيح على شاشات دقة الوضوح العالية، مثل "شاشة شبكية العين" في MacBook Pro، هي تحديد قيم `width` و`height` لتكون نصف قيم الملف الأصلي فقط. إليك مثالًا لصورة تستخدم نصف الارتفاع والعرض الأصليين فقط:

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

حدد `width` و`height` في وسم `img` لتكون نصف القيم الأصلية. في هذه الحالة، كل من `height` و`width` الأصليين هما `200px`.

# --hints--

يجب أن يكون `width` في وسم `img` الخاص بك 100 بكسل.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

يجب أن يكون `height` في وسم `img` الخاص بك 100 بكسل.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
