---
id: 587d78b1367417b2b2512b0c
title: اجعل الطباعة متجاوبة
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

بدلًا من استخدام `em` أو `px` لتحديد حجم النص، يمكنك استخدام وحدات إطار العرض للطباعة المتجاوبة. وحدات إطار العرض، مثل النسب المئوية، هي وحدات نسبية، لكنها تعتمد على عناصر مختلفة. وحدات إطار العرض تعتمد على أبعاد إطار العرض (العرض أو الارتفاع) للجهاز، بينما النسب المئوية تعتمد على حجم عنصر الحاوية الأصلية.

الوحدات الأربعة المختلفة لإطار العرض هي:

<ul><li><code>vw</code> (عرض إطار العرض): <code>10vw</code> تعني 10% من عرض إطار العرض.</li><li><code>vh</code> (ارتفاع إطار العرض): <code>3vh</code> تعني 3% من ارتفاع إطار العرض.</li><li><code>vmin</code> (الحد الأدنى لإطار العرض): <code>70vmin</code> تعني 70% من أصغر بُعد في إطار العرض (ارتفاع أو عرض).</li><li><code>vmax</code> (الحد الأقصى لإطار العرض): <code>100vmax</code> تعني 100% من أكبر بُعد في إطار العرض (ارتفاع أو عرض).</li></ul>

إليك مثالًا يحدد وسم `body` ليكون 30% من عرض إطار العرض.

```css
body {
  width: 30vw;
}
```

# --instructions--

حدد `width` لوسم `h2` ليكون 80% من عرض إطار العرض و`width` للفقرة ليكون 75% من أصغر بُعد في إطار العرض.

# --hints--

يجب أن يكون لوسم `h2` الخاص بك `width` بقيمة 80vw.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

يجب أن يكون لوسم `p` الخاص بك `width` بقيمة 75vmin.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
