---
id: 587d78b1367417b2b2512b09
title: جعل الصورة متجاوبة
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

جعل الصور متجاوبة باستخدام CSS بسيط جدًا. فقط أضف هذه الخصائص إلى الصورة:

```css
img {
  max-width: 100%;
  height: auto;
}
```

تضمن الخاصية `max-width` بقيمة `100%` ألا تصبح الصورة أعرض من الحاوية التي بداخلها، وتُبقي الخاصية `height` بقيمة `auto` على نسبة العرض إلى الارتفاع الأصلية للصورة.

# --instructions--

أضف قواعد التنسيق إلى صنف `responsive-img` لجعل الصورة متجاوبة. يجب ألا تصبح أعرض من الحاوية التي بداخلها (في هذه الحالة، نافذة المعاينة) ويجب أن تحافظ على نسبة العرض إلى الارتفاع الأصلية. بعد إضافة الشفرة، غيّر حجم نافذة المعاينة لترى كيف تتصرف الصور.

# --hints--

يجب أن يحتوي صنف `responsive-img` على خاصية `max-width` مضبوطة على `100%`.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

يجب أن يحتوي صنف `responsive-img` على خاصية `height` مضبوطة على `auto`.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
