---
id: 587d78b0367417b2b2512b08
title: إنشاء استعلام وسائط
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

استعلامات الوسائط هي تقنية جديدة أُدخلت في CSS3 تُغيّر عرض المحتوى بناءً على أحجام إطارات العرض المختلفة. إطار العرض هو المساحة المرئية للمستخدم في صفحة ويب، ويختلف حسب الجهاز المستخدم للوصول إلى الموقع.

تتكون استعلامات الوسائط من نوع وسائط، وإذا تطابق هذا النوع مع نوع الجهاز الذي يُعرض عليه المستند، تُطبق الأنماط. يمكنك وضع عدد غير محدود من المحددات والأنماط داخل استعلام الوسائط الخاص بك.

إليك مثال على استعلام وسائط يعرض المحتوى عندما يكون عرض الجهاز أقل من أو يساوي `100px`:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

واستعلام الوسائط التالي يعرض المحتوى عندما يكون ارتفاع الجهاز أكثر من أو يساوي `350px`:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

تذكّر، يُطبق CSS داخل استعلام الوسائط فقط إذا تطابق نوع الوسائط مع نوع الجهاز المستخدم.

# --instructions--

أضف استعلام وسائط، بحيث يكون للوسم `p` حجم خط `font-size` يساوي `10px` عندما يكون ارتفاع الجهاز أقل من أو يساوي `800px`.

# --hints--

يجب إعلان استعلام `@media` للأجهزة التي يكون فيها `height` أقل من أو يساوي `800px`.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

يجب أن يكون لعنصر `p` الخاص بك `font-size` بقيمة `10px` عندما يكون `height` الجهاز أقل من أو يساوي `800px`.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

يجب أن يكون لعنصر `p` حجم `font-size` ابتدائي قدره `20px` عندما يكون `height` الجهاز أكثر من `800px`.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
