---
id: 5dfb6a35eacea3f48c6300b4
title: ステップ 24
challengeType: 0
dashedName: step-24
---

# --description--

A figure caption (`figcaption`) element is used to add a caption to describe the image contained within the `figure` element.

下記は `A cute cat` というキャプションを持つ `figcaption` 要素の例です。

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

After the image nested in the `figure` element, add a `figcaption` element with text set to:

`Cats love lasagna.`

# --hints--

`figcaption` 要素には開始タグが必要です。 Opening tags have the following syntax: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

`figcaption` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 要素は `figure` 要素内にネストされている必要があります。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

ラザニアの `img` 要素は `figure` 要素内にネストされている必要があります。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

`figure` 要素内にネストされた `figcaption` 要素は、`img` 要素の下に置く必要があります。 You have them in the wrong order.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Your `figcaption` element's text should be `Cats love lasagna.` You have either omitted the text or have a typo.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

