---
id: 5ef9b03c81a63668521804dc
title: ステップ 46
challengeType: 0
dashedName: step-46
---

# --description--

複数の選択肢から回答を 1 つだけ選んでほしい質問には、ラジオボタンを使うことができます。

Here is an example of a radio button with the text set as `cat`:

```html
<input type="radio"> cat
```

`input` 要素は空要素であることを忘れないでください。

Before the text input, add a radio button with the text set as:

`Indoor`

# --hints--

ラジオボタン用の `input` 要素を 1 つ作成する必要があります。 構文を確認してください。

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

`input` 要素には開始タグが必要ですが、終了タグは必要ありません。

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

ラジオボタン用の `input` 要素は 1 つだけ追加してください。 余分なものは削除してください。

```js
assert(document.querySelectorAll('form > input'), 2);
```

新しい `input` 要素は既存の `input` の上にあり、`type` 属性が `text` に設定されている必要があります。 順番に誤りがあるようです。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

新しい `input` 要素に `type` 属性がありません。 開始タグのタグ名の後にスペースがあることを確認してください。

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

新しい `input` 要素には `type` 属性が 1 つだけ必要です。 余分なものは削除してください。

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

新しい `input` 要素には `type` 属性があり、値が `radio` に設定されている必要があります。 値が設定されていないか、誤字脱字があります。 属性値は引用符で囲むことを忘れないでください。

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

新しい `input` 要素の `type` 属性は `radio` に設定されていますが、属性の値は常に引用符で囲むことが推奨されています。

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` ボタンのテキスト `Indoor` は、ラジオボタンの前ではなく、後に配置してください。

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

`Indoor` というテキストは `radio` ボタンの右側に配置してください。 テキストが設定されていないか、誤字脱字があります。

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

