---
id: 5efc4f528d6a74d05e68af74
title: ステップ 58
challengeType: 0
dashedName: step-58
---

# --description--

`input` 要素用のテキストと要素自体を関連付ける方法はもう一つあります。 テキストを `label` 要素にネストして、`input` 要素の `id` 属性と同じ値を持つ `for` 属性を追加することです。

次のような `input` 要素があるとします:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

上記の `input` 要素に `label` 要素を使ってテキストを関連付けると、次のようになります:

```html
<label for="breakfast">Breakfast</label>
```

`Loving` というテキストだけを `label` 要素内にネストして、適切な `for` 属性を追加することにより、テキスト `Loving` とチェックボックスを関連付けてください。

# --hints--

チェックボックスが削除・変更されていないことを確認してください。 チェックボックスのコードは、開始時のコードから変更しないでください。

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

変更前と変わらず、チェックボックスが `id` 属性を持ち、値が `loving` に設定されている必要があります。  期待される値は `--fcc-expected--` ですが、見つかった値は `--fcc-actual--` でした。

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

チェックボックスの要素の後に、新しく `label`要素を追加してください。 開始タグと終了タグが両方あることを確認してください。

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

`label` 要素はチェックボックスの後に配置してください。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

`label` 要素はテキスト `Loving` を持つようにしてください。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

テキスト `Loving` は、新しい `label` 要素内だけに存在する必要があります。 `label` 要素の外側に重複するテキストがないように注意してください。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

`label` 要素に `for` 属性がありません。 `label` の開始タグのタグ名と `for` 属性の間には、必ず半角スペースを入れてください。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

`label` 要素には値が `loving` の `for` 属性が必要です。  期待される値は `--fcc-expected--` ですが、見つかった値は `--fcc-actual--` でした。

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

チェックボックスと新しい `label` 要素の間にはスペースが必要です。

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

