---
id: 671141feba228a35cefba82d
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

テキストの `cute cats` の部分を、`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg` にリンクするアンカー要素にしてください。

# --hints--

1 つ目の `p` 要素内に、新しいアンカー要素があるべきです。

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

アンカーのテキストは `cute cats` であるべきです。

```js
assert.strictEqual(document.querySelector('a')?.innerText, "cute cats");
```

新しいアンカー要素の `href` 属性は `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg` であるべきです。

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

`cute cats` という記述はコード内に 1 つだけにしてください。

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

`p` 要素のテキストは、引き続き `Everyone loves cute cats online!` のままであるべきです。

```js
assert.strictEqual(document.querySelector('p')?.innerText, "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
