---
id: 671141d8e32fe934c26fa1be
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

Add the words `See more` before the anchor element and `in our gallery` after the anchor element.

# --hints--

既存のアンカー要素を使用してください。新しいものを作成しないでください。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

The text inside the anchor element should not change, it must be `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

You should have the words `See more` before the anchor element. Make sure you include the space after the last word.

```js
assert.match(code, /See more <a href=/)
```

アンカー要素の後に `in our gallery` があることを確認してください。 Make sure you include the space before the first word.

```js
assert.match(code, /<\/a> in our gallery/)
```

`See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` がコードに含まれていることを確認してください。

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
