---
id: 671141d8e32fe934c26fa1be
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

アンカー要素の前に `See more` を、アンカー要素の後に `in our gallery` を追加してください。

# --hints--

既存のアンカー要素を使用してください。新しいものを作成しないでください。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

アンカー要素内のテキストは変更しないでください。`cat photos` でなければなりません。

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

アンカー要素の前に `See more` があることを確認してください。 最後の単語の後にスペースを入れてください。

```js
assert.match(code, /See more <a href=/)
```

アンカー要素の後に `in our gallery` があることを確認してください。 最初の単語の前にスペースを入れてください。

```js
assert.match(code, /<\/a> in our gallery/)
```

`See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` がコードに含まれていることを確認してください。

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
