---
id: 66bb6a9c2dd58b73cd759034
title: Construye una Aplicación de Boleta de Calificaciones
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Deberías tener una función llamada `getAverage` que tome un array de puntajes de prueba como parámetro y devuelva el puntaje promedio.
1. You should have a function named `getGrade` that takes a student score as a parameter and returns a string representing a letter grade based on the score. Here are the scores and their corresponding letter grades:

    | Score Range | Grade   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. Deberías tener una función llamada `hasPassingGrade` que tome un puntaje como parámetro y devuelva `true` o `false` dependiendo de si el puntaje corresponde a una calificación aprobatoria.
1. La función `hasPassingGrade` debería usar la función `getGrade` para obtener la calificación con letra y usarla para determinar si la calificación es aprobatoria. Una calificación aprobatoria es cualquier cosa diferente de `"F"`.
1. You should have a function named `studentMsg` that takes an array of scores and a student score as the parameters. The function should return a string with the format:

    - <code>"Promedio de la clase: average-goes-here. Tu calificación: grade-goes-here.  You passed the course."`, if the student passed the course.
    - <code>"Promedio de la clase: average-goes-here. Tu calificación: grade-goes-here.  You failed the course."`, if the student failed the course.

    Replace `average-goes-here` with the average of total scores and `grade-goes-here` with the student's grade.
    Use `getAverage` to get the average score and `getGrade` to get the student's grade.

# --hints--

Deberías tener una función llamada `getAverage`.

```js
assert.isFunction(getAverage);
```

Tu función `getAverage` debe devolver un número.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` debe devolver `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` debe devolver `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` debe devolver `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` debería devolver `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` debería devolver `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Tu función `getAverage`, que calcula el promedio, debería devolver la puntuación media de las evaluaciones.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Deberías tener una función llamada `getGrade`.

```js
assert.isFunction(getGrade);
```

Tu función `getGrade` debería devolver `"A+"` si el puntaje es `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Tu función `getGrade` debería devolver `"A"` si el puntaje está entre `90` y `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Tu función `getGrade` debe devolver `"B"` si la puntuación está entre `80` y `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Tu función `getGrade` debería devolver `"C"` si el puntaje está entre `70` y `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Tu función `getGrade` debe devolver `"D"` si la puntuación está entre `60` y `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Tu función `getGrade` debe devolver `"F"` si la puntuación está entre `0` y `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Deberías tener una función llamada `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Tu función `hasPassingGrade` debe devolver un valor booleano.

```js
assert.isBoolean(hasPassingGrade(100));
```

Tu función `hasPassingGrade` debe devolver `true` si la nota es `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Tu función `hasPassingGrade` debe devolver `false` si la nota es `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Tu función `hasPassingGrade` debe devolver `true` para todos las calificaciones pasadas.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Deberías tener una función llamada `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` debería devolver el siguiente mensaje: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` debería devolver el siguiente mensaje: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` debería devolver el siguiente mensaje: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` debería devolver el siguiente mensaje: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Tu función `studentMsg` debe devolver el mensaje correcto basado en la puntuación del estudiante y el promedio de la clase.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
