---
id: 673b567e3ba535dda140d278
title: Побудова системи голосування
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

У цій лабораторній роботі ви побудуєте систему голосування, яка використовує `Map` для створення опитування та `Set` для запобігання дублюванню голосів.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні ініціалізувати змінну `poll` новим об’єктом `Map`.

2. Ви повинні мати функцію `addOption`, яка приймає параметр `option`.

3. У функції `addOption`:

   - Якщо `option` ще не існує в опитуванні, його слід додати до опитування з порожнім `Set` як значенням для відстеження голосуючих. Також слід повернути повідомлення `Option "<option>" added to the poll.`

   - Якщо `option` вже існує, слід повернути повідомлення `Option "<option>" already exists.`

   - Якщо ви намагаєтеся додати порожній варіант, функція має повернути повідомлення `Option cannot be empty.`

4. Ви повинні мати функцію `vote`, яка приймає два параметри: `option` (варіант для голосування) і `voterId` (унікальний ідентифікатор голосуючого).

5. У функції `vote`:

   - Якщо `option` не існує в опитуванні, функція має повернути повідомлення `Option "<option>" does not exist.`

   - Якщо `option` існує, функція має перевірити, чи `voterId` вже голосував за цей `option`.

   - Якщо голосуючий уже голосував, функція має повернути повідомлення `Voter <voterId> has already voted for "<option>".`

   - Якщо голосуючий не голосував, `voterId` має бути доданий до `Set` `voters` для цього варіанту. Функція має повернути повідомлення `Voter <voterId> voted for "<option>".`

6. У вашому `poll` має бути щонайменше три варіанти.

7. У вашому `poll` має бути щонайменше три голоси.

8. Ви повинні мати функцію `displayResults`, яка повертає результати опитування у такому форматі:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Ви повинні мати змінну `poll`, ініціалізовану новим об’єктом `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Визначте функцію `addOption`, яка приймає параметр `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Визначте функцію `vote`, яка приймає два параметри: `option` і `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Визначте функцію `displayResults` для відображення результатів опитування.

```js
assert.isFunction(displayResults);
```

У вашому `poll` має бути щонайменше три варіанти.

```js
assert.isAtLeast(poll.size, 3);
```

У вашому `poll` має бути щонайменше три голоси.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Переконайтеся, що кожен варіант для голосування відповідає об’єкту `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` має повертати `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Додавання порожнього варіанту має повертати `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Якщо `Turkey` вже додано, `addOption("Turkey")` має повертати `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Якщо `Malaysia` існує серед варіантів для голосування, `vote("Malaysia", "traveler1")` має повертати `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` має оновлювати `Set` голосуючих для варіанту.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Якщо `traveler1` намагається проголосувати за `Algeria` знову, `vote("Algeria", "traveler1")` має повертати `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Дублікати голосів не повинні збільшувати розмір `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Якщо `Nigeria` відсутній серед варіантів для голосування, `vote("Nigeria", "traveler2")` має повертати `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Унікальний варіант може отримувати кілька голосів.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Якщо варіант `Turkey` отримує два голоси, варіант `Morocco` — один голос, а варіант `Spain` не отримує жодного голосу, `displayResults()` має повертати `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` має повертати результати у правильному форматі.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
