---
id: 684aaf9ec670c68d20efd0d0
title: بناء مدير إعدادات المستخدم
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

في هذا المختبر، ستبني مدير إعدادات المستخدم الذي يسمح للمستخدمين بإدارة إعداداتهم مثل السمة، اللغة، والإشعارات. ستنفذ دوالًا لإضافة، تحديث، حذف، وعرض إعدادات المستخدم.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب تعريف دالة باسم `add_setting` بمعلمتين تمثلان قاموسًا للإعدادات وصف مرتب يحتوي على زوج مفتاح-قيمة

1. يجب أن تقوم دالة `add_setting` بما يلي:
    
- تحويل المفتاح والقيمة إلى أحرف صغيرة.
   - إذا كان الإعداد موجودًا، أرجع الرسالة `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - إذا لم يكن الإعداد موجودًا، أضف زوج المفتاح-القيمة إلى القاموس المعطى للإعدادات وأرجع الرسالة `Setting '[key]' added with value '[value]' successfully!`.
   - يجب أن تحتوي الرسائل المرجعة على المفتاح والقيمة بأحرف صغيرة.

1. يجب تعريف دالة باسم `update_setting` بمعلمتين تمثلان قاموسًا للإعدادات وصف مرتب يحتوي على زوج مفتاح-قيمة.

1. يجب أن تقوم دالة `update_setting` بما يلي:
       
- تحويل المفتاح والقيمة إلى أحرف صغيرة.
   - إذا كان الإعداد موجودًا، حدّث قيمته في القاموس المعطى للإعدادات وأرجع الرسالة: `Setting '[key]' updated to '[value]' successfully!`
   - إذا لم يكن الإعداد موجودًا، أرجع الرسالة `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - يجب أن تحتوي الرسائل المرجعة على المفتاح والقيمة بأحرف صغيرة.


1. يجب تعريف دالة باسم `delete_setting` بمعلمتين تمثلان قاموسًا للإعدادات ومفتاحًا.

1. يجب أن تقوم دالة `delete_setting` بما يلي:
    
- تحويل المفتاح الممرر إلى أحرف صغيرة.
   - إذا كان الإعداد موجودًا، أزل زوج المفتاح-القيمة من القاموس المعطى للإعدادات وأرجع الرسالة `Setting '[key]' deleted successfully!`
   - إذا لم يكن الإعداد موجودًا، أرجع الرسالة `Setting not found!`
   - يجب أن تحتوي الرسائل المرجعة على المفتاح بأحرف صغيرة.

1. يجب تعريف دالة باسم `view_settings` بمعلمة واحدة تمثل قاموسًا للإعدادات.

1. يجب أن تقوم دالة `view_settings` بما يلي:
   
- أرجع الرسالة `No settings available.` إذا كان القاموس المعطى للإعدادات فارغًا.
   - إذا احتوى القاموس على أي إعدادات، أرجع سلسلة نصية تعرض الإعدادات. يجب أن تبدأ السلسلة بـ `Current User Settings:` تليها أزواج المفتاح-القيمة، كل منها في سطر جديد وبالمفتاح مكتوبًا بحرف كبير. على سبيل المثال، `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` يجب أن تُرجع:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. لاختبار الشفرة، يجب إنشاء قاموس باسم `test_settings` لتخزين بعض تفضيلات إعدادات المستخدم.

# --hints--

يجب إنشاء قاموس باسم `test_settings` وإضافة بعض القيم إليه.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

يجب تعريف دالة باسم `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

يجب أن تحتوي دالة `add_setting` على معلمتين.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

يجب أن تحول دالة `add_setting` المفتاح إلى أحرف صغيرة.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

يجب أن تحول دالة `add_setting` القيمة إلى أحرف صغيرة.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` يجب أن تُرجع رسالة الخطأ `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` يجب أن تضيف زوج مفتاح-قيمة جديد وتُرجع رسالة النجاح `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

يجب أن تضيف دالة `add_setting` زوج المفتاح-القيمة المعطى إلى القاموس بشكل صحيح.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

يجب تعريف دالة باسم `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

يجب أن تحتوي دالة `update_setting` على معلمتين.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

يجب أن تحول دالة `update_setting` المفتاح إلى أحرف صغيرة.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

يجب أن تحول دالة `update_setting` القيمة إلى أحرف صغيرة.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` يجب أن تحدّث مفتاحًا موجودًا وتُرجع رسالة النجاح `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` يجب أن تُرجع رسالة الخطأ `Setting 'volume' does not exist! Cannot update a non-existing setting.` عندما لا يكون المفتاح موجودًا.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

يجب أن تحدّث دالة `update_setting` زوج المفتاح-القيمة المعطى في القاموس بشكل صحيح.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

يجب تعريف دالة باسم `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

يجب أن تحتوي دالة `delete_setting` على معلمتين.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

يجب أن تحول دالة `delete_setting` المفتاح إلى أحرف صغيرة.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` يجب أن تزيل المفتاح الموجود وتُرجع رسالة النجاح `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

يجب أن تُرجع دالة `delete_setting` رسالة الخطأ `Setting not found!` عندما لا يكون المفتاح موجودًا.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

يجب أن تزيل دالة `delete_setting` المفتاح المعطى من القاموس بشكل صحيح.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

يجب تعريف دالة باسم `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

يجب أن تحتوي دالة `view_settings` على معلمة واحدة.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

يجب أن تُرجع دالة `view_settings` الرسالة `No settings available.` إذا كان القاموس المعطى فارغًا.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

يجب أن تُرجع دالة `view_settings` الإعدادات منسقة للقاموس غير الفارغ.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

يجب أن تكتب دالة `view_settings` الحرف الأول من كل اسم إعداد بحرف كبير.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

يجب أن تعرض دالة `view_settings` النتائج الصحيحة وتنهي بسطر جديد.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
