---
id: 680a3d4c065557260de2ff33
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

正如你可能从之前的课程中回忆的那样，在 Python 中，你通过在赋值操作符（`=`）左侧写变量名，右侧写要赋值的值来声明变量：

```py
variable_name = value
```

创建一个名为 `shift` 的变量，并为你的新变量赋值为 `5`。

# --hints--

你应该声明一个名为 `shift` 的变量。 注意将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你应该将值 `5` 赋值给你的 `shift` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
