---
id: 680a5e7f3d47715196b12ca0
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

现在你已经有了整个移位字母表，移除 `print(shifted_alphabet)` 调用。

# --hints--

你应该从你的代码中移除 `print(shifted_alphabet)`。

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
