---
id: 680a3d4c065557260de2ff33
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将创建一个实现凯撒密码的程序，这是一种通过字母表中移动字母来加密消息的方法。
首先创建一个名为 `shift` 的变量，并将值 `5` 赋值给你的新变量。

# --hints--

你应该声明一个名为 `shift` 的变量。确保将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你应该将值 `5` 赋值给你的 `shift` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
