---
id: 680a549f60f5a3fedfc0edd9
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

正如你从输出中看到的，移位字母表从字母 `f` 开始，因为 `shift` 的值是 `5`。但现在字母表的前五个字母——`a`、`b`、`c`、`d` 和 `e`——在移位字母表中缺失，所以你需要将它们添加到移位字母表的末尾。

`+` 操作符用于将两个或更多字串组合在一起，这个过程称为串接，如下所示：

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

修改现有的 `shifted_alphabet` 变量赋值：使用切片语法提取缺失的 `alphabet` 首部部分，并将其连接到 `alphabet[shift:]`。

作为提醒，`sentence[start:stop]` 返回 `sentence` 中从位置 `start`（包含）到 `stop`（不包含）的字符。

# --hints--

你不应该直接修改 `alphabet` 变量。

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

你应该在声明 `shifted_alphabet` 变量的同一行，将 `alphabet[shift:]` 和缺失的 `alphabet` 开头部分的串接赋值给它。使用 `shift` 指定切片的结束位置。

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
