---
id: 680a431692a5f863f726c438
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在本次研讨会中，你将构建一个 *Caesar cipher*。这是一种最简单的加密文本技术之一，其方法是用字母表中固定位置后面的字母替换明文的每个字母。例如，位移为 5 时，`a` 会被替换为 `f`，`b` 会被替换为 `g`，依此类推。

要实现这个密码，你需要创建一个从移位指示的位置开始的新版本的字母表。正如你在之前的课程中学到的，你可以使用字符串切片提取字串的一部分：

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

创建一个名为 `shifted_alphabet` 的变量，并使用切片语法将 `alphabet` 中从 `shift` 索引开始的部分赋值给它。然后，调用内置的 `print()` 函数在终端打印 `shifted_alphabet` 并查看结果。

# --hints--

你应该声明一个名为 `shifted_alphabet` 的新建变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

你应该将 `alphabet[shift:]` 赋值给 `shifted_alphabet`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

你应该调用 `print()` 函数并传入 `shifted_alphabet` 作为参数。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
