---
id: 680a836235461fce1ff53b3f
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

`translate()` 方法接受由 `str.maketrans()` 生成的翻译表作为参数。它在一个字串上调用，并返回原始字串的拷贝，其中的字符已根据翻译表被替换：

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

在 `text` 上调用 `translate()` 方法，将 `translation_table` 作为参数传入，并将结果赋值给名为 `encrypted_text` 的变量。

# --hints--

你应该有一个名为 `encrypted_text` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

你应该将 `text.translate(translation_table)` 赋值给 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
