---
id: 680a82806c4c6bc3567f2708
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

声明一个名为 `text` 的新建变量，并赋值为字串 `hello world`。这将是要加密的消息。

# --hints--

你应该声明一个名为 `text` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

你应该将字串 `hello world` 赋值给 `text`。记住要用单引号或双引号将字串括起来，并注意字母的大小写。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
