---
id: 680a41ef5bfdc059cfd9d4b5
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

声明另一个名为 `alphabet` 的变量，并将字串 `abcdefghijklmnopqrstuvwxyz` 赋值给该变量。

# --hints--

你应该声明一个名为 `alphabet` 的变量。 注意将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

你应该将字串 `abcdefghijklmnopqrstuvwxyz` 赋值给你的 `alphabet` 变量。记住使用单引号或双引号来引用该字串。

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
