---
id: 680a7e200400089967b8b1f4
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

`str.maketrans()` 方法接受两个长度相等的字串，并返回一个翻译表，该表将第一个字串的每个字符映射到第二个字串的对应字符。翻译表中的每个字符都以 Unicode 序数保存，该序数是唯一标识该字符的数字。

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

你将在课程后面学习更多关于 `str.maketrans()` 返回的结构类型。现在，创建一个将 `alphabet` 中的字符映射到 `shifted_alphabet` 中字符的翻译表，并将其赋值给名为 `translation_table` 的变量。

# --hints--

你应该声明一个名为 `translation_table` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

你应该调用 `str.maketrans()`，传入 `alphabet` 和 `shifted_alphabet` 作为参数，并将结果赋值给 `translation_table`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
