---
id: 680a3d4c065557260de2ff33
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将创建一个名为凯撒密码的程序，这是一种通过字母表中移动字母来加密消息的方法。

首先创建一个名为 `shift` 的变量，并为你的新变量赋值为 `5`。

# --hints--

你应该声明一个名为 `shift` 的变量。 注意将变量名放在行首。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

你应该将值 `5` 赋值给你的 `shift` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
