---
id: 6818b80519e6c6c47a8bb6e8
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

正如你在之前的课程中学到的，函数可以有参数，参数是在函数内可以引用的变量。这里有一个带有两个参数 `num1` 和 `num2` 的 `sum` 函数。

```py
def sum(num1, num2):
    print(num1 + num2)
```

要加密的消息和偏移量仍然在你的函数中硬编码。为你的函数添加两个名为 `text` 和 `shift` 的参数，并从 `caesar` 函数体中删除 `text` 和 `shift` 变量的声明。

# --hints--

你的 `caesar` 函数应有两个参数，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

你不应该在 `caesar` 函数体内声明 `text` 变量。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

你不应该在 `caesar` 函数体内声明 `shift` 变量。

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
