---
id: 6818e9c7c8dc5694e8f64fc3
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在，通过使用字串 `freeCodeCamp` 和数字 `3` 作为参数调用 `caesar` 函数来测试它。将函数调用赋值给名为 `encrypted_text` 的变量。

# --hints--

你应该在全局的作用域中有一个名为 `encrypted_text` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

你应该将 `caesar('freeCodeCamp', 3)` 赋值给 `encrypted_text` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
