---
id: 6818fcb250f34e62cab32c39
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

你可能已经注意到，尽管消息是加密的，但大写字符保持不变。这是因为翻译表格不包含大写字母。

要修复它，你需要使用 `upper()` 方法修改传递给你的 `str.maketrans()` 调用的字串：

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

通过将每个参数自身的大写版本连接到每个参数，更新你的 `str.maketrans()` 调用。

# --hints--

你应该通过将每个参数的大写版本连接到该参数本身，来修改你的 `str.maketrans()` 调用的参数。

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
