---
id: 6819c5c9faf2548e556163d5
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

正如你从之前的课程中回忆的，`not` 操作符是一元逻辑操作符，用于取反一个表达式：

```py
print(not True) # False
print(not False) # True
```

使用 `not` 操作符修正 `if` 语句的条件，以便当 `shift` *不是* 整数时，你的函数返回 `Shift must be an integer value.`。

# --hints--

你应该在你的 `if` 语句中使用 `not` 操作符来否定 `isinstance(shift, int)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
