---
id: 6819be9055828d65c813eec5
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

如果第一个参数是第二个参数的实例，`isinstance()` 函数返回 `True`，否则返回 `False`：

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

将你当前的 `if` 语句条件替换为 `isinstance()` 调用。将 `shift` 作为第一个参数传入，将 `int` 作为第二个参数传入。 

# --hints--

你的 `if` 语句应使用 `isinstance(shift, int)` 作为其条件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
