---
id: 6819bc0637b80256a33adccc
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

既然你已经实现了密码的基本功能，现在是添加一些验证的时候了。为此，你需要一个 `if` 语句。以下是 `if` 语句的语法提醒：

```py
if condition:
    # code to run when condition is true
```

在你的函数体开始处，创建一个 `if` 语句。暂时使用 `True` 作为条件，并在 `if` 语句体内返回字串 `Shift must be an integer value.`

# --hints--

你的 `caesar` 函数中应该有一个 `if` 语句。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

你的 `if` 语句应使用 `True` 作为其条件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

你应该从你的 `if` 语句中返回字串 `Shift must be an integer value.`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
