---
id: 6819c7beb4a8289c2b4f24d5
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

你的函数不应接受负数或零的 `shift`。因此，在你的第一个 `if` 语句之后，创建另一个 `if` 语句，检查 `shift` 是否小于 `1` 并返回字串 `Shift must be a positive integer.`。

# --hints--

你应该有第二个 `if` 语句，用来查看 `shift` 是否小于 `1`。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

你的第二个 `if` 语句应返回 `Shift must be a positive integer.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
