---
id: 6819cf38880d8cc9b9f6af7a
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

Python 允许你为函数中的参数指定默认值，从而创建一个可以用更少或不带参数调用的函数。以下是如何创建一个带有默认值的 `name` 参数的函数：

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

为你的函数添加第三个名为 `encrypt` 的参数，并将其默认值设为 `True`。 

# --hints--

你应该为你的函数添加第三个名为 `encrypt` 的参数，默认值为 `True`。不要修改现有的参数。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
