---
id: 681a1d60e037127e7c8e0230
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

你将使用上一步添加的参数来确定函数是否应该加密传递给它的文本（默认行为，`encrypt=True`），或者是否应该解密加密的消息。

创建一个 `if` 语句，用于查看 `encrypt` 是否不为真。在新的 `if` 语句中，将 `shift` 设置为 `- shift`。这是为了启用相对于加密进程方向相反的位移。

# --hints--

当 `encrypt` 不为真时，你应该将 `shift` 设置为 `- shift`。

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
