---
id: 681a44f4c3235f7d8f428545
title: 步骤 25
challengeType: 20
dashedName: step-25
---

# --description--

现在你将测试 `decrypt` 函数。将赋值给 `encrypted_text` 的值替换为以下字串，该字串表现一个待解密的消息：`Pbhentr vf sbhaq va hayvxryl cynprf.`

然后，声明一个名为 `decrypted_text` 的变量，并将其赋值为调用 `decrypt`，第一个参数为 `encrypted_text`，第二个参数为偏移量 `13`。

最后，在终端上打印 `decrypted_text`。至此，凯撒密码完成。

# --hints--

你应该将 `Pbhentr vf sbhaq va hayvxryl cynprf.` 赋值给 `encrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

你应该有一个名为 `decrypted_text` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

你应该将 `decrypt(encrypted_text, 13)` 赋值给 `decrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

你应该打印 `decrypted_text`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
