---
id: 681a2e35c406ade2f0a5f9b2
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

是时候测试 `encrypt` 函数了。使用相同的参数，将你的 `caesar` 调用替换为对 `encrypt` 的调用。你将在终端看到相同的输出。

# --hints--

你的 `encrypted_text` 变量应具有 `encrypt('freeCodeCamp', 3)` 的值。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = encrypt('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--
print(encrypted_text)
```
