---
id: 681a2a6cf95e82ca30866364
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

声明两个名为 `encrypt` 和 `decrypt` 的函数，均带有 `text` 和 `shift` 参数。

你将使用 `encrypt` 进行加密进程，使用 `decrypt` 进行解密，并用描述性名称标记这两个操作。

从你的两个新建函数中返回一个 `caesar` 调用，传入 `text` 和 `shift`，但确保对 `decrypt` 函数也传入 `False` 作为第三个参数。

# --hints--

你应该有一个名为 `encrypt` 的函数，带有两个参数，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

你的 `encrypt` 函数应返回 `caesar(text, shift)`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

你应该有一个名为 `decrypt` 的函数，带有两个参数，`text` 和 `shift`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

你的 `decrypt` 函数应返回 `caesar(text, shift, False)`。

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
