---
id: 660ef5105b8ba095307a0e50
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер знову виведіть у консоль вашу змінну `character`. Ви повинні побачити рядок `"Hello"`, а потім рядок `"World"` у консолі.

# --hints--

Ви повинні використати `console.log` вдруге.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Ви повинні вивести `character` у консоль вдруге.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


Ваш новий `console.log()` має йти після повторного присвоєння.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
