---
id: 660ee6e3a242da6bd579de69
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Одна з найважливіших концепцій у програмуванні — це змінні. <dfn>Змінна</dfn> вказує на конкретну адресу в пам’яті, де зберігається значення. Змінним дають ім’я, яке можна використовувати у всьому коді для доступу до цього значення.

Оголошення змінної означає присвоєння їй імені. У JavaScript це часто роблять за допомогою ключового слова `let`. Наприклад, ось як ви оголосите змінну `hello`:

```js
let hello;
```

Імена змінних підпорядковуються певним правилам: вони можуть містити літери, числа, знаки долара та підкреслення, але не можуть містити пробілів і не повинні починатися з числа.

Використайте ключове слово `let`, щоб оголосити змінну з ім’ям `character`.

_Примітка_: Загальноприйнято завершувати інструкції в JavaScript крапкою з комою. `;`

# --hints--

Ви повинні використовувати `let` у своєму коді.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Ви повинні використовувати `character` у своєму коді.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Ви повинні використовувати `let`, щоб оголосити змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Ваше оголошення має закінчуватися крапкою з комою.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
