---
id: 660ef31a5be625914a0102cd
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Коли змінну оголошують за допомогою ключового слова `let`, ви можете пізніше <dfn>перепризначити</dfn> (або змінити значення) цієї змінної. У цьому прикладі значення `programmer` змінюється з `"Naomi"` на `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Зверніть увагу, що при перепризначенні змінної ви **не** використовуєте ключове слово `let` знову.

Після вашого `console.log` призначте значення `"World"` змінній `character`.

# --hints--

Ви повинні двічі використати `character` у своєму коді.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Ви повинні використати оператор присвоєння, щоб перепризначити `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Ви повинні використати рядок `"World"` у своєму коді.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Ваша змінна `character` має мати значення `"World"`.

```js
assert.equal(character, "World");
```

Ваше перепризначення не повинно використовувати `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
