---
id: 660ef6355e8f5a9e67fe5f46
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

За допомогою типу даних `number` ви можете виконувати математичні операції, наприклад додавання. Спробуйте надрукувати `count + 1` на консолі.

# --hints--

Отримайте доступ до `console` у коді.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Отримайте доступ до `.log()` з `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Виведіть `count + 1` на консолі.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
