---
id: 660eebd83100d37862268781
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Вашій змінній `character` наразі не призначене значення. Ви можете призначити значення, використовуючи оператор присвоєння <dfn>assignment</dfn> `=`. Наприклад:

```js
let hello = "Hello";
```

Присвоєння значення змінній у момент її оголошення називається <dfn>ініціалізацією</dfn>.

Ініціалізуйте вашу змінну `character`, призначивши їй значення `"Hello"` під час оголошення.

# --hints--

Ви повинні використати оператор присвоєння `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Ви повинні використати рядок `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Ви повинні використати подвійні лапки для вашого рядка `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Ви повинні призначити `"Hello"` вашій змінній `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
