---
id: 660ef6355e8f5a9e67fe5f46
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

За допомогою типу даних `number` ви можете виконувати математичні операції, наприклад додавання. Спробуйте вивести `count + 1` у консоль.

# --hints--

Ви повинні отримати доступ до `console` у своєму коді.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Ви повинні отримати доступ до `.log()` `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Ви повинні вивести `count + 1` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
