---
id: 660ef19b95d3308e7dd31bb6
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>Консоль</dfn> дозволяє друкувати та переглядати JavaScript вивід. Ви можете надіслати інформацію на консоль за допомогою `console.log()`. Наприклад, цей код надрукує на консоль `"Naomi"`:

```js
let developer = "Naomi";
console.log(developer);
```

Наведений вище код отримує доступ до змінної `developer` з її назвою у `console.log()`. Зауважте, що значення в дужках — це значення, яке буде надруковано.

Надрукуйте значення змінної `character` на консолі. Потім, натисніть кнопку «Консоль», щоб бачити JavaScript консоль.

# --hints--

Не змінюйте ініціалізацію змінної `character`.

```js
assert.strictEqual(character, 'Hello');
```

You should access the `console` in your code.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

You should access the `log` property of the `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

You should use parentheses to call the `.log()` method.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

You should print the `character` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
