---
id: 660ef857f2806aa626d29d17
title: Крок 17
challengeType: 1
dashedName: step-17
---

# --description--

Ви також можете виконувати віднімання (`-`), множення (`*`) і ділення (`/`). Не соромтеся експериментувати з операторами та числами у вашому `console.log`. Коли будете готові рухатися далі, видаліть `console.log`.

# --hints--

У вашому коді не має бути `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Значення вашої змінної `character` має залишатися `"Hello"`.

```js
assert.equal(character, "Hello");
```

Значення вашої змінної `count` має залишатися `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
