---
id: 660ef5c1904955978a986a5c
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Використайте `let` для оголошення змінної `count`. Призначте їй числове `8`. Коли ви використовуєте числове значення, не використовуйте лапки. Наприклад:

```js
let money = 100;
```

# --hints--

У вашому коді має бути друге ключове слово `let`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Ви маєте використати `let` для оголошення змінної `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Ви маєте призначити числове значення `8` вашій змінній `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
