---
id: 660ef55dd468079679ee0092
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер ви готові оголосити свою наступну змінну. Видаліть обидві інструкції `console.log` і переназначення `character`.

Також видаліть змінну `secondCharacter`, але залиште ініціалізацію `character` без змін.

# --hints--

У вашому коді не повинно бути інструкцій `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Ви не повинні переназначати значення `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Ваша змінна `character` має мати значення `"Hello"`.

```js
assert.equal(character, "Hello");
```

У вас не повинно бути змінної `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
