---
id: 660ef19b95d3308e7dd31bb6
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

<dfn>Консоль</dfn> дозволяє виводити та переглядати вихідні дані JavaScript. Ви можете надсилати інформацію до консолі, використовуючи `console.log()`. Наприклад, цей код виведе `"Naomi"` у консоль:

```js
let developer = "Naomi";
console.log(developer);
```

Код вище отримує доступ до змінної `developer` за її ім’ям у `console.log()`. Зверніть увагу, що значення між круглими дужками — це значення, яке буде виведено.

Виведіть значення змінної `character` у консоль. Потім натисніть кнопку «Console», щоб переглянути консоль JavaScript.

# --hints--

Ініціалізацію змінної `character` змінювати не слід.

```js
assert.strictEqual(character, 'Hello');
```

У вашому коді слід отримати доступ до `console`.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Вам слід отримати доступ до властивості `log` об’єкта `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Використайте круглі дужки, щоб викликати метод `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Виведіть змінну `character` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
