---
id: 660ef0f7c4b8e68ccd1f0786
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

У JavaScript є сім примітивних типів даних, і рядок (String) — один із них. У JavaScript <dfn>рядок</dfn> представляє послідовність символів і може бути взятий у одинарні (`'`) або подвійні (`"`) лапки.

Зверніть увагу, що рядки є <dfn>незмінними</dfn>, тобто після створення їх не можна змінити. Змінній все одно можна призначити інше значення.

Змініть свій рядок `"Hello"`, щоб він використовував одинарні лапки.

# --hints--

У вашому коді не повинно бути подвійних лапок.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Використовуйте одинарні лапки для вашого рядка `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

У коді все ще слід використовувати `let`.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Ви все ще повинні оголосити змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Значення змінної `character` має залишатися рядком `"Hello"`.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
