---
id: 660f207334fabaeac3269c38
title: Крок 41
challengeType: 1
dashedName: step-41
---

# --description--

Щоб маніпулювати рядком `result`, ви використаєте інший тип циклу. Зокрема, цикл `for...of`, який ітерує по кожному елементу ітерованого об’єкта та тимчасово призначає його змінній.

Синтаксис циклу `for...of` виглядає так:

```js
for (const value of iterable) {

}
```

Зверніть увагу, що можна використовувати `const`, бо змінна існує лише під час однієї ітерації, а не протягом усього циклу.

Створіть цикл `for...of`, щоб ітерувати по вашому масиву `rows`, призначаючи кожне значення змінній `row`.

# --hints--

Вам слід використати інше ключове слово `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Вам слід оголосити змінну `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Вам слід використати `const` для оголошення змінної `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Ваш цикл `for...of` має оголосити змінну `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Ваша змінна `row` має бути вилучена з `rows` за допомогою ключового слова `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Тіло вашого циклу `for...of` має бути порожнім.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
