---
id: 660f540c2176ea5dec01306d
title: Крок 118
challengeType: 1
dashedName: step-118
---

# --description--

Чудова робота! Експериментуйте з різними значеннями для ваших змінних `character`, `count` та `inverted`.

Коли будете готові перейти до наступного проєкту, встановіть `character` у `"!"`, `count` у `10` та `inverted` у `false`, щоб продовжити.

Вітаємо з завершенням вашого першого проєкту на JavaScript!

# --hints--

Ви повинні встановити `character` у `"!"`.

```js
assert.equal(character, "!");
```

Ви повинні встановити `count` у `10`.

```js
assert.equal(count, 10);
```

Ви повинні встановити `inverted` у `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
