---
id: 660f1b6e60bd9edf902c81fd
title: Крок 36
challengeType: 1
dashedName: step-36
---

# --description--

Ваша інструкція <dfn>ітерації</dfn> скаже циклу, що робити з ітератором після кожного запуску.

Коли ви переназначаєте змінну, ви можете використати цю змінну, щоб посилатися на попереднє значення перед переназначенням. Це дозволяє робити такі речі, як додати три до існуючого числа. Наприклад, `bees = bees + 3` збільшить значення `bees` на три.

Використайте цей синтаксис, щоб замінити ваш рядок `"iteration"` на інструкцію переназначення, яка збільшує `i` на одиницю.

# --hints--

Ви повинні додати одиницю до вашої змінної `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Ви повинні призначити `i + 1` назад вашій змінній `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Ваш цикл `for` повинен збільшувати `i` на `1` під час кожної ітерації.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
