---
id: 660f4934fb48f63abd5ae371
title: Крок 95
challengeType: 1
dashedName: step-95
---

# --description--

Використання `done` для відстеження кількості згенерованих рядків працює, але ви насправді можете трохи спростити логіку.

Масиви мають спеціальну властивість `length`, яка дозволяє побачити, скільки значень, або <dfn>елементів</dfn>, міститься в масиві. Ви можете отримати доступ до цієї властивості за допомогою синтаксису на кшталт `myArray.length`.

Зверніть увагу, що `rows.length` у виклику `padRow` призведе до помилки зміщення на одиницю, бо `done` збільшується *перед* викликом.

Оновіть вашу умову, щоб перевіряти, чи `rows.length` менше за `count`.

# --hints--

Ваш цикл `while` має перевіряти, чи `rows.length` менше за `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
