---
id: 660f4a1472f8e63d76162ce5
title: Крок 99
challengeType: 1
dashedName: step-99
---

# --description--

Що як зробити вашу піраміду догори дригом, або <dfn>інвертованою</dfn>? Час спробувати!

Почніть зі створення нового `for` циклу. Оголосіть свій ітератор `i` і призначте йому значення `count`, потім використайте булевий `false` для ваших умовних і ітераційних інструкцій.

# --hints--

Ваш код має містити `for` цикл.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Ваш `for` цикл має ініціалізувати `i` значенням `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Ваш `for` цикл має використовувати `false` як умову.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


Ваш `for` цикл має використовувати `false` як ітерацію.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
