---
id: 660f4cffb1459d45e34902d1
title: Крок 106
challengeType: 1
dashedName: step-106
---

# --description--

Ви насправді можете побудувати перевернуту піраміду, не потрібно робити "зворотній" цикл, як раніше.

Для цього вам потрібно вивчити кілька нових методів масиву. Почніть із використання `const` для оголошення змінної `numbers`. Призначте їй масив з елементами `1`, `2` та `3`. Потім виведіть у консоль масив `numbers`.

# --hints--

Ви повинні використати `const` для оголошення змінної `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Ваша змінна `numbers` має бути масивом.

```js
assert.isArray(numbers);
```

Ваш масив `numbers` повинен містити елементи `1`, `2` та `3` у такому порядку.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Ви повинні вивести у консоль свій масив `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
