---
id: 660f3ce51f70571e1c5227c8
title: Крок 76
challengeType: 1
dashedName: step-76
---

# --description--

Коментарі можуть бути корисними для пояснення, чому ваш код використовує певний підхід, або для залишення нотаток на майбутнє.

У JavaScript ви можете використати `//` для залишення однолінійного коментаря у своєму коді.

Додайте однолінійний коментар над вашим циклом `for`, щоб нагадати собі змінити код на інший тип циклу.

# --hints--

Ви повинні починати однолінійний коментар з `//`.

```js
assert.match(code, /\/\//);
```

Ваш однолінійний коментар має містити щонайменше п’ять символів.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
