---
id: 660f09a2694b59c3a10ee304
title: Крок 23
challengeType: 1
dashedName: step-23
---


# --description--

Поки що видаліть свій перший виклик `console.log` і присвоєння `rows[rows.length - 1]`. Залиште другий виклик `console.log` `rows` на потім.

# --hints--

Ви повинні видалити інструкцію `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Ви повинні видалити переназначення `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Ви не повинні видаляти інструкцію `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
