---
id: 660f20473aef47e9b8c9afc6
title: Крок 40
challengeType: 1
dashedName: step-40
---

# --description--

Додайте інструкцію для виведення значення `result`. Залежно від консолі, яку ви використовуєте, ви можете нічого не побачити у виводі.

# --hints--

Ви повинні додати виклик `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Ви повинні вивести значення вашої змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
