---
id: 660f23b53db70af0f2620e78
title: Крок 43
challengeType: 1
dashedName: step-43
---

# --description--

Тепер усі ваші числа з’являються в одному рядку. Так не вийде створити піраміду.

Вам потрібно додати новий рядок до кожного рядка. Однак натискання клавіші Enter для вставлення розриву рядка між лапками в JavaScript призведе до помилки синтаксису. Замість цього потрібно використати спеціальну керовану послідовність `\n`, яка інтерпретується як новий рядок при виведенні рядка. Наприклад:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Використайте другий оператор додавання, щоб додати новий рядок після значень `result` і `row`.

# --hints--

Ви повинні використати керовану послідовність `\n`. Пам’ятайте, що це має бути рядок, тому вона обгорнута в лапки.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Ви повинні конкатенувати вашу змінну `row` зі змінною `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Ви повинні конкатенувати керовану послідовність `\n` зі змінною `row`.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Ви повинні призначити всю конкатенацію назад змінній `result`.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
