---
id: 660f0ee51d7460ce88cd248d
title: Крок 29
challengeType: 1
dashedName: step-29
---

# --description--

Ви очікували побачити `4` у консолі? `.push()` повертає нову довжину масиву після додавання значення, яке ви йому передаєте.

Важливо розуміти, які значення повертає метод. Приділіть час, щоб поекспериментувати з `.push()` і `.pop()`. Коли будете готові, видаліть усі свої виклики `.push()` і `.pop()`, а також інструкції `console.log`.

# --hints--

У вас не має бути виклику `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

У вас не має бути виклику `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

У вас не має бути жодних інструкцій log.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

У вас не має бути змінної `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

У вас не має бути змінної `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
