---
id: 660f505d02b2bd513a1c3468
title: Крок 111
challengeType: 1
dashedName: step-111
---

# --description--

Ваша піраміда більше не перевернута. Це тому, що ви додаєте нові рядки в **кінець** масиву.

Оновіть тіло циклу, щоб додавати нові рядки на початок масиву.

# --hints--

Ви повинні використовувати метод `unshift` об’єкта `rows`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Ви повинні передати виклик `padRow()` як аргумент для методу `.unshift()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Ви повинні передати `i` як перший аргумент виклику `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Ви повинні передати `count` як другий аргумент виклику `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
