---
id: 660f033cf051ebb50ea3bf48
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

У програмуванні часто потрібно працювати з великою кількістю даних. Існує багато структур даних, які допомагають організувати та керувати вашими даними. Однією з найпростіших структур даних є <dfn>масив</dfn>.

<dfn>Масив</dfn> — це непримітивний тип даних, який може містити послідовність значень. Непримітивні типи даних відрізняються від примітивних тим, що можуть зберігати більш складні дані. Примітивні типи даних, як-от рядки та числа, можуть містити лише одне значення одночасно.

Масиви позначаються за допомогою квадратних дужок (`[]`). Ось приклад змінної зі значенням порожнього масиву:

```js
let array = [];
```

Оголосіть змінну `rows` і призначте їй порожній масив.

# --hints--

Ви повинні мати змінну `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Використайте `let` для оголошення вашої змінної `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Призначте порожній масив вашій змінній `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
