---
id: 660f2b6fd54ac1fc142804dd
title: Крок 47
challengeType: 1
dashedName: step-47
---

# --description--

Логіка форматування цієї піраміди, ймовірно, стане складною, тож це чудовий момент, щоб винести цей код у функцію.

<dfn>Функція</dfn> — це блок коду, який можна повторно використовувати у вашому застосунку. Функції оголошуються за таким синтаксисом:

```js
function name(parameter) {

}
```

Ключове слово `function` повідомляє JavaScript, що змінна `name` буде функцією. `parameter` — це змінна, яка представляє значення, що передається у функцію під час її виклику. Функція може мати стільки <dfn>параметрів</dfn>, скільки ви хочете, або ж не мати їх зовсім. Як і в циклі `for`, простір між фігурними дужками — це <dfn>тіло функції</dfn>.

Оголосіть функцію `padRow`. Поки що не створюйте жодних параметрів. Тіло функції має бути порожнім. Пам’ятайте, що для іменування потрібно використовувати camel case.

# --hints--

Ви повинні використати ключове слово `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Ви повинні оголосити функцію `padRow`.

```js
assert.isFunction(padRow);
```

Ваша функція `padRow()` не повинна мати параметрів.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Тіло вашої функції `padRow()` має бути порожнім.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
