---
id: 660f2fbd45b520046cac68e8
title: Крок 65
challengeType: 1
dashedName: step-65
---

# --description--

Пам’ятаєте, на одному з попередніх кроків ви вивчали повернені значення. Функція може <dfn>повертати</dfn> значення, щоб ваш застосунок міг його окремо обробити.

У функції ключове слово `return` використовується для вказівки поверненого значення. Наприклад, ця функція поверне значення, передане першому параметру:

```js
function name(parameter) {
  return parameter;
}
```

Використайте ключове слово `return`, щоб повернути значення змінної `character`, повторене `rowNumber` разів.

# --hints--

Ви повинні використати метод `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Ви повинні використати метод `.repeat()` для вашої змінної `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Ви повинні передати `rowNumber` у ваш виклик `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Ви повинні використати ключове слово `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Ви повинні повернути результат виклику `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
