---
id: 660f255022991ef34ed0ee88
title: Крок 44
challengeType: 1
dashedName: step-44
---

# --description--

Виведення чисел не створить візуально привабливої піраміди. Тепер, коли ви виводите відформатований вміст свого масиву `rows`, настав час оновити ваш початковий цикл.

Замість того, щоб додавати `i` до масиву, додайте значення вашої змінної `character`.

# --hints--

Вам більше не потрібно додавати вашу змінну `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Вам слід додавати вашу змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
