---
id: 660f280dda5040f707c76b4a
title: Крок 45
challengeType: 1
dashedName: step-45
---

# --description--

Тепер у вас є серія символів `#`, але форма піраміди все ще відсутня. На щастя, змінна `i` представляє поточний номер "рядка" у вашому циклі, що дозволяє використати її для створення структури, схожої на піраміду.

Щоб досягти цього, ви використаєте метод `.repeat()`, доступний для рядків. Цей метод приймає число як аргумент, яке вказує, скільки разів повторити цільовий рядок. Наприклад, використання `.repeat()` для генерації рядка `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Використайте метод `.repeat()` для вашого `character` і передайте йому `i` як число.

# --hints--

Ви повинні використати метод `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Ви повинні використати метод `.repeat()` для вашої змінної `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Ви повинні передати `i` у ваш метод `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Ви повинні використати метод `.repeat()` у методі `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
