---
id: 660f165270622fd4ec0da3f7
title: Крок 31
challengeType: 1
dashedName: step-31
---

# --description--

Оголошення змінної за допомогою ключового слова `let` дозволяє її перепризначити. Це означає, що пізніше можна змінити `character` на зовсім інше значення.

Для цього проєкту ви не захочете змінювати значення цих змінних. Тому натомість використайте `const`, щоб оголосити їх. `const` змінні є спеціальними.

По-перше, `const` змінна не може бути перепризначена як `let` змінна. Цей код кине помилку:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

`const` змінна також не може бути неініціалізованою. Цей код кине помилку:

```js
const firstName;
```

Замініть ключові слова `let` на `const`.

# --hints--

Використайте `const`, щоб оголосити змінну `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Використайте `const`, щоб оголосити змінну `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Використайте `const`, щоб оголосити змінну `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Код не повинен використати `let`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
