---
id: 660f48a419b40238e2b8b4d5
title: Крок 93
challengeType: 1
dashedName: step-93
---

# --description--

Ваш цикл більше не залежить від змінної `continueLoop`. Це робить цю змінну <dfn>непотрібним оголошенням</dfn>. Загалом, варто уникати непотрібних оголошень, щоб уникнути плутанини в майбутньому.

Видаліть змінну `continueLoop`.

# --hints--

У вас більше не повинно бути змінної `continueLoop`.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
