---
id: 660f4f79e2a82a4e92290f44
title: Крок 110
challengeType: 1
dashedName: step-110
---

# --description--

Іноді ви можете захотіти повернути попередній код, який ви закоментували. Ви можете зробити це, видаливши `/*` і `*/` навколо цього коду. Це називається <dfn>розкоментуванням</dfn>.

Розкоментуйте лише свій перший `for` цикл. Залиште однолінійний коментар і два інші багатолінійні коментарі на місці.

# --hints--

Ви повинні розкоментувати свій перший `for` цикл.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Ви не повинні видаляти свій однолінійний коментар.

```js
assert.match(code, /\/\//);
```

Ви не повинні розкоментувати свій `while` цикл.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Ви не повинні розкоментувати свій другий `for` цикл.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
