---
id: 660f4e74f7fd3f4a99ac2e50
title: Крок 108
challengeType: 1
dashedName: step-108
---

# --description--

Масиви також мають метод `.shift()`. Він видалить **перший** елемент масиву, на відміну від `.pop()`, який видаляє останній елемент. Ось приклад методу `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

Масив `numbers` стане `[2, 3]`.

Безпосередньо під вашим масивом `numbers` оголосіть змінну `shifted` і призначте їй результат виклику `.shift()` на масиві `numbers`. На наступному рядку виведіть змінну `shifted` у консоль.

# --hints--

Ви повинні використати `const`, щоб оголосити змінну `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Ви повинні викликати `.shift()` на вашому масиві `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Ви повинні призначити результат виклику `.shift()` змінній `shifted`.

```js
assert.equal(shifted, 1);
```

Ви повинні вивести змінну `shifted` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
