---
id: 660f34626216270c682e2f7b
title: Крок 66
challengeType: 1
dashedName: step-66
---

# --description--

Виклик <dfn>функції</dfn> дозволяє вам фактично використовувати функцію. Можливо, ви не помічали, але методи на кшталт `.push()`, які ви використовували, були викликами функцій.

Функцію викликають, звертаючись до її імені та додаючи `()`. Ось як викликати функцію `test`:

```js
test();
```

Замініть `character.repeat(i + 1)` у вашому виклику `.push()` на виклик функції для вашої функції `padRow`. Поки що не додавайте до неї аргументів.

# --hints--

Ви не повинні використовувати `i + 1` у вашому виклику `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Ви не повинні використовувати `character.repeat` у вашому виклику `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Ви повинні викликати `padRow` у вашому виклику `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

У вашому виклику `padRow` не повинно бути аргументів.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
