---
id: 660f49e32001983c90b75850
title: Крок 98
challengeType: 1
dashedName: step-98
---

# --description--

Це дуже чистий і функціональний цикл. Чудова робота! Але ще є що досліджувати.

Використайте багаторядковий коментар, щоб перетворити на коментарі свій `while` цикл.

# --hints--

Ваш `while` цикл має бути перетворений на коментарі.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Тіло вашого `while` циклу має бути перетворене на коментарі.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
