---
id: 660f17294346b7d69e79db3d
title: Крок 32
challengeType: 1
dashedName: step-32
---

# --description--

Тепер ви готові почати створювати свій генератор піраміди. Ваша змінна `character` слугуватиме будівельним блоком для піраміди.

`"Hello"` може не дуже добре підходити для цього. Змініть значення `character` на символ решітки `"#"`.

# --hints--

Ваша змінна `character` має бути символом решітки.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
