---
id: 660f0a55847d6cc485f29ba5
title: Крок 26
challengeType: 1
dashedName: step-26
---

# --description--

<dfn>Метод</dfn> у JavaScript — це функція, пов’язана з певними значеннями або об’єктами. Приклад, з яким ви вже стикалися, — це метод `.log()`, який є частиною об’єкта `console`.

Масиви мають власні методи, і перший, який ви дослідите, — це метод `.push()`. Він дозволяє «додати» значення в кінець масиву. Ось приклад, як додати число `12` у кінець масиву:

```js
array.push(12);
```

Використайте `.push()`, щоб додати рядок `"freeCodeCamp"` у кінець вашого масиву `rows`. Додайте цей код перед вашим `console.log`, щоб побачити зміни у масиві.

# --hints--

Вам слід використати `.push()` у вашому коді.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Вам слід використати метод `.push()` вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Ви маєте передати рядок `"freeCodeCamp"` у ваш метод `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
