---
id: 660f1bf673487ae0bb25b900
title: Крок 37
challengeType: 1
dashedName: step-37
---

# --description--

Тепер цикл має виконатися вісім разів. Усередині тіла циклу надрукуйте значення ітератора `i` та подивіться, що станеться.

# --hints--

Ви повинні використати `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Зареєструйте значення `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Зареєструйте значення `i` в циклі `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
