---
id: 660f2eccfe3f820304af1b39
title: Крок 64
challengeType: 1
dashedName: step-64
---

# --description--

Щоб знати, як форматувати рядок, вашій функції `padRow` потрібно знати, на якому номері рядка ви зараз і скільки всього рядків генерується.

Найкращий спосіб зробити це — створити для них параметри функції. Додайте вашій функції `padRow` параметри `rowNumber` і `rowCount`. Кілька параметрів розділяються комою:

```js
function name(first, second) {

}
```

# --hints--

Ваша функція `padRow` має мати параметр `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Після параметра `rowNumber` слід додати кому.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Ваша функція `padRow` має мати параметр `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
