---
id: 660f447efc0e722f016c1be0
title: Крок 81
challengeType: 1
dashedName: step-81
---

# --description--

Текст знову з’явився! Це тому, що `"false"` є рядком, який при оцінці як булевий набуває значення `true`. Це означає, що `"false"` є правдивим значенням.

<dfn>Правдиве значення</dfn> — це значення, яке вважається істинним при оцінці як булеве. Більшість значень, з якими ви зустрінетеся в JavaScript, будуть правдивими.

<dfn>Хибне значення</dfn> — це протилежне, значення, яке вважається хибним при оцінці як булеве. У JavaScript є визначений список хибних значень. Деякі з них включають `false`, `0`, `""`, `null`, `undefined` і `NaN`.

Спробуйте змінити умову вашої `if` інструкції на порожній рядок `""`, який є хибним значенням.

# --hints--

Ваша `if` інструкція має мати `""` як умову.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
