---
id: 660f2b6fd54ac1fc142804dd
title: Крок 47
challengeType: 1
dashedName: step-47
---

# --description--

Логіка форматування цієї піраміди, ймовірно, стане складнішою, а це означає, що зараз чудовий час витягти цей код у функцію.

<dfn>Функція</dfn> — це блок коду, який можна повторно використати всюди у програмі. Функції оголошуються з таким синтаксисом:

```js
function name(parameter) {

}
```

Ключове слово `function` каже JavaScript, що змінна `name` буде функцією. `parameter` — це змінна, яка представляє значення, яке передається у функцію під час її використання. Функція може мати так багато або так мало <dfn>параметрів</dfn> як вам завгодно. Подібно до циклу `for`, простір між фігурними дужками є <dfn>тілом функції</dfn>.

Оголосіть функцію `padRow`. Ще не створюйте жодних змінних параметрів. Тіло функції має бути порожнім. Пам’ятайте, що вам потрібно використати верблюдячий регістр для конвенції іменування.

# --hints--

Використайте ключове слово `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Оголосіть функцію `padRow`.

```js
assert.isFunction(padRow);
```

Функція `padRow()` не повинна мати жодного параметра.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Функція `padRow()` повинна мати порожнє тіло.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
