---
id: 660f061d259bbebc37461080
title: Крок 20
challengeType: 1
dashedName: step-20
---

# --description--

Ви можете отримати доступ до значень у масиві, використовуючи <dfn>індекс</dfn> значення. Індекс — це число, яке представляє позицію значення в масиві, починаючи з `0` для першого значення.

Ви можете отримати доступ до значення, використовуючи <dfn>дужкову нотацію</dfn>, наприклад `array[0]`.

Використайте `console.log` і дужкову нотацію, щоб вивести перше значення у вашому масиві `rows`.

# --hints--

У вашому коді має бути інструкція `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Ви маєте отримати доступ до вашого масиву `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Ви маєте використати дужкову нотацію з вашим масивом `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Ви маєте використати дужкову нотацію, щоб отримати перший елемент вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Ви маєте вивести перший елемент вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
