---
id: 660f0f980e98e8cf77f1ce31
title: Крок 30
challengeType: 1
dashedName: step-30
---

# --description--

Змініть своє оголошення `rows`, щоб знову призначити йому порожній масив.

Також змініть свій рядок `'Hello'`, щоб знову використовувати подвійні лапки. Загалом, не має значення, які з двох ви віддаєте перевагу, але варто дотримуватися послідовності у цьому виборі в усьому проєкті.

# --hints--

Ваш масив `rows` має бути порожнім.

```js
assert.empty(rows);
```

Ваш рядок `"Hello"` має використовувати подвійні лапки.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
