---
id: 660f3915b41a441537ec9f5e
title: Крок 72
challengeType: 1
dashedName: step-72
---

# --description--

Оскільки ви збільшуєте `i` лише на `1`, ви можете використати <dfn>оператор збільшення</dfn> `++`. Цей оператор збільшує значення змінної на 1, оновлюючи присвоєння для цієї змінної. Наприклад, `test` стане `8` тут:

```js
let test = 7;
test++;
```

Замініть призначення додавання оператором збільшення для ітерації циклу.

# --hints--

Цикл `for` не повинен використати призначення додавання з `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Цикл `for` має використовувати оператор збільшення на `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
