---
id: 660f3b664421471aa595170f
title: Крок 74
challengeType: 1
dashedName: step-74
---

# --description--

Піраміда зараз виглядає трохи дивно. Оскільки ви починаєте цикл з `1` замість `0`, вам не потрібно додавати одиницю до `i`, коли передаєте його в `padRow`.

Оновіть перший аргумент вашого виклику `padRow`, щоб він був `i`.

# --hints--

Ви не повинні передавати `i + 1` у `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Ви повинні передати `i` у `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
