---
id: 660f0c34aad72dc712b97624
title: Крок 27
challengeType: 1
dashedName: step-27
---

# --description--

Ще один метод, який є важливим для цього проєкту, — це метод `.pop()`. Він видаляє останній елемент із масиву і <dfn>повертає</dfn> цей елемент.

Коли метод повертає значення, можна уявити, що він віддає це значення вам назад, роблячи його доступним для використання в інших частинах вашого коду.

Створіть нову змінну з назвою `popped` і призначте їй результат виклику `rows.pop()`. Потім виведіть у консоль `popped`.

# --hints--

Ви повинні оголосити змінну з назвою `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Ви повинні використати `let` для оголошення змінної з назвою `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Ви повинні викликати метод `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Ви повинні викликати метод `.pop()` для вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Ви повинні вивести у консоль змінну `popped`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
