---
id: 660f17d4e9f227d86e834abd
title: Крок 33
challengeType: 1
dashedName: step-33
---

# --description--

Щоб згенерувати піраміду, потрібно створити кілька рядків. Коли потрібно виконувати завдання повторно, доки не буде виконано певну умову, ви використовуєте <dfn>цикл</dfn>. Існує багато способів написати цикл.

Ви почнете з базового `for` циклу. `for` цикли мають такий синтаксис:

```js
for (iterator; condition; iteration) {
  logic;
}
```

У наступних кроках ви детально розглянете кожен компонент циклу. А поки що створіть `for` цикл, який містить терміни `"iterator"`, `"condition"` і `"iteration"` для трьох компонентів. Тіло циклу, розділ у фігурних дужках `{}`, залиште порожнім.


# --hints--

У вас має бути `for` цикл.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Першим компонентом вашого `for` циклу має бути рядок `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Другим компонентом вашого `for` циклу має бути рядок `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Третім компонентом вашого `for` циклу має бути рядок `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Тіло вашого `for` циклу має бути порожнім.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
