---
id: 660f38c34a4de6141c0c369f
title: Крок 71
challengeType: 1
dashedName: step-71
---

# --description--

Ваш генератор піраміди тепер працює як очікувалося. Але це чудова можливість детальніше дослідити код, який ви написали.

Оператор додавання — не єдиний спосіб додати значення до змінної. Оператор <dfn>додавання з присвоєнням</dfn> можна використовувати як скорочення, що означає «взяти початкове значення змінної, додати це значення і призначити результат назад змінній». Наприклад, ці дві інструкції дадуть однаковий результат:

```js
test = test + 1;
test += 1;
```

Оновіть інструкцію ітерації у циклі `for`, щоб використовувати додавання з присвоєнням.

# --hints--

Ваш цикл `for` не повинен використовувати `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Ваш цикл `for` має використовувати додавання з присвоєнням з `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
