---
id: 660f44f10ea40f300b896a5e
title: Крок 84
challengeType: 1
dashedName: step-84
---

# --description--

Тепер, коли ви попрактикувалися з інструкціями `if...else if...else`, ви можете видалити їх зі свого коду.

Після цього використайте `let` для оголошення змінної `continueLoop` і призначте їй булеве значення `false`. Потім використайте `let` для оголошення змінної `done` і призначте їй значення `0`.

# --hints--

У вас не повинно бути інструкції `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

У вас не повинно бути інструкції `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

У вас не повинно бути інструкції `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Використайте `let` для оголошення змінної `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Ваша змінна `continueLoop` повинна мати значення `false`.

```js
assert.isFalse(continueLoop);
```

Використайте `let` для оголошення змінної `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Ваша змінна `done` повинна мати значення `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
