---
id: 660f0da9bf1035c9097af20a
title: Крок 28
challengeType: 1
dashedName: step-28
---

# --description--

Ви повинні були побачити, що `"freeCodeCamp"` виводиться в консоль. Це тому, що `.pop()` повертає значення, яке було видалене з масиву — а раніше ви додали `"freeCodeCamp"` в кінець масиву.

А що повертає `.push()`? Призначте свій існуючий `rows.push()` новій змінній `pushed` і виведіть її в консоль.

# --hints--

Ви повинні оголосити змінну `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Ви повинні використати `let` для оголошення змінної `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Ви повинні призначити `rows.push("freeCodeCamp")` змінній `pushed`.

```js
assert.equal(pushed, 4);
```

Ви повинні вивести в консоль змінну `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
