---
id: 660f07d231941bc11719f664
title: Крок 21
challengeType: 1
dashedName: step-21
---

# --description--

Масиви особливі тим, що вважаються <dfn>змінними</dfn>. Це означає, що ви можете безпосередньо змінювати значення за індексом.

Наприклад, цей код призначить число `25` другому елементу масиву:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Оновіть **третій** елемент вашого масиву `rows`, щоб він став числом `10`. Потім виведіть масив `rows` у консоль.

# --hints--

Вам слід знову використати дужкову нотацію для масиву `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Вам слід звернутися до третього елемента масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Вам слід використати оператор присвоєння для третього елемента масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Вам слід призначити значення `10` третьому елементу вашого масиву `rows`.

```js
assert.equal(rows[2], 10);
```

У вашому коді має бути друга інструкція `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Вам слід вивести масив `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
