---
id: 660f3ba3cceef11b6ba08b59
title: Крок 75
challengeType: 1
dashedName: step-75
---

# --description--

На жаль, тепер зникла нижня частина піраміди. Це сталося через те, що ви створили ще одну <dfn>помилку зміщення на одиницю</dfn>.

Ваш початковий цикл проходив для значень `i` від `0` до `7`, бо `count` є `8`, а ваша умова вимагає, щоб `i` було менше за `count`. Тепер ваш цикл виконується для значень `i` від `1` до `7`.

Ваш цикл потрібно оновити, щоб він також виконувався, коли `i` дорівнює `8`. З огляду на вашу логіку, це означає, що цикл має виконуватися, коли `i` є <dfn>менше або рівне</dfn> `count`. Для цього можна використати оператор «менше або рівне» `<=`.

Оновіть умову циклу так, щоб він виконувався, поки `i` менше або рівне `count`.

# --hints--

Ваш цикл `for` не повинен перевіряти, чи `i` менше за `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Ваш цикл `for` має перевіряти, чи `i` менше або рівне `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
