---
id: 660f415b76859a2736771607
title: Крок 78
challengeType: 1
dashedName: step-78
---

# --description--

Ваша піраміда знову зникла. Це нормально — так і має бути.

Перед тим як створити новий цикл, вам потрібно дізнатися про інструкції `if`. Інструкція <dfn>`if`</dfn> дозволяє запускати блок коду лише тоді, коли виконується певна умова. Вони мають такий синтаксис:

```js
if (condition) {
  logic
}
```

Створіть інструкцію `if` з булевою умовою `true`. У тілі виведіть рядок `"Condition is true"`.

# --hints--

Ви маєте створити інструкцію `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Ваша інструкція `if` має мати `true` як умову.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

У тілі `if` має бути виведено `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
