---
id: 660f165270622fd4ec0da3f7
title: Крок 31
challengeType: 1
dashedName: step-31
---

# --description--

Оголошення змінної за допомогою ключового слова `let` дозволяє їй бути переназначеною. Це означає, що ви можете пізніше змінити `character` на зовсім інше значення.

Для цього проєкту ви не захочете змінювати ці значення змінних. Тому замість цього слід використовувати `const` для їх оголошення. Змінні `const` є особливими.

По-перше, змінна `const` не може бути переназначена, як змінна `let`. Цей код викличе помилку:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Змінна `const` також не може бути неініціалізованою. Цей код викличе помилку:

```js
const firstName;
```

Замініть свої ключові слова `let` на `const`.

# --hints--

Ви повинні використовувати `const` для оголошення вашої змінної `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Ви повинні використовувати `const` для оголошення вашої змінної `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Ви повинні використовувати `const` для оголошення вашої змінної `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Вам не слід використовувати `let` у вашому коді.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
