---
id: 660f229d2dbe09ef2954a4a1
title: Крок 42
challengeType: 1
dashedName: step-42
---

# --description--

Пам’ятайте, що у попередньому циклі ви використовували оператор додавання `+`, щоб збільшити значення `i` на `1`.

Ви можете зробити подібне зі значенням рядком, додавши новий рядок до наявного рядка. Наприклад, `hello = hello + " World";` додасть рядок `" World"` до наявного рядка, що зберігається в змінній `hello`. Це називається <dfn>конкатенацією</dfn>.

У циклі `for...of` використайте оператор додавання, щоб конкатенувати значення `row` до значення `result`.

# --hints--

Використайте оператор конкатенації на змінній `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Конкатенуйте `row` до змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Призначте результат конкатенації назад до змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
