---
id: 660f1cedf3676fe26122ebf6
title: Крок 38
challengeType: 1
dashedName: step-38
---

# --description--

Ви повинні побачити числа від нуля до семи, виведені у вашій консолі, по одному на кожній лінії. Це стане основою для створення вашої піраміди.

Замініть вашу інструкцію виведення на інструкцію, яка додасть `i` до вашого масиву `rows`.

# --hints--

У вас не має бути виклику `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Ви повинні викликати `.push()` для вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Ви повинні додати `i` до вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Ваш виклик `.push()` має відбуватися у вашому циклі `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
