---
id: 660f46460f9c36330ebc07d8
title: Крок 88
challengeType: 1
dashedName: step-88
---

# --description--

Оператор рівності може призводити до дивної поведінки в JavaScript. Наприклад, `"0" == 0` є true, хоча одне є рядком, а інше — числом.

<dfn>Оператор строгої рівності</dfn> `===` використовується для перевірки, чи два значення рівні і мають однаковий тип. Загалом, це оператор рівності, який ви завжди повинні використовувати. З оператором строгої рівності `"0" === 0` стає false, тому що, хоча вони можуть мати однакове значення нуля, вони не одного типу.

Оновіть вашу умову `done == count`, щоб використовувати оператор строгої рівності.

# --hints--

Ваша умова `if` має використовувати оператор строгої рівності.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
