---
id: 660f229d2dbe09ef2954a4a1
title: Крок 42
challengeType: 1
dashedName: step-42
---

# --description--

Пам’ятайте, у вашому попередньому циклі ви використали оператор додавання `+`, щоб збільшити значення `i` на `1`.

Ви можете зробити подібне з рядком, додавши новий рядок до існуючого. Наприклад, `hello = hello + " World";` додасть рядок `" World"` до існуючого рядка, збереженого у змінній `hello`. Це називається <dfn>конкатенацією</dfn>.

У вашому циклі `for...of` використайте оператор додавання, щоб конкатенувати значення `row` зі значенням `result`.

# --hints--

Ви повинні використати оператор конкатенації для вашої змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Ви повинні конкатенувати `row` до вашої змінної `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Ви повинні призначити результат конкатенації назад змінній `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
