---
id: 660f383d4c772c12ff59904b
title: Крок 70
challengeType: 1
dashedName: step-70
---

# --description--

Ви можете передавати повні вирази як аргумент. Функція отримає результат обчислення цього виразу. Наприклад, ці два виклики функцій, нижче, дадуть однаковий результат. (Note that the order of operations rule PEMDAS—Parenthesis, Exponents, Multiplication, Division, Addition, Subtraction—applies):

```js
test(2 * 3 + 1);
test(7);
```

Знову дивлячись на шаблон:

```md
··#··
·###·
#####
```

Оновіть значення `character`, щоб воно повторювалося `2 * rowNumber - 1` разів.

Знову відкрийте консоль, щоб побачити оновлений результат.

# --hints--

Функція має передати `2 * rowNumber - 1` до виклику `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
