---
id: 660f2a70ad6225fa503e71c3
title: Крок 46
challengeType: 1
dashedName: step-46
---

# --description--

Ви наближаєтеся до мети! На цьому етапі ви стикаєтеся з тим, що називають <dfn>помилкою зміщення на одиницю</dfn> — поширеною проблемою в мовах з індексацією, що починається з нуля, таких як JavaScript.

Перший індекс вашого масиву `rows` — це `0`, тому ви починаєте свій цикл `for` з `i = 0`. Але повторення рядка нуль разів дає нічого для виведення.

Щоб виправити це, додайте `1` до значення `i` у вашому виклику `.repeat()`. Не призначайте це назад у `i`, як ви робили в умовах циклу.

# --hints--

Ви повинні додати `1` до `i` у вашому методі `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
