---
id: 660f374d532dc41189cc9cc2
title: Крок 69
challengeType: 1
dashedName: step-69
---

# --description--

Тепер настав час трохи математики. Розглянемо піраміду з трьох рядів. Якщо ми хочемо, щоб вона була по центру, вона виглядала б приблизно так:

```md
··#··
·###·
#####
```

Порожні пробіли замінені на крапки-посередині для кращої читабельності. Якщо продовжити цей шаблон, можна побачити, що пробіли на початку та в кінці рядка слідують певному шаблону.

Оновіть свої рядки з пробілами так, щоб вони повторювалися `rowCount - rowNumber` разів.

Відкрийте консоль, щоб побачити результат.

# --hints--

Ви повинні викликати `.repeat()` на своїх рядках `" "`, щоб повторити їх `rowCount - rowNumber` разів.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
