---
id: 660f1a00ac619ddc1e259a66
title: Крок 35
challengeType: 1
dashedName: step-35
---

# --description--

<dfn>Умова</dfn> циклу `for` визначає, скільки разів він має ітеруватися. Коли `condition` стане хибою, цикл зупиниться.

У JavaScript булеве значення може бути або `true`, або `false`. Це не рядки — ви дізнаєтеся більше про різницю пізніше.

Поки що ви будете використовувати <dfn>оператор «менше ніж»</dfn> (`<`). Він дозволяє перевірити, чи є значення зліва меншим за значення справа. Наприклад, `count < 3` буде оцінено як `true`, якщо `count` є `2`, і як `false`, якщо `count` є `4`.

Замініть свій рядок `"condition"` на умову, яка перевіряє, чи `i` менше за `count`.

# --hints--

Ви повинні використовувати оператор «менше ніж».

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Ви повинні використовувати оператор «менше ніж», щоб перевірити, чи `i` менше за `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Ваш цикл `for` має використовувати `i < count` як умову.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
