---
id: 660f359af3e32e0f1a6880b7
title: Крок 68
challengeType: 1
dashedName: step-68
---

# --description--

Тепер ви повинні бачити той самий набір символів у своїй консолі. Ваша функція `padRow` робить точно те саме, що й раніше, але тепер вона знаходиться в окремому повторно використовуваному розділі.

Використайте оператор додавання, щоб додати один пробіл `" "` на початок і кінець вашого повторюваного рядка `character`.

Пам’ятайте, що ви можете використовувати оператор `+` для конкатенації рядків ось так:

```js
" " + "string"
```

# --hints--

Ви повинні додати один пробіл на початок повернутого значення.

```js
assert.match(padRow(1, 1), /^\s/);
```

Ви повинні додати один пробіл у кінець повернутого значення.

```js
assert.match(padRow(1, 1), /\s$/);
```

Ваша функція `padRow()` має повертати повторювану серію `character` з пробілом перед і після серії.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
