---
id: 660f50a21fe7645252804f2b
title: Крок 112
challengeType: 1
dashedName: step-112
---

# --description--

Що якби у вас був спосіб перемикатися між перевернутою пірамідою та стандартною пірамідою?

Почніть з оголошення змінної `inverted` та призначення їй значення `true`. Ви не змінюєте цю змінну у своєму коді, але вам потрібно буде використати `let`, щоб наші тести могли змінювати її пізніше.

# --hints--

Ви повинні оголосити змінну `inverted` з `let`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Ви повинні ініціалізувати `inverted` значенням `true`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
