---
id: 660f34e99571070d56d2f231
title: Крок 67
challengeType: 1
dashedName: step-67
---

# --description--

Ваша функція `padRow` має два параметри, які ви визначили. Значення передаються цим параметрам під час виклику функції.

Значення, які ви передаєте під час виклику функції, називаються <dfn>аргументами</dfn>, і ви <dfn>передаєте</dfn> аргументи у виклик функції. Ось приклад виклику функції з `"Hello"`, переданим як аргумент:

```js
test("Hello");
```

Передайте `i + 1` і `count` як аргументи у виклик `padRow`. Як і параметри, аргументи розділяються комою.

# --hints--

Ви повинні передати `i + 1` у виклик `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Після аргументу `i + 1` має бути кома.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Ви повинні передати `count` як другий аргумент.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
