---
id: 660f18f059fe0fda192ce394
title: Крок 34
challengeType: 1
dashedName: step-34
---

# --description--

Вашому циклу тепер потрібен правильний ітератор. <dfn>Ітератор</dfn> — це змінна, яку ви можете оголосити спеціально у вашому `for` циклі, щоб контролювати, як цикл ітерує або проходить через вашу логіку.

Загальною конвенцією є використання `i` як змінної ітератора в циклі. `for` цикл дозволяє оголосити це у круглих дужках `()`. Наприклад, ось `for` цикл, який оголошує змінну `index` і призначає їй значення `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Замініть рядок `"iterator"` на оголошення `let` для змінної `i`. Призначте їй значення `0` на початок. Це дасть змінній `i` значення `0` **першого разу**, коли ваш цикл запуститься.

# --hints--

Ви повинні використати `let` для оголошення змінної `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Ви повинні призначити `0` вашій змінній `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Ваш `for` цикл повинен починати ітератор `i` зі значення `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
