---
id: 6610c77d50636722e5b6be17
title: Крок 60
challengeType: 1
dashedName: step-60
---

# --description--

Під інструкцією `return` виведіть рядок `"This works!"` у консоль.

Після цього ви побачите, що рядок `"This works!"` не відображається в консолі, а лінія `console.log("This works!")` сірого кольору.

Скопіюйте виклик консолі та вставте його над інструкцією `return`. Тепер рядок `"This works!"` має з’явитися в консолі.

Важливо знати про ключове слово `return`, що воно не просто визначає значення, яке повертається з вашої функції, а й зупиняє виконання коду всередині функції або блочної інструкції. Це означає, що будь-який код після інструкції `return` не буде виконаний.

# --hints--

Ваша функція `padRow` має повертати змінну `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Ваш перший виклик `console.log` має йти після ключового слова `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Ваш другий виклик `console.log` має йти перед ключовим словом `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
