---
id: 6610c424b7119919b62932f4
title: Крок 53
challengeType: 1
dashedName: step-53
---

# --description--

Функція не обов’язково має повертати жорстко закодоване значення. Вона може повертати значення, збережене у змінній. Параметри — це спеціальні змінні для функції, тому їх також можна повертати.

Змініть свою `padRow` функцію так, щоб вона `return` параметр `name` безпосередньо.

# --hints--

Ваша `padRow` функція має повертати значення параметра `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
