---
id: 6610c48c4ea0891afa7c4696
title: Крок 54
challengeType: 1
dashedName: step-54
---

# --description--

Якщо ви знову відкриєте консоль, то побачите, що ваша функція `padRow` повертає `undefined`, навіть якщо ви визначили повернене значення! Це тому, що параметрам потрібно передавати значення під час виклику функції.

Коли ви передаєте значення у виклик функції, це значення називається аргументом. Ось приклад виклику функції `demo` з передачею `"Naomi"` як аргументу для параметра `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Передайте своє ім’я як аргумент для параметра `name` у вашому виклику `padRow`. Пам’ятайте, що ваше ім’я — це рядок, тому потрібно використовувати лапки.

# --hints--

Ви повинні передати рядок у ваш виклик `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
