---
id: 6610c16c4fa0df12c0e30675
title: Крок 49
challengeType: 1
dashedName: step-49
---

# --description--

Ви викликаєте свою `padRow` функцію, але нічого не робите з цим викликом. Усі функції в JavaScript <dfn>повертають</dfn> значення, тобто вони надають визначений результат свого виклику, який ви можете використати в іншому місці.

Щоб побачити результат виклику вашої `padRow` функції, оголосіть змінну `call` і призначте їй виклик вашої існуючої `padRow` функції.

# --hints--

Ви повинні оголосити змінну `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Ви повинні використати `const` для оголошення змінної `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Ви повинні призначити `call` результат виклику `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
