---
id: 6610b8f6a98d25f4d485a94d
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

За замовчуванням значення неініціалізованої змінної — `undefined`. Це спеціальний тип даних, який представляє значення, що ще не має визначення.

Ви все одно можете призначити значення неініціалізованій змінній. Ось приклад:

```js
let uninitialized;
uninitialized = "assigned";
```

Призначте рядок `"Test"` вашій змінній `secondCharacter` під оголошенням. Відкрийте консоль, щоб побачити, як змінився ваш лог.

# --hints--

Ви не повинні ініціалізувати `secondCharacter`. Пам’ятайте, що ініціалізація означає присвоєння значення під час оголошення змінної.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Ви повинні використати оператор присвоєння для `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Ви повинні призначити рядок `"Test"` вашій змінній `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
