---
id: 6610c2d8d67563174fcf96dc
title: Крок 52
challengeType: 1
dashedName: step-52
---

# --description--

Коли у вас є значення, яке явно записане у вашому коді, як рядок `"Hello!"` у вашій функції, це вважається <dfn>жорстко закодованим</dfn>. Жорстке кодування значення всередині функції може зробити її менш універсальною, ніж ви хотіли б.

Замість цього ви можете визначити <dfn>параметри</dfn> для функції. Параметри — це спеціальні змінні, яким присвоюють значення під час виклику функції, і які можна використовувати у функції, щоб динамічно змінювати результат коду функції.

Щоб додати параметр до функції, потрібно додати ім’я змінної всередині круглих дужок. Наприклад, ця функція `demo` має параметр `name`:

```js
function demo(name) {

}
```

`name` здається корисним параметром, тож додайте його до вашої функції `padRow`.

# --hints--

Ваша функція `padRow` має мати параметр `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
