---
id: 6610b741b54b90f0c0fb3d58
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Коли імена змінних складаються з кількох слів, існують певні конвенції щодо того, як слід писати слова з великої літери. У JavaScript конвенція, яку слід використовувати, — це <dfn>camel case</dfn>.

Camel case означає, що перше слово в імені повністю у нижньому регістрі, а всі наступні слова починаються з великої літери. Ось кілька прикладів camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Використайте camel case, щоб оголосити нову змінну `secondCharacter`.

# --hints--

Ви повинні оголосити змінну `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Ви повинні використати `let`, щоб оголосити вашу змінну `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Ви не повинні призначати значення вашій змінній `secondCharacter`. Не забудьте крапку з комою.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
