---
id: 6610b9f7619764fad5fd516d
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Ви також можете призначити значення однієї змінної іншій змінній. Наприклад:

```js
let first = "One";
let second = "Two";
second = first;
```

Змінна `second` тепер матиме значення `"One"`.

Щоб побачити це на практиці, змініть присвоєння `secondCharacter` з `"Test"` на вашу змінну `character`.

Потім відкрийте консоль, щоб побачити, що буде виведено.

# --hints--

Ви не повинні призначати значення `"Test"` вашій змінній `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Ви повинні призначити значення змінної `character` вашій змінній `secondCharacter`. Не забудьте крапку з комою.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Ваша змінна `secondCharacter` тепер має мати значення `"World"`.

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
