---
id: 6610c21b3ef82015573ffbbe
title: Крок 51
challengeType: 1
dashedName: step-51
---

# --description--

Ваша змінна `call` має значення `undefined`, хоча ви її визначили! Це тому, що ваша функція `padRow` наразі не повертає значення. За замовчуванням функції повертають `undefined` як своє значення.

Щоб повернути щось інше, потрібно використати ключове слово `return`. Ось приклад функції, яка повертає рядок `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Використайте ключове слово `return`, щоб ваша функція повертала рядок `"Hello!"`.

# --hints--

Ваша функція `padRow` має використовувати ключове слово `return`.

```js
assert.match(padRow.toString(), /return/);
```

Ваша функція `padRow` має повертати рядок `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
