---
id: 6610c6541c82551f95e765ab
title: Крок 58
challengeType: 1
dashedName: step-58
---

# --description--

Змінні також можна оголосити всередині функції. Такі змінні вважаються такими, що знаходяться в <dfn>локальній області видимості</dfn> або <dfn>області видимості блоку</dfn>. Змінна, оголошена всередині функції, може використовуватися лише всередині цієї функції. Якщо ви спробуєте отримати до неї доступ поза межами функції, отримаєте помилку референсу.

Щоб побачити це на практиці, використайте `const` для оголошення змінної `test` у вашій функції `padRow`. Ініціалізуйте її значенням `"Testing"`.

Потім, під вашою функцією, спробуйте вивести `test` у консоль. Ви побачите помилку, бо вона не визначена поза межами локальної області видимості функції. Видаліть цей `console.log`, щоб пройти тести і продовжити.

# --hints--

Ваша функція має оголосити змінну `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Ви маєте ініціалізувати `test` значенням `"Testing"`. Не забудьте крапку з комою.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Ваша змінна `test` має йти перед ключовим словом `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
