---
id: 6610c71600966a2191d3a64a
title: Крок 59
challengeType: 1
dashedName: step-59
---

# --description--

Значення, які повертає функція, використовуються шляхом виклику цієї функції. Ви можете використовувати виклик функції безпосередньо як значення, яке вона повертає, або зберегти повернене значення у змінну. Таким чином, ви можете використовувати значення, призначене локальній змінній, поза межами функції, в якій воно було створене.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Щоб використати своє значення `"Testing"`, поверніть його з функції `padRow`, оновивши інструкцію `return` так, щоб вона повертала лише змінну `test`.


# --hints--

Ваша функція `padRow` має повертати змінну `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
