---
id: 6610c83b52583e245a079217
title: Крок 61
challengeType: 1
dashedName: step-61
---

# --description--

Тепер ваша змінна `call` має значення `"Testing"`. Але ваша функція більше не використовує параметр `name`.

Видаліть параметр `name` з оголошення функції, а потім видаліть рядок `"CamperChan"` з виклику `padRow`.

Також видаліть обидва `console.log` з функції `padRow`.
# --hints--

Ваша функція `padRow` не повинна мати параметр `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

Ви не повинні передавати `"CamperChan"` у виклик `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Ви все ще повинні викликати функцію `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

У вас не повинно бути `console.log` перед ключовим словом `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

У вас не повинно бути `console.log` після ключового слова `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
