---
id: 6610c538372aa61cc0f5b122
title: Крок 57
challengeType: 1
dashedName: step-57
---

# --description--

Змінні в JavaScript доступні в певній <dfn>області видимості</dfn>. Іншими словами, місце оголошення змінної визначає, де у вашому коді її можна використовувати.

Першою є глобальна область видимості. Змінні, які оголошені поза будь-яким «блоком», як-от функція або `for` цикл, знаходяться у <dfn>глобальній області видимості</dfn>. Ваші змінні `character`, `count` і `rows` усі знаходяться у глобальній області видимості.

Коли змінна знаходиться у глобальній області видимості, функція може отримати до неї доступ у своєму визначенні. Ось приклад функції, яка використовує глобальну змінну `title`:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Цей приклад поверне `"Professor Naomi"`. Оновіть вашу функцію `padRow`, щоб вона повертала значення конкатенації вашої змінної `character` на початку параметра `name`.

# --hints--

Ваша функція `padRow` має конкатенувати `character` на початку `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Ваша функція `padRow` має повертати результат `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
