---
id: 6610c87eac0f0b256d7b037e
title: Крок 62
challengeType: 1
dashedName: step-62
---

# --description--

Оскільки ваша функція більше не використовувала параметр, зміна аргументу на неї не вплинула.

Продовжуйте і видаліть оголошення `test` та інструкцію `return` з вашої функції `padRow`, щоб функція знову була порожньою.

# --hints--

У вашій функції `padRow` не має бути змінної `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

Ваша функція `padRow` не повинна повертати значення.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
