---
id: 6610bf6fa14d700beed1b109
title: Крок 87
challengeType: 1
dashedName: step-87
---

# --description--

Оператор <dfn>рівності</dfn> `==` використовується для перевірки, чи два значення рівні. Щоб порівняти два значення, ви б використали інструкцію на кшталт `value == 8`.

Під `done++` всередині вашого циклу додайте інструкцію `if`. Ця інструкція має перевіряти, чи `done` дорівнює `count` за допомогою оператора рівності.

# --hints--

Ви маєте використати інструкцію `if` у своєму циклі. Її слід додати після `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Ваша інструкція `if` має використовувати оператор рівності для порівняння `done` і `count` у умові.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
