---
id: 6610bbed59bc2a0194d85533
title: Крок 22
challengeType: 1
dashedName: step-22
---

# --description--

Зверніть увагу, як значення всередині вашого масиву `rows` було змінено безпосередньо? Це називається <dfn>мутація</dfn>. Коли ви дізнаєтеся більше про масиви, ви навчитеся, коли можна мутувати масив, а коли цього робити не слід.

Перед тим як рухатися далі, це чудова можливість вивчити загальне використання масивів. Наразі ваш код отримує доступ до останнього елемента масиву за допомогою `rows[2]`. Але ви можете не знати, скільки елементів у масиві, коли хочете взяти останній.

Ви можете скористатися властивістю `.length` масиву — вона повертає кількість елементів у масиві. Щоб отримати останній елемент будь-якого масиву, можна використати такий синтаксис:

```js
array[array.length - 1]
```

`array.length` повертає кількість елементів у масиві. Віднімаючи `1`, ви отримуєте індекс останнього елемента масиву. Цю ж ідею можна застосувати до вашого масиву `rows`.

Оновіть свій `rows[2]`, щоб динамічно отримувати доступ до останнього елемента масиву `rows`. Скористайтеся наведеним вище прикладом для допомоги.

У консолі ви не повинні побачити жодних змін.

# --hints--

Ви повинні використати властивість `.length` вашого масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

Ви повинні використати `rows.length` у дужковій нотації.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

Ви повинні відняти `1` від довжини у дужковій нотації.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
