---
id: 6610bbed59bc2a0194d85533
title: Крок 22
challengeType: 1
dashedName: step-22
---

# --description--

Помітьте, як безпосередньо змінено значення всередині масиву `rows`? Це називається <dfn>мутацією</dfn> або <dfn>зміненням</dfn>. Коли ви дізнаєтеся більше про масиви, ви навчитеся, коли потрібно змінювати масив, а коли ні.

Перш ніж рухатися далі, це чудова можливість навчитися загально користуватися масивами. Наразі код отримує доступ до останнього елемента в масиві з `rows[2]`. Але ви можете не знати, скільки елементів у масиві, коли вам потрібен останній.

Ви можете використати властивість масиву `.length` — вона повертає кількість елементів у масиві. Щоб отримати останній елемент будь-якого масиву, ви можете використовувати такий синтаксис:

```js
array[array.length - 1]
```

`array.length` повертає кількість елементів у масиві. Віднімаючи `1`, ви отримаєте індекс останнього елемента в масиві. Ви можете застосувати цю саму концепцію до масиву `rows`.

Оновіть свої `rows[2]`, щоб динамічно отримати доступ до останнього елемента в масиві `rows`. Зверніться до прикладу вище, щоб допомогти вам.

Ви не повинні побачити жодних змін на консолі.

# --hints--

Використайте властивість `.length` масиву `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

Використайте `rows.length` у нотації в дужках.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

Відніміть `1` від довжини у нотації в дужках.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
