---
id: 664599653fcd6e97104f9261
title: Крок 83
challengeType: 1
dashedName: step-83
---

# --description--

Іноді ви захочете запускати різний код, коли всі умови `if...else if` є `false`. Ви можете зробити це, додавши блок `else`.

Блок `else` буде оцінюватися лише якщо умови в блоках `if` та `else if` не виконуються.

Тут блок `else` додано до блоку `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Додайте блок `else` до блоку `else if`. Всередині блоку `else` виведіть рядок `"This is the else block"` у консоль.

Щоб побачити результати в консолі, ви можете вручну змінити `<` у інструкції `else if` на `>`. Це зробить умову `false` істинною, і блок `else` буде виконано.

# --hints--

Ви повинні мати блок `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Ваш блок `else` повинен виводити рядок `"This is the else block"` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
