---
id: 66458f0a05df478aa627629e
title: Крок 82
challengeType: 1
dashedName: step-82
---

# --description--

Текст знову зник! Порожні рядки оцінюються як `false`, що робить їх <dfn>хибним значенням</dfn>. Ви дізнаєтеся більше про істинні та хибні значення в майбутніх проєктах.

Окрім інструкцій `if`, у JavaScript також є інструкції <dfn>else if</dfn>. Інструкції `else if` дозволяють перевіряти кілька умов в одному блоці коду.

Ось синтаксис інструкції `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Якщо перша умова `false`, JavaScript перевірить наступну умову в ланцюжку. Якщо друга умова `false`, JavaScript перевірить третю умову, і так далі.

Під вашою інструкцією `if` додайте інструкцію `else if`, яка перевіряє, чи `5` менше за `10`. Потім у тілі інструкції `else if` виведіть у консоль рядок `"5 is less than 10"`.

Перевірте консоль, щоб побачити результати.

# --hints--

У вас має бути інструкція `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Ваша інструкція `else if` має перевіряти, чи `5` менше за `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Ви маєте вивести рядок `"5 is less than 10"` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
