---
id: 664ed97c55a99e5ffab759dc
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Перед тим як рухатися далі, варто трохи повторити концепції, які ви вивчили.

Використайте ключове слово `let`, щоб оголосити змінну `profession` та змінну `age`. Ініціалізуйте `profession` рядком `"teacher"`, але не ініціалізуйте `age` жодним значенням.

Виведіть обидві змінні в консоль, щоб побачити результати.

# --hints--

Ви повинні оголосити змінну `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Ви повинні використати `let`, щоб оголосити змінну `profession`.

```js
assert.match(code, /let\s+profession/);
```

Ви повинні призначити рядок `"teacher"` змінній `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Ви повинні оголосити змінну `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Ви повинні використати `let`, щоб оголосити змінну `age`.

```js
assert.match(code, /let\s+age/);
```

Ви не повинні призначати значення змінній `age`.

```js
assert.isUndefined(age);
```

Ви повинні вивести `age` у консоль.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Ви повинні вивести `profession` у консоль.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
