---
id: 664ed9344ba3495e1f67f8ea
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Тепер, коли ви повторили оголошення та ініціалізацію, видаліть код, який ви написали для повторення.

Не видаляйте свій `character` змінну.

# --hints--

У вас не повинно бути `profession` змінної.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

У вас не повинно бути `age` змінної.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

У вас не повинно бути `console.log` інструкції для `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

У вас не повинно бути `console.log` інструкції для `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

У вас має залишитися ваша `character` змінна.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
