---
id: 665732da4815b70bb083915e
title: Крок 24
challengeType: 1
dashedName: step-24
---


# --description--

У попередніх кроках ви дізналися все про роботу з масивами. Зробіть паузу, щоб повторити те, що ви вивчили.

Почніть з оголошення змінної `cities` і ініціалізуйте її як масив рядків `"London"`, `"New York"` та `"Mumbai"`. Потім виведіть цю змінну в консоль.

Після виведення змініть останній елемент `cities` на рядок `"Mexico City"`, а потім знову виведіть змінну `cities`.

Якщо все зроблено правильно, у консолі ви побачите такий вихід.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Ви повинні використати `let` для оголошення змінної `cities`.

```js
assert.match(code, /let\s+cities/);
```

Ви повинні призначити змінній `"London"` масив рядків `"New York"`, `"Mumbai"` та `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Ви повинні використати `console.log()` для виведення всього масиву `cities` у консоль.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Ви повинні оновити останній елемент масиву `cities` на рядок `"Mexico City"`. Пам’ятайте, що ви можете отримати доступ до останнього елемента масиву, використовуючи `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

У вашому коді має бути дві інструкції `console.log(cities)`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
