---
id: 665732e62eb8b30c249b5c2c
title: Крок 25
challengeType: 1
dashedName: step-25
---

# --description--

Тепер ви готові перейти до наступної множини уроків з масивами.

Видаліть увесь свій код з попереднього кроку.

# --hints--

Ви повинні видалити всі референси на вашу змінну `cities`.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
