---
id: 66643c9a9e0ad1c787f5fe2b
title: Крок 56
challengeType: 1
dashedName: step-56
---

# --description--

Після швидкого повторення ви маєте видалити свою `addTwoNumbers` функцію, `sum` змінну та оператор виводу.

# --hints--

У вашому коді не повинно бути `addTwoNumbers` функції.

```js
assert.notMatch(code, /addTwoNumbers/);
```

У вашому коді не повинно бути `sum` змінної.

```js
assert.notMatch(code, /sum/);
```

Ви не повинні виводити `sum` змінну.

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
