---
id: 66643c93e05093c728abdbe9
title: Крок 55
challengeType: 1
dashedName: step-55
---

# --description--

Перш ніж рухатися далі, приділіть хвилину повторенню того, як працюють функції.

Оголосіть функцію з іменем `addTwoNumbers`. Ця функція має приймати два аргументи і повертати суму цих двох аргументів.

Ваша функція не повинна використовувати жорстко закодовані значення. Приклад функції з жорстко закодованим значенням може бути таким:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

Ця функція повертала б `"John Doe"` незалежно від аргументів, переданих параметрам `firstName` і `lastName`, тому `"John Doe"` вважається жорстко закодованим значенням.

Оголосіть змінну `sum` і призначте їй значення виклику вашої функції `addTwoNumbers` з аргументами `5` і `10`. Виведіть змінну `sum` у консоль.

# --hints--

У вас має бути функція з іменем `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

Ваша функція `addTwoNumbers` має мати два параметри.

```js
assert.lengthOf(addTwoNumbers, 2);
```

Ваша функція має повертати суму двох параметрів.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

Ваша функція не повинна повертати жорстко закодоване значення. Тобто вона має працювати з будь-якими двома числовими аргументами.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

Вам потрібно оголосити змінну `sum`.

```js
assert.isDefined(sum);
```

Вашій змінній `sum` має бути призначене значення `15`.

```js
assert.strictEqual(sum, 15);
```

Ви маєте призначити `sum` значення, отримане викликом функції `addTwoNumbers` з аргументами `5` і `10`.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

Виведіть у консоль вашу змінну `sum`.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
