---
id: 67fccc2463253b213a000142
title: Construir um Reprodutor Multimídia
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

Nas lições anteriores, você foi apresentado ao trabalho com os elementos `audio` e `video`. Neste laboratório, você irá construir um player multimídia que exibirá uma faixa de `audio` e `video` com uma transcrição.

Para o elemento `audio`, você precisará incluir um elemento `source` que é usado para especificar a mídia que está sendo utilizada. 

Aqui está um exemplo:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

O elemento `source` também pode ser usado no elemento `video` assim:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter um elemento `h1` para o título principal da página.
2. Você deve ter três elementos `section`.
3. Dentro do primeiro elemento `section`, você deve ter um elemento `h2` para o título da música que está tocando.
4. Abaixo do elemento `h2`, você deve ter um elemento `audio` com o atributo `controls` e um atributo `aria-label`.
5. Dentro do elemento `audio`, você deve ter um elemento `source` com um atributo `src` apontando para um arquivo de áudio e um atributo `type`. Você pode usar este URL de áudio se desejar: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. Dentro do segundo elemento `section`, você deve ter um elemento `h2` para o título do vídeo em reprodução.
7. Abaixo do elemento `h2`, você deve ter um elemento `video` com os atributos `controls`, `width` e um atributo `aria-label`.
8. Dentro do elemento `video`, você deve ter um elemento `source` com um atributo `src` apontando para um arquivo de vídeo e um atributo `type`. Você pode usar este URL de vídeo se desejar: `https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4`
9. Abaixo do elemento `source`, você deve ter um elemento `track` com um atributo `src` apontando para um arquivo de legendas e um atributo `kind`, um atributo `srclang` e um atributo `label`.
10. Dentro do terceiro elemento `section`, você deve ter um elemento `h2` para o título da seção, por exemplo, "Transcrição".
11. Abaixo do elemento `h2`, você deve ter um elemento `p` com a transcrição do vídeo.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Seu `h1` deve conter algum texto que represente o título principal da página.  

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```  

Você deve ter três elementos `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

A primeira `section` deve conter um elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Seu `h2` deve conter algum texto que represente o título da música que está tocando.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

A primeira `section` deve conter um elemento `audio`.  

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Seu elemento `audio` deve ter um atributo `controls`.  

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Seu elemento `audio` deve ter um atributo `aria-label` com um texto que descreva o áudio.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Seu elemento `audio` deve conter um elemento `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Seu elemento `source` deve ter um atributo `src` apontando para um arquivo de áudio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Seu elemento `source` deve ter um atributo `type` com um valor especificando o tipo do áudio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

A segunda `section` deve conter um elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Seu `h2` deve conter algum texto que represente o título do vídeo em reprodução.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

A segunda `section` deve conter um elemento `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Seu elemento `video` deve ter um atributo `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Seu elemento `video` deve ter um atributo `width` com um valor que especifique a largura do vídeo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

O elemento `video` deve conter um elemento `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

O elemento `source` deve ter um atributo `src` apontando para um arquivo de vídeo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

O elemento `source` deve ter um atributo `type` com um valor especificando o tipo de vídeo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Seu elemento `video` deve conter um elemento `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Seu elemento `track` deve ter um atributo `src` apontando para um arquivo de legendas.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Seu elemento `track` deve ter um atributo `kind` com um texto que descreva o tipo de arquivo.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Seu elemento `track` deve ter um atributo `srclang` com o texto descrevendo o idioma das legendas.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Seu elemento `track` deve ter um atributo `label` com um texto que descreva o idioma das legendas.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

A terceira `section` deve conter um elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Seu `h2` deve conter algum texto que represente o título da seção.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

A terceira `section` deve conter um elemento `p` para a transcrição.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
