---
id: 669e81368e52b3a5c35a2dc5
title: بناء صفحة تجميع فيديو
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك عنصر `main` كطفل وحيد لعنصر `body`.
1. يجب أن يكون لديك عنصر `h1` بعنوان موضوع صفحتك.
1. يجب أن يكون لديك فقرة تقدم موضوع صفحتك أسفل عنصر `h1`.
1. يجب أن يكون لديك ثلاثة عناصر `section` أسفل الفقرة الأولى.
1. يجب أن يحتوي كل قسم على عنصر `h2` وفقرة وعنصر `iframe` بهذا الترتيب.
1. يجب أن تحتوي عناصر `iframe` الثلاثة على خاصية `src` تشير إلى فيديو صالح.
1. يجب أن يحتوي كل عنصر `iframe` أيضًا على خاصية `title` لوصف المحتوى المضمن، وخاصيتي `height` و`width` لضبط حجم العنصر بشكل مناسب.

# --hints--

يجب أن يكون لديك عنصر `main` داخل عنصر `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

يجب أن يكون عنصر `main` هو الطفل الوحيد لعنصر `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

يجب أن يكون لديك عنصر `h1` بعنوان موضوع صفحتك داخل عنصر `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

يجب أن يكون لديك فقرة تقدم موضوع صفحتك أسفل عنصر `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

يجب أن يكون لديك ثلاثة عناصر `section` أسفل أول عنصر `p` لديك.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

يجب أن يبدأ كل عنصر `section` بعنصر `h2` يعمل كعنوان لذلك القسم.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

يجب أن يحتوي كل عنصر `section` على عنصر `p` لتقديم محتوى الفيديو كطفله الثاني.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

يجب أن يحتوي كل عنصر `section` على عنصر `iframe` كطفله الثالث.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

يجب أن يحتوي أول عنصر `iframe` على خاصية `src` تشير إلى فيديو صالح.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

يجب أن يحتوي أول عنصر `iframe` على خاصية `title` لوصف المحتوى المضمن.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

يجب أن يحتوي أول عنصر `iframe` على خاصية `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

يجب أن يحتوي أول عنصر `iframe` على خاصية `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

يجب أن يحتوي ثاني عنصر `iframe` على خاصية `src` تشير إلى فيديو صالح.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

يجب أن يحتوي ثاني عنصر `iframe` على خاصية `title` لوصف المحتوى المضمن.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

يجب أن يحتوي ثاني عنصر `iframe` على خاصية `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

يجب أن يحتوي ثاني عنصر `iframe` على خاصية `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

يجب أن يحتوي ثالث عنصر `iframe` على خاصية `src` تشير إلى فيديو صالح.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

يجب أن يحتوي ثالث عنصر `iframe` على خاصية `title` لوصف المحتوى المضمن.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

يجب أن يحتوي ثالث عنصر `iframe` على خاصية `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

يجب أن يحتوي ثالث عنصر `iframe` على خاصية `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
