---
id: a103376db3ba46b2d50db289
title: 스파이널 케이스 변환기 구현하기
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

스파이널 케이스는 모든 단어가 소문자로 되어 있고 하이픈으로 구분된 문자열 형식입니다. `"this-is-spinal-tap"`는 스파이널 케이스 문자열의 예입니다.

이 실습에서는 주어진 문자열을 스파이널 케이스로 변환하는 함수를 만드세요.

**목표**: 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `spinalCase`라는 이름의 함수를 만들어야 합니다.
2. `spinalCase` 함수는 하나의 인수, 문자열을 받아야 합니다.
3. `spinalCase` 함수는 스파이널 케이스 형식의 문자열을 반환해야 합니다. 예를 들어, 인수가 `ProductLanding page`라면 함수는 `product-landing-page`를 반환해야 합니다.

# --hints--

`spinalCase`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(spinalCase);
```

`spinalCase`는 하나의 인수를 받아야 합니다.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")`는 문자열 `this-is-spinal-tap`를 반환해야 합니다.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")`는 문자열 `this-is-spinal-tap`를 반환해야 합니다.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")`는 문자열 `the-andy-griffith-show`를 반환해야 합니다.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")`는 문자열 `teletubbies-say-eh-oh`를 반환해야 합니다.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")`는 문자열 `all-the-small-things`를 반환해야 합니다.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
