---
id: 66c06d618d075c7f7f1b890a
title: Crie um adivinhador do futuro
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

Neste laboratório, você definirá cinco previsões e selecionará aleatoriamente uma delas para exibir ao usuário.

Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve inicializar as cinco variáveis `fortune1`, `fortune2`, `fortune3`, `fortune4` e `fortune5` com um valor de string de sua escolha. Você pode usar as opções abaixo, se desejar:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Você deve selecionar um número aleatório entre 1 e 5, incluindo ambos, e atribuí-lo à variável `randomNumber`.

3. Você deve criar uma variável `selectedFortune` e atribuir a previsão apropriada com base nestas regras:

   - Se `randomNumber` for 1, atribua o valor de `fortune1` a `selectedFortune`.
   - Se `randomNumber` for 2, atribua o valor de `fortune2` a `selectedFortune`.
   - Se `randomNumber` for 3, atribua o valor de `fortune3` a `selectedFortune`.
   - Se `randomNumber` for 4, atribua o valor de `fortune4` a `selectedFortune`.
   - Se `randomNumber` for 5, atribua o valor de `fortune5` a `selectedFortune`.

4. Você deve registrar `selectedFortune` no console.

# --hints--

Você deve inicializar `fortune1` com um valor de string.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

You should initialize `fortune2` with a string value.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Você deve inicializar `fortune3` com um valor de string.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Você deve inicializar `fortune4` com um valor de string.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Você deve inicializar `fortune5` com um valor de string.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Você deve usar o método `Math.random()` para gerar um número aleatório.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Você deve gerar um número aleatório entre 1 e 5, incluindo ambos, e atribuí-lo à variável `randomNumber`.

```js
assert.isNotNull(randomNumber);
assert.include([1,2,3,4,5], randomNumber);
```

Você deve ter uma variável `selectedFortune` que é atribuída com base no valor de `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

O `randomNumber` deve corresponder à previsão. Por exemplo, se `randomNumber` for 1, o `selectedFortune` deve ser igual a `fortune1` e assim por diante.

```js
const condition1 = randomNumber === 1 && selectedFortune === fortune1;
const condition2 = randomNumber === 2 && selectedFortune === fortune2;
const condition3 = randomNumber === 3 && selectedFortune === fortune3;
const condition4 = randomNumber === 4 && selectedFortune === fortune4;
const condition5 = randomNumber === 5 && selectedFortune === fortune5;

assert.isTrue(condition1 || condition2 || condition3 || condition4 || condition5);
```

Você deve exibir o `selectedFortune` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
