---
id: 56533eb9ac21ba0edf2244b3
title: 섭씨를 화씨로 변환하기
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

섭씨를 화씨로 변환하는 공식은 섭씨 온도에 `9/5`를 곱하고 `32`를 더하는 것입니다.

섭씨 온도를 나타내는 변수 `celsius`가 주어집니다. 이미 정의된 변수 `fahrenheit`를 사용해 주어진 섭씨 온도에 해당하는 화씨 온도를 할당하세요. 위에 언급한 공식을 사용해 섭씨 온도를 화씨로 변환하세요.

# --hints--

`convertCtoF(0)`는 숫자를 반환해야 합니다

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)`는 `-22` 값을 반환해야 합니다

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)`는 `14` 값을 반환해야 합니다

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)`는 `32` 값을 반환해야 합니다

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)`는 `68` 값을 반환해야 합니다

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)`는 `86` 값을 반환해야 합니다

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
