---
id: 579e2a2c335b9d72dd32e05c
title: 배열 자르기와 붙이기
challengeType: 1
forumTopicId: 301148
dashedName: slice-and-splice
---

# --description--

두 개의 배열과 인덱스가 주어집니다.

첫 번째 배열의 각 요소를 순서대로 두 번째 배열에 복사하세요.

두 번째 배열의 인덱스 `n`부터 요소를 삽입하기 시작하세요.

결과 배열을 반환하세요. 함수 실행 후 입력 배열들은 그대로 유지되어야 합니다.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)`은 `[4, 1, 2, 3, 5]`를 반환해야 합니다.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)`은 `["a", 1, 2, "b"]`를 반환해야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)`는 `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`를 반환해야 합니다.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

첫 번째 배열의 모든 요소는 원래 순서대로 두 번째 배열에 추가되어야 합니다. `frankenSplice([1, 2, 3, 4], [], 0)`은 `[1, 2, 3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

함수 실행 후 첫 번째 배열은 그대로 유지되어야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

함수 실행 후 두 번째 배열은 그대로 유지되어야 합니다.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js
function frankenSplice(arr1, arr2, n) {
  return arr2;
}

frankenSplice([1, 2, 3], [4, 5, 6], 1);
```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  // It's alive. It's alive!
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}

frankenSplice([1, 2, 3], [4, 5], 1);
```
