---
id: a24c1a4622e3c05097f71d67
title: 내가 속할 위치
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

값(두 번째 인수)을 정렬된 배열(첫 번째 인수)에 삽입해야 하는 가장 낮은 인덱스를 반환하세요. 반환값은 숫자여야 합니다.

예를 들어, `getIndexToIns([1,2,3,4], 1.5)`는 `1`을 반환해야 합니다. 왜냐하면 `1`(인덱스 0)보다 크지만 `2`(인덱스 1)보다는 작기 때문입니다.

마찬가지로, `getIndexToIns([20,3,5], 19)`는 `2`를 반환해야 합니다. 배열이 정렬된 후 `[3,5,20]`처럼 보일 것이고 `19`은 `20`(인덱스 2)보다 작고 `5`(인덱스 1)보다는 크기 때문입니다.

# --hints--

`getIndexToIns([10, 20, 30, 40, 50], 35)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

`getIndexToIns([10, 20, 30, 40, 50], 35)`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

`getIndexToIns([10, 20, 30, 40, 50], 30)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

`getIndexToIns([40, 60], 50)`는 `1`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

`getIndexToIns([3, 10, 5], 3)`는 `0`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

`getIndexToIns([5, 3, 20, 3], 5)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)`는 `2`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([2, 20, 10], 19)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

`getIndexToIns([2, 5, 10], 15)`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

`getIndexToIns([2, 5, 10], 15)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

`getIndexToIns([], 1)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(getIndexToIns([], 1));
```

`getIndexToIns([], 1)`는 `0`를 반환해야 합니다.

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
