---
id: a302f7aae1aa3152a5b413bc
title: 숫자의 팩토리얼 구하기
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

주어진 정수의 팩토리얼을 반환하세요.

정수가 `n`로 나타난다면, 팩토리얼은 `n`보다 작거나 같은 모든 양의 정수의 곱입니다.

팩토리얼은 종종 축약 표기법 `n!`로 나타냅니다.

예를 들어: `5! = 1 * 2 * 3 * 4 * 5 = 120`

함수에는 0 이상인 정수만 입력됩니다.

# --hints--

`factorialize(5)`는 숫자를 반환해야 합니다.

```js
assert.isNumber(factorialize(5));
```

`factorialize(5)`는 `120`를 반환해야 합니다.

```js
assert.strictEqual(factorialize(5), 120);
```

`factorialize(10)`는 `3628800`를 반환해야 합니다.

```js
assert.strictEqual(factorialize(10), 3628800);
```

`factorialize(20)`는 `2432902008176640000`를 반환해야 합니다.

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

`factorialize(0)`는 `1`를 반환해야 합니다.

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
