---
id: a6e40f1041b06c996f7b2406
title: 찾는 사람이 임자입니다
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

배열 `arr`을 살펴보고 '진리 검사'를 통과하는 첫 번째 요소를 반환하는 함수를 만드세요. 이는 요소 `x`가 주어졌을 때, `func(x)`가 `true`이면 '진리 검사'를 통과하는 것을 의미합니다. 검사에 통과하는 요소가 없으면 `undefined`를 반환하세요.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })`는 `8`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })`는 `undefined`를 반환해야 합니다.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
