---
id: a77dbc43c33f39daa4429b4f
title: 불리언인지 확인하기
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

값이 불리언 원시형인지 확인하세요. `true` 또는 `false`을 반환하세요.

불리언 원시형은 `true`와 `false`입니다.

# --hints--

`booWho(true)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(booWho(true));
```

`booWho(false)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
