---
id: ab6137d4e35944e21037b769
title: 문장 제목 형식으로 바꾸기
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

주어진 문자열에서 각 단어의 첫 글자를 대문자로 바꾸어 반환하세요. 나머지 글자는 모두 소문자로 만드세요.

이 연습에서는 `the`와 `of` 같은 연결어도 대문자로 처리해야 합니다.

# --hints--

`titleCase("I'm a little tea pot")`는 문자열을 반환해야 합니다.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")`는 문자열 `I'm A Little Tea Pot`를 반환해야 합니다.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")`는 문자열 `Short And Stout`를 반환해야 합니다.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")`는 문자열 `Here Is My Handle Here Is My Spout`를 반환해야 합니다.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
