---
id: acda2fb1324d9b0fa741e6b5
title: 끝나는 부분 확인하기
challengeType: 1
forumTopicId: 16006
dashedName: confirm-the-ending
---

# --description--

문자열(첫 번째 인수, `str`)이 주어진 대상 문자열(두 번째 인수, `target`)로 끝나는지 확인하세요.

이 도전 과제는 ES2015에 도입된 `.endsWith()` 메서드를 사용해 해결할 수 있습니다. 하지만 이 도전 과제에서는 JavaScript의 substring 메서드 중 하나를 사용하시길 바랍니다.

# --hints--

`confirmEnding("Bastian", "n")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing", "mountain")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(
  confirmEnding(
    'If you want to save our world, you must hurry. We dont know how much longer we can withstand the nothing',
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

이 도전 과제를 해결할 때 내장 메서드 `.endsWith()`를 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js
function confirmEnding(str, target) {
  return str;
}

confirmEnding('Bastian', 'n');
```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
