---
id: afcc8d540bea9ea2669306b6
title: 문자열 반복하기 문자열 반복하기
challengeType: 1
forumTopicId: 16041
dashedName: repeat-a-string-repeat-a-string
---

# --description--

주어진 문자열 `str` (첫 번째 인수)를 `num`번 (두 번째 인수) 반복하세요. `num`가 양수가 아니면 빈 문자열을 반환하세요. 이 도전 과제에서는 내장 `.repeat()` 메서드를 사용하지 마세요.

# --hints--

`repeatStringNumTimes("*", 3)`는 문자열 `***`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)`는 문자열 `abcabcabc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)`는 문자열 `abcabcabcabc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)`는 문자열 `abc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)`는 문자열 `********`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)`는 빈 문자열(`""`)을 반환해야 합니다.

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

내장 `repeat()` 메서드는 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

`repeatStringNumTimes("abc", 0)`는 `""`를 반환해야 합니다.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

# --seed--

## --seed-contents--

```js
function repeatStringNumTimes(str, num) {
  return str;
}

repeatStringNumTimes('abc', 3);
```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
