---
id: a103376db3ba46b2d50db289
title: Реалізувати конвертер у spinal case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case — це формат рядка, у якому всі слова в нижньому регістрі та розділені дефісами. `"this-is-spinal-tap"` — приклад рядка у spinal case.

У цій лабораторній роботі ви створите функцію, яка конвертує заданий рядок у spinal case.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з іменем `spinalCase`.
2. Функція `spinalCase` має приймати один аргумент — рядок.
3. Функція `spinalCase` має повертати рядок у форматі spinal case. Наприклад, якщо аргументом є `ProductLanding page`, функція має повернути `product-landing-page`.

# --hints--

Ви повинні створити функцію з іменем `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` має приймати один аргумент.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` має повертати рядок `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` має повертати рядок `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` має повертати рядок `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` має повертати рядок `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` має повертати рядок `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
