---
id: 587d78b1367417b2b2512b09
title: 이미지 반응형 만들기
challengeType: 0
forumTopicId: 301140
dashedName: make-an-image-responsive
---

# --description--

CSS로 이미지를 반응형으로 만드는 것은 사실 매우 간단합니다. 이미지에 다음 속성들을 추가하기만 하면 됩니다:

```css
img {
  max-width: 100%;
  height: auto;
}
```

`max-width`의 `100%`는 이미지가 속한 컨테이너보다 절대 넓어지지 않도록 보장하고, `height`의 `auto`는 이미지가 원래의 가로세로 비율을 유지하도록 합니다.

# --instructions--

`responsive-img` 클래스에 스타일 규칙을 추가하여 반응형으로 만드세요. 이 클래스는 컨테이너(이 경우 미리보기 창)보다 절대 넓어지지 않아야 하며, 원래의 가로세로 비율을 유지해야 합니다. 코드를 추가한 후 미리보기 창 크기를 조절하여 이미지가 어떻게 동작하는지 확인하세요.

# --hints--

`responsive-img` 클래스에는 `max-width` 값으로 `100%`가 설정되어 있어야 합니다.

```js
assert.strictEqual(getComputedStyle(document.querySelector('.responsive-img')).maxWidth, '100%');
```

`responsive-img` 클래스에는 `height` 값으로 `auto`가 설정되어 있어야 합니다.

```js
assert.match(code, /height:\s*?auto;/g);
```

# --seed--

## --seed-contents--

```html
<style>
.responsive-img {


}

img {
  width: 600px;
}
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```

# --solutions--

```html
<style>
  .responsive-img {
    max-width: 100%;
    height: auto;
  }

  img {
    width: 600px;
  }
</style>

<img
  class="responsive-img"
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickerPack.jpg"
  alt="freeCodeCamp stickers set"
/>
```
