---
id: 587d78b1367417b2b2512b0c
title: 반응형 타이포그래피 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/crzN7T8'
forumTopicId: 301141
dashedName: make-typography-responsive
---

# --description--

텍스트 크기를 지정할 때 `em` 또는 `px` 대신 뷰포트 단위를 사용해 반응형 타이포그래피를 만들 수 있습니다. 뷰포트 단위는 백분율처럼 상대 단위이지만, 기준이 되는 대상이 다릅니다. 뷰포트 단위는 기기의 뷰포트 크기(너비 또는 높이)를 기준으로 하고, 백분율은 부모 컨테이너 요소의 크기를 기준으로 합니다.

네 가지 뷰포트 단위는 다음과 같습니다:

<ul><li><code>vw</code> (뷰포트 너비): <code>10vw</code>는 뷰포트 너비의 10%를 의미합니다.</li><li><code>vh</code> (뷰포트 높이): <code>3vh</code>는 뷰포트 높이의 3%를 의미합니다.</li><li><code>vmin</code> (뷰포트 최솟값): <code>70vmin</code>은 뷰포트의 더 작은 쪽(높이 또는 너비)의 70%를 의미합니다.</li><li><code>vmax</code> (뷰포트 최댓값): <code>100vmax</code>는 뷰포트의 더 큰 쪽(높이 또는 너비)의 100%를 의미합니다.</li></ul>

다음은 `body` 태그의 크기를 뷰포트 너비의 30%로 설정한 예시입니다.

```css
body {
  width: 30vw;
}
```

# --instructions--

`width` 태그의 `h2`를 뷰포트 너비의 80%로 설정하고, 문단의 `width`를 뷰포트의 더 작은 쪽 크기의 75%로 설정하세요.

# --hints--

`h2` 태그는 80vw의 `width`를 가져야 합니다.

```js
assert.match(
  __helpers.removeCssComments(code),
  /h2\s*?{\s*?width:\s*?80vw;\s*?}/g
);
```

`p` 태그는 75vmin의 `width`를 가져야 합니다.

```js
assert.match(
  __helpers.removeCssComments(code),
  /p\s*?{\s*?width:\s*?75vmin;\s*?}/g
);
```

# --seed--

## --seed-contents--

```html
<style>

</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  h2 {
    width: 80vw;
  }
  p {
    width: 75vmin;
  }
</style>

<h2>Importantus Ipsum</h2>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
