---
id: 587d78b0367417b2b2512b08
title: 미디어 쿼리 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cqwKrtm'
forumTopicId: 301139
dashedName: create-a-media-query
---

# --description--

미디어 쿼리는 CSS3에서 도입된 새로운 기법으로, 뷰포트 크기에 따라 콘텐츠의 표현 방식을 변경합니다. 뷰포트는 사용자가 볼 수 있는 웹 페이지의 영역이며, 사이트에 접근하는 기기에 따라 다릅니다.

미디어 쿼리는 미디어 유형으로 구성되며, 해당 미디어 유형이 문서가 표시되는 기기의 유형과 일치하면 스타일이 적용됩니다. 미디어 쿼리 안에는 원하는 만큼 선택자와 스타일을 넣을 수 있습니다.

다음은 기기의 너비가 `100px` 이하일 때 콘텐츠를 반환하는 미디어 쿼리 예시입니다:

```css
@media (max-width: 100px) {
  /* CSS Rules */
}
```

다음 미디어 쿼리는 기기의 높이가 `350px` 이상일 때 콘텐츠를 반환합니다:

```css
@media (min-height: 350px) {
  /* CSS Rules */
}
```

미디어 쿼리 안의 CSS는 미디어 유형이 사용 중인 기기와 일치할 때만 적용된다는 점을 기억하세요.

# --instructions--

기기의 높이가 `p` 이하일 때 `font-size` 태그가 `10px`인 `800px`을 갖도록 미디어 쿼리를 추가하세요.

# --hints--

기기의 `@media`이 `height` 이하인 경우를 위한 `800px` 쿼리를 선언해야 합니다.

```js
const media = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(
  media.some(x => x.media?.mediaText?.includes('(max-height: 800px)'))
);
```

기기의 `p`이 `font-size` 이하일 때 `10px` 요소가 `height`인 `800px`을 가져야 합니다.

```js
const rules = new __helpers.CSSHelp(document).getRuleListsWithinMedia(
  '(max-height: 800px)'
);
assert.strictEqual(
  rules?.find(x => x.selectorText === 'p')?.style?.fontSize,
  '10px'
);
```

기기의 `p`이 `font-size` 초과일 때 `20px` 요소가 초기 `height`인 `800px`을 가져야 합니다.

```js
const ifPFirst = new __helpers.CSSHelp(document)
  .getCSSRules()
  ?.find(x => x?.selectorText === 'p' || x?.media);
assert.strictEqual(ifPFirst?.style?.fontSize, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 20px;
  }

  /* Only change code below this line */

  /* Only change code above this line */
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 20px;
  }

  @media (max-height: 800px) {
    p {
      font-size: 10px;
    }
  }
</style>

<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit. Vivamus quis tempus
  massa. Aenean erat nisl, gravida vel vestibulum cursus, interdum sit amet
  lectus. Sed sit amet quam nibh. Suspendisse quis tincidunt nulla. In hac
  habitasse platea dictumst. Ut sit amet pretium nisl. Vivamus vel mi sem.
  Aenean sit amet consectetur sem. Suspendisse pretium, purus et gravida
  consequat, nunc ligula ultricies diam, at aliquet velit libero a dui.
</p>
```
