---
id: 587d78b1367417b2b2512b0a
title: 고해상도 디스플레이용 레티나 이미지 사용하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pzrPu4/cVZ4Rfp'
forumTopicId: 301142
dashedName: use-a-retina-image-for-higher-resolution-displays
---

# --description--

인터넷에 연결된 기기가 늘어나면서 크기와 사양이 다양해지고, 사용하는 디스플레이도 외부와 내부에서 다를 수 있습니다. 픽셀 밀도는 한 기기에서 다른 기기와 다를 수 있는 요소이며, 이 밀도는 인치당 픽셀 수(PPI) 또는 인치당 점 수(DPI)로 알려져 있습니다. 가장 유명한 디스플레이는 최신 Apple MacBook Pro 노트북과 최근 iMac 컴퓨터에 사용되는 "레티나 디스플레이"입니다. "레티나"와 "비레티나" 디스플레이 간 픽셀 밀도 차이 때문에, 고해상도 디스플레이를 염두에 두지 않고 만든 일부 이미지는 고해상도 디스플레이에서 렌더링할 때 "픽셀화"되어 보일 수 있습니다.

MacBook Pro의 "레티나 디스플레이" 같은 고해상도 디스플레이에서 이미지를 제대로 보이게 하는 가장 간단한 방법은 `width`와 `height` 값을 원본 파일 크기의 절반으로 정의하는 것입니다. 다음은 원본 높이와 너비의 절반만 사용하는 이미지 예시입니다.

```html
<style>
  img {
    height: 250px;
    width: 250px;
  }
</style>
<img src="coolPic500x500" alt="A most excellent picture" />
```

# --instructions--

`width` 태그의 `height`와 `img`를 원래 값의 절반으로 설정하세요. 이 경우 원본 `height`와 원본 `width` 모두 `200px`입니다.

# --hints--

`img` 태그는 100픽셀의 `width`를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('img').width, 100);
```

`img` 태그는 100픽셀의 `height`를 가져야 합니다.

```js
assert.strictEqual(document.querySelector('img').height, 100);
```

# --seed--

## --seed-contents--

```html
<style></style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```

# --solutions--

```html
<style>
  img {
    height: 100px;
    width: 100px;
  }
</style>

<img
  src="https://cdn.freecodecamp.org/curriculum/responsive-web-design-principles/FCCStickers-CamperBot200x200.jpg"
  alt="freeCodeCamp sticker that says 'Because CamperBot Cares'"
/>
```
