---
id: 68f1fac259a4afabb478b960
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Puede que hayas notado que no enlazaste al video real. Lo harás en la siguiente fase. Cuando se trata de tipos de archivos de video, hay diferencias en el soporte del navegador. Para adaptarte a esto, puedes usar elementos `source` dentro del elemento `video` y el navegador seleccionará el primer `source` compatible.

Aquí tienes un ejemplo de un elemento `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

El elemento `source` es un elemento vacío, por lo que no tiene una etiqueta de cierre. 

Agrega un elemento `source` dentro de tu elemento `video`. 


# --hints--

Debes tener un elemento `source` dentro de tu elemento `video`.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```
