---
id: 68f18724d71db1dda8d6fec7
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, construirás un reproductor de video HTML. El boilerplate de HTML ha sido proporcionado para ti.

Crea un elemento `h1` y asígnale el texto `Working with the HTML Video Element`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe contener el texto `Working with the HTML Video Element`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
