---
id: 68f1f57723debf9383c75947
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El atributo `loop` reiniciará el video una vez que termine la reproducción. Piensa en un meme de internet que repite la reproducción. Omitir el atributo `loop` hará que el video se reproduzca una sola vez.

El atributo `loop` es un atributo booleano y no necesita un valor.

Agrega el atributo `loop` al elemento `video`.

# --hints--

Tu elemento `video` debe tener el atributo `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
