---
id: 68f1f371fefd448b4f70b898
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

En una lección anterior, aprendiste sobre los diferentes atributos disponibles para el elemento `video`. El atributo `width` determina el ancho del video en píxeles.

Agrega el atributo `width` al elemento `video` con un valor de `640`.

# --hints--

Tu elemento `video` debe tener el atributo `width` con el valor `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
  </body>
</html>
```
