---
id: 68f1f8c248306da204655117
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

El atributo `muted` silenciará el audio en la reproducción inicial. Si tienes `controls` habilitados, el usuario podrá activar el audio. Omitir el atributo `muted` reproducirá el audio en la reproducción inicial.

El atributo `muted` es un atributo booleano y no necesita un valor.

Agrega el atributo `muted` al elemento `video`.

# --hints--

Tu elemento `video` debe tener el atributo `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
