---
id: 68f1f6f97af9749a42260a43
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

El atributo `controls` proporciona controles de reproducción que incluyen reproducción, retroceso y control de volumen para el elemento `video`. 

El atributo `controls` es un atributo booleano y no necesita un valor.

Agrega el atributo `controls` al elemento `video`.

Ahora deberías ver el elemento de video mostrado en la página.

# --hints--

Tu elemento `video` debe tener el atributo `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
