---
id: 68f1eb9ad0a7575e4dbc31d5
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

A continuación, crea un elemento `video` debajo del `h1`. En los siguientes pasos, agregarás los atributos necesarios para que el reproductor de video funcione. 

# --hints--

Debes tener un elemento `video`.

```js
assert.exists(document.querySelector('video'));
```

El nuevo elemento `video` debe estar debajo del `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
