---
id: 68f1f993460f06a64c033061
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

El atributo `poster` es una imagen en miniatura del video. Piensa en los videos que ves en YouTube. Se muestra mientras el video se está descargando. Si se omite el atributo, se muestra el primer fotograma del video durante la fase de descarga.

Ahora, agrega el atributo `poster` con el valor `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` a tu elemento `video`.

# --hints--

Tu elemento `video` debe tener el atributo `poster` con el valor `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
