---
id: 68f26d52b13e0f5b6e2d2e09
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Otro tipo MIME común es el tipo MIME `video/webm`.

WebM es un formato de archivo multimedia audiovisual de código abierto desarrollado por Google, diseñado principalmente para contenido multimedia basado en la web. Soporta códecs de video como VP8, VP9 y AV1, y códecs de audio como Vorbis y Opus, lo que lo convierte en una opción popular para los elementos de video y audio de HTML5.

Debajo de tu primer elemento `source`, agrega otro elemento `source` y asígnale un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` y un atributo `type` con el valor `video/webm`.

# --hints--

Debes tener un segundo elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Tu segundo elemento `source` debe tener un atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Tu segundo elemento `source` debe tener un atributo `src` con un valor de `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Tu segundo elemento `source` debe tener un atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Tu segundo elemento `source` debe tener un atributo `type` con un valor de `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
