---
id: 68f27139bddf9f6bf2821acf
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

El último elemento `source` que agregarás será para el tipo MIME `video/quicktime`.

QuickTime es una arquitectura multimedia extensible creada por Apple, que soporta la reproducción, transmisión, codificación y transcodificación de una variedad de formatos de medios digitales. No es tan popular como el formato MP4, pero puede que la necesites para soporte de aplicaciones heredadas.

Debajo de tu tercer elemento `source`, agrega un cuarto elemento `source` y asígnale un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` y un atributo `type` con el valor `video/quicktime`.

¡Felicidades! Completaste el Taller de Reproductor de Video HTML.

# --hints--

Debes tener un cuarto elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Tu cuarto elemento `source` debe tener un atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Tu cuarto elemento `source` debe tener un atributo `src` con un valor de `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Tu cuarto elemento `source` debe tener un atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Tu cuarto elemento `source` debe tener un atributo `type` con un valor de `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
    type="video/mp4"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
    type="video/webm"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
    type="video/ogg"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
