---
id: 68f268880219c2485bd71d2a
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Para especificar el recurso multimedia del video, necesitarás agregar el atributo `src` al elemento `source`.

Agrega el atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.

# --hints--

Tu elemento `source` debe tener un atributo `src`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('src'));
```

Tu elemento `source` debe tener un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('src'),
'https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    <source>
--fcc-editable-region--
  </video>
</body>
</html>
```
