---
id: 68f2701b2160c765f23b60f8
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Otro tipo MIME común es el tipo MIME `video/ogg`.

Ogg es un formato contenedor de multimedia digital diseñado para proporcionar una transmisión eficiente y manipulación de multimedia digital. Es mantenido por la Fundación Xiph.Org y es libre y abierto, sin restricciones por patentes de software. Su nombre se deriva de "ogging", jerga del juego de computadora Netrek.

Debajo de tu segundo elemento `source`, agrega un tercer elemento `source` y asígnale un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` y un atributo `type` con el valor `video/ogg`.

# --hints--

Debes tener un tercer elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Tu tercer elemento `source` debe tener un atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Tu tercer elemento `source` debe tener un atributo `src` con un valor de `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Tu tercer elemento `source` debe tener un atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Tu tercer elemento `source` debe tener un atributo `type` con un valor de `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
    type="video/mp4"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
    type="video/webm"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```
