---
id: 5f331e55dfab7a896e53c3a1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

The `title` is one of several elements that provide extra information not visible on the web page, but it is useful for search engines or how the page gets displayed.

Dentro del elemento `head`, anida un elemento `meta` con un atributo llamado `charset` con el valor `utf-8` para decirle al navegador cómo se codificarán los caracteres para la página.

# --hints--

Debes tener una etiqueta `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

El elemento `meta` es un elemento vacío, no debe tener una etiqueta final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Your `meta` tag should have a `charset` attribute.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
