---
id: 5f3326b143638ee1a09ff1e3
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

Debes tener una etiqueta `<body>` de apertura.

```js
assert(code.match(/<body>/i));
```

Debes tener una etiqueta `</body>` de cierre.

```js
assert(code.match(/<\/body>/i));
```

No debes modificar el elemento `head`. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

