---
id: 5f33294a6af5e9188dbdb8f3
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

Tu código debe tener una etiqueta de apertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Tu código debe tener una etiqueta de cierre `</main>`.

```js
assert(code.match(/<\/main>/i));
```

No debes cambiar el elemento `body`. Asegúrate de que no has borrado accidentalmente la etiqueta de cierre.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

La etiqueta `main` debe estar dentro de la etiqueta `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
