---
id: 5f33071498eb2472b87ddee4
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS is the language used to style an HTML document. It describes how HTML elements should be displayed on the screen.

As you learned in the last few steps of the Cat Photo App, there is a basic structure needed to start building your web page. Cada documento HTML debe tener una declaración `DOCTYPE` y un elemento `html`. El `DOCTYPE` le indica al navegador qué versión del código HTML está en el documento. Y el elemento `html` representa el elemento raíz que contiene todos los demás elementos.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

Tu Codigo debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Deberías tener un `<html lang="en">` de apertura con el idioma configurado en inglés.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Debes tener una etiqueta `</html>` de cierre.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
