---
id: 5f3477cb2e27333b1ab2b955
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Ahora tienes que vincular el archivo `styles.css` para que los estilos se apliquen de nuevo. Dentro del elemento `head`, añade un elemento `link`. Asigna un atributo `rel` con el valor `"stylesheet"` y un atributo `href` con el valor `"styles.css"`.

Recuerda que el elemento `link` es un elemento vacío.

# --hints--

Tu código debe tener un elemento `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

El elemento `link` es un elemento vacío, no debe tener una etiqueta final `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

No debes modificar el elemento `head` existente. Asegúrate de que no borras la etiqueta de cierre.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Debes tener un elemento `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
