---
id: 5f3c866d5414453fc2d7b480
title: Paso 31
challengeType: 0
dashedName: step-31
---

# --description--

Empezando debajo del par café/precio existente, añade los siguientes cafés y precios, utilizando elementos `article` con dos elementos `p` anidados dentro de cada uno:

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

Al igual que antes, el texto del primer elemento `p` debe ser el sabor al café y el texto del segundo elemento `p` debe ser el precio.

# --hints--

Debes tener cinco elementos `article`.

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

Cada elemento `article` debe tener dos elementos `p`.

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

Tu primer elemento `article` debe tener elementos `p` con el texto `French Vanilla` y `3.00`.

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText, /French Vanilla/i);
assert.match(children?.[1]?.innerText, /3\.00/i);
```

Tu segundo elemento `article` debe tener elementos `p` con el texto `Caramel Macchiato` y `3.75`.

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText, /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText, /3\.75/i);
```

Tu tercer elemento `article` debe tener elementos `p` con el texto `Pumpkin Spice` y `3.50`.

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText, /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText, /3\.50/i);
```

Tu cuarto elemento `article` debe tener elementos `p` con el texto `Hazelnut` y `4.00`.

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText, /Hazelnut/i);
assert.match(children?.[1]?.innerText, /4\.00/i);
```

Tu quinto elemento `article` debe tener elementos `p` con el texto `Mocha` y `4.50`.

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText, /Mocha/i);
assert.match(children?.[1]?.innerText, /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
