---
id: 5dc17d3bf86c76b9248c6eb4
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Crea un elemento `p` sotto il tuo elemento `h2` e dagli il seguente testo:

`Everyone loves cute cats online!`

# --hints--

Il tuo elemento `p` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

I tuoi tag `p` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML sono scritti in minuscolo.

```js
assert.notMatch(code, /<\/?P>/);
```

Il tuo elemento `p` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/p\>/);
```

Il testo del tuo elemento `p` dovrebbe essere `Everyone loves cute cats online!` Hai omesso il testo oppure c’è un errore di battitura.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText?.trim().replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Il tuo elemento `p` dovrebbe essere sotto l’elemento `h2`. Li hai nell’ordine sbagliato.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
