---
id: 5dc174fcf86c76b9248c6eb2
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop continuerai a lavorare con elementi HTML di base come intestazioni, paragrafi e liste costruendo un'app per foto di gatti.

Inizia il workshop aggiungendo un elemento `h1` con il testo `CatPhotoApp`.

# --hints--

Il testo `CatPhotoApp` dovrebbe essere presente nel codice. Potresti voler controllare l'ortografia.

```js
assert.match(code, /catphotoapp/i);
```

Il tuo elemento `h1` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

I tuoi tag `h1` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML si scrivono in minuscolo.

```js
assert.notMatch(code, /<\/?H1>/);
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura. I tag di chiusura hanno questa sintassi: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Il testo del tuo elemento `h1` dovrebbe essere `CatPhotoApp`. Hai o omesso il testo, fatto un errore di battitura o non si trova tra il tag di apertura e il tag di chiusura dell'elemento `h1`.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
