---
id: 5dc1798ff86c76b9248c6eb3
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Sotto l'elemento `h1`, aggiungi un elemento `h2` con questo testo:

`Cat Photos`

# --hints--

Il tuo elemento `h1` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura. I tag di chiusura hanno questa sintassi: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Dovresti avere solo un elemento `h1`. Rimuovi quello in eccesso.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Il testo del tuo elemento `h1` dovrebbe essere 'CatPhotoApp'. Hai omesso il testo oppure c'è un errore di battitura.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```

I tuoi tag `h2` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML sono scritti in minuscolo.

```js
assert.notMatch(code, /<\/?H2>/);
```

Il tuo elemento `h2` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('h2'));
```

Il tuo elemento `h2` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/h2\>/);
```

Il testo del tuo elemento `h2` dovrebbe essere `Cat Photos`. Inserisci solo il testo `Cat Photos` tra i tag di apertura e chiusura `h2`.

```js
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat photos');
```

Il tuo elemento `h2` dovrebbe essere sotto l'elemento `h1`. L'elemento `h1` ha maggiore importanza e deve stare sopra l'elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H1'), collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
