---
id: 5dc17dc8f86c76b9248c6eb5
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

I commenti ti permettono di lasciare messaggi senza influire sulla visualizzazione nel browser. Ti consentono anche di rendere il codice inattivo. Un commento in HTML inizia con `<!--`, contiene un numero qualsiasi di righe di testo e termina con `-->`.

Ecco un esempio di commento con `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Aggiungi un commento sopra l'elemento `p` con questo testo:

`TODO: Add link to cat photos`

# --hints--

Il tuo commento dovrebbe iniziare con `<!--`. Ti manca uno o più caratteri che definiscono l'inizio di un commento.

```js
assert.match(code, /<!--/);
```

Il tuo commento dovrebbe terminare con `-->`. Ti manca uno o più caratteri che definiscono la fine di un commento.

```js
assert.match(code, /-->/);
```

Il tuo codice non dovrebbe avere caratteri di commento di apertura/chiusura in eccesso. Hai un `<!--` o `-->` in più che viene mostrato nel browser.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

Il tuo commento dovrebbe essere sopra l'elemento `p`. Li hai nell'ordine sbagliato.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Il tuo commento dovrebbe contenere il testo `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
