---
id: 5dc2385ff86c76b9248c6eb7
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

In HTML5 ci sono alcuni elementi che identificano diverse aree di contenuto. Questi elementi rendono il tuo HTML più facile da leggere e aiutano con l'ottimizzazione per i motori di ricerca (SEO) e l'accessibilità.

L'elemento `main` è usato per rappresentare il contenuto principale del corpo di un documento HTML. Il contenuto all'interno dell'elemento `main` dovrebbe essere unico nel documento e non dovrebbe essere ripetuto in altre parti del documento.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifica la sezione principale di questa pagina aggiungendo un tag di apertura `<main>` prima dell'elemento `h1` e un tag di chiusura `</main>` dopo l'elemento `p`.

# --hints--

L'elemento `main` dovrebbe avere un tag di apertura. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

L'elemento `main` dovrebbe avere un tag di chiusura. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/main\>/);
```

Il tag di apertura dell'elemento `main` dovrebbe essere al di sotto del tag di apertura dell'elemento `body`. Sono nell'ordine sbagliato.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Il tag di apertura dell'elemento `main` dovrebbe essere sopra l'elemento `h1`. Sono nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Il tag di chiusura dell'elemento `main` dovrebbe essere al di sotto dell'elemento `p`. Sono nell'ordine sbagliato.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>

--fcc-editable-region--
  </body>
</html>
```
