---
id: 5dc24614f86c76b9248c6ebd
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Puoi collegare a un'altra pagina con l'elemento àncora (`a`).

Ecco un esempio che collega a `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Aggiungi un elemento àncora dopo il paragrafo che colleghi a `https://freecatphotoapp.com`. A questo punto, il collegamento non apparirà nell'anteprima.

# --hints--

Il tuo elemento àncora (`a`) dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Il tuo elemento àncora (`a`) dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/a\>/);
```

Il tuo elemento àncora (`a`) dovrebbe essere sotto l'elemento `p`. Li hai nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

Il tuo elemento àncora (`a`) non ha un attributo `href`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelector('a')?.hasAttribute('href'));
```

Il tuo elemento àncora (`a`) dovrebbe collegare a `https://freecatphotoapp.com`. Hai omesso l'URL o c'è un errore di battitura.

```js
assert.equal(
  document.querySelector('a')?.getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Anche se hai impostato l'attributo `a` dell'elemento àncora (`href`) sul collegamento corretto, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
