---
id: 5dc24165f86c76b9248c6ebc
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Tutti gli elementi `img` dovrebbero avere un attributo `alt`. Il testo dell'attributo `alt` viene usato dai lettori di schermo per migliorare l'accessibilità ed è mostrato se l'immagine non riesce a caricarsi.

Ecco un esempio di un elemento `img` con un attributo `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

All'interno dell'elemento `img`, aggiungi un attributo `alt` con questo testo:

`A cute orange cat lying on its back`

# --hints--

Il tuo codice dovrebbe avere un elemento `img`. Hai rimosso l'elemento `img` in un passaggio precedente.

```js
assert.exists(document.querySelector('img'));
```

Il tuo elemento `img` non ha un attributo `alt`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Il valore dell'attributo `img` del tuo elemento `alt` è impostato su qualcosa di diverso da `A cute orange cat lying on its back`. Assicurati che il valore dell'attributo `alt` sia racchiuso tra virgolette.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
