---
id: 5dfa3589eacea3f48c6300ae
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

All'interno del secondo elemento `section`, aggiungi un nuovo elemento `h2` con il testo `Cat Lists`.

# --hints--

Il tuo elemento `section` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Il tuo elemento `h2` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Il tuo elemento `h2` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Il secondo elemento `h2` dovrebbe essere subito sopra il tag di chiusura del secondo elemento `section`. Non è nella posizione corretta.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection?.lastElementChild?.nodeName, 'H2');
```

Il secondo elemento `h2` dovrebbe contenere il testo `Cat Lists`. Hai omesso il testo oppure c'è un errore di battitura.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    ?.lastElementChild?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
