---
id: 5ef9b03c81a63668521804e7
title: Passaggio 35
challengeType: 0
dashedName: step-35
---

# --description--

L'elemento `footer` viene usato per definire un piè di pagina per un documento o una sezione. Un piè di pagina di solito contiene informazioni sull'autore del documento, dati sul copyright, collegamenti ai termini di utilizzo, informazioni di contatto e altro.

Dopo l'elemento `main`, aggiungi un elemento `footer`.

# --hints--

Hai eliminato l'elemento `main` oppure manca il tag di apertura o il tag di chiusura.

```js
assert.exists(document.querySelector('main'));
assert.match(code, /<\/main>/);
```

Il tuo elemento `footer` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('footer'));
```

Il tuo elemento `footer` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/footer\>/);
```

Il tuo elemento `footer` dovrebbe essere sotto il tag di chiusura dell'elemento `main`. Lo hai messo da un'altra parte.

```js
assert.equal(document.querySelector('main')?.nextElementSibling.nodeName, 'FOOTER');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
