---
id: 5ef9b03c81a63668521804ea
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Nota che tutto ciò che hai aggiunto alla pagina finora è all'interno dell'elemento `body`. Tutti i contenuti che devono essere visualizzati nella pagina devono trovarsi all'interno dell'elemento `body`. Tuttavia, altre informazioni importanti vanno all'interno dell'elemento `head`.

L'elemento `head` è usato per contenere metadati sul documento, come il titolo, i collegamenti a fogli di stile e script. I metadati sono informazioni sulla pagina che non vengono direttamente visualizzate nella pagina.

Aggiungi un elemento `head` sopra l'elemento `body`.

# --hints--

Hai eliminato l'elemento `body` oppure manca un tag di apertura o di chiusura.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

L'elemento `head` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<nomeElemento>`.

```js
assert.match(code, /\<head\>/);
```

L'elemento `head` dovrebbe avere un tag di chiusura. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /\<\/head\>/);
```

L'elemento `head` dovrebbe essere al di sopra del tag di apertura dell'elemento `body`. Lo hai inserito nel posto sbagliato.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
