---
id: 5ef9b03c81a63668521804d0
title: Step 37
challengeType: 0
dashedName: step-27
---

# --description--

Per porre maggiore enfasi su una parola o una frase specifica, è possibile utilizzare l'elemento `em`.

Metti in evidenza la parola `love` nell'elemento `figcaption` inserendola in un elemento `em`.

# --hints--

L'elemento `em` dovrebbe avere un tag di apertura. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('em'));
```

L'elemento `em` dovrebbe avere un tag di chiusura. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/em\>/);
```

Hai cancellato l'elemento `figcaption` oppure manca un tag di apertura o chiusura.

```js
assert.exists(document.querySelector('figcaption'));
assert.match(code, /<\/figcaption\>/);
```

L'elemento `em` dovrebbe racchiudere il testo `love`. Hai omesso il testo o hai un refuso.

```js
assert.equal(
  document.querySelector('figcaption > em')?.innerText?.trim().toLowerCase(), 'love'
);
```

Il testo dell'elemento `figcaption`dovrebbe essere `Cats love lasagna`. Controlla eventuali errori e che ci siano gli spazi necessari attorno ai tag di apertura e chiusura dell'elemento `em`.

```js
assert.match(
  document
    .querySelector('figcaption')
    ?.innerText?.trim().replace(/\s+/gi, ' '),
    /^cats love lasagna\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
