---
id: 5ef9b03c81a63668521804ec
title: Passaggio 40
challengeType: 0
dashedName: step-40
---

# --description--

Nota che l'intero contenuto della pagina è annidato all'interno di un elemento `html`. L'elemento `html` è l'elemento radice di una pagina HTML e racchiude tutto il contenuto della pagina.

Puoi anche specificare la lingua della tua pagina aggiungendo l'attributo `lang` all'elemento `html`.

Aggiungi l'attributo `lang` con il valore `en` al tag di apertura `html` per indicare che la lingua della pagina è l'inglese.

# --hints--

Hai eliminato l'elemento `html` oppure manca il tag di apertura o il tag di chiusura.

```js
assert.match(code, /\<html.*?\>/);
assert.match(code, /\<\/html\>/);
```

Il tuo elemento `html` dovrebbe avere un attributo `lang` con il valore `en`. Potresti aver omesso l'attributo/valore oppure aver commesso un errore di battitura.

```js
const extraSpacesRemoved = code.replace(/\s+/g, " ");
assert.match(extraSpacesRemoved, /\<html\s+lang\s*\=\s*("?|'?)en\1\s*\>/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>
--fcc-editable-region--
  <head>
    <title>CatPhotoApp</title>
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
