---
id: 5ef9b03c81a63668521804e9
title: Passaggio 37
challengeType: 0
dashedName: step-37
---

# --description--

Trasforma il testo `freeCodeCamp.org` esistente in un collegamento racchiudendolo in un elemento àncora (`a`). L'attributo `href` dovrebbe essere impostato su `https://www.freecodecamp.org`.

# --hints--

Il tuo elemento àncora (`a`) dovrebbe essere annidato all'interno dell'elemento `p` dentro `footer`. Assicurati di aggiungere un tag di apertura e un tag di chiusura per l'elemento àncora (`a`).

```js
assert.isNotEmpty(document.querySelectorAll('footer > p > a'));
```

Il tuo elemento àncora (`a`) dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert.exists(aElemClosingTags);
assert.lengthOf(aElemClosingTags, 4);
```

Il tuo elemento àncora (`a`) dovrebbe avere un attributo `href` con valore `https://www.freecodecamp.org`. Potresti aver omesso l'attributo/valore o avere un errore di battitura.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(nestedAnchor?.getAttribute('href').toLowerCase(), 'https://www.freecodecamp.org');
```

Il testo del collegamento dovrebbe essere `freeCodeCamp.org`. Potresti aver omesso il testo o avere un errore di battitura.

```js
const nestedAnchor = document.querySelectorAll('footer > p > a')[0];
assert.equal(
  nestedAnchor?.innerText?.trim().toLowerCase(),
    'freecodecamp.org'
);
```

Dopo aver annidato l'elemento àncora (`a`), l'unico contenuto dell'elemento `p` visibile nel browser dovrebbe essere `No Copyright - freeCodeCamp.org`. Controlla di nuovo il testo, gli spazi o la punteggiatura sia dell'elemento `p` sia dell'elemento àncora annidato.

```js
const pText = document.querySelectorAll('footer > p')[0]?.innerText?.trim().toLowerCase().replace(/\s+/g, ' ');
assert.match(pText, /^no copyright - freecodecamp\.org$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```
