---
id: 5ef9b03c81a63668521804d4
title: Passaggio 34
challengeType: 0
dashedName: step-34
---

# --description--

L'elemento `strong` si usa per indicare che un testo ha un'importanza forte o è urgente.

Nel `figcaption` che hai appena aggiunto, indica che `hate` ha un'importanza forte racchiudendolo in un elemento `strong`.

# --hints--

Il tuo elemento `strong` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('strong'));
```

Il tuo elemento `strong` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/strong\>/);
```

Il tuo elemento `strong` dovrebbe circondare la parola `hate` nel testo `Cats hate other cats.` Hai omesso il testo o c'è un errore di battitura.

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    ?.querySelector('strong')
    ?.innerText?.trim().toLowerCase(), 'hate'
);
```

Il testo di `figcaption` dovrebbe essere `Cats hate other cats.` Controlla errori di battitura e che gli spazi necessari siano presenti intorno ai tag di apertura e chiusura dell'elemento `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
   secondFigCaption?.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
          <figcaption>Cats hate other cats.</figcaption>  
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
