---
id: 5f07c98cdb9413cbd4b16750
title: Passo 19
challengeType: 0
dashedName: step-19
---

# --description--

È ora di aggiungere una nuova sezione. Aggiungi un secondo elemento `section` al di sotto dell'elemento `section` esistente.

# --hints--

L'elemento `section` dovrebbe avere un tag di apertura. Opening tags have this syntax: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Dovresti aggiungere un solo tag di apertura `section`. Rimuovere quelli di troppo.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

L'elemento `section` dovrebbe avere un tag di chiusura. Closing tags have a `/` just after the `<` character.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Dovresti aggiungere un solo tag di chiusura `section`. Rimuovi quelli di troppo.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Il secondo elemento `section` non dovrebbe essere annidato nel primo elemento `section`.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Entrambi gli elementi `section` dovrebbero essere tra i tag di apertura e chiusura dell'elemento `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```
