---
id: 5f07be6ef7412fbad0c5626b
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

Prima di aggiungere qualsiasi nuovo contenuto, dovresti usare un elemento `section` per separare il contenuto delle foto di gatti dal contenuto futuro.

L'elemento `section` serve a definire sezioni in un documento, come capitoli, intestazioni, piè di pagina o altre sezioni del documento. È un elemento semantico che aiuta con la SEO e l'accessibilità.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Prendi il tuo elemento `h2`, due elementi `p` e un elemento àncora (`a`) e annidali in un elemento `section`.

# --hints--

Il tuo elemento `section` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

Il tuo elemento `section` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert.match(code, /<\/section\s*>/i);
```

L'intero elemento `section` dovrebbe essere tra il tag di apertura e il tag di chiusura dell'elemento `main`.

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

Gli elementi `h2`, `p` esistenti e l'elemento àncora (`a`) dovrebbero essere tra il tag di apertura e il tag di chiusura dell'elemento `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

L'elemento `h1` non dovrebbe essere annidato nell'elemento `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
