---
id: 671b6e873249bb35c9debfcf
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Aggiungi le parole `See more` prima dell'elemento di ancoraggio e `in our gallery` dopo l'elemento di ancoraggio.

# --hints--

Dovresti utilizzare l'elemento di ancoraggio esistente, non crearne uno nuovo.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Il testo all'interno dell'elemento di ancoraggio non dovrebbe cambiare, deve essere `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

You should have `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` in your code.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
