---
id: 671b6e873249bb35c9debfcf
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Aggiungi le parole `See more ` prima dell'elemento àncora e ` in our gallery` dopo l'elemento àncora.

# --hints--

Dovresti usare l'elemento àncora esistente, non crearne uno nuovo.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Il testo all'interno dell'elemento àncora non deve cambiare, deve essere `cat photos`.

```js
assert.strictEqual(document.querySelector('a')?.innerText?.replace(/\s+/g, ' ').trim(), "cat photos");
```

Dovresti avere `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` nel tuo codice.

```js
assert.match(code, /See more\s+<a href="https:\/\/freecatphotoapp\.com">cat photos<\/a>\s+in our gallery/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
