---
id: 6720b1f2680ec706ed534acf
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、さまざまな疑似クラスと疑似要素を使ってグリーティングカードをデザインする練習をします。HTMLのボイラープレートは用意されています。

`styles.css` ファイルをリンクしてワークショップを始めてください。

# --hints--

`link` 要素は `head` 要素の中にある必要があります。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 要素には、値が `rel` の `stylesheet` 属性を付けてください。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 要素には、値が `href` の `styles.css` 属性を付けてください。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
