---
id: 6720b7f3155a9278d6b6a8b8
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

いよいよグリーティングカードのスタイルを設定する時間です。

`body` 要素のセレクターを追加して、次のことをしてください。

- `font-family` を `Arial` に変更し、その後に一般的な `sans-serif` を続ける、
- 全方向に `40px` のパディングを設定する、
- `text-align` プロパティを `center` に設定する。

# --hints--

`body` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

`font-family` セレクター内で `"Arial", sans-serif` を `body` に設定してください。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

`padding` セレクター内で `40px` を `body` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

`text-align` セレクター内で `center` を `body` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
