---
id: 6720b3191bbafd269bae0ae8
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`div` 要素を作成し、`id` の `greeting-card` と `class` の `card` を持たせてください。

`div` 要素の中に、テキストが `h1` の `Happy Birthday!` を追加してください。次に、`class` という `message` を持ち、テキストが `Wishing you all the happiness and joy on your special day!` の段落要素を追加してください。

# --hints--

`div` 要素があるはずです。

```js
assert.exists(document.querySelector("div"));
```

`div` 要素は `id` の `greeting-card` を持つべきです。

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

`div` 要素は `class` の `card` を持つべきです。

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

`h1` 要素の中に `div` 要素があるはずです。

```js
assert.exists(document.querySelector("div > h1"));
```

`h1` 要素はテキスト `Happy Birthday!` を含むべきです。

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

`p` の中に `div` 要素があるはずです。

```js
assert.exists(document.querySelector("div > p"));
```

`p` 要素は `class` の `message` を持つべきです。

```js
assert.exists(document.querySelector("div > p.message"));
```

`p` 要素はテキスト `Wishing you all the happiness and joy on your special day!` を含むべきです。

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
