---
id: 672b4434dd6d6a1edc405419
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`a` 要素の中に `.card-links` 要素を2つ追加してください。

1つ目はテキストが `Send Card`、`class` が `send-link`、`href` が `#send` である必要があります。

2つ目はテキストが `Share on Social Media`、`class` が `share-link`、`href` が `#share` である必要があります。

# --hints--

`a` の中の最初の `.card-links` 要素はテキストが `Send Card` である必要があります。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

最初の `a` 要素は `class` が `send-link` の値を持つ必要があります。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

最初の `a` 要素は `href` が `#send` の値を持つ必要があります。

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

`a` の中の2つ目の `.card-links` 要素はテキストが `Share on Social Media` である必要があります。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

2つ目の `a` 要素は `class` が `share-link` の値を持つ必要があります。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

2つ目の `a` 要素は `href` が `#share` の値を持つ必要があります。

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
