---
id: 672e44e811e610232fead333
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

テキスト `h2` を含む `#send` を `Sending your card...` に追加し、次にテキスト `p` を含む `Card successfully sent to your recipient!` 要素を追加してください。

# --hints--

`#send` 要素は `h2` 要素を含む必要があります。

```js
assert.exists(document.querySelector("#send > h2"));
```

`h2` 要素はテキスト `Sending your card...` を含む必要があります。

```js
assert.strictEqual(document.querySelector("#send > h2").innerText.trim(), "Sending your card...");
```

`#send` の中の2番目の要素は `p` 要素である必要があります。

```js
assert.strictEqual(document.querySelector("#send").children[1].tagName, "P");
```

`p` 要素はテキスト `Card successfully sent to your recipient!` を持つ必要があります。

```js
assert.strictEqual(document.querySelector("#send").children[1].innerText.trim(), "Card successfully sent to your recipient!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
    <section id="share"></section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
