---
id: 680a7e200400089967b8b1f4
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Die `str.maketrans()`-Methode nimmt zwei Strings gleicher Länge und gibt eine Übersetzungstabelle zurück, die jedes Zeichen des ersten Strings dem entsprechenden Zeichen des zweiten Strings zuordnet. Jedes Zeichen in der Übersetzungstabelle wird als Unicode-Ordinal gespeichert, eine Zahl, die das Zeichen eindeutig identifiziert.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Später im Curriculum lernen Sie mehr über die Art der Struktur, die `str.maketrans()` zurückgibt. Erstellen Sie vorerst eine Übersetzungstabelle, die die Zeichen in `alphabet` den Zeichen in `shifted_alphabet` zuordnet, und weisen Sie sie einer Variablen namens `translation_table` zu.

# --hints--

Sie sollten eine Variable namens `translation_table` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Sie sollten `str.maketrans()` aufrufen und dabei `alphabet` und `shifted_alphabet` als Argumente übergeben und das Ergebnis `translation_table` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
