---
id: 680a5e7f3d47715196b12ca0
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Jetzt, da Sie Ihr gesamtes verschobenes Alphabet haben, entfernen Sie den `print(shifted_alphabet)`-Aufruf.

# --hints--

Sie sollten `print(shifted_alphabet)` aus Ihrem Code entfernen.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
