---
id: 680a836235461fce1ff53b3f
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Die `translate()`-Methode nimmt als Argument die von `str.maketrans()` erzeugte Übersetzungstabelle. Sie wird auf einen String aufgerufen und gibt eine Kopie des ursprünglichen Strings zurück, bei der die Zeichen basierend auf der Übersetzungstabelle ersetzt wurden:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Rufen Sie die `translate()`-Methode auf `text` auf, übergeben Sie `translation_table` als Argument und weisen Sie das Ergebnis einer Variablen namens `encrypted_text` zu.

# --hints--

Sie sollten eine Variable namens `encrypted_text` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Sie sollten `text.translate(translation_table)` `encrypted_text` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
