---
id: 680a549f60f5a3fedfc0edd9
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Wie Sie an der Ausgabe sehen können, beginnt das verschobene Alphabet mit dem Buchstaben `f`, weil `shift` den Wert `5` hat. Aber jetzt fehlen die ersten fünf Buchstaben des Alphabets – `a`, `b`, `c`, `d` und `e` – im verschobenen Alphabet, daher müssen Sie diese am Ende des verschobenen Alphabets hinzufügen. Sie können den Operator `+` verwenden, um den fehlenden Teil des Alphabets anzuhängen.

Ändern Sie die bestehende Zuweisung der Variablen `shifted_alphabet`: Verwenden Sie die Slicing-Syntax, um den fehlenden ersten Teil von `alphabet` zu extrahieren und an `alphabet[shift:]` anzuhängen.

Zur Erinnerung: `sentence[start:stop]` gibt die Zeichen von `sentence` von der Position `start` (einschließlich) bis `stop` (ausschließlich) zurück.

# --hints--

Sie sollten die Variable `alphabet` nicht direkt ändern.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Sie sollten die Verkettung von `alphabet[shift:]` und dem fehlenden ersten Teil von `alphabet` der Variablen `shifted_alphabet` in derselben Zeile zuweisen, in der sie deklariert wurde. Verwenden Sie `shift`, um anzugeben, wo das Slicing enden soll.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
