---
id: 680a41ef5bfdc059cfd9d4b5
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Deklarieren Sie eine weitere Variable namens `alphabet` und weisen Sie dieser Variable den String `abcdefghijklmnopqrstuvwxyz` zu.

# --hints--

Sie sollten eine Variable namens `alphabet` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Sie sollten der Variablen `abcdefghijklmnopqrstuvwxyz` den String `alphabet` zuweisen. Denken Sie daran, entweder einfache oder doppelte Anführungszeichen zu verwenden, um den String einzuschließen.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
