---
id: 680a3d4c065557260de2ff33
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie ein Programm namens caesar cipher, eine Verschlüsselungsmethode, die Buchstaben im Alphabet verschiebt, um Nachrichten zu codieren.

Beginnen Sie damit, eine Variable namens `shift` zu erstellen und weisen Sie Ihrer neuen Variable den Wert `5` zu.

# --hints--

Sie sollten eine Variable namens `shift` deklarieren. Achten Sie darauf, den Variablennamen am Anfang der Zeile zu platzieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Sie sollten den Wert `5` Ihrer `shift`-Variable zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
