---
id: 680a431692a5f863f726c438
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

In diesem Workshop bauen Sie eine *Caesar-Verschlüsselung*. Dies ist eine der einfachsten Techniken, um Text zu verschlüsseln, bei der jeder Buchstabe des Klartexts durch den Buchstaben ersetzt wird, der eine feste Anzahl von Positionen im Alphabet weiter unten steht. Zum Beispiel wird bei einer Verschiebung um 5 `a` durch `f` ersetzt, `b` durch `g` und so weiter.

Um diese Verschlüsselung zu implementieren, müssen Sie eine neue Version Ihres Alphabets erstellen, die an der durch die Verschiebung angegebenen Position beginnt. Wie Sie in einer vorherigen Lektion gelernt haben, können Sie einen Teil eines Strings mit String-Slicing extrahieren:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Erstellen Sie eine Variable namens `shifted_alphabet` und verwenden Sie die Slicing-Syntax, um ihr den Teil von `alphabet` zuzuweisen, der am Index von `shift` beginnt. Rufen Sie dann die eingebaute `print()`-Funktion auf, um `shifted_alphabet` im Terminal auszugeben und das Ergebnis anzusehen.

# --hints--

Sie sollten eine neue Variable namens `shifted_alphabet` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Sie sollten `alphabet[shift:]` `shifted_alphabet` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Sie sollten die `print()`-Funktion aufrufen und `shifted_alphabet` als Argument übergeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
