---
id: 680a82806c4c6bc3567f2708
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Deklarieren Sie eine neue Variable mit dem Namen `text` und weisen Sie ihr den String `hello world` zu. Dies wird die zu verschlüsselnde Nachricht sein.

# --hints--

Sie sollten eine Variable mit dem Namen `text` deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Sie sollten den String `hello world` `text` zuweisen. Denken Sie daran, den String entweder in einfache oder doppelte Anführungszeichen zu setzen und achten Sie auf die Groß- und Kleinschreibung.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
