---
id: 6810891bd6de8a87a833df42
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wie Sie an der Ausgabe sehen können, wurde die Nachricht verschlüsselt. Der nächste Schritt besteht darin, Ihren Code wiederverwendbar zu machen, falls Sie verschiedene Nachrichten verschlüsseln möchten.

Dafür müssen Sie eine Funktion erstellen. Zur Erinnerung: So erstellen Sie eine Funktion namens `spam`, die `Spam!` im Terminal ausgibt:

```py
def spam():
    print('Spam!')
```

Erstellen Sie eine Funktion namens `caesar`. Legen Sie Ihren gesamten bisherigen Code in den Funktionskörper. Achten Sie darauf, für alle Zeilen im Funktionskörper die gleiche Einrückung beizubehalten.

# --hints--

Sie sollten eine Funktion namens `caesar` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Sie sollten den gesamten bisher geschriebenen Code in den Funktionskörper von `caesar` verschieben. Achten Sie darauf, für alle Zeilen im Funktionskörper die gleiche Einrückung beizubehalten...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
