---
id: 6818e9c7c8dc5694e8f64fc3
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Testen Sie nun die Funktion `caesar`, indem Sie sie mit dem String `freeCodeCamp` und der Zahl `3` als Argumente aufrufen. Weisen Sie den Funktionsaufruf einer Variablen mit dem Namen `encrypted_text` zu.

# --hints--

Sie sollten eine Variable mit dem Namen `encrypted_text` im globalen Gültigkeitsbereich haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Sie sollten `caesar('freeCodeCamp', 3)` der Variablen `encrypted_text` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
