---
id: 6818ed1adc249bb91d373729
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, dass die `return`-Anweisung verwendet wird, um einen Wert aus einer Funktion zurückzugeben, damit Sie ihn an anderer Stelle in Ihrem Code verwenden können:

```py
def spam():
    return 'Spam!'
```

Entfernen Sie den `print(encrypted_text)`-Aufruf aus Ihrer Funktion. Löschen Sie dann die Deklaration der `encrypted_text`-Variablen und geben Sie direkt `text.translate(translation_table)` aus Ihrer Funktion zurück.

# --hints--

Sie sollten kein `print(encrypted_text)` innerhalb der `caesar`-Funktion haben.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

Sie sollten keine `encrypted_text`-Variable innerhalb der `caesar`-Funktion haben.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Ihre `caesar`-Funktion sollte `text.translate(translation_table)` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
