---
id: 6818b80519e6c6c47a8bb6e8
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Wie Sie in einer vorherigen Lektion gelernt haben, können Funktionen Parameter haben, das sind Variablen, auf die innerhalb der Funktion verwiesen werden kann. Hier ist eine `sum`-Funktion mit zwei Parametern, `num1` und `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Die Nachricht, die verschlüsselt werden soll, und die Verschiebung sind in Ihrer Funktion noch fest codiert. Geben Sie Ihrer Funktion zwei Parameter mit den Namen `text` und `shift` und löschen Sie die Deklarationen der Variablen `text` und `shift` aus dem Funktionskörper von `caesar`.

# --hints--

Ihre `caesar`-Funktion sollte zwei Parameter haben, `text` und `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Sie sollten keine `text`-Variable innerhalb des Funktionskörpers von `caesar` deklariert haben.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Sie sollten keine `shift`-Variable innerhalb des Funktionskörpers von `caesar` deklariert haben.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
