---
id: 6818eb3f98c0bca5af7c8d03
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Ihr Code ist jetzt wiederverwendbar. Allerdings gibt die `caesar`-Funktion standardmäßig eine Nachricht im Terminal aus und liefert `None` zurück. Um das zu beweisen, geben Sie am Ende Ihres Codes `encrypted_text` aus.

# --hints--

Sie sollten `encrypted_text` außerhalb des Funktionskörpers ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
