---
id: 6818fcb250f34e62cab32c39
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Sie haben vielleicht bemerkt, dass die Großbuchstaben trotz der Verschlüsselung unverändert geblieben sind. Das liegt daran, dass die Übersetzungstabelle keine Großbuchstaben enthält.

Um das zu beheben, müssen Sie die Strings, die an Ihren `str.maketrans()`-Aufruf übergeben werden, mit der `upper()`-Methode ändern:

```py
greet = 'Hello'
print(greet.upper()) # HELLO
```

Aktualisieren Sie Ihren `str.maketrans()`-Aufruf, indem Sie an jedes Argument die Großbuchstabenversion des Arguments anhängen.

# --hints--

Sie sollten die Argumente Ihres `str.maketrans()`-Aufrufs ändern, indem Sie an jedes Argument die Großbuchstabenversion des Arguments anhängen.

```js
({ test: () => runPython(`
_table = _Node(_code).find_function("caesar").find_variable("translation_table")
_first = _table.is_equivalent("translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())")
_second = _table.is_equivalent("translation_table = str.maketrans(alphabet.upper() + alphabet, shifted_alphabet.upper() + shifted_alphabet)")
assert _first or _second
assert caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
