---
id: 6819c5c9faf2548e556163d5
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Wie Sie sich aus vorherigen Lektionen erinnern, ist der `not`-Operator ein unärer logischer Operator, der einen Ausdruck negiert:

```py
print(not True) # False
print(not False) # True
```

Verwenden Sie den `not`-Operator, um die Bedingung der `if`-Anweisung zu korrigieren, sodass Ihre Funktion `Shift must be an integer value.` zurückgibt, wenn `shift` *kein* Integer ist.

# --hints--

Sie sollten den `not`-Operator verwenden, um `isinstance(shift, int)` in Ihrer `if`-Anweisung zu negieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
