---
id: 6819ca464db219ab86f37dc4
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Wie Sie bereits überprüft haben, sollte der zum Verschlüsseln des Texts übergebene Shift ein positiver ganzzahliger Wert sein. Er darf jedoch `25` (den letzten Index von `alphabet`) nicht überschreiten.

Fügen Sie eine zweite Bedingung zur `if`-Anweisung hinzu, die überprüft, ob `shift` größer als `25` ist. Denken Sie daran, dass die logische ODER-Verknüpfung in Python mit dem `or`-Operator umgesetzt wird.

Aktualisieren Sie außerdem die zurückgegebene Meldung auf `Shift must be an integer between 1 and 25.`

# --hints--

Sie sollten den `or`-Operator verwenden, um nach `shift < 1` eine zweite Bedingung hinzuzufügen, die überprüft, ob `shift` größer als `25` ist. Die bestehende Bedingung dürfen Sie nicht verändern.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Ihre zweite `if`-Anweisung sollte `Shift must be an integer between 1 and 25.` zurückgeben

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
