---
id: 6819bc0637b80256a33adccc
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Nachdem Sie die grundlegenden Funktionen der Chiffre implementiert haben, ist es Zeit, eine Validierung hinzuzufügen. Dafür benötigen Sie eine `if`-Anweisung.

Erstellen Sie am Anfang des Funktionskörpers eine `if`-Anweisung. Verwenden Sie vorerst `True` als Bedingung, und geben Sie im Körper der `if`-Anweisung den String `Shift must be an integer value.` zurück.

# --hints--

Sie sollten eine `if`-Anweisung innerhalb Ihrer `caesar`-Funktion haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Ihre `if`-Anweisung sollte `True` als Bedingung verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Sie sollten den String `Shift must be an integer value.` aus Ihrer `if`-Anweisung zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
