---
id: 6819cf38880d8cc9b9f6af7a
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Python ermöglicht es Ihnen, für die Parameter einer Funktion einen Standardwert anzugeben, sodass Sie eine Funktion erstellen können, die mit weniger oder gar keinen Argumenten aufgerufen werden kann. So erstellen Sie eine Funktion mit einem `name`-Parameter, der einen Standardwert hat:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Fügen Sie Ihrer Funktion einen dritten Parameter namens `encrypt` hinzu und geben Sie ihm einen Standardwert von `True`. 

# --hints--

Sie sollten Ihrer Funktion einen dritten Parameter namens `encrypt` mit einem Standardwert von `True` hinzufügen. Ändern Sie die bestehenden Parameter nicht.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
