---
id: 6819be9055828d65c813eec5
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Die `isinstance()`-Funktion gibt `True` zurück, wenn ihr erstes Argument eine Instanz des zweiten Arguments ist, und sonst `False`:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Ersetzen Sie die aktuelle Bedingung Ihrer `if`-Anweisung durch einen `isinstance()`-Aufruf. Übergeben Sie `shift` als erstes Argument und `int` als zweites Argument.

# --hints--

Ihre `if`-Anweisung sollte `isinstance(shift, int)` als Bedingung verwenden.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
