---
id: 681a1d60e037127e7c8e0230
title: Schritt 22
challengeType: 20
dashedName: step-22
---

# --description--

Sie werden den im vorherigen Schritt hinzugefügten Parameter verwenden, um zu bestimmen, ob die Funktion den übergebenen Text verschlüsseln soll (Standardverhalten, `encrypt=True`) oder ob sie eine verschlüsselte Nachricht entschlüsseln soll.

Erstellen Sie eine `if`-Anweisung, die überprüft, ob `encrypt` nicht wahr ist. Innerhalb der neuen `if`-Anweisung setzen Sie `shift` auf `- shift`. Dies ist notwendig, damit die Verschiebung in die entgegengesetzte Richtung zum Verschlüsselungsprozess erfolgen kann.

# --hints--

Sie sollten `shift` auf `- shift` setzen, wenn `encrypt` nicht wahr ist.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
