---
id: 681a2a6cf95e82ca30866364
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Deklarieren Sie zwei Funktionen mit den Namen `encrypt` und `decrypt`, beide mit den Parametern `text` und `shift`.

Sie verwenden `encrypt` für die Verschlüsselung und `decrypt` für die Entschlüsselung, wodurch jede Aktion einen beschreibenden Namen erhält.

Beide neuen Funktionen sollen einen `caesar`-Aufruf zurückgeben, der `text` und `shift` übergibt. In `decrypt` übergeben Sie außerdem `False` als dritten Parameter.

# --hints--

Sie sollten eine Funktion mit dem Namen `encrypt` haben, die zwei Parameter hat, `text` und `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Ihre `encrypt`-Funktion sollte `caesar(text, shift)` zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Sie sollten eine Funktion mit dem Namen `decrypt` haben, die zwei Parameter hat, `text` und `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Ihre `decrypt`-Funktion sollte `caesar(text, shift, False)` zurückgeben.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
