---
id: 681a2e35c406ade2f0a5f9b2
title: Schritt 24
challengeType: 20
dashedName: step-24
---

# --description--

Es ist Zeit, die Funktion `encrypt` zu testen. Ersetzen Sie Ihren `caesar`-Aufruf durch einen Aufruf von `encrypt` mit denselben Argumenten. Sie werden dieselbe Ausgabe im Terminal sehen.

# --hints--

Ihre Variable `encrypted_text` sollte den Wert von `encrypt('freeCodeCamp', 3)` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = encrypt('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--
print(encrypted_text)
```
