---
id: 681a44f4c3235f7d8f428545
title: Schritt 25
challengeType: 20
dashedName: step-25
---

# --description--

Jetzt werden Sie die `decrypt`-Funktion testen. Ersetzen Sie den Wert, der `encrypted_text` zugewiesen ist, durch den folgenden String, der eine zu entschlüsselnde Nachricht darstellt: `Pbhentr vf sbhaq va hayvxryl cynprf.`

Deklarieren Sie dann eine Variable mit dem Namen `decrypted_text` und weisen Sie ihr einen Aufruf von `decrypt` mit `encrypted_text` als erstem Argument und einer Verschiebung von `13` als zweitem Argument zu.

Geben Sie schließlich `decrypted_text` im Terminal aus. Damit ist die Caesar-Verschlüsselung abgeschlossen.

# --hints--

Sie sollten `Pbhentr vf sbhaq va hayvxryl cynprf.` `encrypted_text` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

Sie sollten eine Variable mit dem Namen `decrypted_text` haben.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

Sie sollten `decrypt(encrypted_text, 13)` `decrypted_text` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

Sie sollten `decrypted_text` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
