---
id: 587d78aa367417b2b2512aed
title: 宣告 HTML 文件的 Doctype
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cra98AJ'
forumTopicId: 301095
dashedName: declare-the-doctype-of-an-html-document
---

# --description--

到目前為止的挑戰涵蓋了特定的 HTML 元素及其用途。然而，有一些元素為你的頁面提供整體結構，應該包含在每個 HTML 文件中。

在你的文件頂部，你需要告訴瀏覽器你的頁面使用的是哪個版本的 HTML。HTML 是一種不斷演進的語言，並且定期更新。大多數主要瀏覽器都支援最新的規格，也就是 HTML5。然而，較舊的網頁可能使用較早版本的語言。

你告訴瀏覽器這些資訊，是在第一列添加 `<!DOCTYPE ...>` 頁籤，其中 `...` 部分是 HTML 的版本。對於 HTML5，你使用 `<!DOCTYPE html>`。

`!` 和大寫的 `DOCTYPE` 很重要，尤其是對較舊的瀏覽器而言。`html` 不區分大小寫。

接著，你的其餘 HTML 程式碼需要包裹在 `html` 頁籤中。開啟的 `<html>` 放在 `<!DOCTYPE html>` 行的正下方，關閉的 `</html>` 放在頁面的最後。

這是一個頁面結構的範例。你的 HTML 程式碼應該放在兩個 `html` 標記之間的空間內。

```html
<!DOCTYPE html>
<html>

</html>
```

# --instructions--

為你的空白 HTML 文件在程式碼編輯器頂部添加一個 `DOCTYPE` 頁籤以支援 HTML5。在其下方，添加開啟和關閉的 `html` 頁籤，並將一個 `h1` 元素包裹其中。標題可以包含任何文字。

# --hints--

你的程式碼應該包含 `<!DOCTYPE html>` 頁籤。

```js
assert.match(code,/<!DOCTYPE\s+?html\s*?>/gi);
```

應該有一個 `html` 元素。

```js
assert.lengthOf(document.querySelectorAll('html'), 1);
```

`html` 頁籤應該包裹一個 `h1` 元素。

```js
assert.match(code,/<html>\s*?<h1>\s*?.*?\s*?<\/h1>\s*?<\/html>/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <h1> Hello world </h1>
</html>
```
