---
id: bad87fee1348bd9aedf08833
title: 使用佔位符文字填空
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

網頁開發者傳統上使用 <dfn>lorem ipsum 文字</dfn> 作為佔位文字。這段 lorem ipsum 文字是隨機擷取自古羅馬西塞羅的一段著名文章。

自 16 世紀以來，排版員一直使用 Lorem ipsum 文字作為佔位文字，這個傳統在網路上仍然持續。

好吧，5 世紀已經夠長了。既然我們正在建置一個 CatPhotoApp，讓我們使用一種稱為「kitty ipsum」的文字。

# --instructions--

將你的 `p` 元素內的文字取代為這段 kitty ipsum 文字的前幾個字：`Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

你的 `p` 元素應該包含所提供的「kitty ipsum」文字的前幾個字。

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
