---
id: bad87fee1348bd9aedd08830
title: 為你的表單添加提交按鈕
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cp2Nkhz'
forumTopicId: 16627
dashedName: add-a-submit-button-to-a-form
---

# --description--

讓我們為你的表單添加一個 `submit` 按鈕。點擊此按鈕將會將你的表單中的資料傳送到你在表單的 `action` 屬性中指定的 URL。

這裡有一個範例提交按鈕：

```html
<button type="submit">this button submits the form</button>
```

# --instructions--

將一個型別為 `submit`、文字為 `Submit` 的按鈕添加為你的 `form` 元素的最後一個元素。

# --hints--

你的 `form` 應該在裡面有一個 `button`。

```js
const form = document.querySelector('form');
const children = form.querySelectorAll(`:scope ${"button"}`); 
assert.isNotEmpty(children); 
```

你的提交按鈕應該有屬性 `type` 設定為 `submit`。

```js
assert.strictEqual(document.querySelector('button').getAttribute('type'), 'submit');
```

你的提交按鈕應該只包含文字 `Submit`。

```js
const text = document.querySelector('button').textContent;
assert.match(text,/^\s*submit\s*$/gi);
```

你的 `button` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.strictEqual(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
