---
id: bad88fee1348bd9aedf08816
title: 使用錨點元素連結到頁面內部區段
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

`a` (*anchor*) 元素也可以用來創建內部連結，以跳轉到網頁中的不同區段。

要創建內部連結，你需要將連結的 `href` 屬性指定為井字號 `#` 加上你想要內部連結的元素的 `id` 屬性值，通常該元素位於頁面較下方。接著，你需要為你要連結的元素添加相同的 `id` 屬性。`id` 是用來唯一描述元素的屬性。

以下是一個內部錨點連結及其標的元素的範例：

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

當使用者點擊 `Contacts` 聯結時，將會被帶到網頁中帶有 **Contacts** 標題元素的區段。

# --instructions--

將你的外部聯結改為內部聯結，方法是將 `href` 屬性改為 `#footer`，並將文字從 `cat photos` 改為 `Jump to Bottom`。

從錨點標記中移除 `target="_blank"` 屬性，因為這會導致連結的文件在新的視窗頁籤中開啟。

然後為頁面底部的 `<footer>` 元素添加一個值為 `footer` 的 `id` 屬性。

# --hints--

你的頁面上應該只有一個 anchor 標籤。

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

你的頁面上應該只有一個 `footer` 標籤。

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

`a` 頁籤應該有一個設定為 "#footer" 的 `href` 屬性。

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

`a` 頁籤不應該有 `target` 屬性。

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

`a` 文字應該是「跳到底部」。

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

`footer` 頁籤應該設定 `id` 屬性為 "footer"。

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
