---
id: bad87fee1348bd9aedf08834
title: 創建一組圓鈕
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

你可以為你想讓使用者從多個選項中只給出一個答案的問題使用 <dfn>圓鈕、單選按鈕</dfn>。

圓鈕是 `input` 的一種型別。

你的每個圓鈕都可以巢狀在自己的 `label` 元素中。將 `input` 元素包裹在 `label` 元素內時，會自動將圓鈕輸入與包圍它的標籤元素關聯起來。

所有相關的圓鈕應該具有相同的 `name` 屬性以創建圓鈕群組。透過創建圓鈕群組，選擇任何單一圓鈕時，將自動取消選取同一群組內的其他按鈕，確保使用者只提供一個答案。

以下是圓鈕、單選按鈕的範例：

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

最佳做法是在 `label` 元素上設定 `for` 屬性，其值與 `input` 元素的 `id` 屬性值相符。這讓輔助技術能夠創建標籤與相關 `input` 元素之間的聯結關係。例如：

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

我們也可以將 `input` 元素巢狀在 `label` 標籤內：

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

為你的表單添加一對圓鈕，每個圓鈕都巢狀在自己的 `label` 元素中。一個應該有 `indoor` 選項，另一個應該有 `outdoor` 選項。兩者都應該共用 `name` 屬性為 `indoor-outdoor`，以創建一個圓鈕群組。

# --hints--

你的頁面應該有兩個 `radio` 按鈕元素。

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

你的圓鈕應該被賦予 `name` 屬性為 `indoor-outdoor`。

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

你的兩個圓鈕元素中的每一個都應該巢狀在自己的 `label` 元素中。

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

你的每個 `label` 元素都應該有一個結束標籤。

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

你的其中一個圓鈕應該有標籤 `indoor`。

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

你的其中一個圓鈕應該有標籤 `outdoor`。

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

你的每個圓鈕元素都應該添加在 `form` 標記內。

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
