---
id: bad87fee1348bd9aedf08812
title: 為你的網站添加影像
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

你可以使用 `img` 元素為你的網站添加影像，並使用 `src` 屬性指向特定影像的 URL。

這方面的一個範例是：

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

請注意 `img` 是一個空元素。

所有 `img` 元素**必須**有一個 `alt` 屬性。`alt` 屬性內的文字用於螢幕閱讀器以提升存取性，且當影像載入失敗時會顯示該文字。

**注意：** 如果影像純粹是裝飾用途，使用空的 `alt` 屬性是最佳做法。

理想情況下，`alt` 屬性不應包含特殊字元，除非必要。

讓我們為上面的 `img` 範例添加一個 `alt` 屬性：

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

讓我們嘗試為我們的網站添加一張影像。

在現有的 `main` 元素內，於現有的 `p` 元素之前插入一個 `img` 元素。

現在設定 `src` 屬性，使其指向網址 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`。

最後，別忘了為你的 `img` 元素添加帶有適用文字的 `alt` 屬性。

# --hints--

你的頁面應該有一個影像元素。

```js
assert.exists(document.querySelector('img'));
```

你的影像應該有一個指向小貓影像的 `src` 屬性。

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

你的影像元素的 `alt` 屬性不應該是空的。

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
