---
id: bad87fee1348bd9aedf08827
title: 創建項目符號無序列表
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML 有一個用於建立 <dfn>無序列表</dfn> 或項目符號樣式列表的特殊元素。

無序列表以開啟的 `<ul>` 元素開始，接著是任意數量的 `<li>` 元素。最後，無序列表以 `</ul>` 結束。

例如：

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

會創建一個以項目符號樣式呈現 `milk` 和 `cheese` 的列表。

# --instructions--

移除最後兩個 `p` 元素，並在頁面底部創建一個包含貓咪喜愛的三件事的無序列表。

# --hints--

創建一個 `ul` 元素。

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

你的 `ul` 元素中應該有三個 `li` 元素。

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

你的 `ul` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

你的 `li` 元素應該有結束標籤。

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

你的 `li` 元素不應包含空字串或僅有空白字元。

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
