---
id: bad87fee1348bd9aedf08802
title: 取消註解 HTML
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

註解是一種讓你能在程式碼中為其他開發者留下評論的方法，而不會影響顯示給終端用戶的最終輸出。

註解也是讓程式碼變成非作用中狀態而不必完全刪除它的方便方法。

HTML 中的註解以 `<!--` 開始並以 `-->` 結束。

# --instructions--

取消註解你的 `h1`、`h2` 和 `p` 元素。

# --hints--

你的 `h1` 元素應該透過取消註解在頁面上可見。

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

你的 `h2` 元素應該透過取消註解在頁面上可見。

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

你的 `p` 元素應該透過取消註解在頁面上可見。

```js
assert.notEmpty(document.querySelectorAll('p'));
```

頁面上不應該顯示任何尾隨的註解標籤（例如 `-->`）。

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
