---
id: bad87fee1348bd9aedf08804
title: 將 HTML 註解掉
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

請記住，為了開始一個註解，你需要使用 `<!--`，而為了結束一個註解，你需要使用 `-->`。

在你的 `h2` 元素開始之前，你需要結束註解。

# --instructions--

將你的 `h1` 元素和你的 `p` 元素註解掉，但不要註解你的 `h2` 元素。

# --hints--

你的 `h1` 元素應該被註解掉，這樣它在頁面上就不會顯示。

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

你的 `h2` 元素不應該被註解掉，以便它能在頁面上顯示。

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

你的 `p` 元素應該被註解掉，這樣它才不會在頁面上顯示。

```js
assert.isEmpty(document.querySelectorAll('p'));
```

你的每個註解都應該以 `-->` 結尾。

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

你不應該更改程式碼中 `h1`、`h2` 或 `p` 元素的順序。

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
