---
id: bad87fee1348bd9aecf08801
title: HTML5 元素介紹
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 介紹了更多具描述性的 HTML 頁籤，包括 `main`、`header`、`footer`、`nav`、`video`、`article`、`section` 等。

這些頁籤為你的 HTML 提供描述性的結構，使你的 HTML 更容易閱讀，並有助於搜尋引擎優化（SEO）和存取。`main` HTML5 頁籤幫助搜尋引擎和其他開發者找到你頁面的主要內容。

範例用法，一個 `main` 元素內巢狀兩個子元素：

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**注意：** 許多新的 HTML5 頁籤及其利益已在應用無障礙章節中涵蓋。

# --instructions--

創建第二個 `p` 元素，並使用以下 kitty ipsum 文字：`Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

然後，創建一個 `main` 元素，並且只將兩個 `p` 元素巢狀在 `main` 元素內。

# --hints--

你應該有 2 個帶有 Kitty Ipsum 文字的 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

你的每個 `p` 元素都應該有一個結束標籤。

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

你的 `p` 元素應包含所提供的額外 `kitty ipsum` 文字的前幾個字。

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

你的程式碼應該有一個 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` 元素應該有兩個段落元素作為子元素。

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

開啟的 `main` 頁籤應該位於第一個段落頁籤之前。

```js
assert.match(code,(/<main>\s*?<p>/g));
```

關閉的 `main` 頁籤應該在第二個關閉的段落頁籤之後。

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
