---
id: bad87fed1348bd9aedf08833
title: 刪除 HTML 元素
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ckK73C9'
forumTopicId: 17559
dashedName: delete-html-elements
---

# --description--

我們的手機沒有太多垂直空間。

讓我們移除不必要的元素，這樣我們就可以開始建置我們的 CatPhotoApp。

# --instructions--

刪除你的 `h1` 元素，以便我們可以簡化視圖。

# --hints--

你的 `h1` 元素應該被刪除。

```js
assert.notMatch(code,/<h1>/gi);
assert.notMatch(code,/<\/h1>/gi);
```

你的 `h2` 元素應該在頁面上。

```js
assert.match(code,/<h2>[\w\W]*<\/h2>/gi);
```

你的 `p` 元素應該在頁面上。

```js
assert.match(code,/<p>[\w\W]*<\/p>/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2><p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
