---
id: bad87fee1348bd9aedf08801
title: 使用段落元素進行通知
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

`p` 元素是在網站上用於段落文字的首選元素。`p` 是 "paragraph" 的縮寫。

你可以這樣創建一個段落元素：

```html
<p>I'm a p tag!</p>
```

# --instructions--

在你的 `h2` 元素下方創建一個 `p` 元素，並給它文字 `Hello Paragraph`。

**注意：** 作為慣例，所有 HTML 頁籤皆以小寫字母書寫，例如 `<p></p>` 而非 `<P></P>`。

# --hints--

你的程式碼應該有一個有效的 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

你的 `p` 元素應該有文字 `Hello Paragraph`。

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

你的 `p` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
