---
id: bad87fee1348bd9aedf08830
title: 為你的 `Text Field` 添加 Placeholder 文字
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

佔位文字是在你的 `input` 元素中，使用者尚未輸入任何內容前所顯示的文字。

你可以這樣創建佔位符文字：

```html
<input type="text" placeholder="this is placeholder text">
```

**注意：** 請記得 `input` 是一個空元素。

# --instructions--

將你的文字 `input` 的 `placeholder` 值設為 "cat photo URL"。

# --hints--

你應該為現有的文字 `input` 元素添加一個 `placeholder` 屬性。

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

你應該將你的 `placeholder` 屬性的值設為 `cat photo URL`。

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

完成的 `input` 元素不應該有結束標籤。

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

完成的 `input` 元素應該具有有效的語法。

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
