---
id: bad87fee1348bd9aedf08816
title: 使用錨點元素連結到外部頁面
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

你可以使用 `a` (*anchor*) 元素來連結到你的網頁外的內容。

`a` 元素需要一個稱為 `href` 屬性的目標網頁位址。它們也需要錨點文字。以下是一個範例：

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

然後你的瀏覽器會將文字 `this links to freecodecamp.org` 顯示為你可以點擊的連結。而該連結會帶你到網路位址 `https://www.freecodecamp.org`。

# --instructions--

創建一個 `a` 元素，該元素聯結到 `https://www.freecatphotoapp.com` 並且其錨點文字為 "cat photos"。

# --hints--

你的 `a` 元素應該有 `cat photos` 的錨點文字。

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

你需要一個 `a` 元素，該元素聯結到 `https://www.freecatphotoapp.com`

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

你的 `a` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
