---
id: bad87fee1348bd9aedf08828
title: 創建有序列表
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cQ3B8TM'
forumTopicId: 16824
dashedName: create-an-ordered-list
---

# --description--

HTML 有另一個特殊元素用於建立 <dfn>有序列表</dfn>，或編號列表。

有序列表以開啟的 `<ol>` 元素開始，接著是任意數量的 `<li>` 元素。最後，有序列表以 `</ol>` 頁籤關閉。

例如：

```html
<ol>
  <li>Garfield</li>
  <li>Sylvester</li>
</ol>
```

會創建一個包含 `Garfield` 和 `Sylvester` 的編號列表。

# --instructions--

創建一個包含貓咪最討厭的前三件事的有序列表。

# --hints--

你應該有一個有序的列表，用於 `Top 3 things cats hate:`

```js
const previousElement = document.querySelector('ol').previousElementSibling; 
assert.match(previousElement.textContent,/Top 3 things cats hate:/i);
```

你應該為 `Things cats love:` 準備一個無序列表。

```js
const previousElement = document.querySelector('ul').previousElementSibling; 
assert.match(previousElement.textContent,/Things cats love:/i);
```

你應該只有一個 `ul` 元素。

```js
assert.lengthOf(document.querySelectorAll('ul'), 1);
```

你應該只有一個 `ol` 元素。

```js
assert.lengthOf(document.querySelectorAll('ol'), 1);
```

你的 `ul` 元素中應該有三個 `li` 元素。

```js
assert.lengthOf(document.querySelectorAll('ul li'), 3);
```

你的 `ol` 元素中應該有三個 `li` 元素。

```js
assert.lengthOf(document.querySelectorAll('ol li'), 3);
```

你的 `ul` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/ul>/g);
assert.strictEqual(code.match(/<\/ul>/g).length ,code.match(/<ul>/g).length);
```

你的 `ol` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/ol>/g);
assert.strictEqual(code.match(/<\/ol>/g).length ,code.match(/<ol>/g).length);
```

你的 `li` 元素應該有一個結束標籤。

```js
assert.match(code,/<\/li>/g);
assert.match(code,/<li>/g);
assert.strictEqual(code.match(/<\/li>/g).length ,code.match(/<li>/g).length);
```

你無序列表中的 `li` 元素不應該是空的。

```js
[...document.querySelectorAll('ul li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

你有序列表中的 `li` 元素不應該是空的。

```js
[...document.querySelectorAll('ol li')].forEach((val) =>
  assert.isNotEmpty(__helpers.removeWhiteSpace(val.textContent))
);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>

</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>hate 1</li>
    <li>hate 2</li>
    <li>hate 3</li>
  </ol>
</main>
```
