---
id: bad87fee1348bd9aede08817
title: 將錨點元素巢狀於段落中
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

你可以在其他文字元素中巢狀嵌套聯結。

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

讓我們拆解範例。一般文字包裹在 `p` 元素中：

```html
<p> Here's a ... for you to follow. </p>
```

接下來是 *anchor* 元素 `<a>`（需要閉合標籤 `</a>`）：  

```html
<a> ... </a>
```

`target` 是一個指定在哪裡開啟聯結的錨點標籤屬性。值 `_blank` 指定在新分頁中開啟聯結。`href` 是一個包含聯結 URL 位址的錨點標籤屬性：  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

`a` 元素內的文字 `link to www.freecodecamp.org` 稱為 <dfn>錨點文字</dfn>，並會顯示可點擊的連結：

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

範例的最終輸出將會是這樣：  

Here's a <a href="https://www.freecodecamp.org" target="_blank">link to www.freecodecamp.org</a> for you to follow.

# --instructions--

將現有的 `a` 元素巢狀於新的 `p` 元素中。不要創建新的錨點標籤。新的段落應該有文字顯示 `View more cat photos`，其中 `cat photos` 是一個連結，其餘部分是純文字。

# --hints--

你應該只擁有一個 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

`a` 元素應該聯結到 "`https://www.freecatphotoapp.com`"。

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

你的 `a` 元素應該有 `cat photos` 的錨點文字

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

你應該創建一個新的 `p` 元素。你的 HTML 程式碼中應該至少有 3 個 `p` 頁籤。

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

你的 `a` 元素應該巢狀在你的新 `p` 元素內。

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

你的 `p` 元素應該有文字 `View more `（後面有一個空格）。

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

你的 `a` 元素 <em>不應該</em> 有文字 `View more`。

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

你的每個 `p` 元素都應該有一個結束標籤。

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

你的每個 `a` 元素都應該有一個結束標籤。

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
