---
id: 680a82806c4c6bc3567f2708
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Tangaza kigezo kipya kinachoitwa `text` na weka thamani yake kuwa mfuatano wa herufi `hello world`. Hii itakuwa ujumbe wa kuficha.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `hello world` kwa `text`. Kumbuka kufunga mfuatano wa herufi kwa alama za nukuu moja au alama za nukuu maradufu na zingatia herufi kubwa na ndogo.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
