---
id: 680a41ef5bfdc059cfd9d4b5
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Tangaza kigezo kingine kinachoitwa `alphabet` na weka thamani ya mfuatano wa herufi `abcdefghijklmnopqrstuvwxyz` kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `alphabet`. Zingatia kuweka jina la kigezo mwanzoni mwa mstari.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `abcdefghijklmnopqrstuvwxyz` kwa kigezo chako cha `alphabet`. Kumbuka kutumia alama za nukuu moja au alama za nukuu maradufu kufunga mfuatano wa herufi.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
