---
id: 680a3d4c065557260de2ff33
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda programu inayoitwa caesar cipher ambayo ni njia ya usimbaji inayobadilisha herufi katika alfabeti ili kuficha ujumbe.

Anza kwa kutangaza kigezo kinachoitwa `shift` na weka thamani `5` kwa kigezo chako kipya.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `shift`. Zingatia kuweka jina la kigezo mwanzoni mwa mstari.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Unapaswa kuweka thamani `5` kwa kigezo chako cha `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
