---
id: 680a431692a5f863f726c438
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Katika warsha hii, utaunda *Caesar cipher*. Hii ni mojawapo ya mbinu rahisi zaidi za kusimba maandishi, ambayo inahusisha kubadilisha kila herufi ya maandishi ya kawaida na herufi inayopatikana kwa nafasi fulani iliyowekwa chini katika alfabeti. Kwa mfano, kwa mabadiliko ya 5, `a` itabadilishwa kuwa `f`, `b` kuwa `g` na kadhalika.

Ili kutekeleza cipher hii, utahitaji kuunda toleo jipya la alfabeti yako ambalo linaanza katika nafasi inayotangazwa na mabadiliko. Kama ulivyojifunza katika funzo la awali, unaweza kutoa sehemu ya mfuatano wa herufi kwa kutumia kukata mfuatano wa herufi (string slicing):

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Tengeneza kigezo kinachoitwa `shifted_alphabet` na tumia sintaksia ya kukata mfuatano wa herufi ili kuweka sehemu ya `alphabet` inayotoka kwenye kielezo cha `shift`. Kisha, fanya wito kwa kitendakazi iliyo ndani (built-in) cha `print()` ili uchapishe `shifted_alphabet` kwenye terminal na uangalie matokeo.

# --hints--

Unapaswa kutangaza kigezo kipya kinachoitwa `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Unapaswa kuweka thamani ya `alphabet[shift:]` kwa `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Unapaswa kufanya wito kwa kitendakazi cha `print()` ukipitisha `shifted_alphabet` kama hoja.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
