---
id: 680a5e7f3d47715196b12ca0
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa baada ya kuwa na alfabeti yako yote iliyohamishwa, ondoa wito la `print(shifted_alphabet)`.

# --hints--

Unapaswa kuondoa `print(shifted_alphabet)` kutoka kwenye msimbo wako.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
