---
id: 680a549f60f5a3fedfc0edd9
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Kama unavyoona kutoka kwenye matokeo, alfabeti iliyosogezwa inaanza kwa herufi `f` kwa sababu `shift` ina thamani `5`. Lakini sasa herufi tano za kwanza za alfabeti – `a`, `b`, `c`, `d` na `e` – zinakosekana katika alfabeti iliyosogezwa, hivyo utahitaji kuziongeza mwishoni mwa alfabeti iliyosogezwa. Unaweza kutumia opereta `+` kuunganisha sehemu iliyokosekana ya alfabeti.

Badilisha mgawo wa thamani uliopo wa `shifted_alphabet`: tumia sintaksia ya kukata (slicing) ili kutoa sehemu ya kwanza iliyokosekana ya `alphabet` na kuiongeza kwenye `alphabet[shift:]`.

Kama ukumbusho, `sentence[start:stop]` hurejesha herufi za `sentence` kuanzia nafasi `start` (ikiwa ni pamoja) hadi `stop` (isipokuwa).

# --hints--

Haupaswi kubadilisha moja kwa moja mgawo wa thamani wa `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Unapaswa kuweka thamani ya muunganiko wa `alphabet[shift:]` na sehemu ya kwanza iliyokosekana ya `alphabet` kwenye mgawo wa thamani wa `shifted_alphabet` kwenye mstari uleule ulipotangazwa. Tumia `shift` kuonyesha mpaka wa kukata.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
