---
id: 680a836235461fce1ff53b3f
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Njia ya `translate()` inachukua kama hoja jedwali la tafsiri lililotengenezwa na `str.maketrans()`. Inaitwa kwenye mfuatano wa herufi na inarudisha nakala ya mfuatano wa herufi wa awali ambapo herufi zimebadilishwa kulingana na jedwali la tafsiri:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Fanya wito wa njia ya `translate()` kwenye `text` ukipitisha `translation_table` kama hoja na weka thamani ya matokeo kwenye kigezo kinachoitwa `encrypted_text`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Unapaswa kuweka `text.translate(translation_table)` kwa `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
