---
id: 680a3d4c065557260de2ff33
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaunda programu inayotekeleza Caesar cipher, njia ya usimbaji inayobadilisha herufi katika alfabeti ili kuficha ujumbe.
Anza kwa kuunda kigezo kinachoitwa `shift` na weka thamani `5` kwa kigezo chako kipya.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `shift`. Hakikisha jina la kigezo lipo mwanzoni mwa mstari.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Unapaswa kuweka thamani `5` kwa kigezo chako cha `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
