---
id: 680a7e200400089967b8b1f4
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Njia ya `str.maketrans()` inachukua mfuatano wa herufi miwili wenye urefu sawa na kurudisha jedwali la tafsiri linaloonyesha kila herufi ya mfuatano wa kwanza kwa herufi inayolingana katika mfuatano wa pili. Kila herufi katika jedwali la tafsiri huhifadhiwa kama nambari ya Unicode, nambari inayotambulisha herufi hiyo kwa kipekee.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Utajifunza zaidi kuhusu aina ya muundo unaorejeshwa na `str.maketrans()` baadaye katika mtaala. Kwa sasa, tengeneza jedwali la tafsiri linaloonyesha herufi katika `alphabet` kwa herufi katika `shifted_alphabet` na uweke kwenye kigezo kinachoitwa `translation_table`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Unapaswa kuita `str.maketrans()` ukipitisha `alphabet` na `shifted_alphabet` kama hoja, na weka matokeo kwenye `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
