---
id: 6810891bd6de8a87a833df42
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Kama unavyoona kutoka kwenye matokeo, ujumbe umefichwa kwa usimbaji. Hatua inayofuata itakuwa kufanya msimbo wako uweze kutumika tena ikiwa ungetaka kusimba ujumbe tofauti.

Kwa ajili ya hilo, utahitaji kuunda kitendakazi. Kama ukumbusho, hivi ndivyo unavyounda kitendakazi kinachoitwa `spam` kinachochapisha `Spam!` kwenye terminal:

```py
def spam():
    print('Spam!')
```

Unda kitendakazi kinachoitwa `caesar`. Weka msimbo wote uliouandika ndani ya mwili wa kitendakazi hicho. Zingatia kuweka kiwango sawa cha kuingizwa kwa mistari yote ndani ya mwili wa kitendakazi.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

Unapaswa kuhamisha msimbo wote uliouandika hadi sasa ndani ya mwili wa kitendakazi cha `caesar`. Hakikisha unaweka kiwango sawa cha kuingizwa kwa mistari yote ndani ya mwili wa kitendakazi...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
