---
id: 6818fcb250f34e62cab32c39
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Huenda umeona kwamba, ingawa ujumbe umefichwa, herufi kubwa hazijagusiwa. Hii hutokea kwa sababu jedwali la tafsiri halijumuishi herufi kubwa.

Ili kurekebisha hili, utahitaji kubadilisha mfuatano wa herufi unaotumwa kwa wito wako wa `str.maketrans()` kwa kutumia njia ya `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Sahihisha wito wako wa `str.maketrans()` kwa kuunganisha kwenye kila hoja toleo la herufi kubwa la hoja hiyo yenyewe.

# --hints--

Unapaswa kubadilisha hoja za wito wako wa `str.maketrans()` kwa kuunganisha kwenye kila hoja toleo la herufi kubwa la hoja hiyo yenyewe.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
