---
id: 6818b80519e6c6c47a8bb6e8
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Kama ulivyojifunza katika funzo lililopita, vitendakazi vinaweza kuwa na vigezo, ambavyo ni vigezo vinavyoweza kurejea ndani ya kitendakazi. Hapa kuna kitendakazi cha `sum` chenye vigezo viwili, `num1` na `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Ujumbe wa kuficha na mabadiliko bado umeandikwa moja kwa moja ndani ya kitendakazi chako. Toa vigezo viwili kwa kitendakazi chako vinavyoitwa `text` na `shift`, na futa tamko la vigezo vya `text` na `shift` kutoka ndani ya mwili wa kitendakazi cha `caesar`.

# --hints--

Kitendakazi chako cha `caesar` kinapaswa kuwa na vigezo viwili, `text` na `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Haupaswi kuwa na kigezo cha `text` kilichotangazwa ndani ya mwili wa kitendakazi cha `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Haupaswi kuwa na kigezo cha `shift` kilichotangazwa ndani ya mwili wa kitendakazi cha `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
