---
id: 6818e9c7c8dc5694e8f64fc3
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa, jaribu kitendakazi cha `caesar` kwa kukiwita na mfuatano wa herufi `freeCodeCamp` na nambari `3` kama hoja. Weka wito wa kitendakazi hicho katika kigezo kinachoitwa `encrypted_text`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `encrypted_text` katika wigo wa ulimwengu.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Unapaswa kuweka thamani ya `caesar('freeCodeCamp', 3)` katika kigezo cha `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
