---
id: 6818eb3f98c0bca5af7c8d03
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Sasa msimbo wako unaweza kutumika tena. Hata hivyo, kitendakazi cha `caesar` kinachapisha ujumbe kwenye terminal na kinarejesha `None` kwa chaguo-msingi. Ili kuthibitisha hilo, chapisha `encrypted_text` mwishoni mwa msimbo wako.

# --hints--

Unapaswa kuchapisha `encrypted_text` nje ya mwili wa kitendakazi.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
