---
id: 6819be9055828d65c813eec5
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Kitendakazi cha `isinstance()` hurudisha `True` ikiwa hoja yake ya kwanza ni mfano wa hoja ya pili, na `False` vinginevyo:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Badilisha sharti la sasa la tamko lako la `if` kwa wito wa `isinstance()`. Pitia `shift` kama hoja ya kwanza, na `int` kama hoja ya pili.

# --hints--

Tamko lako la `if` linapaswa kutumia `isinstance(shift, int)` kama sharti lake.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
