---
id: 6819c5c9faf2548e556163d5
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Kama unavyokumbuka kutoka kwa mafundisho ya awali, opereta `not` ni opereta wa mantiki wa unary ambaye hukataa usemi wa kihesabu:

```py
print(not True) # False
print(not False) # True
```

Tumia opereta `not` kurekebisha hali ya tamko la `if` ili kitendakazi chako kirudishe `Shift must be an integer value.` wakati `shift` si nambari kamili.

# --hints--

Unapaswa kutumia opereta `not` kukataa `isinstance(shift, int)` katika tamko lako la `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
