---
id: 681a44f4c3235f7d8f428545
title: Hatua 25
challengeType: 20
dashedName: step-25
---

# --description--

Sasa utaanza kujaribu kitendakazi cha `decrypt`. Badilisha thamani imepewa `encrypted_text` na mfuatano wa herufi ufuatao, unaowakilisha ujumbe wa kufungua: `Pbhentr vf sbhaq va hayvxryl cynprf.`

Kisha, tangaza kigezo kinachoitwa `decrypted_text` na kimupe wito wa `decrypt` ukiwa na `encrypted_text` kama hoja yake ya kwanza na mabadiliko ya `13` kama hoja ya pili.

Mwishowe, chapisha `decrypted_text` kwenye terminali. Kwa kufanya hivyo, usimbaji wa Caesar utakuwa umekamilika.

# --hints--

Unapaswa kuweka `Pbhentr vf sbhaq va hayvxryl cynprf.` kuwa `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

Unapaswa kuwa na kigezo kinachoitwa `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

Unapaswa kuweka `decrypt(encrypted_text, 13)` kuwa `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

Unapaswa kuchapisha `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
