---
id: 63b607af6fcdb119aae9b16a
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Katika fomu yako, watumizi wanapaswa kuwa na uwezo wa kuongeza aina mbalimbali za milo pamoja na idadi ya kalori zao.

Unda kipengele cha `fieldset` chenye `id` iliyowekwa kuwa `breakfast`.

Ndani ya kipengele hicho, unda `legend` yenye maandishi ya `Breakfast`, na `div` tupu yenye `class` iliyowekwa kuwa `input-container`.

# --hints--

Unapaswa kuunda kipengele cha `fieldset` ndani ya `form` yako.

```js
assert.exists(document.querySelector('form fieldset'));
```

Kipengele chako cha `fieldset` kinapaswa kuja baada ya kipengele chako cha `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

Kipengele chako cha `fieldset` kinapaswa kuwa na `id` iliyowekwa kuwa `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Unapaswa kuunda kipengele cha `legend` ndani ya `fieldset` yako.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

Kipengele chako cha `legend` kinapaswa kuwa na maandishi ya `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Unapaswa kuunda kipengele cha `div` ndani ya `fieldset` yako.

```js
assert.exists(document.querySelector('form fieldset div'));
```

Kipengele chako cha `div` kinapaswa kuja baada ya kipengele chako cha `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

Kipengele chako cha `div` kinapaswa kuwa na `class` iliyowekwa kuwa `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
