---
id: 63b60aaaa65f8922bfce6b7e
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Katika kipengele chako cha `span`, tengeneza kipengele cha `label` kwa ajili ya `entry-dropdown` na kipa maandishi ya `Add food or exercise:`. Kisha tengeneza kipengele cha `select` chenye sifa ya `id` iliyowekwa kuwa `entry-dropdown` na sifa ya `name` iliyowekwa kuwa `options`. Chini yake, ongeza kipengele cha `button` chenye sifa ya `id` iliyowekwa kuwa `add-entry` na maandishi ya `Add Entry`.

Toa kipengele chako cha `button` sifa ya kipengele ya `type` iliyowekwa kuwa `button` ili kuzuia uwasilishaji wa fomu kiotomatiki.

# --hints--

Unapaswa kuongeza kipengele cha `label` kwenye kipengele chako cha `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Kipengele chako kipya cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` iliyowekwa kuwa `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Kipengele chako kipya cha `label` kinapaswa kuwa na maandishi ya `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Unapaswa kuongeza kipengele cha `select` kwenye kipengele chako cha `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Kipengele chako cha `select` kinapaswa kuja baada ya kipengele chako cha `label`.

```js
assert.equal(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName, 'LABEL');
```

Kipengele chako kipya cha `select` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Kipengele chako kipya cha `select` kinapaswa kuwa na sifa ya kipengele ya `name` iliyowekwa kuwa `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Kipengele chako cha `select` kinapaswa kuwa tupu na kuwa na lebo ya kufunga ya `</select>`.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Unapaswa kuongeza kipengele cha `button` kwenye kipengele chako cha `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Kipengele chako cha `button` kinapaswa kuja baada ya kipengele chako cha `select`.

```js
assert.equal(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName, 'SELECT');
```

Kipengele chako kipya cha `button` kinapaswa kuwa na sifa ya kipengele ya `id` iliyowekwa kuwa `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Kipengele chako kipya cha `button` kinapaswa kuwa na maandishi ya `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
