---
id: 6976db7ecfa770bf21307b20
title: Erstellen Sie eine Remix-Playlist-Engine
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

In diesem Labor erstellen Sie ein Programm, das aus mehreren von Zuhörern eingereichten Playlists eine einzige Remix-Playlist erstellt.

Jeder Zuhörer liefert eine Liste von Songs, die er hören möchte. Einige Songs können mehrfach vorkommen, und einige Künstler können zu oft vertreten sein. Ihre Aufgabe ist es, diese Playlists Schritt für Schritt zu verarbeiten: Sie kombinieren sie zu einer Liste, bewerten jeden Song, entfernen doppelte Songs, begrenzen, wie oft derselbe Künstler erscheint, und erstellen dann eine endgültige Wiedergabereihenfolge.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `flattenPlaylists` erstellen, die ein Array von Playlists akzeptiert, wobei jede Playlist ein Array von Objekten mit den folgenden Eigenschaften ist: `trackId`, `artist`, `title`, `votes`, `bpm`. Wenn die Eingabe kein Array ist, sollte `flattenPlaylists` ein leeres Array zurückgeben. Ein Beispiel für eine Playlist wurde für Sie bereitgestellt. Sie können dieses Beispiel verwenden, um Ihre Funktion zu testen.

2. `flattenPlaylists` sollte ein flaches Array von Track-Objekten zurückgeben, wobei jedes Objekt alle ursprünglichen Track-Eigenschaften plus eine `source`-Eigenschaft enthält, die ein Array mit dem Playlist-Index und dem Track-Index angibt, woher der Track stammt.

3. Sie sollten eine Funktion namens `scoreTracks` erstellen, die ein Array von Track-Objekten akzeptiert, wie es von `flattenPlaylists` zurückgegeben wird (jedes mit den Eigenschaften `trackId`, `artist`, `title`, `votes`, `bpm` und `source`) und ein neues Array von Track-Objekten zurückgibt, wobei jedem eine `score`-Eigenschaft hinzugefügt wird, die mit der Formel `votes * 10 - Math.abs(bpm - 120)` berechnet wird.

4. Sie sollten eine Funktion namens `dedupeTracks` erstellen, die ein Array von Track-Objekten akzeptiert, wie es von `scoreTracks` zurückgegeben wird, und ein neues Array zurückgibt, in dem doppelte `trackId`-Einträge entfernt wurden, wobei nur das erste Vorkommen jedes Eintrags erhalten bleibt.

5. Sie sollten eine Funktion namens `enforceArtistQuota` erstellen, die ein Array von Track-Objekten akzeptiert, wie es von `dedupeTracks` zurückgegeben wird, und eine Zahl, die die maximal erlaubte Anzahl von Vorkommen pro Künstler angibt. Die Funktion sollte ein neues Array zurückgeben, in dem kein Künstler öfter als die angegebene Anzahl erscheint, wobei die frühesten Vorkommen erhalten bleiben.

6. Sie sollten eine Funktion namens `buildSchedule` erstellen, die ein Array von Track-Objekten akzeptiert, wie es von `enforceArtistQuota` zurückgegeben wird, und ein neues Array von `{ slot, trackId }`-Objekten zurückgibt, wobei `slot` ein 1-basierter Index ist, der die Position jedes Tracks in der Sendereihenfolge darstellt.

7. Sie sollten eine Funktion namens `remixPlaylist` erstellen, die ein Array von Playlists und die maximale Anzahl erlaubter Vorkommen pro Künstler akzeptiert. Die Funktion sollte den endgültigen Sendungsplan als Array von `{ slot, trackId }`-Objekten zurückgeben, indem sie `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` und `buildSchedule` der Reihe nach aufruft.

# --hints--

Sie sollten eine Funktion namens `flattenPlaylists` haben.

```js
assert.isFunction(flattenPlaylists);
```

Sie sollten von `flattenPlaylists` ein leeres Array zurückgeben, wenn die Eingabe kein Array ist.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Jeder Track, der von `flattenPlaylists` zurückgegeben wird, sollte ein `source`-Feld enthalten, das ein Array mit dem Playlist-Index und dem Track-Index ist.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Sie sollten eine Funktion namens `scoreTracks` haben.

```js
assert.isFunction(scoreTracks)
```

Jeder Track, der von `scoreTracks` zurückgegeben wird, sollte ein numerisches `score`-Feld enthalten.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Sie sollten `score` mit einem Ziel-BPM von `120` und folgender Formel berechnen: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Sie sollten eine Funktion namens `dedupeTracks` haben.

```js
assert.isFunction(dedupeTracks)
```

Wenn doppelte `trackId`-Werte existieren, sollte `dedupeTracks` nur das erste Vorkommen des Tracks behalten.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Sie sollten eine Funktion namens `enforceArtistQuota` haben.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` sollte sicherstellen, dass kein Künstler mehr als `maxPerArtist` Mal erscheint, indem zusätzliche Tracks entfernt werden, während die frühesten erhalten bleiben.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Sie sollten eine Funktion namens `buildSchedule` haben.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` sollte ein Array von Objekten mit der Form `{ slot, trackId }` zurückgeben, wobei `slot` bei `1` beginnt.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Sie sollten eine Funktion namens `remixPlaylist` haben.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` sollte die Hilfsfunktionen der Reihe nach aufrufen, um den endgültigen Plan zu erstellen.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
