---
id: 68de40f53fcaf106f51cc20f
title: Jenga mchezaji wa sauti na video wa HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele cha `h1` kwa kichwa kikuu cha ukurasa.
2. Unapaswa kuwa na vipengele viwili vya `section`.
3. Ndani ya kipengele cha kwanza cha `section`, unapaswa kuwa na kipengele cha `h2` kwa kichwa cha video inayochezwa.
4. Chini ya kipengele cha `h2`, unapaswa kuwa na kipengele cha `video` chenye sifa za `controls` na `width`. Sifa ya `width` inapaswa kuwekwa kuwa `640`.
5. Ndani ya kipengele cha `video`, unapaswa kuwa na kipengele cha `source` chenye sifa ya `src` inayorejelea faili la video na sifa ya `type`.
   - Unaweza kutumia `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Ndani ya kipengele cha pili cha `section`, unapaswa kuwa na kipengele cha `h2` kwa kichwa cha wimbo unaochezwa.
7. Chini ya kipengele cha `h2`, unapaswa kuwa na kipengele cha `audio` chenye sifa za `controls` na `loop`, na sifa ya `src` inayorejelea faili la sauti.
   - Unaweza kutumia `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - Au `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Unapaswa kuwa na kipengele cha `h1` kwa kichwa kikuu cha ukurasa.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

`h1` yako haipaswi kuwa tupu.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Unapaswa kuwa na vipengele viwili vya `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Ndani ya kipengele cha kwanza cha `section`, unapaswa kuwa na kipengele cha `h2` kwa kichwa cha video inayochezwa.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Chini ya kipengele cha `h2`, unapaswa kuwa na kipengele cha `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

Kipengele cha `video` kinapaswa kuwa na sifa ya `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

Kipengele cha `video` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Ndani ya kipengele cha `video`, unapaswa kuwa na kipengele cha `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

Kipengele cha `source` kinapaswa kuwa na sifa ya `src` inayowakilisha faili la video.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

Kipengele cha `source` kinapaswa kuwa na sifa ya kipengele ya `type`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

Sifa ya `type` inapaswa kuendana na aina ya media ya kiendelezi cha faili cha sifa ya `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Ndani ya kipengele cha pili cha `section`, unapaswa kuwa na kipengele cha `h2` kwa kichwa cha wimbo unaochezwa.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Vipengele vyote vya `h2` vinapaswa kuwa na maandishi fulani.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Chini ya kipengele cha `h2`, unapaswa kuwa na kipengele cha `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

Kipengele cha `audio` kinapaswa kuwa na sifa ya `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

Kipengele cha `audio` kinapaswa kuwa na sifa ya `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

Kipengele cha `audio` kinapaswa kuwa na sifa ya `src` inayorejelea faili la sauti.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
