---
id: 6814d8e1516e86b171929de4
title: "Challenge 1: Vowel Balance"
challengeType: 29
dashedName: challenge-1
---

# --description--

문자열이 주어졌을 때, 문자열 앞부분 절반의 모음 개수가 뒷부분 절반의 모음 개수와 같은지 확인하세요.

- 문자열에는 어떤 문자도 포함될 수 있습니다.
- `a`, `e`, `i`, `o`, `u` 문자는 대소문자 구분 없이 모음으로 간주합니다.
- 문자열 길이가 홀수일 경우 가운데 문자는 무시하세요.

# --hints--

`is_balanced("racecar")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("racecar"), True)`)
}})
```

`is_balanced("Lorem Ipsum")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Lorem Ipsum"), True)`)
}})
```

`is_balanced("Kitty Ipsum")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("Kitty Ipsum"), False)`)
}})
```

`is_balanced("string")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("string"), False)`)
}})
```

`is_balanced(" ")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced(" "), True)`)
}})
```

`is_balanced("abcdefghijklmnopqrstuvwxyz")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("abcdefghijklmnopqrstuvwxyz"), False)`)
}})
```

`is_balanced("123A#b!E&*456-o.U")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_balanced("123A#b!E&*456-o.U"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_balanced(s):

    return s
```

# --solutions--

```py
def is_balanced(s):
    vowels = set("aeiouAEIOU")
    n = len(s)
    half = n // 2

    first_half = s[:half]
    second_half = s[-half:]

    def count_vowels(sub):
        return sum(1 for char in sub if char in vowels)

    return count_vowels(first_half) == count_vowels(second_half)
```
