---
id: 681cb1b0dab50c87ddb2e518
title: "Challenge 7: Targeted Sum"
challengeType: 29
dashedName: challenge-7
---

# --description--

숫자 배열과 정수 목표값이 주어졌을 때, 배열에서 두 개의 고유한 숫자를 찾아 그 합이 목표값이 되도록 하세요. 두 숫자의 인덱스를 배열로 반환하거나, 두 숫자의 합이 목표값이 되지 않으면 `"Target not found"`를 반환하세요.

- 반환하는 배열은 인덱스가 오름차순이어야 합니다.

# --hints--

`find_target([2, 7, 11, 15], 9)` should return `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` should return `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` should return `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` should return `'Target not found'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
