---
id: 681cb1afdab50c87ddb2e516
title: "Challenge 5: Jbelmud Text"
challengeType: 29
dashedName: challenge-5
---

# --description--

문자열이 주어지면, 각 단어가 다음 조건에 따라 변환된 문자열을 반환하세요:

- 단어의 첫 글자와 마지막 글자는 그대로 유지합니다.
- 첫 글자와 마지막 글자 사이의 모든 글자는 알파벳 순서로 정렬합니다.
- 입력 문자열에는 구두점이 없으며 모두 소문자로만 이루어져 있습니다.

# --hints--

`jbelmu("hello world")` should return `"hello wlord"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("hello world"), "hello wlord")`)
}})
```

`jbelmu("i love jumbled text")` should return `"i love jbelmud text"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("i love jumbled text"), "i love jbelmud text")`)
}})
```

`jbelmu("freecodecamp is my favorite place to learn to code")` should return `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe")`)
}})
```

`jbelmu("the quick brown fox jumps over the lazy dog")` should return `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def jbelmu(text):

    return text
```

# --solutions--

```py
def jbelmu(text):
    words = text.split()
    jumbled = []

    for word in words:
        if len(word) <= 3:
            jumbled.append(word)
        else:
            first = word[0]
            last = word[-1]
            middle = ''.join(sorted(word[1:-1]))
            jumbled.append(first + middle + last)
    
    return ' '.join(jumbled)
```
