---
id: 681cb05adab50c87ddb2e513
title: "Challenge 2: Base Check"
challengeType: 29
dashedName: challenge-2
---

# --description--

숫자를 나타내는 문자열과 2부터 36까지의 정수 진법이 주어졌을 때, 해당 숫자가 그 진법에서 유효한지 판단하세요.

- 문자열에는 정수와 대문자 또는 소문자가 포함될 수 있습니다.
- 검사는 대소문자를 구분하지 않아야 합니다.
- 진법은 2에서 36 사이의 어떤 숫자도 될 수 있습니다.
- 모든 문자가 주어진 진법에서 유효한 숫자이면 숫자는 유효합니다.
- 진법별 유효 숫자 예시:
  - 진법 2: 0-1
  - 진법 8: 0-7
  - 진법 10: 0-9
  - 진법 16: 0-9 및 A-F
  - 진법 36: 0-9 및 A-Z

# --hints--

`is_valid_number("10101", 2)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10101", 2), True)`);
}})
```

`is_valid_number("10201", 2)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10201", 2), False)`)
}})
```

`is_valid_number("76543210", 8)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("76543210", 8), True)`)
}})
```

`is_valid_number("9876543210", 8)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 8), False)`)
}})
```

`is_valid_number("9876543210", 10)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 10), True)`)
}})
```

`is_valid_number("ABC", 10)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 10), False)`)
}})
```

`is_valid_number("ABC", 16)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 16), True)`)
}})
```

`is_valid_number("Z", 36)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("Z", 36), True)`)
}})
```

`is_valid_number("ABC", 20)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 20), True)`)
}})
```

`is_valid_number("4B4BA9", 16)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("4B4BA9", 16), True)`)
}})
```

`is_valid_number("5G3F8F", 16)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 16), False)`)
}})
```

`is_valid_number("5G3F8F", 17)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 17), True)`)
}})
```

`is_valid_number("abc", 10)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 10), False)`)
}})
```

`is_valid_number("abc", 16)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 16), True)`)
}})
```

`is_valid_number("AbC", 16)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("AbC", 16), True)`)
}})
```

`is_valid_number("z", 36)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("z", 36), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_number(n, base):

    return n
```

# --solutions--

```py
def is_valid_number(n, base):
    allChars = "0123456789abcdefghijklmnopqrstuvwxyz"
    newN = n.lower()

    availableChars = allChars[0:base]

    for char in newN:
        if char not in availableChars:
            return False

    return True
```
