---
id: 68216ef80f957572e7c340c5
title: "Challenge 12: Message Decoder"
challengeType: 29
dashedName: challenge-12
---

# --description--

비밀 메시지 문자열과 메시지를 인코딩하기 위해 문자를 이동시킨 횟수를 나타내는 정수가 주어지면, 디코딩된 문자열을 반환하세요.

- 양수는 메시지가 알파벳에서 앞으로 이동되었음을 의미합니다.
- 음수는 메시지가 알파벳에서 뒤로 이동되었음을 의미합니다.
- 대소문자를 구분하며, 디코딩된 문자는 인코딩된 문자와 같은 대소문자를 유지해야 합니다.
- 알파벳이 아닌 문자는 디코딩하지 않아야 합니다.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` should return `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` should return `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` should return `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```
