---
id: 6821ebda237de8297eaee792
title: "Challenge 16: Reverse Parenthesis"
challengeType: 29
dashedName: challenge-16
---

# --description--

올바르게 중첩된 소괄호를 포함하는 문자열이 주어졌을 때, 다음 규칙을 사용해 문자열의 디코딩된 버전을 반환하세요:

- 각 소괄호 쌍 안의 모든 문자는 뒤집혀야 합니다.
- 최종 결과에는 소괄호가 제거되어야 합니다.
- 소괄호가 중첩된 경우, 가장 안쪽 쌍부터 먼저 뒤집고 그 결과를 바깥 쌍의 뒤집기에 포함해야 합니다.
- 모든 소괄호가 균형 있게 맞춰져 있고 올바르게 중첩되어 있다고 가정하세요.

# --hints--

`decode("(f(b(dc)e)a)")` should return `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` should return `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
