---
id: 6821ebee237de8297eaee796
title: "Challenge 20: Array Duplicates"
challengeType: 29
dashedName: challenge-20
---

# --description--

정수 배열이 주어지면, 처음 배열에서 두 번 이상 나타나는 정수들로 이루어진 배열을 오름차순으로 정렬하여 반환하세요. 두 번 이상 나타나는 값이 없으면 빈 배열을 반환하세요.

- 반환하는 배열에는 각 값이 한 번만 포함되어야 합니다.

# --hints--

`find_duplicates([1, 2, 3, 4, 5])` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 5]), [])`)
}})
```

`find_duplicates([1, 2, 3, 4, 1, 2])` should return `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 1, 2]), [1, 2])`)
}})
```

`find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` should return `[-6, 0, 2, 4, 5, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_duplicates(arr):

    return arr
```

# --solutions--

```py
def find_duplicates(arr):
    seen = set()
    duplicates = set()

    for num in arr:
        if num in seen:
            duplicates.add(num)
        else:
            seen.add(num)

    return sorted(duplicates)
```
