---
id: 6821ebce237de8297eaee790
title: "Challenge 14: Character Battle"
challengeType: 29
dashedName: challenge-14
---

# --description--

당신의 군대와 상대 군대를 나타내는 두 개의 문자열이 주어졌을 때, 각 군대의 같은 위치에 있는 캐릭터들이 다음 규칙에 따라 전투를 벌입니다:

- 문자 `a-z`는 각각 `1-26`의 힘을 가집니다.
- 문자 `A-Z`는 각각 `27-52`의 힘을 가집니다.
- 숫자 `0-9`는 자신의 숫자 값만큼 힘을 가집니다.
- 그 외 모든 문자는 힘이 0입니다.
- 각 캐릭터는 한 번의 전투만 할 수 있습니다.

각 전투에서 더 강한 캐릭터가 승리합니다. 승리 횟수가 더 많은 군대가 전쟁에서 이깁니다. 다음 값을 반환하세요:

- 당신의 군대가 상대 군대보다 캐릭터 수가 더 많으면 `"Opponent retreated"`를 반환하세요.
- 상대 군대가 당신의 군대보다 캐릭터 수가 더 많으면 `"We retreated"`를 반환하세요.
- 당신의 군대가 더 많은 전투에서 이겼으면 `"We won"`를 반환하세요.
- 상대 군대가 더 많은 전투에서 이겼으면 `"We lost"`를 반환하세요.
- 두 군대가 같은 수의 전투에서 이겼으면 `"It was a tie"`를 반환하세요.

# --hints--

`battle("Hello", "World")` should return `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` should return `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` should return `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
