---
id: 6821ebdf237de8297eaee793
title: "Challenge 17: Unorder of Operations"
challengeType: 29
dashedName: challenge-17
---

# --description--

정수 배열과 문자열 연산자 배열이 주어지면, 왼쪽에서 오른쪽으로 순서대로 숫자에 연산을 적용하세요. 모든 숫자가 사용될 때까지 필요에 따라 연산을 반복하세요. 최종 결과를 반환하세요.

예를 들어 `[1, 2, 3, 4, 5]`와 `['+', '*']`이 주어지면, 표준 연산 순서를 무시하고 왼쪽에서 오른쪽으로 `1 + 2 * 3 + 4 * 5`를 평가한 결과를 반환하세요.

- 유효한 연산자는 `+`, `-`, `*`, `/`, 그리고 `%`입니다.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` should return `3`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3)`)
}})
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` should return `38`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38)`)
}})
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` should return `60`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60)`)
}})
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` should return `30`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30)`)
}})
```

`evaluate([33, 11, 29, 13], ['/', '-'])` should return `-2`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(evaluate([33, 11, 29, 13], ['/', '-']), -2)`)
}})
```

# --seed--

## --seed-contents--

```py
def evaluate(numbers, operators):

    return numbers
```

# --solutions--

```py
def do_math(a, b, operator):
    if operator == '+':
        return a + b
    elif operator == '-':
        return a - b
    elif operator == '*':
        return a * b
    elif operator == '/':
        return a / b
    else:
        return a % b

def evaluate(numbers, operators):
    total = numbers[0]

    for i in range(1, len(numbers)):
        operator = operators[(i - 1) % len(operators)]
        total = do_math(total, numbers[i], operator)

    return total
```
