---
id: 6821ebea237de8297eaee795
title: "Challenge 19: Candlelight"
challengeType: 29
dashedName: challenge-19
---

# --description--

처음 가진 촛불의 개수를 나타내는 정수와, 새 촛불을 만들기 위해 필요한 사용한 촛불의 개수를 나타내는 정수가 주어질 때, 만들고 태울 수 있는 만큼 모두 태운 후 사용한 촛불의 총 개수를 반환하세요.

예를 들어, 7개의 촛불이 있고 새 촛불을 만들기 위해 2개의 사용한 촛불이 필요하다면:

1. 7개의 촛불을 태워 7개의 사용한 촛불을 얻습니다,
2. 6개의 사용한 촛불을 재활용해 3개의 새 촛불을 만듭니다 (1개의 사용한 촛불이 남음),
3. 3개의 촛불을 태워 3개의 사용한 촛불을 더 얻습니다 (총 4개),
4. 4개의 사용한 촛불을 재활용해 2개의 새 촛불을 만듭니다,
5. 2개의 촛불을 태워 2개의 사용한 촛불을 얻습니다,
6. 2개의 사용한 촛불을 재활용해 1개의 새 촛불을 만듭니다,
7. 1개의 촛불을 태웁니다.

이 예시에서는 총 13개의 촛불을 태운 셈입니다.

# --hints--

`burn_candles(7, 2)` should return `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` should return `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` should return `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` should return `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` should return `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
