---
id: 68adce01c0e1144d0a902956
title: "도전 과제 25: 모음 반복기"
challengeType: 29
dashedName: challenge-25
---

# --description--

문자열이 주어지면, 각 모음이 이전에 만난 모음보다 한 번 더 반복되는 새로운 문자열 버전을 반환하세요. 예를 들어, 문장에서 첫 번째 모음은 변경하지 않습니다. 두 번째 모음은 연속으로 두 번 나타나야 합니다. 세 번째 모음은 연속으로 세 번 나타나야 하며, 이런 식으로 계속됩니다.

- `a`, `e`, `i`, `o`, `u` 문자는 대문자든 소문자든 모음으로 간주합니다.
- 원래 모음은 대소문자를 유지하세요.
- 반복된 모음은 소문자로 만드세요.
- 모음이 아닌 모든 문자는 원래 대소문자를 유지하세요.

# --hints--

`repeat_vowels("hello world")`는 `"helloo wooorld"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")`는 `"freeeCooodeeeeCaaaaamp"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")`는 `"AEeIiiOoooUuuuu"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")`는 `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
