---
id: 68adce01c0e1144d0a902958
title: "도전 과제 26: IPv4 유효성 검사기"
challengeType: 29
dashedName: challenge-26
---

# --description--

문자열이 주어졌을 때, 유효한 IPv4 주소인지 확인하세요. 유효한 IPv4 주소는 점(`.`)으로 구분된 네 개의 정수로 구성됩니다. 각 숫자는 다음 조건을 만족해야 합니다:

- 0부터 255 사이여야 합니다(포함).
- 선행 0이 없어야 합니다(예: 0은 허용되지만 01은 허용되지 않습니다).
- 숫자 문자만 허용됩니다.

# --hints--

`is_valid_ipv4("192.168.1.1")`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.1"), True)`)
}})
```

`is_valid_ipv4("0.0.0.0")`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("0.0.0.0"), True)`)
}})
```

`is_valid_ipv4("255.255.255.255")`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.255.255.255"), True)`)
}})
```

`is_valid_ipv4("208.67.222.222")`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("208.67.222.222"), True)`)
}})
```

`is_valid_ipv4("255.01.50.111")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.01.50.111"), False)`)
}})
```

`is_valid_ipv4("255.00.50.111")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("255.00.50.111"), False)`)
}})
```

`is_valid_ipv4("256.101.50.115")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("256.101.50.115"), False)`)
}})
```

`is_valid_ipv4("192.168.101.")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.101."), False)`)
}})
```

`is_valid_ipv4("192168145213")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192168145213"), False)`)
}})
```

`is_valid_ipv4("192.168.1.a")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.a"), False)`)
}})
```

`is_valid_ipv4("192.168.o.1")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.o.1"), False)`)
}})
```

`is_valid_ipv4("192.168.1.1a")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_ipv4("192.168.1.1a"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_ipv4(ipv4):

    return ipv4
```

# --solutions--

```py
def is_valid_ipv4(ipv4):
    parts = ipv4.split(".")    

    if len(parts) != 4:
        return False

    for part in parts:
        if not part.isdigit():
            return False

        num = int(part)
        if num < 0 or num > 255:
            return False

        if len(part) > 1 and part.startswith("0"):
            return False

    return True
```
