---
id: 68af0687ef34c76c28ffa54b
title: "Challenge 32: Reverse Sentence"
challengeType: 29
dashedName: challenge-32
---

# --description--

단어들로 이루어진 문자열이 주어지면, 단어들의 순서를 거꾸로 하여 새로운 문자열을 반환하세요. 예를 들어, 첫 번째 단어는 반환된 문자열의 끝에 와야 하고, 마지막 단어는 반환된 문자열의 시작에 와야 합니다.

- 주어진 문자열에서 단어들은 하나 이상의 공백으로 구분될 수 있습니다.
- 반환된 문자열에서는 단어 사이에 공백이 하나만 있어야 합니다.

# --hints--

`reverse_sentence("world hello")` should return `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` should return `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` should return `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` should return `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
