---
id: 68af0687ef34c76c28ffa54d
title: "Challenge 33: Screen Time"
challengeType: 29
dashedName: challenge-33
---

# --description--

일주일 동안의 시간을 나타내는 7개의 정수로 이루어진 입력 배열이 주어집니다. 각 정수는 해당 요일에 휴대폰 사용 시간(시간 단위)을 나타냅니다. 다음 조건에 따라 화면 사용 시간이 너무 많은지 판단하세요:

- 하루라도 10시간 이상 사용하면 너무 많습니다.
- 연속된 3일의 평균 사용 시간이 8시간 이상이면 너무 많습니다.
- 7일 전체 평균 사용 시간이 6시간 이상이면 너무 많습니다.

# --hints--

`too_much_screen_time([1, 2, 3, 4, 5, 6, 7])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 4, 5, 6, 7]), False)`)
}})
```

`too_much_screen_time([7, 8, 8, 4, 2, 2, 3])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([7, 8, 8, 4, 2, 2, 3]), False)`)
}})
```

`too_much_screen_time([5, 6, 6, 6, 6, 6, 6])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([5, 6, 6, 6, 6, 6, 6]), False)`)
}})
```

`too_much_screen_time([1, 2, 3, 11, 1, 3, 4])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 11, 1, 3, 4]), True)`)
}})
```

`too_much_screen_time([1, 2, 3, 10, 2, 1, 0])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([1, 2, 3, 10, 2, 1, 0]), True)`)
}})
```

`too_much_screen_time([3, 3, 5, 8, 8, 9, 4])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 3, 5, 8, 8, 9, 4]), True)`)
}})
```

`too_much_screen_time([3, 9, 4, 8, 5, 7, 6])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(too_much_screen_time([3, 9, 4, 8, 5, 7, 6]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def too_much_screen_time(hours):

    return hours
```

# --solutions--

```py
def too_much_screen_time(hours):
    for h in hours:
        if h >= 10:
            return True

    for i in range(5):
        avg3 = sum(hours[i:i+3]) / 3
        if avg3 >= 8:
            return True

    weekly_avg = sum(hours) / 7
    if weekly_avg >= 6:
        return True

    return False
```
