---
id: 68b06e589bf2273243814771
title: "Challenge 36: Thermostat Adjuster"
challengeType: 29
dashedName: challenge-36
---

# --description--

현재 방 온도와 목표 온도가 주어졌을 때, 다음 조건에 따라 방 온도를 어떻게 조절해야 하는지 나타내는 문자열을 반환하세요:

- 현재 온도가 목표 온도보다 낮으면 `"heat"`를 반환하세요.
- 현재 온도가 목표 온도보다 높으면 `"cool"`를 반환하세요.
- 현재 온도가 목표 온도와 같으면 `"hold"`를 반환하세요.

# --hints--

`adjust_thermostat(68, 72)` should return `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` should return `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` should return `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` should return `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` should return `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` should return `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
