---
id: 68b06e589bf227324381476f
title: "Challenge 35: Word Frequency"
challengeType: 29
dashedName: challenge-35
---

# --description--

문단이 주어지면 가장 자주 등장하는 단어 세 개를 배열로 반환하세요.

- 문단의 단어들은 공백으로 구분됩니다.
- 문단에서 대소문자를 구분하지 마세요. 예를 들어 `Hello`와 `hello`을 같은 단어로 처리하세요.
- 문단에서 구두점은 무시하세요. 구두점은 쉼표(`,`), 마침표(`.`), 느낌표(`!`)로 구성됩니다.
- 반환하는 배열의 모든 단어는 소문자여야 합니다.
- 반환하는 배열은 가장 자주 등장하는 단어가 첫 번째로 오도록 내림차순이어야 합니다.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` should return `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` should return `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` should return `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```
