---
id: 68b06e589bf2273243814777
title: "Challenge 39: Fill The Tank"
challengeType: 29
dashedName: challenge-39
---

# --description--

연료 탱크의 크기, 현재 연료량, 갤런당 가격이 주어졌을 때, 탱크를 가득 채우는 데 드는 비용을 반환하세요.

- `tankSize`는 탱크의 총 용량(갤런 단위)입니다.
- `fuelLevel`는 탱크에 현재 남아 있는 연료량(갤런 단위)입니다.
- `pricePerGallon`는 연료 1갤런당 가격입니다.
- 반환값은 소수점 둘째 자리까지 반올림한 형식인 `"$d.dd"`이어야 합니다.

# --hints--

`cost_to_fill(20, 0, 4.00)` should return `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` should return `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` should return `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` should return `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` should return `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
