---
id: 68b06e589bf2273243814775
title: "Challenge 38: Slug Generator"
challengeType: 29
dashedName: challenge-38
---

# --description--

문자열이 주어지면 다음 조건을 사용해 URL에 적합한 버전의 문자열을 반환하세요:

- 모든 문자는 소문자여야 합니다.
- 문자, 숫자, 공백이 아닌 모든 문자는 제거해야 합니다.
- 모든 공백은 URL 인코딩된 공백 코드 `%20`로 교체해야 합니다.
- 연속된 공백은 하나의 `%20`로 교체해야 합니다.
- 반환된 문자열은 앞뒤에 `%20`가 없어야 합니다.

# --hints--

`generate_slug("helloWorld")` should return `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` should return `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` should return `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` should return `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` should return `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
