---
id: 68b1f72371a5ac895ac70a08
title: "Challenge 43: Digits vs Letters"
challengeType: 29
dashedName: challenge-43
---

# --description--

문자열이 주어지면, 숫자가 문자보다 많으면 `"digits"`을, 문자가 숫자보다 많으면 `"letters"`을, 숫자와 문자의 개수가 같으면 `"tie"`를 반환하세요.

- 숫자는 `0-9`로 구성됩니다.
- 문자는 대문자 또는 소문자의 `a-z`로 구성됩니다.
- 그 외의 문자는 무시하세요.

# --hints--

`digits_or_letters("abc123")` should return `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123"), "tie")`)
}})
```

`digits_or_letters("a1b2c3d")` should return `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("a1b2c3d"), "letters")`)
}})
```

`digits_or_letters("1a2b3c4")` should return `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("1a2b3c4"), "digits")`)
}})
```

`digits_or_letters("abc123!@#DEF")` should return `"letters"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("abc123!@#DEF"), "letters")`)
}})
```

`digits_or_letters("H3110 W0R1D")` should return `"digits"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("H3110 W0R1D"), "digits")`)
}})
```

`digits_or_letters("P455W0RD")` should return `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(digits_or_letters("P455W0RD"), "tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def digits_or_letters(s):

    return s
```

# --solutions--

```py
def digits_or_letters(s):
    digit_count = 0
    letter_count = 0

    for char in s:
        if char.isdigit():
            digit_count += 1
        elif char.isalpha():
            letter_count += 1

    if digit_count > letter_count:
        return "digits"
    if letter_count > digit_count:
        return "letters"
    return "tie"
```
