---
id: 68b1f72371a5ac895ac70a0a
title: "Challenge 44: String Mirror"
challengeType: 29
dashedName: challenge-44
---

# --description--

두 개의 문자열이 주어졌을 때, 두 번째 문자열이 첫 번째 문자열의 거울인지 확인하세요.

- 문자열이 거울이라고 간주되려면, 같은 문자를 역순으로 포함해야 합니다.
- 대문자와 소문자는 구분하세요.
- 알파벳이 아닌 모든 문자는 무시하세요.

# --hints--

`is_mirror("helloworld", "helloworld")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("helloworld", "helloworld"), False)`)
}})
```

`is_mirror("Hello World", "dlroW olleH")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW olleH"), True)`)
}})
```

`is_mirror("RaceCar", "raCecaR")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "raCecaR"), True)`)
}})
```

`is_mirror("RaceCar", "RaceCar")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("RaceCar", "RaceCar"), False)`)
}})
```

`is_mirror("Mirror", "rorrim")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Mirror", "rorrim"), False)`)
}})
```

`is_mirror("Hello World", "dlroW-olleH")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "dlroW-olleH"), True)`)
}})
```

`is_mirror("Hello World", "!dlroW !olleH")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror("Hello World", "!dlroW !olleH"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror(str1, str2):

    return str1
```

# --solutions--

```py
def is_mirror(str1, str2):
    clean1 = "".join(c for c in str1 if c.isalpha())
    clean2 = "".join(c for c in str2 if c.isalpha())
    return clean1[::-1] == clean2
```
