---
id: 68b7687dded630607aceccb1
title: "Challenge 48: Spam Detector"
challengeType: 29
dashedName: challenge-48
---

# --description--

다음과 같은 형식의 전화번호가 주어졌을 때 `"+A (BBB) CCC-DDDD"`, 각 문자는 다음과 같이 숫자를 나타냅니다:

- `A`은 국가 번호를 나타내며, 숫자 자릿수는 제한이 없습니다.
- `BBB`는 지역 번호를 나타내며 항상 세 자리 숫자입니다.
- `CCC`과 `DDDD`는 지역 번호의 나머지 부분으로 각각 세 자리와 네 자리 숫자입니다.

다음 기준에 따라 스팸 번호인지 판단하세요:

- 국가 번호가 2자리보다 길거나 0(`0`)으로 시작하지 않는 경우.
- 지역 번호가 900보다 크거나 200보다 작은 경우.
- 지역 번호 첫 세 자리 숫자의 합이 지역 번호 마지막 네 자리 숫자 안에 포함된 경우.
- 번호에 같은 숫자가 네 번 이상 연속으로 반복되는 경우(형식 문자는 무시).

# --hints--

`is_spam("+0 (200) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
