---
id: 68b7cadffed0e75a517da677
title: "Challenge 54: P@ssw0rd Str3ngth!"
challengeType: 29
dashedName: challenge-54
---

# --description--

비밀번호 문자열이 주어지면, 비밀번호 강도에 따라 `"weak"`, `"medium"`, 또는 `"strong"`를 반환하세요.

비밀번호는 다음 규칙에 따라 평가됩니다:

- 최소 8자 이상이어야 합니다.
- 대문자와 소문자를 모두 포함해야 합니다.
- 숫자를 최소 하나 포함해야 합니다.
- 다음 집합에서 최소 하나의 특수 문자를 포함해야 합니다: `!`, `@`, `#`, `$`, `%`, `^`, `&`, 또는 `*`.

규칙을 두 개 미만으로 충족하면 `"weak"`를 반환하세요.
규칙을 2개 또는 3개 충족하면 `"medium"`를 반환하세요.
모든 4개 규칙을 충족하면 `"strong"`를 반환하세요.

# --hints--

`check_strength("123456")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("123456"), "weak")`)
}})
```

`check_strength("pass!!!")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("pass!!!"), "weak")`)
}})
```

`check_strength("Qwerty")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("Qwerty"), "weak")`)
}})
```

`check_strength("PASSWORD")` should return `"weak"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD"), "weak")`)
}})
```

`check_strength("PASSWORD!")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PASSWORD!"), "medium")`)
}})
```

`check_strength("PassWord%^!")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("PassWord%^!"), "medium")`)
}})
```

`check_strength("qwerty12345")` should return `"medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("qwerty12345"), "medium")`)
}})
```

`check_strength("S3cur3P@ssw0rd")` should return `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("S3cur3P@ssw0rd"), "strong")`)
}})
```

`check_strength("C0d3&Fun!")` should return `"strong"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(check_strength("C0d3&Fun!"), "strong")`)
}})
```

# --seed--

## --seed-contents--

```py
def check_strength(password):

    return password
```

# --solutions--

```py
import re
def check_strength(password):
    rules_met = 0

    if len(password) >= 8:
        rules_met += 1
    if re.search(r'[a-z]', password) and re.search(r'[A-Z]', password):
        rules_met += 1
    if re.search(r'\d', password):
        rules_met += 1
    if re.search(r'[!@#$%^&*]', password):
        rules_met += 1

    if rules_met < 2:
        return "weak"
    elif rules_met < 4:
        return "medium"
    else:
        return "strong"
```
