---
id: 68b7687dded630607aceccaf
title: "Challenge 47: Caught Speeding"
challengeType: 29
dashedName: challenge-47
---

# --description--

차량들이 관찰된 속도를 나타내는 배열과 제한 속도를 나타내는 숫자가 주어졌을 때, 과속한 차량 수와 그 차량들의 제한 속도 초과 평균 속도를 포함하는 배열을 반환하세요.

- 과속한 차량이 없으면 `[0, 0]`를 반환하세요.

# --hints--

`speeding([50, 60, 55], 60)` should return `[0, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([50, 60, 55], 60), [0, 0])`)
}})
```

`speeding([58, 50, 60, 55], 55)` should return `[2, 4]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([58, 50, 60, 55], 55), [2, 4])`)
}})
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` should return `[4, 8.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5])`)
}})
```

`speeding([100, 105, 95, 102], 100)` should return `[2, 3.5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([100, 105, 95, 102], 100), [2, 3.5])`)
}})
```

`speeding([40, 45, 44, 50, 112, 39], 55)` should return `[1, 57]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57])`)
}})
```

# --seed--

## --seed-contents--

```py
def speeding(speeds, limit):

    return speeds
```

# --solutions--

```py
def speeding(speeds, limit):
    speeding = 0
    total_over = 0

    for speed in speeds:
        if speed > limit:
            speeding += 1
            total_over += (speed - limit)
    
    if speeding == 0:
        return [0, 0]
    
    return [speeding, total_over / speeding]
```
