---
id: 68b7cadffed0e75a517da673
title: "Challenge 52: Binary to Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

이진수를 나타내는 문자열이 주어지면, 그 십진수 값을 숫자로 반환하세요.

이진수는 숫자 `0`와 `1`만 사용하여 어떤 수든 나타냅니다. 이진수를 십진수로 변환하려면 각 숫자에 `2`의 거듭제곱을 곱한 후 모두 더하세요. 가장 오른쪽 숫자에 `2^0`을 곱하고, 그 왼쪽 숫자에는 `2^1`를 곱하는 식으로 진행합니다. 모든 숫자에 `2`의 거듭제곱을 곱한 후 결과를 모두 더하세요.

예를 들어, 이진수 `101`는 십진수로 `5`와 같습니다. 이유는 다음과 같습니다:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` should return `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
