---
id: 68b7cadffed0e75a517da671
title: "Challenge 51: Phone Number Formatter"
challengeType: 29
dashedName: challenge-51
---

# --description--

11자리 숫자로 이루어진 문자열이 주어지면, 이 문자열을 `"+D (DDD) DDD-DDDD"` 형식의 전화번호로 반환하세요.

# --hints--

`format_number("05552340182")` should return `"+0 (555) 234-0182"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_number("05552340182"), "+0 (555) 234-0182")`)
}})
```

`format_number("15554354792")` should return `"+1 (555) 435-4792"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_number("15554354792"), "+1 (555) 435-4792")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_number(number):

    return number
```

# --solutions--

```py
def format_number(number):
    country = number[0]
    area = number[1:4]
    prefix = number[4:7]
    line = number[7:]

    return f"+{country} ({area}) {prefix}-{line}"
```
