---
id: 68c1a929005bf54d342aa8d3
title: "Challenge 56: Space Week Day 2: Exoplanet Search"
challengeType: 29
dashedName: challenge-56
---

# --description--

우주 주간 2일차에는 각 문자가 별의 광도 측정값을 나타내는 문자열이 주어집니다. 통과법(transit method)을 사용해 외계 행성이 탐지되었는지 판단하세요. 통과법은 행성이 별 앞을 지나가면서 관측된 광도가 줄어드는 현상입니다.

- 광도 측정값은 `0-9`와 `A-Z` 문자로만 구성되며, 각 측정값은 다음 수치에 해당합니다:
- `0-9` 문자는 `0-9` 광도 수준에 해당합니다.
- `A-Z` 문자는 `10-35` 광도 수준에 해당합니다.

별에 외계 행성이 있다고 간주하려면, 단일 측정값이 모든 측정값 평균의 80% 이하인 경우가 있어야 합니다. 예를 들어, 별의 평균 광도가 10이라면, 단일 측정값이 8 이하일 때 외계 행성이 있다고 판단합니다.

# --hints--

`has_exoplanet("665544554")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
