---
id: 68c1a929005bf54d342aa8d4
title: "Challenge 57: Space Week Day 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

우주 주간 3일차에는 통신 라우트에서 자신, 위성, 그리고 고향 행성 사이의 거리(킬로미터 단위)를 나타내는 배열이 주어집니다. 다음 조건을 사용해 메시지가 라우트를 통해 목적지 행성에 도달하는 데 걸리는 시간을 계산하세요:

- 배열의 첫 번째 값은 현재 위치에서 첫 번째 위성까지의 거리입니다.
- 마지막 값을 제외한 각 값은 다음 위성까지의 거리입니다.
- 배열의 마지막 값은 이전 위성에서 고향 행성까지의 거리입니다.
- 메시지는 초당 300,000 km 속도로 이동합니다.
- 메시지가 **통과하는** 각 위성은 0.5초의 전송 지연을 추가합니다.
- 소수점 넷째 자리까지 반올림한 숫자를 반환하되, 끝의 0은 제거하세요.

# --hints--

`send_message([300000, 300000])` should return `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` should return `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` should return `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` should return `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` should return `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` should return `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
