---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Goldilocks Zone"
challengeType: 29
dashedName: challenge-59
---

# --description--

우주 주간 5일차에는 별의 "골디락스 존"을 계산합니다. 골디락스 존은 별 주위에서 액체 상태의 물이 존재하기에 조건이 "딱 맞는" 영역입니다.

별의 질량이 주어지면, 천문 단위로 골디락스 존의 시작 거리와 끝 거리를 배열로 반환하세요.

골디락스 존을 계산하는 방법:

1. 별의 광도를 질량의 3.5제곱으로 구하세요.
2. 존의 시작 거리는 광도의 제곱근에 0.95를 곱한 값입니다.
3. 존의 끝 거리는 광도의 제곱근에 1.37을 곱한 값입니다.

- 거리는 소수점 둘째 자리까지 반올림해서 반환하세요.

예를 들어, 질량이 `1`일 때 `[0.95, 1.37]`을 반환하세요.

# --hints--

`goldilocks_zone(1)` should return `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` should return `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` should return `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` should return `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` should return `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
