---
id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Landing Spot"
challengeType: 29
dashedName: challenge-58
---

# --description--

우주 주간 4일차에는 로버의 잠재적 착륙 지점을 나타내는 숫자 행렬(배열의 배열)이 주어집니다. 다음 규칙에 따라 가장 안전한 착륙 지점을 찾으세요:

- 행렬의 각 지점은 `0-9`부터 포함하는 숫자를 가집니다.
- `0`은 잠재적 착륙 지점을 나타냅니다.
- `0`가 아닌 숫자는 착륙하기에 너무 위험합니다. 숫자가 클수록 위험도가 높습니다.
- 가장 안전한 지점은 주변 셀(대각선 제외, 최대 4개 이웃)의 위험도 합이 가장 낮은 `0` 셀로 정의됩니다.
- 경계 밖 이웃은 무시하세요(모서리와 가장자리 셀은 이웃이 적습니다).
- 가장 안전한 착륙 지점의 인덱스를 반환하세요. 항상 가장 안전한 지점은 하나뿐입니다.

예를 들어, 다음이 주어지면:

```js
[
  [1, 0],
  [2, 0]
]
```

첫 번째 배열에서 `[0, 1]`의 인덱스인 `0`를 반환하세요.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` should return `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` should return `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` should return `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` should return `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
