---
id: 68c497f3aaefc9fd9f1b0e26
title: "Challenge 62: Hex to Decimal"
challengeType: 29
dashedName: challenge-62
---

# --description--

16진수(기수 16)를 나타내는 문자열이 주어지면, 10진수(기수 10) 값으로 정수를 반환하세요.

16진수는 16개의 숫자를 사용하는 숫자 체계입니다:

- `0-9`는 `0`부터 `9`까지의 값을 나타냅니다.
- `A-F`는 `10`부터 `15`까지의 값을 나타냅니다.

부분 변환 표는 다음과 같습니다:

| 16진수 | 10진수 |
|:-------:|:-------:|
| 0       | 0       |
| 1       | 1       |
| ...     | ...     |
| 9       | 9       |
| A       | 10      |
| ...     | ...     |
| F       | 15      |
| 10      | 16      |
| ...     | ...     |
| 9F      | 159     |
| A0      | 160     |
| ...     | ...     |
| FF      | 255     |
| 100     | 256     |

- 문자열에는 `0–9`와 `A–F` 문자만 포함됩니다.

# --hints--

`hex_to_decimal("A")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A"), 10)`)
}})
```

`hex_to_decimal("15")` should return `21`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("15"), 21)`)
}})
```

`hex_to_decimal("2E")` should return `46`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("2E"), 46)`)
}})
```

`hex_to_decimal("FF")` should return `255`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("FF"), 255)`)
}})
```

`hex_to_decimal("A3F")` should return `2623`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(hex_to_decimal("A3F"), 2623)`)
}})
```

# --seed--

## --seed-contents--

```py
def hex_to_decimal(hex):

    return hex
```

# --solutions--

```py
def hex_to_decimal(hex):

    return int(hex, 16)
```
