---
id: 68c497f3aaefc9fd9f1b0e24
title: "Challenge 60: Space Week Day 6: Moon Phase"
challengeType: 29
dashedName: challenge-60
---

# --description--

우주 주간 6일차에는 `"YYYY-MM-DD"` 형식의 날짜가 주어지고, 다음 규칙을 사용해 그날의 달 위상을 결정해야 합니다:

28일로 단순화한 월령 주기를 사용하며, 네 개의 같은 길이 위상으로 나눕니다:

- `"New"`: 1일차부터 7일차까지
- `"Waxing"`: 8일차부터 14일차까지
- `"Full"`: 15일차부터 21일차까지
- `"Waning"`: 22일차부터 28일차까지

28일이 지나면 주기가 1일차인 초승달부터 다시 반복됩니다.

- 주어진 날짜의 위상을 결정하기 위해 `"2000-01-06"`를 기준 초승달(주기 1일차)로 사용하세요.
- 기준 날짜 이전의 날짜는 주어지지 않습니다.
- 올바른 위상을 문자열로 반환하세요.

**참고:** 1일차는 초승달이 뜨는 날로, 마지막 초승달 이후 0일이 지난 상태를 의미합니다.

# --hints--

`moon_phase("2000-01-12")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` should return `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` should return `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` should return `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
