---
id: 68cae5b538ff798bbd4da00a
title: "Challenge 72: Thermostat Adjuster 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

화씨로 된 현재 방 온도와 섭씨로 된 목표 온도가 주어졌을 때, 다음 조건에 따라 방 온도를 어떻게 조절해야 하는지 나타내는 문자열을 반환하세요:

- 현재 온도가 목표 온도보다 낮으면 `"Heat: X degrees Fahrenheit"`를 반환하세요. 여기서 `X`은 목표에 도달하기 위해 방을 가열해야 하는 화씨 온도 차이로, 소수점 첫째 자리까지 반올림한 값입니다.
- 현재 온도가 목표 온도보다 높으면 `"Cool: X degrees Fahrenheit"`를 반환하세요. 여기서 `X`은 목표에 도달하기 위해 방을 냉각해야 하는 화씨 온도 차이로, 소수점 첫째 자리까지 반올림한 값입니다.
- 현재 온도가 목표 온도와 같으면 `"Hold"`를 반환하세요.

섭씨를 화씨로 변환하려면 섭씨 온도에 1.8을 곱하고 결과에 32를 더하세요 (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` should return `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` should return `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` should return `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
