---
id: 68cae5b538ff798bbd4da001
title: "Challenge 63: Battle of Words"
challengeType: 29
dashedName: challenge-63
---

# --description--

팀과 상대 팀을 나타내는 두 문장이 주어질 때, 각 팀의 단어가 상대 팀의 해당 단어와 대결하여 다음 규칙에 따라 어느 팀이 이기는지 결정하세요:

- 주어진 문장은 항상 같은 수의 단어를 포함합니다.
- 단어는 한 칸의 공백으로 구분되며, 오직 알파벳 문자만 포함합니다.
- 각 단어의 값은 그 단어를 구성하는 글자들의 합입니다.
- `a`부터 `z`까지의 글자는 `1`부터 `26`까지의 값에 대응합니다. 예를 들어, `a`는 `1`이고, `z`는 `26`입니다.
- 대문자는 글자의 값을 두 배로 만듭니다. 예를 들어, `A`는 `2`이고, `Z`는 `52`입니다.
- 단어들은 순서대로 대결합니다: 팀의 첫 번째 단어가 상대 팀의 첫 번째 단어와 대결하고, 그 다음 단어들도 같은 방식입니다.
- 단어의 값이 상대 단어의 값보다 크면 그 단어가 이깁니다.
- 더 많은 단어가 이긴 팀이 승리합니다.

팀이 이기면 `"We win"`를, 팀이 지면 `"We lose"`를, 두 팀이 이긴 단어 수가 같으면 `"Draw"`를 반환하세요.

# --hints--

`battle("hello world", "hello word")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
