---
id: 68cae5b538ff798bbd4da009
title: "Challenge 71: Tip Calculator"
challengeType: 29
dashedName: challenge-71
---

# --description--

식사 가격과 사용자 지정 팁 비율이 주어지면 15%, 20%, 그리고 사용자 지정 금액의 세 가지 팁 값을 배열로 반환하세요.

- 가격은 `"$N.NN"` 형식으로 제공됩니다.
- 사용자 지정 팁 비율은 `"25%"` 형식으로 제공됩니다.
- 금액은 동일한 `"$N.NN"` 형식으로 소수점 둘째 자리까지 반올림하여 반환하세요.

예를 들어, `"$10.00"` 식사 가격과 `"25%"` 사용자 지정 팁 값이 주어지면 `["$1.50", "$2.00", "$2.50"]`를 반환하세요.

# --hints--

`calculate_tips("$10.00", "25%")` should return `["$1.50", "$2.00", "$2.50"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"])`)
}})
```

`calculate_tips("$89.67", "26%")` should return `["$13.45", "$17.93", "$23.31"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"])`)
}})
```

`calculate_tips("$19.85", "9%")` should return `["$2.98", "$3.97", "$1.79"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_tips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_tips(meal_price, custom_tip):

    return meal_price
```

# --solutions--

```py
def calculate_tips(meal_price, custom_tip):
    meal = float(meal_price[1:])
    custom_percent = float(custom_tip[:-1])

    tips_percent = [15, 20, custom_percent]
    tips = [f"${meal * p / 100:.2f}" for p in tips_percent]

    return tips
```
