---
id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Missing Socks"
challengeType: 29
dashedName: challenge-69
---

# --description--

처음 가지고 있던 양말 쌍의 수를 나타내는 정수와 세탁한 횟수를 나타내는 또 다른 정수가 주어질 때, 다음 조건을 사용해 현재 가지고 있는 완전한 양말 쌍의 수를 반환하세요:

- 세탁 2회마다 양말 한 짝을 잃습니다.
- 세탁 3회마다 사라진 양말 한 짝을 찾습니다.
- 세탁 5회마다 양말 한 짝이 닳아 버려야 합니다.
- 세탁 10회마다 양말 한 켤레를 새로 삽니다.
- 총 양말 수는 절대 0보다 작을 수 없습니다.
- 규칙은 겹칠 수 있습니다. 예를 들어, 세탁 10회에는 양말 한 짝을 잃고, 한 짝을 버리며, 새 양말 한 켤레를 삽니다.
- 완전한 양말 쌍의 수를 반환하세요.

# --hints--

`sock_pairs(2, 5)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
