---
id: 68d2ba1468508398389487d0
title: "Challenge 75: Hidden Treasure"
challengeType: 29
dashedName: challenge-75
---

# --description--

숨겨진 보물이 포함된 해저 지도를 나타내는 2D 배열과 보물 탐색을 위한 다음 잠수 좌표(`[row, column]`)가 주어질 때, 다음 규칙에 따라 `"Empty"`, `"Found"`, 또는 `"Recovered"`을 반환하세요:

- 주어진 2D 배열에는 정확히 하나의 회수되지 않은 보물이 있으며, 이 보물은 여러 칸을 차지합니다.
- 2D 배열의 각 칸에는 다음 값 중 하나가 들어 있습니다:
  - `"-"`: 보물이 없음.
  - `"O"`: 아직 발견되지 않은 보물의 일부.
  - `"X"`: 이미 발견된 보물의 일부.

- 잠수 위치에 보물이 없으면 `"Empty"`을 반환하세요.
- 잠수 위치에서 보물을 찾았지만 보물의 다른 일부가 아직 발견되지 않았다면 `"Found"`을 반환하세요.
- 잠수 위치에서 보물의 마지막 발견되지 않은 부분을 찾았다면 `"Recovered"`을 반환하세요.

예를 들어, 다음이 주어졌을 때:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

잠수 위치 좌표가 `[2, 1]`이면, 잠수로 보물의 마지막 발견되지 않은 부분을 찾았으므로 `"Recovered"`을 반환하세요.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` should return `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` should return `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` should return `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` should return `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` should return `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` should return `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```
