---
id: 68d2ba1468508398389487ce
title: "Challenge 73: Speak Wisely, You Must"
challengeType: 29
dashedName: challenge-73
---

# --description--

문장이 주어지면, 다음 규칙을 사용해 현명한 스승의 조언처럼 들리는 문장 버전을 반환하세요:

- 단어들은 한 칸의 공백으로 구분됩니다.
- 문장에서 다음 단어들 중 첫 번째로 나타나는 단어를 찾으세요: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- 그 단어를 포함해 그 앞에 있는 모든 단어를 문장 끝으로 옮기고:
  - 옮길 때 단어들의 순서를 유지하세요.
  - 모두 소문자로 만드세요.
  - 그리고 그 앞에 쉼표와 공백을 추가하세요.
- 새 문장의 첫 단어 첫 글자는 대문자로 만드세요.
- 모든 주어진 문장은 한 개의 구두점으로 끝납니다. 원래 문장의 구두점을 유지하고 새 문장 끝으로 옮기세요.
- 새 문장을 반환할 때, 단어 사이에 한 칸의 공백만 있고 문장 시작과 끝에는 공백이 없도록 하세요.

예를 들어, `"You must speak wisely."`가 주어지면 `"Speak wisely, you must."`를 반환하세요.

# --hints--

`wise_speak("You must speak wisely.")` should return `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` should return `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` should return `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` should return `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` should return `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
