---
id: 68d30fc57588d97fd3027b30
title: "Challenge 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

1994년 10월 28일, Netscape Navigator가 출시되어 수백만 명이 초기 웹을 탐험하는 데 도움을 주었습니다.

Netscape Navigator에서 실행한 브라우저 명령 배열이 주어질 때, 다음 규칙을 사용해 모든 명령을 실행한 후 현재 위치한 페이지를 반환하세요:

- 항상 `"Home"` 페이지에서 시작하며, 이 페이지는 명령 배열에 포함되지 않습니다.
- 유효한 명령은 다음과 같습니다:
  - `"Visit Page"`: `"Page"`가 방문할 페이지 이름입니다. 예를 들어, `"Visit About"`는 `"About"` 페이지로 이동합니다. 새 페이지를 방문할 때는 앞으로 가기 기록을 모두 삭제하세요.
  - `"Back"`: 이전 페이지로 이동하거나 이전 페이지가 없으면 현재 페이지에 머뭅니다.
  - `"Forward"`: 앞으로 가기 기록의 다음 페이지로 이동하거나 없으면 현재 페이지에 머뭅니다.

예를 들어, `["Visit About Us", "Back", "Forward"]`가 주어지면 `"About Us"`을 반환하세요.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` should return `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` should return `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` should return `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` should return `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
