---
id: 68d30845cc08266018fc46bc
title: "Challenge 76: Complementary DNA"
challengeType: 29
dashedName: challenge-76
---

# --description--

DNA 수열을 나타내는 문자열이 주어지면 다음 규칙을 사용해 보색 가닥을 반환하세요:

- DNA는 `"A"`, `"C"`, `"G"`, `"T"` 문자로 구성됩니다.
- `"A"`와 `"T"` 문자는 서로 보색입니다.
- `"C"`와 `"G"` 문자는 서로 보색입니다.

예를 들어, `"ACGT"`가 주어지면 `"TGCA"`를 반환하세요.

# --hints--

`complementary_dna("ACGT")` should return `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` should return `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` should return `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` should return `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```
