---
id: 68e39ed6106dac2f0a98fd66
title: "Challenge 84: Infected"
challengeType: 29
dashedName: challenge-84
---

# --description--

1988년 11월 2일, 최초의 주요 인터넷 웜이 배포되어 단 하루 만에 인터넷에 연결된 컴퓨터 약 10%가 감염되었습니다.

이 도전 과제에서는 인터넷 웜이 배포된 이후 경과한 일 수가 주어지고, 다음 규칙에 따라 감염된 컴퓨터 수를 계산해야 합니다:

- 0일째에는 첫 번째 컴퓨터가 감염됩니다.
- 그 다음 날부터는 감염된 컴퓨터 수가 매일 두 배로 증가합니다.
- 매 3일째에는 바이러스가 퍼진 후 패치가 적용되어 감염된 컴퓨터 수가 20% 줄어듭니다. 패치된 컴퓨터 수는 소수점 올림하여 가장 가까운 정수로 만드세요.

예를 들어:

- 0일째: 총 1대의 컴퓨터가 감염됩니다.
- 1일째: 총 2대의 컴퓨터가 감염됩니다.
- 2일째: 총 4대의 컴퓨터가 감염됩니다.
- 3일째: 총 8대의 컴퓨터가 감염됩니다. 그 후 패치를 적용합니다: 8대 * 20% = 1.6대가 패치됨. 1.6을 올림하여 2대가 패치됩니다. 8대 감염 - 2대 패치 = 3일째 후 총 6대가 감염됩니다.

주어진 일 수가 지난 후 감염된 총 컴퓨터 수를 반환하세요.

# --hints--

`infected(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` should return `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` should return `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` should return `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
