---
id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--

표준 카드 한 벌에는 각 무늬마다 13장의 고유 카드가 있습니다. 카드 덱에서 뽑을 카드 수를 나타내는 정수가 주어지면, 뽑을 수 있는 고유 카드 조합의 수를 반환하세요.

- 순서는 중요하지 않습니다. 카드 A를 먼저 뽑고 카드 B를 뽑는 것은 카드 B를 먼저 뽑고 카드 A를 뽑는 것과 같습니다.

예를 들어 `52`가 주어지면 `1`를 반환하세요. 52장 카드 덱에서 52장을 뽑는 조합은 단 하나뿐입니다. 그리고 `2`가 주어지면 `1326`를 반환하세요. 카드 덱에서 2장을 뽑을 때 가능한 카드 조합은 1326가지입니다.

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
