---
id: 68f6587287ad1f4ad39b0c7c
title: "Challenge 90: Character Limit"
challengeType: 29
dashedName: challenge-90
---

# --description--

이 도전 과제에서는 문자열이 주어지고, 그것이 소셜 미디어 게시물에 들어맞는지 판단해야 합니다. 주어진 규칙에 따라 다음 문자열을 반환하세요:

- 40자 제한 내에 들어맞으면 `"short post"`를 반환하세요.
- 40자를 초과하고 80자 제한 내에 들어맞으면 `"long post"`를 반환하세요.
- 두 제한 모두에 들어맞지 않으면 `"invalid post"`를 반환하세요.

# --hints--

`can_post("Hello world")` should return `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` should return `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` should return `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
