---
id: 68f6587287ad1f4ad39b0c83
title: "Challenge 97: GCD"
challengeType: 29
dashedName: challenge-97
---

# --description--

두 개의 양의 정수가 주어지면, 그들의 최대공약수(GCD)를 반환하세요.

- 두 정수의 최대공약수는 나머지를 남기지 않고 두 수 모두를 나누는 가장 큰 수입니다.

예를 들어, `4`의 약수는 `1`, `2`, 그리고 `4`입니다. `6`의 약수는 `1`, `2`, `3`, 그리고 `6`입니다. 따라서 `4`와 `6`이 주어지면, 두 약수 집합 모두에 나타나는 가장 큰 수인 `2`을 반환하세요.

# --hints--

`gcd(4, 6)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(4, 6), 2)`)
}})
```

`gcd(20, 15)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(20, 15), 5)`)
}})
```

`gcd(13, 17)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(13, 17), 1)`)
}})
```

`gcd(654, 456)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(654, 456), 6)`)
}})
```

`gcd(3456, 4320)` should return `864`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(gcd(3456, 4320), 864)`)
}})
```

# --seed--

## --seed-contents--

```py
def gcd(x, y):

    return x
```

# --solutions--

```py
def gcd(x, y):
    while y != 0:
        x, y = y, x % y
    return x
```
