---
id: 68ffb91507a5b645769328cc
title: "Challenge 109: What's My Age Again?"
challengeType: 29
dashedName: challenge-109
---

# --description--

`YYYY-MM-DD` 형식으로 주어진 생일 날짜를 기준으로 2025년 11월 27일 현재 나이를 반환하세요.

- 모든 생일은 2025년 11월 27일 이전의 유효한 날짜라고 가정하세요.
- 나이는 정수로 반환하세요.
- 2025년에 이미 생일이 지났는지 여부를 반드시 고려하세요.

# --hints--

`calculate_age("2000-11-20")` should return `25`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-11-20"), 25)`)
}})
```

`calculate_age("2000-12-01")` should return `24`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2000-12-01"), 24)`)
}})
```

`calculate_age("2014-10-25")` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("2014-10-25"), 11)`)
}})
```

`calculate_age("1994-01-06")` should return `31`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-01-06"), 31)`)
}})
```

`calculate_age("1994-12-14")` should return `30`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_age("1994-12-14"), 30)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_age(birthday):

    return birthday
```

# --solutions--

```py
from datetime import datetime
def calculate_age(birthday):
    today = datetime(2025, 11, 27)
    birth_date = datetime.strptime(birthday, "%Y-%m-%d")

    age = today.year - birth_date.year

    if (today.month, today.day) < (birth_date.month, birth_date.day):
        age -= 1

    return age
```
