---
id: 68ffb91507a5b645769328c4
title: "Challenge 101: Markdown Heading Converter"
challengeType: 29
dashedName: challenge-101
---

# --description--

마크다운 제목을 나타내는 문자열이 주어지면, 이에 상응하는 HTML 제목을 반환하세요.

유효한 마크다운 제목은 다음 조건을 만족해야 합니다:

- 0개 이상의 공백으로 시작하고,
- 1개에서 6개까지 연속된 해시 문자(`#`)가 나오며,
- 최소 한 칸 이상의 공백이 있고, 마지막으로
- 제목 텍스트가 나옵니다.

해시 기호의 개수가 제목 레벨을 결정합니다. 예를 들어, 해시 기호가 하나면 `h1` 태그에 해당하고, 여섯 개면 `h6` 태그에 해당합니다.

주어진 문자열이 위 형식과 정확히 일치하지 않으면 `"Invalid format"`을 반환하세요.

예를 들어, `"# My level 1 heading"`가 주어지면 `"<h1>My level 1 heading</h1>"`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 태그를 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`convert("# My level 1 heading")` should return `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` should return `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` should return `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` should return `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
