---
id: 68ffb91507a5b645769328c6
title: "Challenge 103: LCM"
challengeType: 29
dashedName: challenge-103
---

# --description--

두 정수가 주어지면, 두 수의 최소공배수(LCM)를 반환하세요.

두 수의 최소공배수는 두 수 모두의 배수인 가장 작은 양의 정수입니다. 예를 들어 `4`와 `6`이 주어지면 `12`를 반환하세요. 이유는 다음과 같습니다:

- `4`의 배수는 `4`, `8`, `12` 등입니다.
- `6`의 배수는 `6`, `12`, `18` 등입니다.
- `12`은 두 수 모두의 배수인 가장 작은 수입니다.

# --hints--

`lcm(4, 6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` should return `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` should return `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
