---
id: 69162d64f96574d9bb629eff
title: "Challenge 114: Camel to Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

카멜 표기법으로 된 문자열이 주어지면, 다음 규칙을 사용해 스네이크 표기법 버전의 문자열을 반환하세요:

- 입력 문자열은 오직 문자(`A-Z`와 `a-z`)만 포함하며 항상 소문자로 시작합니다.
- 카멜 표기법 문자열의 모든 대문자는 새로운 단어의 시작입니다.
- 모든 문자를 소문자로 변환하세요.
- 단어는 밑줄(`_`)로 구분하세요.

# --hints--

`to_snake("helloWorld")` should return `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` should return `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` should return `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
