---
id: 69162d64f96574d9bb629f03
title: "Challenge 118: Date Formatter"
challengeType: 29
dashedName: challenge-118
---

# --description--

`"Month day, year"` 형식으로 주어진 날짜를 `"YYYY-MM-DD"` 형식으로 반환하세요.

- 주어진 월은 영어로 된 전체 월 이름입니다. 예를 들어: `"January"`, `"February"` 등입니다.
- 반환값에서는 월과 일이 두 자리 숫자가 되도록 필요하면 앞에 0을 채우세요.

예를 들어, `"December 6, 2025"`가 주어지면 `"2025-12-06"`를 반환하세요.

# --hints--

`format_date("December 6, 2025")` should return `"2025-12-06"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("December 6, 2025"), "2025-12-06")`)
}})
```

`format_date("January 1, 2000")` should return `"2000-01-01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("January 1, 2000"), "2000-01-01")`)
}})
```

`format_date("November 11, 1111")` should return `"1111-11-11"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("November 11, 1111"), "1111-11-11")`)
}})
```

`format_date("September 7, 512")` should return `"512-09-07"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("September 7, 512"), "512-09-07")`)
}})
```

`format_date("May 4, 1950")` should return `"1950-05-04"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("May 4, 1950"), "1950-05-04")`)
}})
```

`format_date("February 29, 1992")` should return `"1992-02-29"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_date("February 29, 1992"), "1992-02-29")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_date(date_string):

    return date_string
```

# --solutions--

```py
def format_date(date_string):
    months = {
        "January": "01",
        "February": "02",
        "March": "03",
        "April": "04",
        "May": "05",
        "June": "06",
        "July": "07",
        "August": "08",
        "September": "09",
        "October": "10",
        "November": "11",
        "December": "12"
    }

    month, day, year = date_string.replace(",", "").split(" ")
    day = day.zfill(2)

    return f"{year}-{months[month]}-{day}"
```
