---
id: 69162d64f96574d9bb629efe
title: "Challenge 113: Miles to Kilometers"
challengeType: 29
dashedName: challenge-113
---

# --description--

숫자로 주어진 마일 거리를 킬로미터로 환산하여 반환하세요.

- 입력은 항상 0 이상의 숫자입니다.
- 1마일은 1.60934킬로미터입니다.
- 결과는 소수점 둘째 자리까지 반올림하세요.
- 반올림한 결과에서 불필요한 끝의 0은 제거하세요.

# --hints--

`convert_to_km(1)` should return `1.61`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(1), 1.61)`)
}})
```

`convert_to_km(21)` should return `33.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(21), 33.8)`)
}})
```

`convert_to_km(3.5)` should return `5.63`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(3.5), 5.63)`)
}})
```

`convert_to_km(0)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0), 0)`)
}})
```

`convert_to_km(0.621371)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_km(0.621371), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_km(miles):

    return miles
```

# --solutions--

```py
def convert_to_km(miles):
    km = miles * 1.60934
    return round(km, 2)
```
