---
id: 69162d64f96574d9bb629f01
title: "Challenge 116: Permutation Count"
challengeType: 29
dashedName: challenge-116
---

# --description--

문자열이 주어지면, 그 문자열의 문자들로 만들 수 있는 서로 다른 순열의 개수를 반환하세요.

- 순열은 문자열의 문자들을 재배열한 모든 경우입니다.
- 중복된 순열은 세지 마세요.
- 문자열에 반복된 문자가 있으면, 반복된 배열은 한 번만 세야 합니다.
- 문자열은 오직 문자(`A-Z`, `a-z`)만 포함합니다.

예를 들어, `"abb"`가 주어지면 `3`을 반환하세요. 문자를 배열하는 고유한 방법이 세 가지이기 때문입니다: `"abb"`, `"bab"`, 그리고 `"bba"`.

# --hints--

`count_permutations("abb")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` should return `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
