---
id: 69162d64f96574d9bb629efd
title: "Challenge 112: AI Detector"
challengeType: 29
dashedName: challenge-112
---

# --description--

오늘의 도전 과제는 2022년 11월 30일 ChatGPT 출시에서 영감을 받았습니다.

하나 이상의 문장으로 이루어진 문자열이 주어졌을 때, 다음 규칙을 사용해 AI가 생성했을 가능성이 있는지 판단하세요:

- 두 개 이상의 대시(`-`)가 포함되어 있습니다.
- 두 개 이상의 괄호 세트(`()`)가 포함되어 있습니다. 괄호 안에 텍스트가 있을 수 있습니다.
- 7자 이상인 단어가 세 개 이상 포함되어 있습니다.

- 단어는 한 칸의 공백으로 구분되며 오직 문자(`A-Z`)로만 구성됩니다. 구두점이나 기타 문자가 단어의 일부로 포함되지 않도록 하세요.

주어진 문장이 위 규칙 중 하나라도 충족하면 `"AI"`를 반환하고, 그렇지 않으면 `"Human"`를 반환하세요.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
