---
id: 691b559495c5cb5a37b9b482
title: "Challenge 122: Markdown Bold Parser"
challengeType: 29
dashedName: challenge-122
---

# --description--

마크다운에서 일부 굵게 표시된 텍스트가 포함될 수 있는 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

- 마크다운에서 굵게 표시된 텍스트는 두 개의 별표(`**`) 또는 두 개의 밑줄(`__`)로 시작하고 끝나는 모든 텍스트입니다.
- 텍스트와 별표 또는 밑줄 사이에는 공백이 있을 수 없지만, 텍스트 자체에는 공백이 있을 수 있습니다.
- 모든 굵게 표시된 부분을 HTML `b` 태그로 변환하여 문자열을 반환하세요.

예를 들어, `"**This is bold**"`가 주어지면 `"<b>This is bold</b>"`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 태그를 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parse_bold("**This is bold**")` should return `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` should return `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` should return `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` should return `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` should return `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
