---
id: 691b559495c5cb5a37b9b489
title: "Challenge 129: Markdown Blockquote Parser"
challengeType: 29
dashedName: challenge-129
---

# --description--

마크다운에 블록인용문이 포함된 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

마크다운에서 블록인용문은 다음 조건을 만족하는 모든 줄입니다:

- 0개 이상의 공백으로 시작하고
- 그다음에 부등호 기호(`>`)
- 이어서 1개 이상의 공백이 나오고
- 마지막으로 블록인용문 텍스트가 옵니다.

블록인용문 텍스트를 여는 HTML `blockquote` 태그와 닫는 태그로 감싸서 반환하세요.

예를 들어, `"> This is a quote"`가 주어지면 `<blockquote>This is a quote</blockquote>`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 태그를 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parse_blockquote("> This is a quote")` should return `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` should return `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` should return `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
