---
id: 691b559495c5cb5a37b9b480
title: "Challenge 120: Pounds to Kilograms"
challengeType: 29
dashedName: challenge-120
---

# --description--

숫자로 주어진 파운드 단위 무게를 받아 문자열 `"(lbs) pounds equals (kgs) kilograms."`을 반환하세요.

- `"(lbs)"`을 입력 숫자로 바꾸세요.
- `"(kgs)"`를 킬로그램으로 변환한 값으로 바꾸되, 소수점 둘째 자리까지 반올림하고 항상 두 자리 소수점을 포함하세요.
- 1파운드는 0.453592킬로그램입니다.
- 입력이 `1`일 경우 `"pound"` 대신 `"pounds"`를 사용하세요.
- 변환된 값이 `1`일 경우 `"kilogram"` 대신 `"kilograms"`를 사용하세요.

# --hints--

`convert_to_kgs(1)` should return `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` should return `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` should return `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` should return `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` should return `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
