---
id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--

마크다운 이미지 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

마크다운 이미지는 다음 형식을 가집니다: `"![alt text](image_url)"`. 여기서:

- `alt text` 은 이미지 설명(`alt` 속성 값)입니다.
- `image_url` 은 이미지의 소스 URL(`src` 속성 값)입니다.

이미지 URL에 `img` 을, 대체 텍스트에 `src` 을 설정한 HTML `alt` 태그 문자열을 반환하세요.

예를 들어, `"![Cute cat](cat.png)"` 이 주어지면 `'<img src="cat.png" alt="Cute cat">'` 을 반환하세요.

- 태그, 속성 순서, 공백, 따옴표 사용이 위와 동일한지 꼭 확인하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 표시하지 않을 수 있으니, 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
