---
id: 6925e2068081f40f549ced1c
title: "Challenge 138: Sum of Divisors"
challengeType: 29
dashedName: challenge-138
---

# --description--

양의 정수가 주어지면, 그 수의 모든 약수의 합을 반환하세요.

- 약수는 그 수를 나누었을 때 나머지가 `0`인 모든 정수입니다.
- 각 약수는 한 번만 계산하세요.

예를 들어, `6`이 주어지면 `12`를 반환하세요. `6`의 약수는 `1`, `2`, `3`, `6`이고, 이들의 합은 `12`입니다.

# --hints--

`sum_divisors(6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` should return `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` should return `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` should return `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` should return `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
