---
id: 69272dcf1c24b44fd79137c4
title: "Challenge 141: Takeoff Fuel"
challengeType: 29
dashedName: challenge-141
---

# --description--

현재 비행기에 있는 연료 갤런 수와 목적지까지 필요한 연료 리터 수가 주어졌을 때, 추가로 몇 갤런의 연료를 넣어야 하는지 계산하세요.

- 1갤런은 3.78541리터입니다.
- 비행기에 이미 충분한 연료가 있다면 `0`를 반환하세요.
- 연료는 정수 갤런 단위로만 추가할 수 있습니다.
- 반환 값에 소수점은 포함하지 마세요.

# --hints--

`fuel_to_add(0, 1)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(0, 1), 1)`)
}})
```

`fuel_to_add(5, 40)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(5, 40), 6)`)
}})
```

`fuel_to_add(10, 30)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(10, 30), 0)`)
}})
```

`fuel_to_add(896, 20500)` should return `4520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(896, 20500), 4520)`)
}})
```

`fuel_to_add(1000, 50000)` should return `12209`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fuel_to_add(1000, 50000), 12209)`)
}})
```

# --seed--

## --seed-contents--

```py
def fuel_to_add(current_gallons, required_liters):

    return current_gallons
```

# --solutions--

```py
import math
def fuel_to_add(current_gallons, required_liters):
    liters_per_gallon = 3.78541
    current_liters = current_gallons * liters_per_gallon
    if current_liters >= required_liters:
        return 0
    liters_needed = required_liters - current_liters
    return math.ceil(liters_needed / liters_per_gallon)
```
