---
id: 69272dcf1c24b44fd79137c5
title: "Challenge 142: Sum the String"
challengeType: 29
dashedName: challenge-142
---

# --description--

숫자와 다른 문자가 섞여 있는 문자열이 주어질 때, 문자열에 포함된 모든 숫자의 합을 반환하세요.

- 연속된 숫자는 하나의 숫자로 처리하세요. 예를 들어, `"13"`는 1 + 3이 아니라 13으로 계산합니다.
- 숫자가 아닌 문자는 무시하세요.

# --hints--

`string_sum("3apples2bananas")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` should return `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` should return `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` should return `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` should return `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```
