---
id: 69272dcf1c24b44fd79137c3
title: "Challenge 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

변수 이름을 나타내는 문자열이 주어지면, 해당 변수 이름을 SCREAMING_SNAKE_CASE로 변환해 반환하세요.

주어진 변수 이름은 다음 형식 중 하나로 작성됩니다:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

위 형식에서 단어들은 밑줄(`_`), 하이픈(`-`), 또는 새 단어가 대문자로 시작하는 방식으로 구분됩니다.

SCREAMING_SNAKE_CASE로 변환하려면:

- 모든 문자를 대문자로 만드세요
- 단어 사이를 밑줄(`_`)로 구분하세요

# --hints--

`to_screaming_snake_case("userEmail")` should return `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` should return `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` should return `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` should return `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` should return `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

`to_screaming_snake_case("my_variable_name")` should return `"MY_VARIABLE_NAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("my_variable_name"), "MY_VARIABLE_NAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
