---
id: 69272dcf1c24b44fd79137c6
title: "Challenge 143: Markdown Italic Parser"
challengeType: 29
dashedName: challenge-143
---

# --description--

마크다운에서 일부 이탤릭 텍스트가 포함될 수 있는 문자열이 주어지면, 이에 상응하는 HTML 문자열을 반환하세요.

- 마크다운에서 이탤릭 텍스트는 단일 별표(`*`) 또는 단일 밑줄(`_`)로 시작하고 끝나는 모든 텍스트입니다.
- 텍스트와 별표 또는 밑줄 사이에는 공백이 있을 수 없지만, 텍스트 자체에는 공백이 있을 수 있습니다.
- 모든 이탤릭 표현을 HTML `i` 태그로 변환하고 문자열을 반환하세요.

예를 들어, `"*This is italic*"`가 주어지면 `"<i>This is italic</i>"`를 반환하세요.

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 태그를 표시하지 않을 수 있습니다. 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parse_italics("*This is italic*")` should return `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` should return `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` should return `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` should return `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` should return `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
