---
id: 69306364df283fcaff2e1ad7
title: "Challenge 146: Left-Handed Seat at the Table"
challengeType: 29
dashedName: challenge-146
---

# --description--

4x2 행렬(배열의 배열)로 표현된 식사 자리 배치에서 왼손잡이가 앉을 수 있는 좌석 수를 구하세요.

- 왼손잡이는 바로 왼쪽에 앉은 사람이 오른손잡이인 자리에 앉을 수 없습니다.

주어진 행렬에서:

- `"R"`는 오른손잡이가 앉은 자리입니다.
- `"L"`는 왼손잡이가 앉은 자리입니다.
- `"U"`는 비어 있는 자리입니다.
- 앉을 수 있는 자리는 비어 있는 자리뿐입니다.
- 윗줄 좌석은 "아래쪽"을 향하고, 아랫줄 좌석은 "위쪽"을 향하고 있어(테이블처럼) 왼쪽과 오른쪽은 좌석 방향에 따라 상대적입니다.
- 모서리 좌석은 옆에 한 자리만 있습니다.

예를 들어, 주어진 행렬에서:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

왼쪽 위 좌석은 바로 왼쪽에 오른손잡이가 있어 앉을 수 없습니다. 나머지 두 빈 좌석은 바로 왼쪽에 오른손잡이가 없으므로 앉을 수 있습니다.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
