---
id: 69306364df283fcaff2e1ada
title: "Challenge 149: vOwElcAsE"
challengeType: 29
dashedName: challenge-149
---

# --description--

문자열이 주어지면 모든 모음을 대문자로 변환하고 나머지 알파벳 문자는 소문자로 변환한 새 문자열을 반환하세요.

- 모음은 `"a"`, `"e"`, `"i"`, `"o"`, `"u"`이며 대소문자 구분 없이 포함됩니다.
- 알파벳이 아닌 문자는 변경하지 않고 그대로 둡니다.

# --hints--

`vowel_case("vowelcase")` should return `"vOwElcAsE"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("vowelcase"), "vOwElcAsE")`)
}})
```

`vowel_case("coding is fun")` should return `"cOdIng Is fUn"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("coding is fun"), "cOdIng Is fUn")`)
}})
```

`vowel_case("HELLO, world!")` should return `"hEllO, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HELLO, world!"), "hEllO, wOrld!")`)
}})
```

`vowel_case("git cherry-pick")` should return `"gIt chErry-pIck"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("git cherry-pick"), "gIt chErry-pIck")`)
}})
```

`vowel_case("HEAD~1")` should return `"hEAd~1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(vowel_case("HEAD~1"), "hEAd~1")`)
}})
```

# --seed--

## --seed-contents--

```py
def vowel_case(s):

    return s
```

# --solutions--

```py
def vowel_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char in vowels:
            result += char.upper()
        elif char.isalpha():
            result += char.lower()
        else:
            result += char

    return result
```
