---
id: 69373793f5a867f769cde13a
title: "Challenge 155: Plant the Crop"
challengeType: 29
dashedName: challenge-155
---

# --description--

농장 밭 크기를 나타내는 정수와 농장 밭 크기의 단위를 나타내는 `"acres"` 또는 `"hectares"`, 그리고 작물 종류가 주어질 때, 해당 밭에 심을 수 있는 그 작물의 개수를 구하세요.

- 1에이커는 4046.86 제곱미터입니다.
- 1헥타르는 10,000 제곱미터입니다.

다음은 입력으로 주어질 작물과 한 그루가 차지하는 공간입니다:

|작물|한 그루당 공간|
|-|-|
|`"corn"`|1 제곱미터|
|`"wheat"`|0.1 제곱미터|
|`"soybeans"`|0.5 제곱미터|
|`"tomatoes"`|0.25 제곱미터|
|`"lettuce"`|0.2 제곱미터|

밭에 심을 수 있는 작물 개수를 소수점 이하는 버리고 정수로 반환하세요.

# --hints--

`get_number_of_plants(1, "acres", "corn")` should return `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` should return `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` should return `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` should return `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` should return `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
