---
id: 69373793f5a867f769cde139
title: "Challenge 154: Par for the Hole"
challengeType: 29
dashedName: challenge-154
---

# --description--

두 개의 정수가 주어집니다. 하나는 골프 홀의 기준 타수(par)이고, 다른 하나는 골퍼가 그 홀에서 친 획 수입니다. 골프 용어를 사용해 골퍼의 점수를 반환하세요.

Return:

- 한 획을 쳤다면 `"Hole in one!"`을 반환하세요.
- 기준 타수보다 두 획 적게 쳤다면 `"Eagle"`을 반환하세요.
- 기준 타수보다 한 획 적게 쳤다면 `"Birdie"`를 반환하세요.
- 기준 타수와 같은 획 수를 쳤다면 `"Par"`을 반환하세요.
- 기준 타수보다 한 획 더 쳤다면 `"Bogey"`를 반환하세요.
- 기준 타수보다 두 획 더 쳤다면 `"Double bogey"`를 반환하세요.

# --hints--

`golf_score(3, 3)` should return `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` should return `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` should return `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` should return `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` should return `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` should return `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
