---
id: 6939b873185d8e00d453563d
title: "Challenge 158: Array Swap"
challengeType: 29
dashedName: challenge-158
---

# --description--

두 개의 값이 있는 배열이 주어지면, 값이 서로 바뀐 배열을 반환하세요.

예를 들어, `["A", "B"]`가 주어지면 `["B", "A"]`을 반환하세요.

# --hints--

`array_swap(["A", "B"])` should return `["B", "A"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["A", "B"]), ["B", "A"])`)
}})
```

`array_swap([25, 20])` should return `[20, 25]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([25, 20]), [20, 25])`)
}})
```

`array_swap([True, False])` should return `[False, True]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap([True, False]), [False, True])`)
}})
```

`array_swap(["1", 1])` should return `[1, "1"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_swap(["1", 1]), [1, "1"])`)
}})
```

# --seed--

## --seed-contents--

```py
def array_swap(arr):

    return arr
```

# --solutions--

```py
def array_swap(arr):

    return [arr[1], arr[0]]
```
