---
id: 6939b873185d8e00d453563c
title: "Challenge 157: Markdown Link Parser"
challengeType: 29
dashedName: challenge-157
---

# --description--

마크다운 링크 문자열이 주어지면, 이에 해당하는 HTML 문자열을 반환하세요.

마크다운 이미지는 다음 형식을 가집니다: `"[link_text](link_url)"`. `a` 태그 문자열을 반환하되, `href`는 `link_url`로 설정하고 `link_text`는 태그 내용으로 사용하세요.

예를 들어, `"[freeCodeCamp](https://freecodecamp.org/)"`가 주어지면 `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`를 반환하세요;

참고: 콘솔은 문자열 내 HTML 태그를 로그로 출력할 때 표시하지 않을 수 있으니, 태그가 포함된 로그를 보려면 브라우저 콘솔을 확인하세요.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` should return `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` should return `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` should return `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
