---
id: 694596b0585c11170ac7c7fa
title: "Challenge 162: Energy Consumption"
challengeType: 29
dashedName: challenge-162
---

# --description--

운동 중 소모한 칼로리 수와 운동 중 전자 기기가 사용한 와트시 수가 주어졌을 때, 어느 쪽이 더 많은 에너지를 사용했는지 판단하세요.

두 값을 비교하려면 다음 변환을 사용해 둘 다 줄 단위로 변환하세요:

- 1 칼로리는 4184 줄입니다.
- 1 와트시는 3600 줄입니다.

Return:

- 운동이 더 많은 에너지를 사용했다면 `"Workout"`.
- 기기가 더 많은 에너지를 사용했다면 `"Devices"`.
- 둘 다 같은 양의 에너지를 사용했다면 `"Equal"`.

# --hints--

`compare_energy(250, 50)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
