---
id: 696655d24b614176d4c9b78d
title: "Challenge 170: Odd or Even Day"
challengeType: 29
dashedName: challenge-170
---

# --description--

타임스탬프(Unix epoch 이후 밀리초 수)가 주어지면 다음을 반환하세요:

- 해당 타임스탬프의 날짜가 홀수일 경우 `"odd"`를 반환하세요.
- 해당 타임스탬프의 날짜가 짝수일 경우 `"even"`를 반환하세요.

예를 들어, 2026년 1월 27일의 타임스탬프인 `1769472000000`가 주어지면, 날짜 숫자(`"odd"`)가 홀수이므로 `27`를 반환하세요.

참고: 타임스탬프는 밀리초 단위이며, 로컬 시간이 아닌 UTC 시간대를 기준으로 날짜를 사용해야 합니다.

# --hints--

`odd_or_even_day(1769472000000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
