---
id: 696655d24b614176d4c9b78c
title: "Challenge 169: FizzBuzz Mini"
challengeType: 29
dashedName: challenge-169
---

# --description--

정수가 주어지면 다음 규칙에 따라 문자열을 반환하세요:

- 숫자가 3으로 나누어떨어지면 `"Fizz"`를 반환하세요.
- 숫자가 5로 나누어떨어지면 `"Buzz"`를 반환하세요.
- 숫자가 3과 5 모두로 나누어떨어지면 `"FizzBuzz"`를 반환하세요.
- 그렇지 않으면 주어진 숫자를 문자열로 반환하세요.

# --hints--

`fizz_buzz_mini(3)` should return `"Fizz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(3), "Fizz")`)
}})
```

`fizz_buzz_mini(4)` should return `"4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(4), "4")`)
}})
```

`fizz_buzz_mini(35)` should return `"Buzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(35), "Buzz")`)
}})
```

`fizz_buzz_mini(75)` should return `"FizzBuzz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(75), "FizzBuzz")`)
}})
```

`fizz_buzz_mini(98)` should return `"98"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_mini(98), "98")`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_mini(n):

    return n
```

# --solutions--

```py
def fizz_buzz_mini(n):
    if n % 3 == 0 and n % 5 == 0:
        return "FizzBuzz"
    if n % 3 == 0:
        return "Fizz"
    if n % 5 == 0:
        return "Buzz"
    return str(n)
```
