---
id: 696655d24b614176d4c9b78b
title: "Challenge 168: Scaled Image"
challengeType: 29
dashedName: challenge-168
---

# --description--

이미지의 너비와 높이를 나타내는 문자열과 이미지를 확대할 숫자가 주어지면, 확대된 너비와 높이를 반환하세요.

- 입력 문자열은 `"WxH"` 형식입니다. 예를 들어, `"800x600"`입니다.
- 배율은 너비와 높이에 곱할 숫자입니다.

확대된 크기를 같은 `"WxH"` 형식으로 반환하세요.

# --hints--

`scale_image("800x600", 2)` should return `"1600x1200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("800x600", 2), "1600x1200")`)
}})
```

`scale_image("100x100", 10)` should return `"1000x1000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("100x100", 10), "1000x1000")`)
}})
```

`scale_image("1024x768", 0.5)` should return `"512x384"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("1024x768", 0.5), "512x384")`)
}})
```

`scale_image("300x200", 1.5)` should return `"450x300"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(scale_image("300x200", 1.5), "450x300")`)
}})
```

# --seed--

## --seed-contents--

```py
def scale_image(size, scale):

    return size
```

# --solutions--

```py
def scale_image(size, scale):
    width, height = map(int, size.split("x"))

    new_width = round(width * scale)
    new_height = round(height * scale)

    return f"{new_width}x{new_height}"
```
