---
id: 69738771fb5a7b8b24cca2a1
title: "Challenge 175: Digital Detox"
challengeType: 29
dashedName: challenge-175
---

# --description--

로그인 기록 배열이 주어졌을 때, 디지털 디톡스 목표를 달성했는지 판단하세요.

각 기록은 `"YYYY-MM-DD HH:mm:ss"` 형식의 문자열입니다.

다음 두 조건이 모두 참일 때 디지털 디톡스 목표를 달성한 것입니다:

- 4시간 이내에 로그인한 횟수가 1회를 넘지 않아야 합니다.
- 하루에 로그인한 횟수가 2회를 넘지 않아야 합니다.

# --hints--

`digital_detox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-02-01 08:00:00", "2026-02-01 12:30:00"]), True)`)
}})
```

`digital_detox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-02-01 04:00:00", "2026-02-01 07:30:00"]), False)`)
}})
```

`digital_detox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-01-31 08:21:30", "2026-01-31 14:30:00", "2026-02-01 08:00:00", "2026-02-01 12:30:00"]), True)`)
}})
```

`digital_detox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-01-31 10:40:21", "2026-01-31 15:19:41", "2026-01-31 21:49:50", "2026-02-01 09:30:00"]), False)`)
}})
```

`digital_detox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 09:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]), True)`)
}})
```

`digital_detox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(digital_detox(["2026-02-05 10:00:00", "2026-02-01 09:00:00", "2026-02-03 22:15:00", "2026-02-02 12:10:00", "2026-02-02 07:15:00", "2026-02-04 01:45:00", "2026-02-01 16:50:00", "2026-02-03 09:30:00"]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def digital_detox(logs):

    return logs
```

# --solutions--

```py
from datetime import datetime

def digital_detox(logs):
    if not logs:
        return True

    times = sorted(
        datetime.strptime(log, "%Y-%m-%d %H:%M:%S")
        for log in logs
    )

    FOUR_HOURS = 4 * 60 * 60

    for i in range(1, len(times)):
        diff = (times[i] - times[i - 1]).total_seconds()
        if diff < FOUR_HOURS:
            return False

    daily_counts = {}
    for time in times:
        day = time.date()
        daily_counts[day] = daily_counts.get(day, 0) + 1
        if daily_counts[day] > 2:
            return False

    return True
```
