---
id: 69738771fb5a7b8b24cca2a4
title: "Challenge 178: Truncate the Text"
challengeType: 29
dashedName: challenge-178
---

# --description--

문자열이 주어지면, 20자 이하면 그대로 반환하세요. 20자보다 길면 처음 17자만 남기고 `"..."`를 끝에 붙여서(총 20자가 되도록) 결과를 반환하세요.

# --hints--

`truncate_text("Hello, world!")` should return `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` should return `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` should return `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` should return `"...................."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```
