---
id: 69738771fb5a7b8b24cca2a2
title: "Challenge 176: Groundhog Day"
challengeType: 29
dashedName: challenge-176
---

# --description--

오늘은 그라운드호그 데이로, 그라운드호그가 자신의 그림자를 보았는지 여부에 따라 날씨를 예측합니다.

그라운드호그의 모습을 나타내는 값을 받아 올바른 예측을 반환하세요:

- 주어진 값이 불리언 `true` (그라운드호그가 그림자를 보았다면)일 경우 `"Looks like we'll have six more weeks of winter."`을 반환하세요.
- 값이 불리언 `false` (그라운드호그가 그림자를 보지 못했다면)일 경우 `"It's going to be an early spring."`을 반환하세요.
- 값이 그 외의 경우(그라운드호그가 나타나지 않았다면) `"No prediction this year."`를 반환하세요.

# --hints--

`groundhog_day_prediction(True)` should return `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` should return `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
