---
id: 69738771fb5a7b8b24cca29f
title: "Challenge 173: Valid Pawn Moves"
challengeType: 29
dashedName: challenge-173
---

# --description--

체스판에서 폰 하나의 위치가 주어지면, 그 폰이 이동할 수 있는 모든 유효한 칸을 오름차순으로 배열에 담아 반환하세요.

표준 체스판은 8x8 크기로, 열은 왼쪽에서 오른쪽으로 `A`부터 `H`까지 레이블이 붙어 있고, 행은 아래에서 위로 `1`부터 `8`까지 레이블이 붙어 있습니다. 다음과 같습니다:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

이 도전 과제에서:

- 당신은 체스판 아래쪽에 있는 플레이어입니다.
- 폰은 한 칸만 "위쪽"으로 이동할 수 있습니다.
- 단, 폰이 시작 행(2행)에 있으면 한 칸 또는 두 칸 위로 이동할 수 있습니다.

예를 들어, `"D4"`가 주어지면, 폰이 이동할 수 있는 유일한 칸인 `["D5"]`를 반환하세요.
`"B2"`가 주어지면, 시작 행에 있으므로 `["B3", "B4"]`를 반환하고, 오름차순으로 정렬하세요.

# --hints--

`find_pawn_moves("D4")` should return `["D5"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("D4"), ["D5"])`)
}})
```

`find_pawn_moves("B2")` should return `["B3", "B4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("B2"), ["B3", "B4"])`)
}})
```

`find_pawn_moves("A7")` should return `["A8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("A7"), ["A8"])`)
}})
```

`find_pawn_moves("G2")` should return `["G3", "G4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("G2"), ["G3", "G4"])`)
}})
```

`find_pawn_moves("E3")` should return `["E4"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_pawn_moves("E3"), ["E4"])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_pawn_moves(position):

    return position
```

# --solutions--

```py
def find_pawn_moves(position):
    column = position[0].upper()
    row = int(position[1])
    moves = []

    if row >= 8:
        return moves

    moves.append(f"{column}{row + 1}")

    if row == 2:
        moves.append(f"{column}{row + 2}")

    return moves
```
