---
id: 69738771fb5a7b8b24cca2a5
title: "Challenge 179: Pocket Change"
challengeType: 29
dashedName: challenge-179
---

# --description--

주머니에 있는 동전을 나타내는 정수 배열이 주어집니다. 각 정수는 센트 단위 동전의 가치를 나타내며, `"$D.CC"` 형식으로 총 금액을 반환하세요.

- 100센트는 1달러와 같습니다.
- 반환값에는 1달러 미만 금액에 대해 선행 0을 포함하고 센트는 항상 정확히 두 자리 숫자로 표시하세요.

# --hints--

`count_change([25, 10, 5, 1])` should return `"$0.41"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1]), "$0.41")`)
}})
```

`count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` should return `"$1.43"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43")`)
}})
```

`count_change([100, 25, 100, 1000, 5, 500, 2000, 25])` should return `"$37.55"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55")`)
}})
```

`count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` should return `"$0.70"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70")`)
}})
```

`count_change([1])` should return `"$0.01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([1]), "$0.01")`)
}})
```

`count_change([25, 25, 25, 25])` should return `"$1.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_change([25, 25, 25, 25]), "$1.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def count_change(change):

    return change
```

# --solutions--

```py
def count_change(change):
    total_cents = sum(change)
    dollars = total_cents // 100
    cents = total_cents % 100
    return f"${dollars}.{cents:02d}"
```
