---
id: 697a49e6ff50d756c9b6935f
title: "Challenge 182: 2026 Winter Games Day 3: Biathlon"
challengeType: 29
dashedName: challenge-182
---

# --description--

각 값이 바이애슬론 한 라운드에서 맞힌 표적 수를 나타내는 정수 배열이 주어질 때, 선수에게 부과되는 총 벌칙 거리(스키를 더 타야 하는 거리)를 반환하세요.

- 각 라운드는 5개의 표적으로 구성됩니다.
- 놓친 표적 하나당 150미터 벌칙 루프가 부과됩니다.

# --hints--

`calculate_penalty_distance([4, 4])` should return `300`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` should return `450`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` should return `150`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` should return `1500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
