---
id: 697a49e6ff50d756c9b6935e
title: "Challenge 181: 2026 Winter Games Day 2: Snowboarding"
challengeType: 29
dashedName: challenge-181
---

# --description--

스노보더의 시작 자세와 도 단위 회전이 주어졌을 때, 착지 자세를 결정하세요.

- 스노보더의 자세는 `"Regular"` 또는 `"Goofy"`입니다.
- 묘기 회전은 90도의 배수입니다. 양수는 시계 방향 회전을, 음수는 반시계 방향 회전을 나타냅니다.
- 착지 자세는 180도 회전할 때마다 뒤집힙니다.

예를 들어, `"Regular"`와 `90`가 주어지면 `"Regular"`를 반환하세요. `"Regular"`와 `180` 도가 주어지면 `"Goofy"`를 반환하세요.

# --hints--

`get_landing_stance("Regular", 90)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` should return `"Goofy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` should return `"Regular"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
