---
id: 697a49e6ff50d756c9b69360
title: "Challenge 183: 2026 Winter Games Day 4: Ski Jumping"
challengeType: 29
dashedName: challenge-183
---

# --description--

거리 점수, 스타일 점수, 바람 보정 값, K-포인트 보너스 값을 받아 스키 점프 점수를 계산하고 메달 획득 여부를 결정하세요.

- 점수는 위 네 가지 값을 모두 더해서 계산합니다.

다른 점프 선수들의 현재 총점은 다음과 같습니다:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- 점수가 가장 높으면 `"Gold"`를 반환하세요
- 두 번째로 높으면 `"Silver"`를 반환하세요
- 세 번째로 높으면 `"Bronze"`를 반환하세요
- 그렇지 않으면 `"No Medal"`를 반환하세요

# --hints--

`ski_jump_medal(125.0, 58.0, 0.0, 6.0)` should return `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(125.0, 58.0, 0.0, 6.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 50.0, 1.0, 4.0)` should return `"Bronze"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 50.0, 1.0, 4.0), "Bronze")`)
}})
```

`ski_jump_medal(122.0, 52.0, -1.0, 4.0)` should return `"Silver"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(122.0, 52.0, -1.0, 4.0), "Silver")`)
}})
```

`ski_jump_medal(118.0, 50.5, -1.5, 4.0)` should return `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(118.0, 50.5, -1.5, 4.0), "No Medal")`)
}})
```

`ski_jump_medal(124.0, 50.5, 2.0, 5.0)` should return `"Gold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(124.0, 50.5, 2.0, 5.0), "Gold")`)
}})
```

`ski_jump_medal(119.0, 49.5, 0.0, 3.0)` should return `"No Medal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(ski_jump_medal(119.0, 49.5, 0.0, 3.0), "No Medal")`)
}})
```

# --seed--

## --seed-contents--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):

    return distance_points
```

# --solutions--

```py
def ski_jump_medal(distance_points, style_points, wind_comp, k_point_bonus):
    my_score = distance_points + style_points + wind_comp + k_point_bonus
    other_scores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0]
    all_scores = other_scores + [my_score]
    all_scores.sort(reverse=True)
    rank = all_scores.index(my_score) + 1

    if rank == 1:
        return "Gold"
    elif rank == 2:
        return "Silver"
    elif rank == 3:
        return "Bronze"
    else:
        return "No Medal"
```
