---
id: 697a49e9860d24853adef67d
title: "Challenge 192: 2026 Winter Games Day 13: Nordic Combined"
challengeType: 29
dashedName: challenge-192
---

# --description--

선수들의 점프 점수 배열이 주어졌을 때, 노르딕 복합의 크로스컨트리 구간 출발 지연 시간을 계산하세요.

가장 높은 점프 점수를 받은 선수가 가장 먼저 출발합니다(지연 시간 0초). 나머지 선수들은 최고 점수와의 차이에 따라 더 늦게 출발합니다.

각 선수의 지연 시간을 계산하려면, 최고 점수에서 해당 선수의 점프 점수를 빼고 그 결과에 1.5를 곱하세요. 지연 시간은 가장 가까운 정수로 올림하세요.

# --hints--

`calculate_start_delays([120, 110, 125])` should return `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` should return `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` should return `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` should return `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```
