---
id: 697a49e6ff50d756c9b69361
title: "Challenge 184: 2026 Winter Games Day 5: Cross-Country Skiing"
challengeType: 29
dashedName: challenge-184
---

# --description--

크로스컨트리 스키 경주의 완주 시간 배열이 주어질 때, 이를 우승자와의 시간 차이로 변환하세요.

- 주어진 시간은 `"H:MM:SS"` 형식의 문자열입니다.
- 주어진 시간은 가장 빠른 순서부터 느린 순서로 정렬되어 있습니다.
- 우승자의 시간(가장 빠른 시간)은 `"0"`에 해당해야 합니다.
- 나머지 시간들은 `"+M:SS"` 형식으로 우승자와의 시간 차이를 보여야 합니다.

예를 들어, `["1:25:32", "1:26:10", "1:27:05"]`가 주어지면 `["0", "+0:38", "+1:33"]`를 반환하세요.

# --hints--

`get_relative_results(["1:25:32", "1:26:10", "1:27:05"])` should return `["0", "+0:38", "+1:33"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:25:32", "1:26:10", "1:27:05"]), ["0", "+0:38", "+1:33"])`)
}})
```

`get_relative_results(["1:00:01", "1:00:05", "1:00:10"])` should return `["0", "+0:04", "+0:09"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:00:01", "1:00:05", "1:00:10"]), ["0", "+0:04", "+0:09"])`)
}})
```

`get_relative_results(["1:10:06", "1:10:23", "1:10:48", "1:12:11"])` should return `["0", "+0:17", "+0:42", "+2:05"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["1:10:06", "1:10:23", "1:10:48", "1:12:11"]), ["0", "+0:17", "+0:42", "+2:05"])`)
}})
```

`get_relative_results(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"])` should return `["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["0:49:13", "0:49:15", "0:50:14", "0:51:30", "0:51:58", "0:52:16", "0:53:12", "0:53:31", "0:56:19", "1:02:20"]), ["0", "+0:02", "+1:01", "+2:17", "+2:45", "+3:03", "+3:59", "+4:18", "+7:06", "+13:07"])`)
}})
```

`get_relative_results(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"])` should return `["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_relative_results(["2:01:15", "2:10:45", "2:10:53", "2:11:04", "2:11:55", "2:13:27", "2:14:30", "2:15:10"]), ["0", "+9:30", "+9:38", "+9:49", "+10:40", "+12:12", "+13:15", "+13:55"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_relative_results(results):

    return results
```

# --solutions--

```py
def get_relative_results(results):
    def time_to_seconds(time_str):
        hours, minutes, seconds = map(int, time_str.split(':'))
        return hours * 3600 + minutes * 60 + seconds
    
    def seconds_to_time_format(secs):
        mins = secs // 60
        secs_remainder = secs % 60
        return f"+{mins}:{secs_remainder:02d}"
    
    winner_seconds = time_to_seconds(results[0])
    
    relative_results = []
    for i, time in enumerate(results):
        if i == 0:
            relative_results.append('0')
        else:
            current_seconds = time_to_seconds(time)
            difference = current_seconds - winner_seconds
            relative_results.append(seconds_to_time_format(difference))
    
    return relative_results
```
