---
id: 697a49e6ff50d756c9b69366
title: "Challenge 189: 2026 Winter Games Day 10: Alpine Skiing"
challengeType: 29
dashedName: challenge-189
---

# --description--

스키장의 수직 낙차, 수평 거리, 그리고 유형이 주어졌을 때, 경사의 난이도 등급을 결정하세요.

난이도 등급을 결정하려면:

- 낙차를 거리로 나누어 경사의 가파름을 계산하세요.
- 그런 다음, 경사 가파름에 경사 유형에 따른 조정 계수를 곱하세요:
  - `"Downhill"`: 경사 가파름에 1.2를 곱하세요
  - `"Slalom"`: 경사 가파름에 0.9를 곱하세요
  - `"Giant Slalom"`: 경사 가파름에 1.0을 곱하세요

Return:

- 조정된 경사 가파름이 0.1 이하이면 `"Green"`을 반환하세요
- 조정된 경사 가파름이 0.1 초과 0.25 이하이면 `"Blue"`를 반환하세요
- 조정된 경사 가파름이 0.25 초과이면 `"Black"`를 반환하세요


# --hints--

`get_hill_rating(95, 900, "Slalom")` should return `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` should return `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` should return `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` should return `"Green"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` should return `"Blue"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` should return `"Black"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
