---
id: 698a1a73ade5ac0e19180fa9
title: "Challenge 206: Playing Card Values"
challengeType: 29
dashedName: challenge-206
---

# --description--

주어진 카드 배열에서 각 카드의 숫자 값을 담은 새 배열을 반환하세요.

Card Values:

- 에이스(`"A"`)는 값이 1입니다.
- 숫자 카드(`"2"` - `"10"`)는 각각 2부터 10까지 해당 숫자 값을 가집니다.
- 그림 카드: 잭(`"J"`), 퀸(`"Q"`), 킹(`"K"`)은 각각 10의 값입니다.

Suits:

- 각 카드는 스페이드(`"S"`), 클로버(`"C"`), 다이아몬드(`"D"`), 하트(`"H"`) 중 하나의 무늬를 가집니다.

Card Format:

- 각 카드는 문자열로 나타냅니다: `"valueSuit"`. 예를 들어, `"AS"`은 스페이드 에이스, `"10H"`는 하트 10, `"QC"`은 클로버 퀸입니다.

# --hints--

`card_values(["3H", "4D", "5S"])` should return `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` should return `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` should return `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` should return `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` should return `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```
