---
id: 698a1a73ade5ac0e19180fa5
title: "Challenge 202: Add Punctuation"
challengeType: 29
dashedName: challenge-202
---

# --description--

마침표가 빠진 문장 문자열이 주어졌을 때, 다음 위치에 마침표(`"."`)를 추가하세요:

- 대문자 바로 앞에 오는 공백 앞에
- 그리고 문자열 끝에

결과 문자열을 반환하세요.

# --hints--

`add_punctuation("Hello world")` should return `"Hello world."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world"), "Hello world.")`)
}})
```

`add_punctuation("Hello world It's nice today")` should return `"Hello world. It's nice today."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Hello world It's nice today"), "Hello world. It's nice today.")`)
}})
```

`add_punctuation("JavaScript is great Sometimes")` should return `"JavaScript is great. Sometimes."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.")`)
}})
```

`add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` should return `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.")`)
}})
```

`add_punctuation("Wait.. For it")` should return `"Wait... For it."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(add_punctuation("Wait.. For it"), "Wait... For it.")`)
}})
```

# --seed--

## --seed-contents--

```py
def add_punctuation(sentences):

    return sentences
```

# --solutions--

```py
def add_punctuation(sentences):
    result = ""
    length = len(sentences)

    for i, c in enumerate(sentences):
        if c == " " and i + 1 < length and sentences[i + 1].isupper():
            result += "."
        result += c

    if not result.endswith("."):
        result += "."

    return result
```
