---
id: 698a1a863194f1f4e63f645e
title: "Challenge 207: Smallest Gap"
challengeType: 29
dashedName: challenge-207
---

# --description--

문자열이 주어지면, 두 개의 동일한 문자 사이에 있는 가장 작은 문자 수(가장 작은 간격)를 가진 부분 문자열을 반환하세요.

- 항상 일치하는 문자 쌍이 최소 한 쌍 이상 있습니다.
- 반환하는 부분 문자열에는 일치하는 문자는 포함하지 않습니다.
- 두 개 이상의 간격이 같은 길이일 경우, 첫 번째 간격에 해당하는 문자를 반환하세요.

예를 들어 `"ABCDAC"`가 주어지면 `"DA"`를 반환하는데, 그 이유는:

- 문자열에서 `"A"`와 `"C"`만 반복됩니다.
- 두 `"A"` 문자 사이의 문자 수는 3이고, 두 `"C"` 문자 사이의 문자 수는 2입니다.
- 따라서 두 `"C"` 문자 사이의 문자열을 반환합니다.

# --hints--

`smallest_gap("ABCDAC")` should return `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` should return `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` should return `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` should return `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` should return `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
