---
id: 698a1a73ade5ac0e19180fa2
title: "Challenge 199: Sequential Difference"
challengeType: 29
dashedName: challenge-199
---

# --description--

숫자 배열이 주어지면, 각 숫자에서 다음 숫자로 가기 위해 필요한 값을 포함하는 새 배열을 반환하세요.

- 마지막 숫자의 경우 다음 숫자가 없으므로 `0`를 사용하세요.

예를 들어, `[1, 2, 4, 7]`이 주어지면 `[1, 2, 3, 0]`를 반환하세요.

# --hints--

`find_differences([1, 2, 4, 7])` should return `[1, 2, 3, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 4, 7]), [1, 2, 3, 0])`)
}})
```

`find_differences([10, 15, 19, 22, 24, 25])` should return `[5, 4, 3, 2, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([10, 15, 19, 22, 24, 25]), [5, 4, 3, 2, 1, 0])`)
}})
```

`find_differences([25, 20, 16, 13, 11, 10])` should return `[-5, -4, -3, -2, -1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([25, 20, 16, 13, 11, 10]), [-5, -4, -3, -2, -1, 0])`)
}})
```

`find_differences([0, 1, 2, 2, 3, 3, 4, 5])` should return `[1, 1, 0, 1, 0, 1, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([0, 1, 2, 2, 3, 3, 4, 5]), [1, 1, 0, 1, 0, 1, 1, 0])`)
}})
```

`find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1])` should return `[1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_differences([1, 2, 5, 12, 34, -15, -1, 41, 113, -222, -99, -40, 10, -18, -6, -2, -1]), [1, 3, 7, 22, -49, 14, 42, 72, -335, 123, 59, 50, -28, 12, 4, 1, 0])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_differences(arr):

    return arr
```

# --solutions--

```py
def find_differences(arr):
    result = []

    for i in range(len(arr)):
        if i == len(arr) - 1:
            result.append(0)
        else:
            result.append(arr[i + 1] - arr[i])

    return result
```
