---
id: 698a1a73ade5ac0e19180fa8
title: "Challenge 205: Perfect Cube Count"
challengeType: 29
dashedName: challenge-205
---

# --description--

두 정수가 주어졌을 때, 두 수를 포함한 범위 내에 존재하는 완전 세제곱수의 개수를 구하세요.

- 더 작은 수가 첫 번째 인수라는 보장은 없습니다.
- 어떤 수가 완전 세제곱수라는 것은 정수 `n`가 존재하여 `n * n * n = number`인 경우입니다. 예를 들어, 27은 `3 * 3 * 3 = 27`이기 때문에 완전 세제곱수입니다.

# --hints--

`count_perfect_cubes(3, 30)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(3, 30), 2)`)
}})
```

`count_perfect_cubes(1, 30)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(1, 30), 3)`)
}})
```

`count_perfect_cubes(30, 0)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(30, 0), 4)`)
}})
```

`count_perfect_cubes(-64, 64)` should return `9`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(-64, 64), 9)`)
}})
```

`count_perfect_cubes(9214, -8127)` should return `41`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_perfect_cubes(9214, -8127), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_perfect_cubes(a, b):

    return a
```

# --solutions--

```py
def count_perfect_cubes(a, b):
    start, end = min(a, b), max(a, b)
    count = 0

    n = 0
    while n ** 3 <= end:
        if n ** 3 >= start:
            count += 1
        n += 1

    n = -1
    while n ** 3 >= start:
        if n ** 3 <= end:
            count += 1
        n -= 1

    return count
```
