---
id: 698a1a863194f1f4e63f6460
title: "Challenge 209: Element Size"
challengeType: 29
dashedName: challenge-209
---

# --description--

주어진 창 크기, 뷰포트 너비 `"vw"` 단위로 표현된 요소의 너비, 뷰포트 높이 `"vh"` 단위로 표현된 요소의 높이를 기준으로 요소의 크기를 픽셀 단위로 계산하세요.

- 주어진 창 크기와 반환할 요소 크기는 `"width x height"`, `"1200 x 800"` 형식의 문자열입니다.

- `"vw"` 단위는 창 너비의 백분율입니다. 예를 들어 `"50vw"` 는 창 너비의 50%를 의미합니다.
- `"vh"` 단위는 창 높이의 백분율입니다. 예를 들어 `"50vh"` 는 창 높이의 50%를 의미합니다.

# --hints--

`get_element_size("1200 x 800", "50vw", "50vh")` should return `"600 x 400"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "50vw", "50vh"), "600 x 400")`)
}})
```

`get_element_size("320 x 480", "25vw", "50vh")` should return `"80 x 240"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("320 x 480", "25vw", "50vh"), "80 x 240")`)
}})
```

`get_element_size("1000 x 500", "7vw", "3vh")` should return `"70 x 15"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1000 x 500", "7vw", "3vh"), "70 x 15")`)
}})
```

`get_element_size("1920 x 1080", "95vw", "100vh")` should return `"1824 x 1080"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1920 x 1080", "95vw", "100vh"), "1824 x 1080")`)
}})
```

`get_element_size("1200 x 800", "0vw", "0vh")` should return `"0 x 0"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1200 x 800", "0vw", "0vh"), "0 x 0")`)
}})
```

`get_element_size("1440 x 900", "100vw", "114vh")` should return `"1440 x 1026"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_element_size("1440 x 900", "100vw", "114vh"), "1440 x 1026")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_element_size(window_size, element_vw, element_vh):

    return window_size
```

# --solutions--

```py
def get_element_size(window_size, element_vw, element_vh):
    window_width, window_height = map(int, window_size.split(" x "))

    vw = int(element_vw.replace("vw", ""))
    vh = int(element_vh.replace("vh", ""))

    width_px = round((vw / 100) * window_width)
    height_px = round((vh / 100) * window_height)

    return f"{width_px} x {height_px}"
```
