---
id: 699c8e045ee7cb94ed2322dd
title: "Challenge 222: Equinox Shadows"
challengeType: 29
dashedName: challenge-222
---

# --description--

오늘은 춘분으로, 태양이 적도 바로 위에 있고 정오에 머리 위에 있습니다. 주어진 시간에 4피트 높이의 수직 기둥이 만드는 그림자를 구하세요.

- 시간은 `"HH:MM"` 24시간 형식의 문자열입니다 (예: `"15:00"`은 오후 3시입니다).
- 시간은 30분 단위로만 주어집니다.

Rules:

- 태양은 오전 6시에 `"east"`에서 떠서 오후 6시에 `"west"`에서 집니다.
- 그림자는 항상 태양과 반대 방향을 가리킵니다.
- 그림자의 길이(피트 단위)는 정오에서 떨어진 시간(시간 단위)의 세제곱입니다.
- 일출 전(오전 6시 이전), 일몰 후(오후 6시 이후), 그리고 정오에는 그림자가 없습니다.

Return:

- 그림자가 있으면 `"(length)ft (direction)"`를 반환하세요. 예를 들어 `"8ft west"`입니다.
- 그렇지 않으면 `"No shadow"`를 반환하세요.

예를 들어 `"10:00"`이 주어지면, 오전 10시는 정오에서 2시간 떨어져 있으므로 2<sup>3</sup> = 8피트이고, 태양이 동쪽에 있으므로 그림자는 서쪽을 가리켜 `"8ft west"`를 반환합니다.

# --hints--

`get_shadow("10:00")` should return `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` should return `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` should return `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` should return `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` should return `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
