---
id: 699c8e045ee7cb94ed2322d5
title: "Challenge 214: Domino Chain Validator"
challengeType: 29
dashedName: challenge-214
---

# --description--

도미노 수열을 나타내는 2차원 배열이 주어졌을 때, 유효한 체인을 이루는지 판단하세요.

- 배열의 각 요소는 도미노를 나타내며, 1부터 6까지(포함) 두 숫자로 이루어진 배열입니다.
- 체인이 유효하려면 각 도미노의 두 번째 숫자가 다음 도미노의 첫 번째 숫자와 일치해야 합니다.
- 첫 번째 도미노의 첫 숫자와 마지막 도미노의 마지막 숫자는 아무것과도 일치할 필요가 없습니다.

# --hints--

`is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

`is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

`is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

`is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

`is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
