---
id: 699c8e045ee7cb94ed2322d6
title: "Challenge 215: Parking Fee Calculator"
challengeType: 29
dashedName: challenge-215
---

# --description--

차를 주차한 시간과 차를 찾은 시간을 나타내는 두 개의 문자열이 주어질 때, 주차 요금을 계산하세요.

- 주어진 문자열은 24시간 시계를 사용하여 `"HH:MM"` 형식으로 되어 있습니다. 예를 들어 `"14:00"`은 오후 2시입니다.
- 첫 번째 문자열은 차를 주차한 시간이고, 두 번째는 차를 찾은 시간입니다.
- 찾은 시간이 주차 시간보다 이르면, 주차가 자정을 넘어 다음 날까지 이어졌다는 뜻입니다.

Fee rules:

- 주차한 시간당 3달러입니다.
- 부분 시간은 다음 한 시간으로 올림합니다.
- 주차가 자정을 넘어 밤새 이어지면 추가로 10달러의 야간 요금이 부과됩니다.
- 총 요금이 5달러 미만일 경우 최솟값으로 5달러가 적용됩니다.

총 요금을 `"$cost"`, `"$5"` 형식으로 반환하세요.

# --hints--

`calculate_parking_fee("09:00", "11:00")` should return `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` should return `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` should return `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` should return `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` should return `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` should return `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
