---
id: 69a890af247de743333bd4cc
title: "Challenge 223: QR Decoder"
challengeType: 29
dashedName: challenge-223
---

# --description--

6x6 행렬(배열의 배열)로 주어진 QR 코드를 나타내는 이진 데이터 문자열을 반환하세요.

- QR 코드는 90도 단위로 회전된 어떤 방향으로도 주어질 수 있습니다.
- 올바르게 방향이 맞춰진 코드는 왼쪽 아래, 왼쪽 위, 오른쪽 위 모서리에 2x2 크기의 `1` (방향 표시자) 그룹이 있습니다.
- 세 개의 2x2 방향 표시자는 이진 데이터에 포함되지 않습니다.
- QR 코드가 올바르게 방향이 맞춰졌을 때 이진 데이터는 왼쪽에서 오른쪽, 위에서 아래 순서로(책을 읽듯이) 읽습니다.
- 코드는 항상 정확히 하나의 유효한 방향을 가집니다.

예를 들어, 다음이 주어졌을 때:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

또는 같은 코드가 다른 방향으로 주어졌을 때:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

세 개의 2x2 방향 표시자를 제외한 모든 이진 데이터인 `"000000000000000000000001"`를 반환하세요.

# --hints--

`decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"])` should return `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001")`)
}})
```

`decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"])` should return `"000000000000000000000001"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001")`)
}})
```

`decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"])` should return `"001101000011000000110100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100")`)
}})
```

`decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"])` should return `"010001000100010101010110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110")`)
}})
```

`decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"])` should return `"010000100100100101001110"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_qr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_qr(qr_code):

    return qr_code
```

# --solutions--

```py
def is_oriented(code):
    marker_indices = [
        (0, 0),
        (0, 1),
        (0, 4),
        (0, 5),
        (1, 0),
        (1, 1),
        (1, 4),
        (1, 5),
        (4, 0),
        (4, 1),
        (5, 0),
        (5, 1)
    ]

    for row, col in marker_indices:
        if code[row][col] != '1':
            return False

    return True


def rotate(code):
    new_code = ['', '', '', '', '', '']
    for i in range(6):
        for j in range(5, -1, -1):
            new_code[i] += code[j][i]
    return new_code


def get_data(code):
    data = ''
    data += code[0][2:4]
    data += code[1][2:4]
    data += code[2]
    data += code[3]
    data += code[4][2:]
    data += code[5][2:]
    return data


def decode_qr(qr_code):
    code = qr_code[:]

    while not is_oriented(code):
        code = rotate(code)

    return get_data(code)
```
