---
id: 69a890af247de743333bd4d0
title: "Challenge 227: Cooldown Time"
challengeType: 29
dashedName: challenge-227
---

# --description--

두 개의 타임스탬프가 주어집니다. 첫 번째는 사용자가 시험을 마친 시점을 나타내고, 두 번째는 현재 시간을 나타냅니다. 사용자가 다시 시험을 볼 수 있는지 판단하세요.

- 두 타임스탬프는 모두 `"YYYY-MM-DDTHH:MM:SS"` 형식으로 주어집니다. 예를 들어 `"2026-03-25T14:00:00"`입니다. 시간은 24시간제로 표시됩니다.
- 사용자는 시험을 다시 보기 전에 최소 48시간을 기다려야 합니다.

# --hints--

`can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T08:00:00", "2026-03-25T14:00:00"), True)`)
}})
```

`can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-24T14:00:00", "2026-03-25T10:00:00"), False)`)
}})
```

`can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-23T09:25:00", "2026-03-25T09:25:00"), True)`)
}})
```

`can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_retake("2026-03-25T11:50:00", "2026-03-23T11:49:59"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_retake(finish_time, current_time):

    return finish_time
```

# --solutions--

```py
from datetime import datetime, timedelta

def can_retake(finish_time, current_time):
    fmt = "%Y-%m-%dT%H:%M:%S"
    finished_dt = datetime.strptime(finish_time, fmt)
    current_dt = datetime.strptime(current_time, fmt)
    diff = current_dt - finished_dt
    cooldown = timedelta(hours=48)

    return diff >= cooldown
```
