---
id: 69a890af247de743333bd4d2
title: "도전 과제 229: 텍스트 자르기 2"
challengeType: 29
dashedName: challenge-229
---

# --description--

문자열이 주어지면, 문자들의 총 너비가 50 단위를 넘지 않도록 잘라낸 새 문자열을 반환하세요.

각 문자는 특정 너비를 가집니다:

| Letters | Width |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

위 표는 모든 대문자와 소문자를 포함합니다. 추가로:

- 공백(`" "`)은 너비가 2입니다
- 마침표(`"."`)는 너비가 1입니다

- 주어진 문자열이 50 단위 이하라면, 문자열을 그대로 반환하세요. 그렇지 않으면
- 문자열을 자르고 끝에 마침표 세 개(`"..."`)를 추가하세요. 마침표 세 개를 포함한 총 너비가 50 단위를 넘지 않으면서 최대한 가까워야 합니다.

# --hints--

`truncate_text("The quick brown fox")` should return `"The quick brown f..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The quick brown fox"), "The quick brown f...")`)
}})
```

`truncate_text("The silky smooth sloth")` should return `"The silky smooth s..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The silky smooth sloth"), "The silky smooth s...")`)
}})
```

`truncate_text("THE LOUD BRIGHT BIRD")` should return `"THE LOUD BRIG..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...")`)
}})
```

`truncate_text("The fast striped zebra")` should return `"The fast striped z..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The fast striped zebra"), "The fast striped z...")`)
}})
```

`truncate_text("The big black bear")` should return `"The big black bear"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("The big black bear"), "The big black bear")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(s):

    return s
```

# --solutions--

```py
def truncate_text(s):
    MAX_WIDTH = 50
    ELLIPSIS = "..."
    ELLIPSIS_WIDTH = 3
    TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH

    char_groups = {
        "ilI.": 1,
        "fjrt ": 2,
        "abcdeghkmnopqrstuvwxyzJL": 3,
        "ABCDEFGHKMNOPQRSTUVWXYZ": 4
    }

    def get_char_width(char):
        for key, width in char_groups.items():
            if char in key:
                return width
        return 3

    def string_width(string):
        total = 0
        for char in string:
            total += get_char_width(char)
        return total

    if string_width(s) <= MAX_WIDTH:
        return s

    result = ""
    total_width = 0

    for char in s:
        char_width = get_char_width(char)

        if total_width + char_width > TRUNCATE_LIMIT:
            break

        result += char
        total_width += char_width

    return result + ELLIPSIS
```
