---
id: 69b1028d6e265413d0198a2a
title: "Challenge 231: ISBN-10 Validator"
challengeType: 29
dashedName: challenge-231
---

# --description--

문자열이 주어지면, 그것이 유효한 ISBN-10인지 판단하세요.

ISBN-10은 하이픈(`"-"`)과 10개의 다른 문자로 구성됩니다. 하이픈을 제거한 후(`"-"`):

- 처음 9개 문자는 숫자여야 하며,
- 마지막 문자는 숫자이거나 숫자 10을 나타내는 문자 `"X"`일 수 있습니다.

유효성을 검사하려면:

- 각 숫자(또는 값)에 그 위치를 곱하세요(첫 번째 숫자는 1을 곱하고, 두 번째는 2를 곱하는 식으로).
- 모든 결과를 더하세요.
- 합계가 11로 나누어떨어지면 유효합니다.

# --hints--

`is_valid_isbn10("0-306-40615-2")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-2"), True)`)
}})
```

`is_valid_isbn10("0-306-40615-1")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-306-40615-1"), False)`)
}})
```

`is_valid_isbn10("0-8044-2957-X")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-8044-2957-X"), True)`)
}})
```

`is_valid_isbn10("X-306-40615-2")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("X-306-40615-2"), False)`)
}})
```

`is_valid_isbn10("0-6822-2589-4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn10("0-6822-2589-4"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn10(s):

    return s
```

# --solutions--

```py
def is_valid_isbn10(s):
    isbn = s.replace("-", "")

    if len(isbn) != 10:
        return False

    total = 0

    for i in range(9):
        if not isbn[i].isdigit():
            return False
        total += int(isbn[i]) * (i + 1)

    last = isbn[9]

    if last == "X":
        total += 10 * 10
    elif last.isdigit():
        total += int(last) * 10
    else:
        return False

    return total % 11 == 0
```
