---
id: 69b1028d6e265413d0198a2c
title: "Challenge 233: Wake-Up Alarm"
challengeType: 29
dashedName: challenge-233
---

# --description--

알람을 맞춘 시간을 나타내는 문자열과 실제로 일어난 시간을 나타내는 문자열이 주어질 때, 일찍 일어났는지, 제시간에 일어났는지, 아니면 늦게 일어났는지 판단하세요.

- 두 시간 모두 `"HH:MM"` 24시간 형식으로 주어집니다.

Return:

- 알람 시간보다 일찍 일어났으면 `"early"`을 반환하세요.
- 알람 시간에 일어났거나 알람 시간 이후 10분 이내의 스누즈 시간 내에 일어났으면 `"on time"`를 반환하세요.
- 알람 시간보다 10분 이상 늦게 일어났으면 `"late"`을 반환하세요.

두 시간은 같은 날에 해당합니다.

# --hints--

`alarm_check("07:00", "06:45")` should return `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("07:00", "06:45"), "early")`)
}})
```

`alarm_check("06:30", "06:30")` should return `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("06:30", "06:30"), "on time")`)
}})
```

`alarm_check("08:10", "08:15")` should return `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:10", "08:15"), "on time")`)
}})
```

`alarm_check("09:30", "09:45")` should return `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("09:30", "09:45"), "late")`)
}})
```

`alarm_check("08:15", "08:25")` should return `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:15", "08:25"), "on time")`)
}})
```

`alarm_check("05:45", "05:56")` should return `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("05:45", "05:56"), "late")`)
}})
```

`alarm_check("04:30", "04:00")` should return `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("04:30", "04:00"), "early")`)
}})
```

# --seed--

## --seed-contents--

```py
def alarm_check(alarm_time, wake_time):

    return alarm_time
```

# --solutions--

```py
def alarm_check(alarm_time, wake_time):
    alarm_h, alarm_m = map(int, alarm_time.split(":"))
    wake_h, wake_m = map(int, wake_time.split(":"))

    alarm_minutes = alarm_h * 60 + alarm_m
    wake_minutes = wake_h * 60 + wake_m
    diff = wake_minutes - alarm_minutes

    if diff < 0:
        return "early"
    if diff <= 10:
        return "on time"
    return "late"
```
