---
id: 69b1028d6e265413d0198a2e
title: "Challenge 235: Capitalized Fibonacci"
challengeType: 29
dashedName: challenge-235
---

# --description--

문자열이 주어지면, 인덱스가 피보나치 수인 문자는 대문자로, 그렇지 않은 문자는 소문자로 변환한 새 문자열을 반환하세요.

피보나치 수열은 각 숫자가 앞의 두 숫자의 합인 수들의 수열입니다. 수열의 처음 10개 숫자는 `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`입니다.

- 첫 번째 문자의 인덱스는 `0`입니다.
- 문자가 아닌 문자의 인덱스가 피보나치 수이면, 변경하지 않고 그대로 두세요.

# --hints--

`capitalize_fibonacci("hello world")` should return `"HELLo woRld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize_fibonacci("hello world"), "HELLo woRld")`)
}})
```

`capitalize_fibonacci("HELLO WORLD")` should return `"HELLo woRld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize_fibonacci("HELLO WORLD"), "HELLo woRld")`)
}})
```

`capitalize_fibonacci("hello, world!")` should return `"HELLo, wOrld!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize_fibonacci("hello, world!"), "HELLo, wOrld!")`)
}})
```

`capitalize_fibonacci("The quick brown fox jumped over the lazy dog.")` should return `"THE qUicK broWn fox jUmped over thE lazy dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize_fibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.")`)
}})
```

`capitalize_fibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` should return `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(capitalize_fibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.")`)
}})
```

# --seed--

## --seed-contents--

```py
def capitalize_fibonacci(s):

    return s
```

# --solutions--

```py
def capitalize_fibonacci(s):
    fibs = {0, 1}
    a, b = 0, 1
    while b < len(s):
        a, b = b, a + b
        fibs.add(b)

    result = []
    for i, char in enumerate(s):
        if not char.isalpha():
            result.append(char)
        elif i in fibs:
            result.append(char.upper())
        else:
            result.append(char.lower())

    return "".join(result)
```
