---
id: 69b1028d6e265413d0198a2d
title: "Challenge 234: Prank Number"
challengeType: 29
dashedName: challenge-234
---

# --description--

하나를 제외한 모든 숫자가 일정한 규칙을 따르는 숫자 배열이 주어질 때, 규칙을 따르지 않는 숫자 하나를 수정한 새 배열을 반환하세요.

규칙은 다음 중 하나입니다:

- 숫자가 일정한 값만큼 증가합니다(덧셈).
- 숫자가 일정한 값만큼 감소합니다(뺄셈).

예를 들어, `[2, 4, 7, 8, 10]`가 주어지면 `[2, 4, 6, 8, 10]`을 반환하세요.

# --hints--

`fix_prank_number([2, 4, 7, 8, 10])` should return `[2, 4, 6, 8, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10])`)
}})
```

`fix_prank_number([10, 10, 8, 7, 6])` should return `[10, 9, 8, 7, 6]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6])`)
}})
```

`fix_prank_number([12, 24, 36, 48, 61, 72, 84, 96])` should return `[12, 24, 36, 48, 60, 72, 84, 96]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96])`)
}})
```

`fix_prank_number([4, 1, -2, -5, -8, -5])` should return `[4, 1, -2, -5, -8, -11]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11])`)
}})
```

`fix_prank_number([0, 100, 200, 300, 150, 500])` should return `[0, 100, 200, 300, 400, 500]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500])`)
}})
```

`fix_prank_number([400, 425, 400, 375, 350, 325, 300])` should return `[450, 425, 400, 375, 350, 325, 300]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300])`)
}})
```

`fix_prank_number([-5, 5, 10, 15, 20])` should return `[0, 5, 10, 15, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20])`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_prank_number(arr):

    return arr
```

# --solutions--

```py
def fix_prank_number(arr):
    diffs = [arr[i + 1] - arr[i] for i in range(len(arr) - 1)]

    step = diffs[0]
    for i in range(len(diffs) - 1):
        if diffs[i] == diffs[i + 1]:
            step = diffs[i]
            break

    result = arr[:]

    for i in range(len(result) - 1):
        if result[i + 1] - result[i] != step:
            if i == 0 and result[2] - result[1] == step:
                result[0] = result[1] - step
            else:
                result[i + 1] = result[i] + step
            break

    return result
```
