---
id: 69b1028d6e265413d0198a2f
title: "도전 과제 236: 브라우저 기록"
challengeType: 29
dashedName: challenge-236
---

# --description--

브라우저 명령어 배열이 주어졌을 때, 두 개의 값을 가진 배열을 반환하세요: URL 배열로 된 기록과 현재 페이지의 인덱스입니다.

유효한 명령어는 다음과 같습니다:

- `"URL"` - 여기서 URL은 웹 주소입니다 (예: `"freecodecamp.org"`). 주어진 URL로 이동하며, 기록의 다음 위치에 추가하고 앞으로의 기록은 삭제합니다.
- `"Back"` - 기록에서 이전 페이지로 이동하거나, 이전 페이지가 없으면 현재 페이지에 머뭅니다.
- `"Forward"` - 기록에서 다음 페이지로 이동하거나, 다음 페이지가 없으면 현재 페이지에 머뭅니다.

예를 들어, `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`가 주어지면 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`를 반환하세요.

# --hints--

`get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"])`는 `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])`는 `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"])`는 `[["example.com", "example.com/contact", "example.com/blog"], 1]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])`는 `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Back", "Back"])`는 `[["example.com", "example.com/about"], 0]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0])`)
}})
```

`get_browser_history(["example.com", "example.com/about", "Forward"])`는 `[["example.com", "example.com/about"], 1]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_browser_history(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_browser_history(commands):

    return commands
```

# --solutions--

```py
def get_browser_history(commands):
    stack = []
    index = -1

    for command in commands:
        if command == "Back":
            if index > 0:
                index -= 1
        elif command == "Forward":
            if index < len(stack) - 1:
                index += 1
        else:
            del stack[index + 1:]
            stack.append(command)
            index += 1

    return [stack, index]
```
