---
id: 69b58ce40693f140c84c8559
title: "Challenge 240: Palindrome Characters"
challengeType: 29
dashedName: challenge-240
---

# --description--

문자열이 주어졌을 때, 그것이 회문인지 판단하고 가운데 문자를 반환하세요(길이가 홀수면 가운데 한 문자, 짝수면 가운데 두 문자).

- 회문은 앞에서 읽으나 뒤에서 읽으나 같은 문자열입니다.
- 회문이 아니면 `"none"`를 반환하세요.

# --hints--

`palindrome_locator("racecar")` should return `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("racecar"), "e")`)
}})
```

`palindrome_locator("level")` should return `"v"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("level"), "v")`)
}})
```

`palindrome_locator("freecodecamp")` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("freecodecamp"), "none")`)
}})
```

`palindrome_locator("noon")` should return `"oo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("noon"), "oo")`)
}})
```

`palindrome_locator("11100111")` should return `"00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(palindrome_locator("11100111"), "00")`)
}})
```

# --seed--

## --seed-contents--

```py
def palindrome_locator(s):

    return s
```

# --solutions--

```py
def palindrome_locator(s):
    if s != s[::-1]:
        return "none"
    mid = len(s) // 2
    return s[mid] if len(s) % 2 == 1 else s[mid - 1] + s[mid]
```
