---
id: 69b559d2903b9e4afe9075f7
title: "Challenge 237: Equation Validation"
challengeType: 29
dashedName: challenge-237
---

# --description--

수학 방정식을 나타내는 문자열이 주어졌을 때, 그것이 올바른지 판단하세요.

- 왼쪽에는 최대 세 개의 양의 정수와 연산자 `+`, `-`, `*`, `/`가 포함될 수 있습니다.
- 방정식은 `"number operator number = number"` 형식으로 주어집니다 (왼쪽에 두 개 또는 세 개의 숫자가 있음). 예를 들어 `"2 + 2 = 4"` 또는 `"2 + 3 - 1 = 4"`가 있습니다.
- 오른쪽은 항상 단일 정수입니다.

표준 연산 순서를 따르세요: 곱셈과 나눗셈은 덧셈과 뺄셈보다 먼저 왼쪽에서 오른쪽으로 평가됩니다.

# --hints--

`is_valid_equation("2 + 2 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 2 = 4"), True)`)
}})
```

`is_valid_equation("2 + 3 - 1 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 3 - 1 = 4"), True)`)
}})
```

`is_valid_equation("8 / 2 = 4")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("8 / 2 = 4"), True)`)
}})
```

`is_valid_equation("10 * 5 = 50")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 * 5 = 50"), True)`)
}})
```

`is_valid_equation("2 - 2 = 0")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 - 2 = 0"), True)`)
}})
```

`is_valid_equation("2 + 9 / 3 = 5")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 9 / 3 = 5"), True)`)
}})
```

`is_valid_equation("20 - 2 * 3 = 14")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("20 - 2 * 3 = 14"), True)`)
}})
```

`is_valid_equation("2 + 5 = 6")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("2 + 5 = 6"), False)`)
}})
```

`is_valid_equation("10 - 2 * 3 = 24")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("10 - 2 * 3 = 24"), False)`)
}})
```

`is_valid_equation("3 + 9 / 3 = 4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_equation("3 + 9 / 3 = 4"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_equation(equation):

    return equation
```

# --solutions--

```py
def is_valid_equation(equation):
    left, right = equation.split(" = ")
    tokens = left.split(" ")

    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        if op in ("*", "/"):
            a = int(tokens[i - 1])
            b = int(tokens[i + 1])
            computed = a * b if op == "*" else a // b
            tokens[i - 1:i + 2] = [str(computed)]
        else:
            i += 2

    result = int(tokens[0])
    i = 1
    while i < len(tokens) - 1:
        op = tokens[i]
        val = int(tokens[i + 1])
        result = result + val if op == "+" else result - val
        i += 2

    return result == int(right)
```
