---
id: 69b58ce40693f140c84c855a
title: "Challenge 241: FizzBuzz Validator"
challengeType: 29
dashedName: challenge-241
---

# --description--

연속된 정수 배열이 주어지고, 3과 5의 배수가 교체된 경우, 이것이 유효한 FizzBuzz 수열인지 판단하세요.

유효한 FizzBuzz 수열에서는:

- 3의 배수는 "Fizz"로 교체됩니다.
- 5의 배수는 "Buzz"로 교체됩니다.
- 3과 5의 공배수는 "FizzBuzz"로 교체됩니다.
- 나머지 숫자는 정수로 그대로 남습니다.

# --hints--

`is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([13, 14, "FizzBuzz", 16, 17])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([13, 14, "FizzBuzz", 16, 17]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", 4, 5])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", 4, 5]), False)`)
}})
```

`is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]), True)`)
}})
```

`is_fizz_buzz([1, 2, "Fizz", "Buzz", 5])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([1, 2, "Fizz", "Buzz", 5]), False)`)
}})
```

`is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]), False)`)
}})
```

`is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_fizz_buzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_fizz_buzz(arr):

    return arr
```

# --solutions--

```py
def is_fizz_buzz(arr):
    start = None
    for i, val in enumerate(arr):
        if isinstance(val, int):
            start = val - i
            break

    if start is None:
        return False

    for i in range(len(arr)):
        n = start + i
        if n % 15 == 0:
            expected = "FizzBuzz"
        elif n % 3 == 0:
            expected = "Fizz"
        elif n % 5 == 0:
            expected = "Buzz"
        else:
            expected = n
        if arr[i] != expected:
            return False

    return True
```
