---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

양의 정수가 주어졌을 때, 그 수 또는 그 수의 어떤 회전이 숫자 자리수로 나누어떨어지는지 확인하세요.

회전은 첫 번째 숫자를 끝으로 옮기는 것을 의미합니다. 예를 들어, 1회전 후 123은 231이 됩니다.

- 먼저 `0` (주어진 숫자)의 회전을 확인하세요.
- 주어진 숫자에는 0이 포함되지 않습니다.
- 나누어떨어지는 회전 숫자가 있으면 첫 번째 회전 숫자를 반환하고, 없으면 `"none"`을 반환하세요.

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
