---
id: 69bc6cb30c1d112a2e110a07
title: "도전 과제 250: 숨겨진 키"
challengeType: 29
dashedName: challenge-250
---

# --description--

250번째 일일 도전 과제에 오신 것을 환영합니다!

암호화된 문자열이 주어지면, 암호 키를 사용해 이를 해독하고 반환하세요.

키를 찾으려면:

- 오늘까지의 일일 도전 과제 중 도전 과제 번호가 25의 배수인 모든 도전 과제를 살펴보세요(이번 도전 과제 포함).
- 그 도전 과제 제목들의 첫 글자를 각각 가져와 문자열로 만드세요. 제목이 문자로 시작하지 않으면 첫 번째 문자를 찾으세요.

메시지를 해독하려면, 암호화된 메시지의 각 글자를 다음과 같이 처리하세요:

  1. 키의 해당 글자를 보세요(메시지가 키보다 길면 키를 반복해서 사용).
  2. 키 글자를 숫자로 변환하세요: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. 암호화된 글자를 알파벳에서 그 숫자만큼 뒤로 이동시키세요.
  4. 이동이 `"A"`보다 앞서면 `"Z"`로 돌아가세요.

예를 들어, 암호화된 메시지가 `"Y"`로 시작하고 첫 번째 키 글자가 `"V"`(22)라면, `"Y"`을 22칸 뒤로 이동시켜 `"C"`을 얻습니다. 이 과정을 각 글자에 반복해 전체 메시지를 해독하세요.

- 글자만 이동시키고, 공백은 그대로 반환하세요.
- 주어진 글자와 반환하는 글자는 모두 대문자입니다.

# --hints--

`decode("YAVJYNXE")`는 `"CONGRATS"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")`는 `"CODING LEGEND"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")`는 `"YOU ARE AWESOME"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")`는 `"KEEP GOING"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")`는 `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
