---
id: 69bc6cb30c1d112a2e110a03
title: "Challenge 246: Name Initials"
challengeType: 29
dashedName: challenge-246
---

# --description--

전체 이름이 문자열로 주어질 때, 이니셜을 반환하세요.

- 이니셜을 만들 이름들은 공백으로 구분됩니다.
- 이니셜은 대문자로 만들어야 합니다.
- 이니셜은 점으로 구분해야 합니다.

예를 들어, `"Tommy Millwood"`는 `"T.M."`를 반환합니다.

# --hints--

`get_initials("Tommy Millwood")` should return `"T.M."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Tommy Millwood"), "T.M.")`)
}})
```

`get_initials("Savanna Puddlesplash")` should return `"S.P."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Savanna Puddlesplash"), "S.P.")`)
}})
```

`get_initials("Frances Cowell Conrad")` should return `"F.C.C."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Frances Cowell Conrad"), "F.C.C.")`)
}})
```

`get_initials("Dragon")` should return `"D."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dragon"), "D.")`)
}})
```

`get_initials("Dorothy Vera Clump Haverstock Norris")` should return `"D.V.C.H.N."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_initials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_initials(name):

    return name
```

# --solutions--

```py
def get_initials(name):
    return ''.join(word[0].upper() + '.' for word in name.split())
```
