---
id: 69bc6cb30c1d112a2e110a06
title: "Challenge 249: String Math"
challengeType: 29
dashedName: challenge-249
---

# --description--

숫자와 다른 문자들이 섞인 문자열이 주어졌을 때, 숫자 사이에 있는 숫자가 아닌 문자들의 개수에 따라 숫자들에 수학 연산을 수행하세요.

- 두 숫자를 구분하는 문자의 개수가 짝수이면 덧셈을 사용하세요.
- 홀수이면 뺄셈을 사용하세요.
- 연속된 숫자들은 하나의 숫자로 간주하세요.
- 연산은 왼쪽에서 오른쪽 순서로 적용하세요.
- 숫자가 아닌 선행 및 후행 문자는 무시하세요.

예를 들어 `"3ab10c8"`가 주어지면 `5`을 반환하세요. 두 숫자 사이에 짝수 개의 문자가 있으므로 3과 10을 더해 13이 됩니다. 그다음 결과인 13과 8 사이에 홀수 개의 문자가 있으므로 8을 빼서 최종 결과를 구합니다.

# --hints--

`do_math("3ab10c8")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("3ab10c8"), 5)`)
}})
```

`do_math("6MINUS4")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("6MINUS4"), 2)`)
}})
```

`do_math("9plus3")` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("9plus3"), 12)`)
}})
```

`do_math("5fkwo#10i#%.<>15P=@20!#B/25")` should return `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("5fkwo#10i#%.<>15P=@20!#B/25"), 15)`)
}})
```

`do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` should return `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67)`)
}})
```

# --seed--

## --seed-contents--

```py
def do_math(s):

    return s
```

# --solutions--

```py
import re

def do_math(s):
    tokens = re.findall(r'\d+|\D+', s)
    result = None
    pending_op = None
    for token in tokens:
        if re.match(r'^\d+$', token):
            if result is None:
                result = int(token)
            elif pending_op == 'add':
                result += int(token)
            else:
                result -= int(token)
        else:
            pending_op = 'add' if len(token) % 2 == 0 else 'subtract'
    return result
```
