---
id: 69c5f3d787b1725d5f00c8ba
title: "도전 과제 256: 가장 가까운 시간 방향"
challengeType: 29
dashedName: challenge-256
---

# --description--

두 시간이 주어졌을 때, 첫 번째 시간에서 두 번째 시간으로 앞으로 가는 것이 더 빠른지, 뒤로 가는 것이 더 빠른지 판단하세요.

- 시간은 24시간 형식으로 주어집니다 (`"HH:MM"`)
- 시계는 순환합니다 (앞으로 이동할 때 23:59에서 00:00으로, 뒤로 이동할 때 00:00에서 23:59로 넘어갑니다)

반환값:

- 앞으로 가는 것이 더 짧으면 `"forward"` 반환
- 뒤로 가는 것이 더 짧으면 `"backward"` 반환
- 두 방향 모두 같은 시간이 걸리면 `"equal"` 반환

# --hints--

`get_direction("10:00", "12:00")`는 `"forward"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("10:00", "12:00"), "forward")`)
}})
```

`get_direction("11:00", "05:00")`는 `"backward"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("11:00", "05:00"), "backward")`)
}})
```

`get_direction("00:00", "12:00")`는 `"equal"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("00:00", "12:00"), "equal")`)
}})
```

`get_direction("15:45", "01:10")`는 `"forward"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("15:45", "01:10"), "forward")`)
}})
```

`get_direction("03:30", "19:50")`는 `"backward"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("03:30", "19:50"), "backward")`)
}})
```

`get_direction("06:30", "18:30")`는 `"equal"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_direction("06:30", "18:30"), "equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_direction(time1, time2):

    return time1
```

# --solutions--

```py
def get_direction(time1, time2):
    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    start = to_minutes(time1)
    end = to_minutes(time2)
    forward = (end - start + 1440) % 1440
    backward = (start - end + 1440) % 1440

    if forward < backward:
        return 'forward'
    if backward < forward:
        return 'backward'
    return 'equal'
```
