---
id: 69c6ff713a52713463aa7926
title: "Challenge 260: Word Score"
challengeType: 29
dashedName: challenge-260
---

# --description--

단어가 주어지면, 표준 문자 값 표를 사용해 점수를 반환하세요:

| Letter | Value |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- 대문자와 소문자는 같은 값을 가집니다.

# --hints--

`get_word_score("hi")` should return `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hi"), 17)`)
}})
```

`get_word_score("hello")` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hello"), 52)`)
}})
```

`get_word_score("hippopotamus")` should return `169`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("hippopotamus"), 169)`)
}})
```

`get_word_score("freeCodeCamp")` should return `94`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_word_score("freeCodeCamp"), 94)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_word_score(word):

    return word
```

# --solutions--

```py
def get_word_score(word):
    return sum(ord(c) - 96 for c in word.lower())
```
