---
id: 69cfca90e8a0a6d4d6871c54
title: "Challenge 270: Longest Common Substring"
challengeType: 29
dashedName: challenge-270
---

# --description--

문자열이 주어지면 두 번 이상 나타나는 가장 긴 부분 문자열을 반환하세요.

- 부분 문자열은 겹칠 수 있습니다.

# --hints--

`get_longest_substring("abracadabra")` should return `"abra"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("abracadabra"), "abra")`)
}})
```

`get_longest_substring("hello world hello")` should return `"hello"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("hello world hello"), "hello")`)
}})
```

`get_longest_substring("mississippi")` should return `"issi"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("mississippi"), "issi")`)
}})
```

`get_longest_substring("ha ha ha ha ha ha ha")` should return `"ha ha ha ha ha ha"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha")`)
}})
```

`get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` should return `"the quick brown fox jumped over"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_longest_substring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_longest_substring(s):

    return s
```

# --solutions--

```py
def get_longest_substring(s):
    longest = ''

    for length in range(len(s) - 1, 0, -1):
        for i in range(len(s) - length + 1):
            sub = s[i:i + length]
            if s.find(sub) != s.rfind(sub):
                return sub

    return longest
```
