---
id: 69cfca90e8a0a6d4d6871c50
title: "Challenge 266: Good Day"
challengeType: 29
dashedName: challenge-266
---

# --description--

`"HH:MM"` 형식(24시간제)으로 된 시간 문자열이 주어지면 다음을 반환하세요:

- `"Good morning"`: `05:00`부터 `11:59`까지의 시간에 대해
- `"Good afternoon"`: `12:00`부터 `17:59`까지의 시간에 대해
- `"Good evening"`: `18:00`부터 `21:59`까지의 시간에 대해
- `"Good night"`: `22:00`부터 `04:59`까지의 시간에 대해

# --hints--

`get_greeting("06:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("06:30"), "Good morning")`)
}})
```

`get_greeting("12:00")` should return `"Good afternoon"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("12:00"), "Good afternoon")`)
}})
```

`get_greeting("21:59")` should return `"Good evening"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("21:59"), "Good evening")`)
}})
```

`get_greeting("00:01")` should return `"Good night"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("00:01"), "Good night")`)
}})
```

`get_greeting("11:30")` should return `"Good morning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_greeting("11:30"), "Good morning")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_greeting(s):

    return s
```

# --solutions--

```py
def get_greeting(s):
    hours, minutes = map(int, s.split(':'))
    total = hours * 60 + minutes

    if 300 <= total < 720:
        return 'Good morning'
    if 720 <= total < 1080:
        return 'Good afternoon'
    if 1080 <= total < 1320:
        return 'Good evening'
    return 'Good night'
```
