---
id: 69cfca90e8a0a6d4d6871c52
title: "도전 과제 268: 나르시시스트 수"
challengeType: 29
dashedName: challenge-268
---

# --description--

양의 정수가 주어졌을 때, 그것이 나르시시스트 수인지 판단하세요.

- 숫자가 나르시시스트 수라는 것은 각 자리 숫자를 전체 자리수만큼 거듭제곱한 값을 모두 더했을 때 원래 숫자와 같다는 뜻입니다.

예를 들어, 153은 3자리 숫자이고, 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153이므로 나르시시스트 수입니다.

# --hints--

`is_narcissistic(153)`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

`is_narcissistic(154)`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

`is_narcissistic(371)`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

`is_narcissistic(512)`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

`is_narcissistic(9)`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

`is_narcissistic(11)`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

`is_narcissistic(9474)`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

`is_narcissistic(6549)`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```
