---
id: 69cfca90e8a0a6d4d6871c4f
title: "도전 과제 265: 가장 깊은 괄호"
challengeType: 29
dashedName: challenge-265
---

# --description--

균형 잡힌 괄호가 포함된 문자열이 주어질 때, 가장 깊게 중첩된 괄호 안의 내용을 반환하세요.

- 괄호는 `()`, `[]`, `{}` 세 가지 유형 중 하나일 수 있습니다.
- 입력에는 항상 하나의 가장 깊은 그룹이 존재합니다.

예를 들어, `"(hello (world))"`가 주어지면 `"world"`를 반환하세요.

# --hints--

`get_deepest_brackets("(hello (world))")`는 `"world"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("(hello (world))"), "world")`)
}})
```

`get_deepest_brackets("[outer [inner] outer]")`는 `"inner"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("[outer [inner] outer]"), "inner")`)
}})
```

`get_deepest_brackets("{a{b}c{d{e}f}g}")`는 `"e"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("{a{b}c{d{e}f}g}"), "e")`)
}})
```

`get_deepest_brackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")`는 `"fox"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox")`)
}})
```

`get_deepest_brackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")`는 `"C"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_deepest_brackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_deepest_brackets(s):

    return s
```

# --solutions--

```py
def get_deepest_brackets(s):
    max_depth = 0
    depth = 0
    content = ''
    start = 0

    for i, ch in enumerate(s):
        if ch in '([{':
            depth += 1
            if depth > max_depth:
                max_depth = depth
                start = i + 1
        elif ch in ')]}':
            if depth == max_depth:
                content = s[start:i]
            depth -= 1

    return content
```
