---
id: 69cfca90e8a0a6d4d6871c4e
title: "Challenge 264: Anagram Groups"
challengeType: 29
dashedName: challenge-264
---

# --description--

단어 배열이 주어지면, 애너그램으로 묶인 단어들의 2차원 배열을 반환하세요.

- 단어들은 같은 글자를 어떤 순서로든 포함하면 애너그램입니다.
- 각 단어는 정확히 하나의 그룹에 속합니다.
- 반환 순서는 중요하지 않습니다.

예를 들어 `["listen", "silent", "hello", "enlist", "world"]`가 주어지면 `[["listen", "silent", "enlist"], ["hello"], ["world"]]`을 반환하세요.

# --hints--

`group_anagrams(["listen", "silent", "hello", "enlist", "world"])` should return `[["listen", "silent", "enlist"], ["hello"], ["world"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["listen", "silent", "hello", "enlist", "world"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["enlist", "listen", "silent"], ["hello"], ["world"]])`)
}})
```

`group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])` should return `[["ate", "eat", "tea"], ["bat"], ["nat", "tan"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["eat", "tea", "tan", "ate", "nat", "bat"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["ate", "eat", "tea"], ["bat"], ["nat", "tan"]])`)
}})
```

`group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])` should return `[["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["care", "race", "acre", "pots", "stop", "tops", "opts", "post", "spot", "evil", "vile", "live", "veil"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["acre", "care", "race"], ["evil", "live", "veil", "vile"], ["opts", "post", "pots", "spot", "stop", "tops"]])`)
}})
```

`group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])` should return `[["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
groups = group_anagrams(["algorithms", "logarithms", "education", "cautioned", "auctioned", "triangle", "integral", "alerting", "relating"])
sorted_groups = sorted([sorted(g) for g in groups], key=lambda g: g[0])
TestCase().assertEqual(sorted_groups, [["alerting", "integral", "relating", "triangle"], ["algorithms", "logarithms"], ["auctioned", "cautioned", "education"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def group_anagrams(words):

    return words
```

# --solutions--

```py
def group_anagrams(words):
    groups = {}
    for word in words:
        key = ''.join(sorted(word))
        if key not in groups:
            groups[key] = []
        groups[key].append(word)
    return list(groups.values())
```
