---
id: 69d03e549613fbdbe21d11e0
title: "도전 과제 271: 약 복용 알림"
challengeType: 29
dashedName: challenge-271
---

# --description--

약 목록과 현재 시간을 나타내는 문자열이 주어질 때, 다음에 복용해야 할 약과 복용까지 남은 시간을 반환하세요.

- 각 약은 `[name, lastTaken]` 형식이며, 여기서 `name`은 약 이름이고 `lastTaken`는 마지막 복용 시간입니다.
- 모든 시간은 `"HH:MM"` (24시간) 형식입니다.

다음 약 복용 일정을 사용하세요:

| 이름 | 일정 |
| - | - |
| Deployxitrin | 08:00, 16:00 |
| Debuggamanizole | 07:00, 13:00, 21:00 |
| Mergeflictamine | 4시간마다 |

`"{name} in Hh Mm"` 형식의 문자열을 반환하세요. 예를 들어, `"Debuggamanizole in 2h 0m"` 또는 `"Deployxitrin in 1h 5m"`입니다.

# --hints--

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")`는 `"Debuggamanizole in 2h 0m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")`는 `"Deployxitrin in 1h 5m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")`는 `"Mergeflictamine in 0h 45m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")`는 `"Debuggamanizole in 0h 1m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")`는 `"Debuggamanizole in 0h 5m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m")`)
}})
```

`medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")`는 `"Mergeflictamine in 3h 30m"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(medication_reminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m")`)
}})
```

# --seed--

## --seed-contents--

```py
def medication_reminder(medications, current_time):

    return medications
```

# --solutions--

```py
def medication_reminder(medications, current_time):
    schedule = {
        'Deployxitrin': {'type': 'fixed', 'times': ['08:00', '16:00']},
        'Debuggamanizole': {'type': 'fixed', 'times': ['07:00', '13:00', '21:00']},
        'Mergeflictamine': {'type': 'interval', 'hours': 4},
    }

    def to_minutes(t):
        h, m = map(int, t.split(':'))
        return h * 60 + m

    current = to_minutes(current_time)
    earliest = float('inf')
    next_med = ''

    for name, last_taken in medications:
        med = schedule[name]
        if med['type'] == 'fixed':
            future_times = [to_minutes(t) for t in med['times'] if to_minutes(t) > current]
            if not future_times:
                continue
            next_due = min(future_times)
        else:
            next_due = to_minutes(last_taken) + med['hours'] * 60

        diff = next_due - current
        if 0 < diff < earliest:
            earliest = diff
            next_med = name

    h = earliest // 60
    m = earliest % 60
    return f"{next_med} in {h}h {m}m"
```
