---
id: 69e2383af7832c8032603b92
title: "도전 과제 276: 잘못된 요소"
challengeType: 29
dashedName: challenge-276
---

# --description--

오름차순으로 정렬된 정수 배열이 있는데, 단 하나의 잘못된 위치에 있는 요소가 있습니다. 그 요소의 인덱스를 반환하세요.

- 만약 여러 요소가 잘못된 위치로 간주될 수 있다면, 첫 번째 요소의 인덱스를 반환하세요.

# --hints--

`find_offender([1, 6, 2, 3, 4, 5])`는 `1`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 6, 2, 3, 4, 5]), 1)`)
}})
```

`find_offender([1, 2, 3, 5, 4, 5])`는 `3`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([1, 2, 3, 5, 4, 5]), 3)`)
}})
```

`find_offender([2, 1])`는 `0`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 1]), 0)`)
}})
```

`find_offender([2, 4, 1, 6, 8])`는 `2`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([2, 4, 1, 6, 8]), 2)`)
}})
```

`find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])`는 `6`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_offender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_offender(arr):

    return arr
```

# --solutions--

```py
def find_offender(arr):
    for i in range(len(arr)):
        without = arr[:i] + arr[i+1:]
        if all(without[j-1] <= without[j] for j in range(1, len(without))):
            return i
```
