---
id: 69e2383af7832c8032603b8f
title: "Challenge 273: ISBN-13 Validator"
challengeType: 29
dashedName: challenge-273
---

# --description--

문자열이 주어졌을 때, 그것이 유효한 ISBN-13 번호인지 확인하세요.

A valid ISBN-13:

- 숫자와 하이픈만 포함합니다
- 하이픈을 제거한 후 정확히 13자리 숫자가 있습니다
- 다음 검사를 통과합니다:
  1. 각 숫자에 1과 3을 번갈아 곱합니다 (첫 번째 숫자는 1, 두 번째는 3, 세 번째는 1, 이런 식으로).
  2. 결과들의 합이 10으로 나누어떨어져야 합니다.

# --hints--

`is_valid_isbn_13("9780306406157")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
