---
id: 69e2383af7832c8032603b94
title: "도전 과제 278: 커피 주문 파서"
challengeType: 29
dashedName: challenge-278
---

# --description--

커피 주문 문자열이 주어지면, 메뉴 항목을 식별하고 형식화된 주문을 반환하세요.

다음 메뉴 항목과 가격을 사용하세요:

| 항목 | 가격 |
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

일치하는 항목들을 `" + "`로 연결한 문자열을 반환하고, 그 뒤에 콜론과 공백(`": "`)을 붙이며, 총 가격을 표시하세요.

예를 들어, `"I'd like an oat latte with vanilla syrup and an extra shot please."`가 주어지면 `"oat latte + vanilla syrup + extra shot: $6.25"`를 반환하세요.

항목은 메뉴에 나타난 순서대로 표시되어야 하며, 총 가격은 항상 소수점 둘째 자리까지 표시해야 합니다.

# --hints--

`format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please.")`는 `"oat latte + vanilla syrup + extra shot: $6.25"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25")`)
}})
```

`format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")`는 `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85")`)
}})
```

`format_coffee_order("Can I get a cold brew with some cream and an extra shot.")`는 `"cold brew + extra shot + cream: $5.75"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75")`)
}})
```

`format_coffee_order("Just an espresso please.")`는 `"espresso: $3.00"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("Just an espresso please."), "espresso: $3.00")`)
}})
```

`format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")`는 `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format_coffee_order("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60")`)
}})
```

# --seed--

## --seed-contents--

```py
def format_coffee_order(order):

    return order
```

# --solutions--

```py
def format_coffee_order(order):
    menu = [
        ("cold brew", 450),
        ("oat latte", 500),
        ("cappuccino", 475),
        ("espresso", 300),
        ("vanilla syrup", 75),
        ("caramel drizzle", 60),
        ("extra shot", 50),
        ("oat milk", 75),
        ("cream", 75),
    ]
    found = [(name, price) for name, price in menu if name in order.lower()]
    total = sum(price for _, price in found)
    items = " + ".join(name for name, _ in found)
    return f"{items}: ${total / 100:.2f}"
```
