---
id: 69f35a5bb823ed620fcb7cb8
title: "도전 과제 279: 가장 긴 도미노 연결 고리"
challengeType: 29
dashedName: challenge-279
---

# --description--

도미노 세트를 나타내는 2D 배열이 주어질 때, 가장 긴 유효한 연결 고리를 반환하세요.

- 각 도미노는 0부터 6까지의 숫자 쌍으로 이루어집니다. 예: `[3, 2]`.
- 연결 고리는 각 도미노의 두 번째 숫자가 다음 도미노의 첫 번째 숫자와 일치할 때 유효합니다.
- 첫 번째 도미노의 첫 숫자와 마지막 도미노의 두 번째 숫자는 아무것과도 일치할 필요가 없습니다.
- 어떤 도미노든 뒤집을 수 있으므로 `[3, 2]`은 `[2, 3]`로 사용할 수 있습니다.
- 항상 정확히 하나의 가장 긴 유효한 연결 고리가 존재합니다.

예를 들어, `[[1, 2], [4, 5], [2, 3]]`이 주어지면 `[[1, 2], [2, 3]]`를 반환하세요.

# --hints--

`get_longest_chain([[1, 2], [4, 5], [2, 3]])`는 `[[1, 2], [2, 3]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[1, 2], [4, 5], [2, 3]])
TestCase().assertIn(result, [[[1, 2], [2, 3]], [[3, 2], [2, 1]]])`)
}})
```

`get_longest_chain([[2, 1], [4, 3], [5, 3]])`는 `[[4, 3], [3, 5]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[2, 1], [4, 3], [5, 3]])
TestCase().assertIn(result, [[[4, 3], [3, 5]], [[5, 3], [3, 4]]])`)
}})
```

`get_longest_chain([[1, 2], [3, 4], [2, 3], [4, 0]])`는 `[[1, 2], [2, 3], [3, 4], [4, 0]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[1, 2], [3, 4], [2, 3], [4, 0]])
TestCase().assertIn(result, [[[1, 2], [2, 3], [3, 4], [4, 0]], [[0, 4], [4, 3], [3, 2], [2, 1]]])`)
}})
```

`get_longest_chain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])`는 `[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])
TestCase().assertIn(result, [[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]], [[5, 6], [6, 6], [6, 1], [1, 1], [1, 4]]])`)
}})
```

`get_longest_chain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])`는 `[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
result = get_longest_chain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])
TestCase().assertIn(result, [[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]], [[6, 5], [5, 5], [5, 4], [4, 4], [4, 0], [0, 3], [3, 3]]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_longest_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def get_longest_chain(dominoes):
    def search(chain, remaining):
        best = chain
        last = chain[-1][1]
        for i, (a, b) in enumerate(remaining):
            rest = remaining[:i] + remaining[i+1:]
            if a == last:
                result = search(chain + [[a, b]], rest)
                if len(result) > len(best):
                    best = result
            if b == last and a != b:
                result = search(chain + [[b, a]], rest)
                if len(result) > len(best):
                    best = result
        return best

    best = []
    for i, (a, b) in enumerate(dominoes):
        rest = dominoes[:i] + dominoes[i+1:]
        r1 = search([[a, b]], rest)
        if len(r1) > len(best):
            best = r1
        if a != b:
            r2 = search([[b, a]], rest)
            if len(r2) > len(best):
                best = r2
    return best
```
