---
id: 69f35a5bb823ed620fcb7cbd
title: "Challenge 284: I Before E"
challengeType: 29
dashedName: challenge-284
---

# --description--

단어나 문장이 주어지면, 모든 단어가 "C 다음이 아니면 I가 E보다 먼저 온다" 규칙을 따르도록 수정된 버전을 반환하세요.

- 단어에 `"ei"`가 `"c"` 앞에 오지 않고 포함되어 있으면 `"ie"`로 바꾸세요.
- 단어에 `"ie"`가 `"c"` 앞에 오면 `"ei"`로 바꾸세요.
- 나머지 단어는 변경하지 않고 그대로 두세요.

# --hints--

`i_before_e("beleive")` should return `"believe"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("beleive"), "believe")`)
}})
```

`i_before_e("recieve")` should return `"receive"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("recieve"), "receive")`)
}})
```

`i_before_e("we recieved a breif")` should return `"we received a brief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved a breif"), "we received a brief")`)
}})
```

`i_before_e("she beleived the friendly niece could percieve the greif")` should return `"she believed the friendly niece could perceive the grief"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief")`)
}})
```

`i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit")` should return `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(i_before_e("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit")`)
}})
```

# --seed--

## --seed-contents--

```py
def i_before_e(sentence):

    return sentence
```

# --solutions--

```py
import re

def i_before_e(sentence):
    words = []
    for word in sentence.split(" "):
        word = re.sub(r"([^c])ei", r"\1ie", word)
        word = re.sub(r"cie", "cei", word)
        words.append(word)
    return " ".join(words)
```
