---
id: 69f35a5bb823ed620fcb7cba
title: "Challenge 281: Bingo Range"
challengeType: 29
dashedName: challenge-281
---

# --description--

빙고 글자가 주어지면, 그 글자에 해당하는 숫자 범위를 반환하세요.

| 글자 | 숫자 범위 |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

가장 작은 수부터 가장 큰 수까지 범위 내 모든 숫자를 포함하는 배열을 반환하세요.

# --hints--

`get_bingo_range("B")` should return `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_range("B"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15])`)
}})
```

`get_bingo_range("I")` should return `[16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_range("I"), [16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30])`)
}})
```

`get_bingo_range("N")` should return `[31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_range("N"), [31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45])`)
}})
```

`get_bingo_range("G")` should return `[46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_range("G"), [46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60])`)
}})
```

`get_bingo_range("O")` should return `[61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_bingo_range("O"), [61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_bingo_range(letter):

    return letter
```

# --solutions--

```py
def get_bingo_range(letter):
    ranges = {"B": (1, 15), "I": (16, 30), "N": (31, 45), "G": (46, 60), "O": (61, 75)}
    start, end = ranges[letter]
    return list(range(start, end + 1))
```
