---
id: 69f35a5bb823ed620fcb7cbc
title: "Challenge 283: String Zipper"
challengeType: 29
dashedName: challenge-283
---

# --description--

두 개의 문자열이 주어지면, 각 문자를 한 번에 하나씩 번갈아가며 합친 새 문자열을 반환하세요. 한 문자열이 더 길면 남은 문자를 끝에 덧붙이세요.

첫 번째 문자열의 첫 문자부터 시작하세요.

# --hints--

`zip_strings("abc", "123")` should return `"a1b2c3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("abc", "123"), "a1b2c3")`)
}})
```

`zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz")` should return `"abcdefghijklmnopqrstuvwxyz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz")`)
}})
```

`zip_strings("day", "night")` should return `"dnaiyght"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("day", "night"), "dnaiyght")`)
}})
```

`zip_strings("python", "javascript")` should return `"pjyatvhaosncript"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("python", "javascript"), "pjyatvhaosncript")`)
}})
```

`zip_strings("feCdCm", "reoeap")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(zip_strings("feCdCm", "reoeap"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def zip_strings(a, b):

    return a
```

# --solutions--

```py
def zip_strings(a, b):
    result = ""
    for i in range(max(len(a), len(b))):
        if i < len(a):
            result += a[i]
        if i < len(b):
            result += b[i]
    return result
```
