---
id: 69f35a5bb823ed620fcb7cbb
title: "Challenge 282: Sleep Debt"
challengeType: 29
dashedName: challenge-282
---

# --description--

오늘까지 매일 밤 잔 수면 시간을 배열로 주고, 목표 수면 시간을 알려줄 때, 수면 부채를 없애기 위해 오늘 밤 몇 시간 자야 하는지 반환하세요.

- 오늘 밤 자야 할 시간도 총 필요한 시간에 포함하세요.
- 오늘 밤 목표 수면 시간 이상 잤다면 `0`를 반환하세요.

# --hints--

`sleep_debt([6, 6, 6, 6, 6, 6], 8)` should return `20`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 6, 6, 6, 6, 6], 8), 20)`)
}})
```

`sleep_debt([6, 7, 8, 4, 8, 6], 7)` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 7, 8, 4, 8, 6], 7), 10)`)
}})
```

`sleep_debt([10, 10, 9, 10, 9, 11], 9)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([10, 10, 9, 10, 9, 11], 9), 4)`)
}})
```

`sleep_debt([8, 7, 6, 7, 6, 8], 6)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 7, 6, 7, 6, 8], 6), 0)`)
}})
```

`sleep_debt([8, 9, 10, 9, 10, 7], 7)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 9, 10, 9, 10, 7], 7), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sleep_debt(hours_slept, target_hours):

    return hours_slept
```

# --solutions--

```py
def sleep_debt(hours_slept, target_hours):
    debt = (len(hours_slept) + 1) * target_hours - sum(hours_slept)
    return max(0, debt)
```
