---
id: 69f8c998d78ad3171a0713bd
title: "도전 과제 290: 피자 파티"
challengeType: 29
dashedName: challenge-290
---

# --description--

오늘 각 사람이 일한 시간을 담은 배열이 주어졌을 때, 피자 파티를 위해 주문할 피자 개수를 반환하세요.

- 각 사람이 일한 시간을 3으로 나누어 조각 수를 구하세요.
- 조각을 부분적으로 먹을 수 없으니, 각 사람의 조각 수를 가장 가까운 정수로 올림하세요.
- 각 사람은 최소 두 조각을 받습니다.
- 피자 한 판은 8조각입니다. 피자 총 개수는 가장 가까운 정수로 올림하세요.

# --hints--

`get_pizzas_to_order([8, 8, 8])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
