---
id: 69f8c998d78ad3171a0713bc
title: "도전 과제 289: 차이의 합"
challengeType: 29
dashedName: challenge-289
---

# --description--

숫자 배열이 주어지면, 각 숫자와 그 다음 숫자 사이의 차이들의 합을 반환하세요.

예를 들어, `[1, 3, 4]`가 주어지면 `3` (2 + 1)을 반환하세요.

# --hints--

`sum_of_differences([1, 3, 4])`는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([1, 3, 4]), 3)`)
}})
```

`sum_of_differences([5, -3, 3, 9, 10])`는 `5`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([5, -3, 3, 9, 10]), 5)`)
}})
```

`sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7])`는 `-16`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([9, 6, 15, -20, 33, 14, 25, 16, -7]), -16)`)
}})
```

`sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75])`는 `25`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_differences([50, 102, -46, 82, -49, 29, 71, 902, -237, 111, -61, 75]), 25)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_differences(arr):

    return arr
```

# --solutions--

```py
def sum_of_differences(arr):
    return sum(arr[i + 1] - arr[i] for i in range(len(arr) - 1))
```
