---
id: 69f8c998d78ad3171a0713bf
title: "도전 과제 292: 더 넓은 가로세로 비율"
challengeType: 29
dashedName: challenge-292
---

# --description--

서로 다른 이미지 크기를 나타내는 두 문자열이 주어질 때, 가로 세로 비율이 더 큰 이미지의 가로세로 비율을 반환하세요.

- 주어진 문자열은 `"WxH"` 형식이며, 예를 들어 `"1920x1080"`와 같습니다.
- 가로세로 비율은 가로와 세로의 비율을 가장 작은 정수로 줄인 값입니다. 예를 들어 `"1920x1080"`는 `"16:9"`로 줄어듭니다.
- `"W:H"` 형식의 문자열을 반환하세요. 예를 들어 `"16:9"`와 같습니다.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")`는 `"16:9"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")`는 `"4:3"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")`는 `"2:1"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")`는 `"9:16"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")`는 `"43:18"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")`는 `"2089:8539"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
