---
id: 69f90c1329a94b37e2a2086d
title: "Challenge 294: Parentheses Combinations"
challengeType: 29
dashedName: challenge-294
---

# --description--

정수 `n`가 주어지면, `n` 쌍의 소괄호로 만들 수 있는 유효한 조합의 개수를 반환하세요.

- 유효한 조합이란 모든 여는 소괄호가 대응하는 닫는 소괄호를 가지며, 닫는 소괄호가 그에 대응하는 여는 소괄호보다 먼저 나오지 않는 문자열입니다.

예를 들어, `2`가 주어지면 유효한 조합은 2가지입니다:

```json
(())
()()
```

# --hints--

`get_combinations(2)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` should return `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` should return `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` should return `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
