---
id: 69f90c1329a94b37e2a2086c
title: "도전 과제 293: 최고의 족보"
challengeType: 29
dashedName: challenge-293
---

# --description--

플레이 카드 다섯 장을 나타내는 문자열 배열이 주어질 때, 최고의 족보 이름을 반환하세요.

- 각 카드는 두 글자 문자열로 나타내며, 등급 다음에 무늬가 옵니다. `"2h"` 예를 들어,
  - 등급은 낮은 것부터 높은 것까지 `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"`, `"A"`입니다.
  - 무늬는 `"h"`, `"d"`, `"c"`, `"s"`입니다.
- 에이스(`"A"`)는 스트레이트에서 높거나 낮은 값으로 사용할 수 있습니다.

족보는 나쁜 것부터 좋은 순서로 다음과 같습니다:

| 이름 | 설명 |
| - | - |
| `"High Card"` | 페어나 그 이상의 족보가 없음 |
| `"Pair"` | 한 등급이 두 장 |
| `"Two Pair"` | 한 등급이 두 장, 다른 등급이 두 장 |
| `"Three of a Kind"` | 한 등급이 세 장 |
| `"Straight"` | 연속된 다섯 등급 |
| `"Flush"` | 같은 무늬 다섯 장 |
| `"Full House"` | 한 등급이 세 장, 다른 등급이 두 장 |
| `"Four of a Kind"` | 한 등급이 네 장 |
| `"Straight Flush"` | 같은 무늬로 연속된 다섯 등급 |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` 같은 무늬 |

최고의 족보 이름을 반환하세요.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])`는 `"Three of a Kind"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])`는 `"Full House"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])`는 `"Flush"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])`는 `"Four of a Kind"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])`는 `"Two Pair"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])`는 `"Straight Flush"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])`는 `"High Card"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])`는 `"Straight"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])`는 `"Pair"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])`는 `"Royal Flush"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
