---
id: 6a0dcd03ee4e68698080ef6c
title: "도전 과제 308: 신용카드 유효성 검사기"
challengeType: 29
dashedName: challenge-308
---

# --description--

신용카드 번호로 된 숫자 문자열이 주어졌을 때, 다음 메서드를 사용해 유효한 카드 번호인지 확인하세요:

- 끝에서 두 번째 숫자부터 시작해 왼쪽으로 한 칸씩 건너뛰며 모든 숫자를 두 배로 만드세요.
- 두 배로 만든 숫자가 9보다 크면 9를 빼세요.
- 두 배로 만든 숫자와 그렇지 않은 숫자 모두 합하세요.
- 합계가 10으로 나누어떨어지면 유효한 번호입니다.

# --hints--

`is_valid_card("4532015112830366")`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")`는 `True`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")`는 `False`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
