---
id: 6a0dcd03ee4e68698080ef6b
title: "도전 과제 307: 구역 규정"
challengeType: 29
dashedName: challenge-307
---

# --description--

도시 건물 배치를 나타내는 2D 그리드(배열의 배열)가 주어졌을 때, 구역 규정을 위반하는 모든 건물의 좌표를 반환하세요.

그리드의 각 셀에는 아래 표에 있는 레이블 중 하나가 들어 있습니다. 건물은 최대 4개의 이웃 셀(수평 또는 수직)이 인접할 수 없는 유형일 경우 위반 상태입니다.

| 레이블 | 유형 | 인접할 수 없는 유형 |
| - | - | - |
| `"i"` | 산업 | `"R"`, `"I"` |
| `"A"` | 농업 | `"C"` |
| `"R"` | 주거 | `"i"`, `"C"` |
| `"I"` | 기관 | `"i"` |
| `"C"` | 상업 | `"R"`, `"A"` |
| `""` (빈 문자열) | 미개발 | 제한 없음 |

위반하는 모든 셀의 좌표를 `[row, col]` 쌍의 배열로, 순서는 상관없이 반환하세요. 위반이 없으면 빈 배열을 반환하세요.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])`는 `[[0, 0], [0, 1]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])`는 `[[0, 1], [1, 1]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])`는 `[]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])`는 `[[0, 1], [0, 2], [0, 3]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])`는 `[[2, 3], [2, 4], [3, 1], [3, 2]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
