---
id: 6a0dcd03ee4e68698080ef68
title: "도전 과제 304: 여행 일정 배치"
challengeType: 29
dashedName: challenge-304
---

# --description--

하루 여행을 위한 최소 두 개의 선택적 정지가 포함된 배열이 주어질 때, 유효한 여행 일정 배치의 수를 반환하세요.

여행 일정에는 항상 `"breakfast"`, `"lunch"`, `"dinner"`가 포함되며, 이들은 인수로 전달되지 않습니다. 선택적 정지는 다음 규칙에 따라 여행 일정 어디에나 배치할 수 있습니다:

- `"breakfast"`는 항상 첫 번째이며, `"lunch"` 앞에 적어도 한 개의 정지가 있어야 합니다.
- `"lunch"`는 `"dinner"`보다 먼저 나타나야 하며, 그 사이에 적어도 한 개의 정지가 있어야 합니다.
- `"dinner"` 이후에는 최대 한 개의 선택적 정지만 나타날 수 있습니다.

유효한 배치의 수를 반환하세요.

# --hints--

`get_itinerary_count(["library", "park"])`는 `2`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])`는 `18`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])`는 `120`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])`는 `840`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])`는 `55440`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
