---
id: 6a0dcc730cb92a616f86f0c0
title: "도전 과제 296: 스키마 검증기 2부"
challengeType: 29
dashedName: challenge-296
---

# --description--

객체(JavaScript) 또는 딕셔너리(Python)가 다음 스키마와 일치하는지 확인하세요:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- 추가 키는 허용됩니다

# --hints--

`is_valid_schema({"username": "alice", "posts": 10, "verified": False})`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "alice", "posts": 10, "verified": False}), True)`)
}})
```

`is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25})`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25}), True)`)
}})
```

`is_valid_schema({"username": "frank", "posts": "21", "verified": True})`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "frank", "posts": "21", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "sam", "posts": 17, "verified": "false"})`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sam", "posts": 17, "verified": "false"}), False)`)
}})
```

`is_valid_schema({"username": "bill", "verified": True})`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bill", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fred", "verified": True})`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fred", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": 5, "posts": 10, "verified": True})`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 5, "posts": 10, "verified": True}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool)
    )
```
