---
id: 6a15cadf5f240d05a2649558
title: "도전 과제 314: 소인수분해"
challengeType: 29
dashedName: challenge-314
---

# --description--

1보다 큰 정수가 주어지면, 오름차순으로 정렬된 숫자 배열로 소인수분해 결과를 반환하세요.

소인수분해는 주어진 정수를 곱해서 만드는 소수들의 집합입니다. 각 정수는 정확히 하나의 소인수분해 집합을 가집니다. 예를 들어, 20의 소인수분해는 `[2, 2, 5]`입니다. 왜냐하면 2 * 2 * 5 = 20이기 때문입니다.

주어진 정수가 소수라면, 단일 요소 배열로 반환하세요.

# --hints--

`prime_factorization(20)`는 `[2, 2, 5]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)`는 `[17]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)`는 `[3, 5]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)`는 `[5, 7]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)`는 `[3, 3, 3, 37]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)`는 `[2, 2, 2, 3, 3, 5]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)`는 `[2, 3, 5, 7, 11, 13, 17]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
