---
id: 6a151d32d772271dd2448c2d
title: "도전 과제 311: 주문 시전자"
challengeType: 29
dashedName: challenge-311
---

# --description--

시전하는 주문 코드 문자열이 주어질 때, 총 점수를 계산하세요.

문자열의 각 문자는 하나의 주문을 나타냅니다:

| 코드 | 주문 | 분류 | 기본 점수 |
| - | - | - | - |
| `"f"` | Fire | Destruction | 3 |
| `"l"` | Lightning | Destruction | 3 |
| `"i"` | Ice | Control | 2 |
| `"w"` | Wind | Control | 2 |
| `"h"` | Heal | Restoration | 1 |
| `"s"` | Shield | Restoration | 1 |

연속으로 서로 다른 분류의 주문을 시전한 횟수에 따라 콤보 배수가 적용됩니다:

- 첫 번째 주문은 항상 기본 점수를 받습니다.
- 이전 주문과 다른 분류의 주문을 연속으로 시전할 때마다 배수가 1씩 증가합니다.
- 이전 주문과 같은 분류의 주문을 시전하면 배수가 다시 1로 초기화됩니다.
- 각 주문의 점수는 기본 점수에 현재 배수를 곱한 값입니다.

주문 수열에서 총 점수를 반환하세요.

# --hints--

`cast("fihwl")`는 `33`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")`는 `45`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")`는 `37`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")`는 `50`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")`는 `101`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
