---
id: 6a15cadf5f240d05a2649557
title: "도전 과제 313: 대여 비용"
challengeType: 29
dashedName: challenge-313
---

# --description--

대여 시각, 반납 시각, 대여 등급이 주어지면 연체료를 포함한 총 대여 비용을 반환하세요.

- 주어진 시각은 UTC ISO 문자열입니다. 예: `"2026-06-18T18:30:00Z"`.
- 대여 등급은 반납 기한까지 남은 일수로 `1`, `3`, 또는 `7`입니다.
- 대여는 대여 기간 마지막 날 UTC 기준 오후 12시까지 또는 그 이전에 반납해야 합니다. 예를 들어, 3월 15일 언제든 대여한 1일 대여는 3월 16일 UTC 기준 오후 12시까지 반납해야 합니다.
- 반납 기한과 시간을 넘긴 매일마다 연체료가 부과됩니다.

가격은 다음과 같습니다:

| 등급 | 기본 비용 | 하루 연체료 |
| - | - | - |
| 1일 | $4.99 | $3.99 |
| 3일 | $3.99 | $2.99 |
| 7일 | $2.99 | $0.99 |

총 비용을 소수점 둘째 자리까지 반올림하여 `"$D.CC"` 형식으로 반환하세요.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)`는 `"$4.99"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)`는 `"$12.97"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)`는 `"$3.99"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)`는 `"$9.97"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)`는 `"$2.99"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)`는 `"$358.40"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
