---
id: 6a15cadf5f240d05a264955c
title: "도전 과제 318: DNA 변이"
challengeType: 29
dashedName: challenge-318
---

# --description--

길이가 같은 두 DNA 가닥이 주어지면, 가닥이 다른 위치(변이)의 인덱스를 배열로 반환하세요.

- DNA 가닥은 `"A"`, `"T"`, `"C"`, `"G"` 문자로 이루어진 문자열입니다
- 인덱스는 오름차순으로 반환하세요
- 변이가 없으면 빈 배열을 반환하세요

# --hints--

`detect_mutations("ATCG", "ATGG")`는 `[2]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATCG", "ATGG"), [2])`)
}})
```

`detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")`는 `[4, 9, 11]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11])`)
}})
```

`detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")`는 `[]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), [])`)
}})
```

`detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")`는 `[15, 16, 17, 18]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18])`)
}})
```

`detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")`는 `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28])`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_mutations(strand1, strand2):

    return strand1
```

# --solutions--

```py
def detect_mutations(strand1, strand2):
    return [i for i, (a, b) in enumerate(zip(strand1, strand2)) if a != b]
```
