---
id: 6a15cadf5f240d05a264955d
title: "도전 과제 319: Frontmatter 파서"
challengeType: 29
dashedName: challenge-319
---

# --description--

frontmatter 블록을 나타내는 문자열이 주어지면, 이를 파싱해서 키와 값을 가진 객체(JavaScript) 또는 딕셔너리(Python)를 반환하세요.

Frontmatter는 `---` 구분자로 감싸져 있고, 그 안에 한 줄에 하나씩 `key: value` 쌍이 들어 있습니다. 예를 들어:

```md
---
title: My Post
draft: false
views: 100
---
```

다음과 같이 반환해야 합니다:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- 숫자, 불리언, 문자열은 각각의 타입으로 반환해야 합니다.
- 주어진 문자열은 줄바꿈 문자(`"\n"`)로 줄이 구분됩니다. 위 예시는 `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`로 주어집니다.

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")`는 `{ title: "My Post", draft: False, views: 100 }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")`는 `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")`는 `{ version: "1.0.0", url: "https://example.com", private: True }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")`는 `{ rating: 4.5, price: 9.99 }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
