---
id: 6a15cadf5f240d05a2649556
title: "도전 과제 312: 스트리밍 비용"
challengeType: 29
dashedName: challenge-312
---

# --description--

스트리밍 서비스 장바구니에 있는 영화를 나타내는 배열과 구독 등급을 나타내는 문자열이 주어지면, 영화들의 총 비용을 반환하세요.

장바구니의 각 항목은 `"format"`(`"HD"` 또는 `"4K"`)과 `"type"`(`"rent"` 또는 `"buy"`)을 가진 객체입니다. 비용은 다음과 같습니다:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

다음 구독 등급 할인율을 적용하세요:

- `"none"`: 정가
- `"basic"`: 10% 할인
- `"premium"`: 25% 할인

총 비용을 소수점 둘째 자리까지 반올림하여 `"$D.CC"` 형식으로 반환하세요.

# --hints--

`get_streaming_bill([{ format: "HD", type: "rent" }], "none")`는 `"$3.99"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}], "none"), "$3.99")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "premium")`는 `"$12.73"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "premium"), "$12.73")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")`는 `"$18.87"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "HD", "type": "rent"}, {"format": "HD", "type": "buy"}], "basic"), "$18.87")`)
}})
```

`get_streaming_bill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")`는 `"$29.98"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "4K", "type": "buy"}, {"format": "4K", "type": "buy"}], "premium"), "$29.98")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")`는 `"$42.96"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}], "none"), "$42.96")`)
}})
```

`get_streaming_bill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }, { format: "HD", type: "buy" }], "basic")`는 `"$50.36"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_streaming_bill([{"format": "HD", "type": "rent"}, {"format": "4K", "type": "rent"}, {"format": "HD", "type": "buy"}, {"format": "4K", "type": "buy"}, {"format": "HD", "type": "buy"}], "basic"), "$50.36")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_streaming_bill(cart, subscription):

    return cart
```

# --solutions--

```py
def get_streaming_bill(cart, subscription):
    pricing = {
        "HD": {"rent": 3.99, "buy": 12.99},
        "4K": {"rent": 5.99, "buy": 19.99},
    }
    discounts = {"none": 1, "basic": 0.9, "premium": 0.75}
    multiplier = discounts[subscription]
    total = sum(pricing[item["format"]][item["type"]] * multiplier for item in cart)
    return f"${total:.2f}"
```
