---
id: 6a15cadf5f240d05a264955f
title: "도전 과제 321: 주기율표 철자법"
challengeType: 29
dashedName: challenge-321
---

# --description--

단어가 주기율표의 원소 기호를 사용해 철자할 수 있는지 확인하세요.

- 단어를 철자할 때 대소문자를 무시하세요. `"neon"`은 기호 `"Ne"`, `"O"`, `"N"`으로 철자할 수 있습니다.

다음은 원소 기호의 전체 목록입니다:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

단어를 철자하는 데 사용된 원소들을 원래 대소문자와 철자 순서대로 배열로 반환하세요. 철자할 수 없으면 빈 배열을 반환하세요.

# --hints--

`get_periodic_spelling("neon")`는 `["Ne", "O", "N"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("neon"), ["Ne", "O", "N"])`)
}})
```

`get_periodic_spelling("rational")`는 `["Ra", "Ti", "O", "N", "Al"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("rational"), ["Ra", "Ti", "O", "N", "Al"])`)
}})
```

`get_periodic_spelling("yarn")`는 `["Y", "Ar", "N"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("yarn"), ["Y", "Ar", "N"])`)
}})
```

`get_periodic_spelling("carbon")`는 `["C", "Ar", "B", "O", "N"]` 또는 `["Ca", "Rb", "O", "N"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("carbon"), [["C", "Ar", "B", "O", "N"], ["Ca", "Rb", "O", "N"]])`)
}})
```

`get_periodic_spelling("noisy")`는 `["N", "O", "I", "S", "Y"]` 또는 `["No", "I", "S", "Y"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("noisy"), [["N", "O", "I", "S", "Y"], ["No", "I", "S", "Y"]])`)
}})
```

`get_periodic_spelling("bicycles")`는 `["B", "I", "C", "Y", "Cl", "Es"]` 또는 `["Bi", "C", "Y", "Cl", "Es"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("bicycles"), [["B", "I", "C", "Y", "Cl", "Es"], ["Bi", "C", "Y", "Cl", "Es"]])`)
}})
```

`get_periodic_spelling("optics")`는 `["O", "P", "Ti", "C", "S"]`, `["O", "P", "Ti", "Cs"]`, `["O", "Pt", "I", "C", "S"]`, 또는 `["O", "Pt", "I", "Cs"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("optics"), [["O", "P", "Ti", "C", "S"], ["O", "P", "Ti", "Cs"], ["O", "Pt", "I", "C", "S"], ["O", "Pt", "I", "Cs"]])`)
}})
```

`get_periodic_spelling("value")`는 `[]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("value"), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_periodic_spelling(word):

    return word
```

# --solutions--

```py
def get_periodic_spelling(word):
    elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"]
    lower = {e.lower(): e for e in elements}

    def spell(w, i):
        if i == len(w):
            return []
        two = w[i:i + 2]
        one = w[i:i + 1]
        if two in lower:
            rest = spell(w, i + 2)
            if rest is not None:
                return [lower[two]] + rest
        if one in lower:
            rest = spell(w, i + 1)
            if rest is not None:
                return [lower[one]] + rest
        return None

    result = spell(word.lower(), 0)
    return result if result is not None else []
```
