---
id: 6a15cadf5f240d05a264955b
title: "도전 과제 317: BMI 계산기"
challengeType: 29
dashedName: challenge-317
---

# --description--

파운드 단위의 몸무게와 인치 단위의 키가 주어지면, 소수점 첫째 자리까지 반올림한 BMI(체질량지수)를 반환하세요.

BMI를 구하려면: 몸무게를 키의 제곱으로 나눈 후, 그 결과에 703을 곱하세요.

# --hints--

`calculate_bmi(180, 70)`는 `25.8`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)`는 `24.0`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)`는 `19.5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)`는 `39.1`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)`는 `22.8`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
