---
id: 6a15cadf5f240d05a264955e
title: "도전 과제 320: 혈액 은행"
challengeType: 29
dashedName: challenge-320
---

# --description--

혈액 은행의 재고 배열과 환자의 혈액형 요청 배열이 주어질 때, `"X of Y patients served"` 형식의 문자열을 반환하세요. 여기서 `X`는 은행 재고에서 혈액을 받을 수 있는 최대 환자 수이고, `Y`는 전체 환자 수입니다.

두 배열의 각 항목은 다음 혈액형 중 하나입니다: `"AB"`, `"A"`, `"B"`, 또는 `"O"`.

호환성 규칙:

- `"AB"`는 모든 혈액형으로부터 받을 수 있습니다.
- `"A"`는 `"A"`와 `"O"`로부터 받을 수 있습니다.
- `"B"`는 `"B"`와 `"O"`로부터 받을 수 있습니다.
- `"O"`는 오직 `"O"`로부터만 받을 수 있습니다.

주어진 배열에서 중복된 항목은 수량을 나타냅니다.

# --hints--

`triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])`는 `"4 of 4 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])`는 `"3 of 5 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served")`)
}})
```

`triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])`는 `"4 of 5 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served")`)
}})
```

`triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])`는 `"4 of 4 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served")`)
}})
```

`triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])`는 `"8 of 11 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served")`)
}})
```

`triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])`는 `"13 of 14 patients served"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_blood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_blood(bank, patients):

    return bank
```

# --solutions--

```py
def triage_blood(bank, patients):
    b = {}
    p = {}
    for t in bank:
        b[t] = b.get(t, 0) + 1
    for t in patients:
        p[t] = p.get(t, 0) + 1

    saved = 0

    def serve(patient, donors):
        nonlocal saved
        for donor in donors:
            n = min(p.get(patient, 0), b.get(donor, 0))
            p[patient] = p.get(patient, 0) - n
            b[donor] = b.get(donor, 0) - n
            saved += n

    serve("O", ["O"])
    serve("A", ["A", "O"])
    serve("B", ["B", "O"])
    serve("AB", ["AB", "A", "B", "O"])
    return f"{saved} of {len(patients)} patients served"
```
