---
id: 6a19b1062d1b153d8ac76d72
title: "도전 과제 324: 중복 문자 개수 세기"
challengeType: 29
dashedName: challenge-324
---

# --description--

두 개의 문자열이 주어지면, 두 번째 문자열에서 첫 번째 문자열에서 찾을 수 있는 문자의 개수를 반환하세요.

- 두 번째 문자열에 중복된 문자는 각각 별도로 계산하세요.

# --hints--

`duplicate_character_count("aloha", "hei")` 는 `1`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("aloha", "hei"), 1)`)
}})
```

`duplicate_character_count("jambo", "bonjour")` 는 `4`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("jambo", "bonjour"), 4)`)
}})
```

`duplicate_character_count("hello", "hola")` 는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello", "hola"), 3)`)
}})
```

`duplicate_character_count("ola", "hej")` 는 `0`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ola", "hej"), 0)`)
}})
```

`duplicate_character_count("ciao", "konnichiwa")` 는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ciao", "konnichiwa"), 5)`)
}})
```

`duplicate_character_count("merhaba", "xin chao")` 는 `2`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("merhaba", "xin chao"), 2)`)
}})
```

`duplicate_character_count("hello world", "hello to everyone around the world")` 는 `26`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello world", "hello to everyone around the world"), 26)`)
}})
```

# --seed--

## --seed-contents--

```py
def duplicate_character_count(str1, str2):

    return str1
```

# --solutions--

```py
def duplicate_character_count(str1, str2):
    return sum(1 for char in str2 if char in str1)
```
