---
id: 6a19dd69388e5a6e59e4f2ef
title: "도전 과제 326: 최대 이익"
challengeType: 29
dashedName: challenge-326
---

# --description--

일일 주가 배열과 예산(달러 단위)이 주어질 때, 주어진 기간 동안 주식을 사고팔아 얻을 수 있는 최대 이익을 계산하세요.

- 주식을 사기 전에는 팔 수 없습니다.
- 주식은 온전한 주 단위로만 살 수 있습니다.
- 최대 가능한 이익을 문자열로 반환하되, 소수점 둘째 자리까지 내림하여 포맷하세요.

# --hints--

`get_max_profit([5, 6], 50)`는 `"10.00"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([5, 6], 50), "10.00")`)
}})
```

`get_max_profit([8, 2, 5, 10], 20)`는 `"80.00"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([8, 2, 5, 10], 20), "80.00")`)
}})
```

`get_max_profit([4, 5, 3, 6], 20)`는 `"18.00"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([4, 5, 3, 6], 20), "18.00")`)
}})
```

`get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200)`는 `"8.31"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([54.40, 51.22, 53.99, 50.28, 53.01, 52.84], 200), "8.31")`)
}})
```

`get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80)`는 `"0.00"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([15.38, 15.01, 14.99, 14.62, 14.28], 80), "0.00")`)
}})
```

`get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25)`는 `"73.80"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_max_profit([121.45, 126.82, 122.91, 124.65, 128.83, 128.83, 127.33], 1230.25), "73.80")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_max_profit(prices, budget):

    return prices
```

# --solutions--

```py
def get_max_profit(prices, budget):
    max_profit = 0

    for buy in range(len(prices) - 1):
        for sell in range(buy + 1, len(prices)):
            shares = int(budget / prices[buy])
            profit = shares * (prices[sell] - prices[buy])
            if profit > max_profit:
                max_profit = profit

    return f"{int(max_profit * 100) / 100:.2f}"
```
