---
id: 6a1d9f98e819ed70a0e994de
title: "도전 과제 332: 이슈 분류"
challengeType: 29
dashedName: challenge-332
---

# --description--

이슈에 마지막으로 게시된 이후 경과한 밀리초 수와 이슈에 마지막으로 게시된 메시지를 기준으로 다음 규칙에 따라 이슈에 대해 무엇을 해야 하는지 결정하세요:

- 마지막 메시지가 7일 미만이면 `"leave it"`을 반환하세요
- 마지막 메시지가 7일 이상이고 내용에 대소문자 구분 없이 `"bump"`가 포함되어 있으면 `"close it"`을 반환하세요
- 그렇지 않으면 `"bump it"`을 반환하세요

# --hints--

`triage_issue(86400000, "Lets fix it")`는 `"leave it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")`는 `"bump it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")`는 `"close it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")`는 `"bump it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")`는 `"close it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")`는 `"leave it"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
