---
id: 6a1d9f98e819ed70a0e994df
title: "도전 과제 333: 이슈 분류 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

이슈 제목과 현재 레이블 배열이 주어지면, 다음 규칙에 따라 업데이트된 레이블 배열을 반환하세요:

이슈에 레이블이 없으면 다음을 추가하세요:

- 제목에 `"error"` 또는 `"bug"`가 포함되어 있으면 `"bug"`와 `"needs triage"`를 추가하세요
- 제목에 `"feature"` 또는 `"add"`가 포함되어 있으면 `"enhancement"`와 `"discussing"`을 추가하세요

그렇지 않고 주어진 레이블에 다음이 포함되어 있으면:

- `"needs triage"`가 있고 제목에 `"simple"` 또는 `"easy"`가 포함되어 있으면 `"needs triage"`를 제거하고 `"good first issue"`를 추가하세요
- `"discussing"`이 있고 제목에 `"planned"` 또는 `"next"`가 포함되어 있으면 `"discussing"`을 제거하고 `"on the roadmap"`를 추가하세요
- 그렇지 않고 `"needs triage"` 또는 `"discussing"`가 있으면 이를 제거하고 `"help wanted"`를 추가하세요

제목에 다음이 포함되어 있으면:

- `"security"`가 있으면 `"critical"` 레이블을 추가하세요

# --hints--

`triage_issue("app crashes with error", [])`는 `["bug", "needs triage"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])`는 `["bug", "help wanted"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])`는 `["enhancement", "discussing"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])`는 `["enhancement", "help wanted"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])`는 `["bug", "needs triage", "critical"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])`는 `["critical"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])`는 `["bug", "good first issue"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])`는 `["enhancement", "on the roadmap"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])`는 `["enhancement", "help wanted", "critical"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
