---
id: 6a1d9f98e819ed70a0e994db
title: "도전 과제 329: 버킷 채우기"
challengeType: 29
dashedName: challenge-329
---

# --description--

2D 그리드와 시작 위치(``[row, col]``), 그리고 새 값을 주면, 시작 위치와 같은 값을 가진 모든 연결된 셀을 새 값으로 바꾸세요.

- 셀은 수평 또는 수직으로 인접해 있을 때 연결된 것으로 간주합니다(대각선은 제외).

업데이트된 그리드를 반환하세요.

# --hints--

`bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B")`는 `[["R", "B"], ["R", "B"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "G"], ["R", "G"]], [0, 1], "B"), [["R", "B"], ["R", "B"]])`)
}})
```

`bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B")`는 `[["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["Y", "G", "G"], ["Y", "Y", "Y"], ["B", "Y", "R"]], [1, 2], "B"), [["B", "G", "G"], ["B", "B", "B"], ["B", "B", "R"]])`)
}})
```

`bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R")`는 `[["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["O", "O", "P"], ["P", "O", "O"], ["P", "P", "O"]], [2, 0], "R"), [["O", "O", "P"], ["R", "O", "O"], ["R", "R", "O"]])`)
}})
```

`bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y")`는 `[["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["T", "T", "R", "T"], ["R", "T", "R", "T"], ["R", "T", "R", "T"], ["T", "T", "T", "T"]], [0, 3], "Y"), [["Y", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["R", "Y", "R", "Y"], ["Y", "Y", "Y", "Y"]])`)
}})
```

`bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G")`는 `[["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "B", "G", "B"], ["R", "B", "B", "G"], ["B", "G", "B", "R"], ["B", "G", "G", "B"]], [2, 2], "G"), [["G", "G", "G", "B"], ["R", "G", "G", "G"], ["B", "G", "G", "R"], ["B", "G", "G", "B"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, pos, new_value):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, pos, new_value):
    row, col = pos
    target = grid[row][col]
    if target == new_value:
        return grid

    def fill(r, c):
        if r < 0 or r >= len(grid):
            return
        if c < 0 or c >= len(grid[0]):
            return
        if grid[r][c] != target:
            return
        grid[r][c] = new_value
        fill(r - 1, c)
        fill(r + 1, c)
        fill(r, c - 1)
        fill(r, c + 1)

    fill(row, col)
    return grid
```
