---
id: 6a1d9f98e819ed70a0e994e1
title: "도전 과제 335: 다섯 개 주사위"
challengeType: 29
dashedName: challenge-335
---

# --description--

값이 1-6인 다섯 개 주사위 배열이 주어지면, 가능한 최고의 족보를 반환하세요.

족보는 낮은 순서부터 높은 순서로 다음과 같습니다:

| 족보 | 설명 |
|------|-------------|
| `"no pair"` | 한 쌍도 없는 경우 또는 그 이상 |
| `"pair"` | 같은 값의 주사위 두 개 |
| `"two pair"` | 서로 다른 두 쌍 |
| `"three of a kind"` | 같은 값의 주사위 세 개 |
| `"small straight"` | 연속된 네 개 값 |
| `"large straight"` | 연속된 다섯 개 값 |
| `"full house"` | 세 개 같은 값과 한 쌍 |
| `"four of a kind"` | 같은 값의 주사위 네 개 |
| `"five of a kind"` | 같은 값의 주사위 다섯 개 |

# --hints--

`five_dice([1, 1, 1, 1, 1])`는 `"five of a kind"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 1, 1, 1, 1]), "five of a kind")`)
}})
```

`five_dice([5, 5, 5, 6, 5])`는 `"four of a kind"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([5, 5, 5, 6, 5]), "four of a kind")`)
}})
```

`five_dice([2, 5, 6, 4, 3])`는 `"large straight"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 5, 6, 4, 3]), "large straight")`)
}})
```

`five_dice([4, 3, 3, 3, 1])`는 `"three of a kind"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 3, 3, 3, 1]), "three of a kind")`)
}})
```

`five_dice([4, 6, 2, 6, 5])`는 `"pair"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([4, 6, 2, 6, 5]), "pair")`)
}})
```

`five_dice([1, 4, 5, 6, 2])`는 `"no pair"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 4, 5, 6, 2]), "no pair")`)
}})
```

`five_dice([1, 3, 4, 6, 2])`는 `"small straight"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([1, 3, 4, 6, 2]), "small straight")`)
}})
```

`five_dice([2, 2, 5, 2, 5])`는 `"full house"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([2, 2, 5, 2, 5]), "full house")`)
}})
```

`five_dice([6, 4, 5, 6, 4])`는 `"two pair"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(five_dice([6, 4, 5, 6, 4]), "two pair")`)
}})
```

# --seed--

## --seed-contents--

```py
def five_dice(dice):

    return dice
```

# --solutions--

```py
def five_dice(dice):
    from collections import Counter
    counts = Counter(dice)
    values = sorted(counts.values(), reverse=True)
    keys = sorted(counts.keys())

    is_large_straight = len(keys) == 5 and keys[-1] - keys[0] == 4
    is_small_straight = not is_large_straight and any(
        keys[i + 3] - keys[i] == 3 and len(set(keys[i:i + 4])) == 4
        for i in range(len(keys) - 3)
    )

    if values[0] == 5: return "five of a kind"
    if values[0] == 4: return "four of a kind"
    if values[0] == 3 and values[1] == 2: return "full house"
    if is_large_straight: return "large straight"
    if is_small_straight: return "small straight"
    if values[0] == 3: return "three of a kind"
    if values[0] == 2 and values[1] == 2: return "two pair"
    if values[0] == 2: return "pair"
    return "no pair"
```
