---
id: 6a1d9f98e819ed70a0e994d9
title: "도전 과제 327: 데이터베이스 마이그레이션"
challengeType: 29
dashedName: challenge-327
---

# --description--

두 개의 데이터베이스 객체가 주어지면, 첫 번째 객체에 없는 프로프를 두 번째 객체에 채워서 반환하세요.

- 이미 레코드에 존재하는 필드는 덮어쓰지 마세요.

# --hints--

`migrate_record({ "username": "", "posts": 0 }, { "verified": True })`는 `{ "username": "", "posts": 0, "verified": True }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "verified": True }), { "username": "", "posts": 0, "verified": True })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 })`는 `{ "username": "camper", "posts": 5 }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "posts": 5 }), { "username": "camper", "posts": 5 })`)
}})
```

`migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" })`는 `{ "username": "camper", "posts": 0, "verified": False }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0, "verified": False }, { "username": "camper" }), { "username": "camper", "posts": 0, "verified": False })`)
}})
```

`migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" })`는 `{ "username": "camper", "role": "admin", "posts": 0 }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "posts": 0 }, { "username": "camper", "role": "admin" }), { "username": "camper", "role": "admin", "posts": 0 })`)
}})
```

`migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" })`는 `{ "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False }`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({ "username": "", "email": "", "posts": 0, "verified": False, "role": "user", "banned": False }, { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin" }), { "username": "camper", "email": "camper@freecodecamp.org", "role": "admin", "posts": 0, "verified": False, "banned": False })`)
}})
```

# --seed--

## --seed-contents--

```py
def migrate_record(schema, record):

    return schema
```

# --solutions--

```py
def migrate_record(schema, record):
    return {**schema, **record}
```
