---
id: 6a1d9f98e819ed70a0e994da
title: "도전 과제 328: Kaprekar의 루틴"
challengeType: 29
dashedName: challenge-328
---

# --description--

4자리 숫자가 주어지면 6174에 도달할 때까지 Kaprekar의 루틴을 적용해야 하는 횟수를 반환하세요.

Kaprekar의 루틴은 다음과 같이 작동합니다:

- 숫자를 내림차순으로 정렬해 가장 큰 수를 만드세요
- 숫자를 오름차순으로 정렬해 가장 작은 수를 만드세요 (필요하면 앞에 0을 채우세요)
- 작은 수를 큰 수에서 빼세요
- 새 숫자로 이 과정을 반복하세요

# --hints--

`kaprekar(1234)`는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)`는 `6`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)`는 `4`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)`는 `7`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)`는 `2`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
