---
id: 6a1d9f98e819ed70a0e994dd
title: "도전 과제 331: 가장 가까운 배수"
challengeType: 29
dashedName: challenge-331
---

# --description--

두 개의 정수가 주어지면, 첫 번째 정수를 두 번째 정수의 가장 가까운 배수로 반올림하세요.

# --hints--

`round_to_nearest_multiple(5, 3)`는 `6`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(5, 3), 6)`)
}})
```

`round_to_nearest_multiple(17, 4)`는 `16`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(17, 4), 16)`)
}})
```

`round_to_nearest_multiple(43, 5)`는 `45`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(43, 5), 45)`)
}})
```

`round_to_nearest_multiple(38, 11)`는 `33`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(38, 11), 33)`)
}})
```

`round_to_nearest_multiple(93, 12)`는 `96`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(round_to_nearest_multiple(93, 12), 96)`)
}})
```

# --seed--

## --seed-contents--

```py
def round_to_nearest_multiple(num, multiple):

    return num
```

# --solutions--

```py
def round_to_nearest_multiple(num, multiple):
    return round(num / multiple) * multiple
```
