---
id: 6a2037a68a0bc2aef0006002
title: "도전 과제 340: 피그 라틴 변환기"
challengeType: 29
dashedName: challenge-340
---

# --description--

문자열이 주어지면 다음 규칙을 사용해 피그 라틴으로 변환하세요:

- 단어가 모음(`"a"`, `"e"`, `"i"`, `"o"`, 또는 `"u"`)으로 시작하면 끝에 `"way"`를 추가하세요. 예를 들어, `"universe"`는 `"universeway"`로 변환됩니다.
- 단어가 하나 이상의 자음으로 시작하면 그 자음들을 끝으로 옮기고 `"ay"`를 추가하세요. 예를 들어, `"hello"`는 `"ellohay"`로 변환됩니다.
- 첫 글자의 대소문자를 유지하세요. 예를 들어, `"Hello"`는 `"Ellohay"`로 변환됩니다.

# --hints--

`pig_latin("universe")`는 `"universeway"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")`는 `"ellohay"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")`는 `"ellohay universeway"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")`는 `"Ellohay universeway"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")`는 `"Igpay Atinlay isway unfay"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")`는 `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
