---
id: 6a2037a68a0bc2aef0006007
title: "도전 과제 345: 단어 블렌더"
challengeType: 29
dashedName: challenge-345
---

# --description--

두 단어가 주어지면 첫 번째 단어의 앞부분과 두 번째 단어의 뒷부분을 합쳐 새 단어를 반환하세요.

- 홀수 길이 단어의 경우, 앞부분이 더 짧은 절반입니다.

# --hints--

`blend_words("turtle", "toucan")`는 `"turcan"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("turtle", "toucan"), "turcan")`)
}})
```

`blend_words("chipmunk", "flamingo")`는 `"chipingo"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("chipmunk", "flamingo"), "chipingo")`)
}})
```

`blend_words("falcon", "pelican")`는 `"falican"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("falcon", "pelican"), "falican")`)
}})
```

`blend_words("hyena", "iguana")`는 `"hyana"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("hyena", "iguana"), "hyana")`)
}})
```

`blend_words("scorpion", "gorilla")`는 `"scorilla"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("scorpion", "gorilla"), "scorilla")`)
}})
```

`blend_words("platypus", "wolverine")`는 `"platerine"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(blend_words("platypus", "wolverine"), "platerine")`)
}})
```

# --seed--

## --seed-contents--

```py
def blend_words(word1, word2):

    return word1
```

# --solutions--

```py
def blend_words(word1, word2):
    first_half = word1[:len(word1) // 2]
    second_half = word2[len(word2) // 2:]
    return first_half + second_half
```
