---
id: 6a2037a68a0bc2aef0005fff
title: "도전 과제 337: 집계 카운터"
challengeType: 29
dashedName: challenge-337
---

# --description--

집계 표시 문자열이 주어지면, 나타내는 총 개수를 반환하세요.

- 각 파이프 `"|"`는 하나의 개수를 나타내요.
- 다섯 번째 표시는 슬래시 `"/"`로 나타내며, 다섯 개 그룹을 완성해요 (`"||||/"`).
- 그룹들은 공백으로 구분돼요.

# --hints--

`get_tally_count("||||")`는 `4`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")`는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")`는 `8`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")`는 `17`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")`는 `41`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
