---
id: 6a2037a68a0bc2aef0006000
title: "도전 과제 338: 반려동물 나이 계산기"
challengeType: 29
dashedName: challenge-338
---

# --description--

반려동물 종류와 사람 나이를 주면, 다음 변환 표를 사용해 반려동물 나이로 환산한 값을 반환하세요:

| 반려동물 | 곱셈값 |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`pet_years("dog", 5)`는 `35`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

`pet_years("cat", 9)`는 `54`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

`pet_years("rabbit", 3)`는 `24`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

`pet_years("hamster", 4)`는 `120`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

`pet_years("guinea pig", 5)`는 `60`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

`pet_years("goldfish", 2)`는 `12`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

`pet_years("bird", 1)`는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
