---
id: 6a2037a68a0bc2aef0006001
title: "도전 과제 339: 배열 청크"
challengeType: 29
dashedName: challenge-339
---

# --description--

배열과 청크 크기가 주어지면, 해당 크기로 나눈 하위 배열들로 분할된 배열을 반환하세요.

- 배열이 정확히 나누어지지 않으면 마지막 청크는 더 작을 수 있습니다.

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)`는 `[[1, 2, 3], [4, 5, 6]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)`는 `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)`는 `[[1, 2, 3], [4, 5]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)`는 `[["a"], ["b"], ["c"], ["d"], ["e"]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)`는 `[[1, 2, 3]]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
