---
id: 6a2037a68a0bc2aef0006008
title: "도전 과제 346: 돼지 저금통"
challengeType: 29
dashedName: challenge-346
---

# --description--

돼지 저금통을 나타내는 객체가 주어지면, 총 가치를 `"$D.CC"` 형식의 문자열로 반환하세요.

객체는 다음 중 어느 것이든 포함할 수 있습니다:

| 동전 | 가치 |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6})`는 `"$1.98"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 3, "nickels": 5, "dimes": 2, "quarters": 6}), "$1.98")`)
}})
```

`piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1})`는 `"$0.41"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 1, "nickels": 1, "dimes": 1, "quarters": 1}), "$0.41")`)
}})
```

`piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5})`는 `"$2.25"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"nickels": 8, "dimes": 6, "quarters": 5}), "$2.25")`)
}})
```

`piggy_bank({})`는 `"$0.00"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({}), "$0.00")`)
}})
```

`piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19})`는 `"$6.76"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(piggy_bank({"pennies": 146, "nickels": 11, "dimes": 0, "quarters": 19}), "$6.76")`)
}})
```

# --seed--

## --seed-contents--

```py
def piggy_bank(coins):

    return coins
```

# --solutions--

```py
def piggy_bank(coins):
    values = {"pennies": 0.01, "nickels": 0.05, "dimes": 0.10, "quarters": 0.25}
    total = sum(values[coin] * count for coin, count in coins.items())
    return f"${total:.2f}"
```
