---
id: 6a2037a68a0bc2aef0006004
title: "도전 과제 342: 주사위 확률"
challengeType: 29
dashedName: challenge-342
---

# --description--

여섯 면체 주사위 여러 개를 굴리고 목표 합계가 주어질 때, 그 합계가 나올 확률을 `"1 in X"` 형식의 문자열로 반환하세요.

- 주사위 개수는 1개에서 6개 사이입니다.
- 주어진 주사위 개수로 목표 합계는 항상 나올 수 있습니다.
- `"X"`는 가장 가까운 정수로 반올림하세요.

# --hints--

`get_odds(1, 5)`는 `"1 in 6"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(1, 5), "1 in 6")`)
}})
```

`get_odds(2, 4)`는 `"1 in 12"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(2, 4), "1 in 12")`)
}})
```

`get_odds(3, 10)`는 `"1 in 8"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(3, 10), "1 in 8")`)
}})
```

`get_odds(4, 7)`는 `"1 in 65"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(4, 7), "1 in 65")`)
}})
```

`get_odds(5, 26)`는 `"1 in 111"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(5, 26), "1 in 111")`)
}})
```

`get_odds(6, 35)`는 `"1 in 7776"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(6, 35), "1 in 7776")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odds(dice, target):

    return dice
```

# --solutions--

```py
def get_odds(dice, target):
    valid_combos = 0
    total_combos = 6 ** dice

    def roll(dice_left, total):
        nonlocal valid_combos
        if dice_left == 0:
            if total == target:
                valid_combos += 1
            return
        for face in range(1, 7):
            roll(dice_left - 1, total + face)

    roll(dice, 0)
    return f"1 in {round(total_combos / valid_combos)}"
```
