---
id: 6a22d77ddf034bc4e35b1d5d
title: "도전 과제 356: 마방진 해결기"
challengeType: 29
dashedName: challenge-356
---

# --description--

하나의 숫자가 빠진 3x3 격자(`0`으로 나타냄)가 주어지면, 마방진을 완성하는 빠진 숫자를 반환하세요. 유효한 숫자가 없으면 `"impossible"`을 반환하세요.

마방진은 모든 행, 열, 대각선의 합이 같은 숫자인 격자입니다.

# --hints--

`solve_magic_square([[2, 7, 6], [9, 0, 1], [4, 3, 8]])`는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[2, 7, 6], [9, 0, 1], [4, 3, 8]]), 5)`)
}})
```

`solve_magic_square([[0, 14, 12], [18, 10, 2], [8, 6, 16]])`는 `4`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[0, 14, 12], [18, 10, 2], [8, 6, 16]]), 4)`)
}})
```

`solve_magic_square([[12, 17, 16], [19, 0, 10], [14, 13, 18]])`는 `"impossible"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[12, 17, 16], [19, 0, 10], [14, 13, 18]]), "impossible")`)
}})
```

`solve_magic_square([[15, 35, 31], [43, 27, 11], [23, 19, 0]])`는 `39`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[15, 35, 31], [43, 27, 11], [23, 19, 0]]), 39)`)
}})
```

`solve_magic_square([[26, 41, 14], [47, 35, 0], [32, 29, 44]])`는 `"impossible"`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[26, 41, 14], [47, 35, 0], [32, 29, 44]]), "impossible")`)
}})
```

# --seed--

## --seed-contents--

```py
def solve_magic_square(grid):

    return grid
```

# --solutions--

```py
def solve_magic_square(grid):
    flat = [n for row in grid for n in row]
    zero_index = flat.index(0)
    row = zero_index // 3
    col = zero_index % 3

    target = next(sum(r) for r in grid if 0 not in r)

    missing = target - sum(grid[row])
    filled = [r[:] for r in grid]
    filled[row][col] = missing

    def check(arr):
        return sum(arr) == target

    for i in range(3):
        if not check(filled[i]):
            return "impossible"
        if not check([filled[r][i] for r in range(3)]):
            return "impossible"

    if not check([filled[0][0], filled[1][1], filled[2][2]]):
        return "impossible"
    if not check([filled[0][2], filled[1][1], filled[2][0]]):
        return "impossible"

    return missing
```
