---
id: 6a22d77ddf034bc4e35b1d5c
title: "도전 과제 355: 모스 부호"
challengeType: 29
dashedName: challenge-355
---

# --description--

모스 부호 문자열이 주어지면, 다음 표를 사용해 해독된 메시지를 반환하세요:

| 부호 | 문자 | 부호 | 문자 |
|------|--------|------|--------|
| `.-` | `A` | `-.` | `N` |
| `-...` | `B` | `---` | `O` |
| `-.-.` | `C` | `.--.` | `P` |
| `-..` | `D` | `--.-` | `Q` |
| `.` | `E` | `.-.` | `R` |
| `..-.` | `F` | `...` | `S` |
| `--.` | `G` | `-` | `T` |
| `....` | `H` | `..-` | `U` |
| `..` | `I` | `...-` | `V` |
| `.---` | `J` | `.--` | `W` |
| `-.-` | `K` | `-..-` | `X` |
| `.-..` | `L` | `-.--` | `Y` |
| `--` | `M` | `--..` | `Z` |

- 글자는 한 칸 띄어쓰기로 구분합니다
- 단어는 세 칸 띄어쓰기로 구분합니다

# --hints--

`decode_morse("--..")`는 `"Z"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("--.."), "Z")`)
}})
```

`decode_morse("... --- ...")`는 `"SOS"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("... --- ..."), "SOS")`)
}})
```

`decode_morse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--.")`는 `"FREECODECAMP"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--."), "FREECODECAMP")`)
}})
```

`decode_morse(".... . .-.. .-.. ---   .-- --- .-. .-.. -..")`는 `"HELLO WORLD"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse(".... . .-.. .-.. ---   .-- --- .-. .-.. -.."), "HELLO WORLD")`)
}})
```

`decode_morse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --.")`는 `"THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode_morse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --."), "THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode_morse(code):

    return code
```

# --solutions--

```py
def decode_morse(code):
    table = {
        '.-': 'A', '-...': 'B', '-.-.': 'C', '-..': 'D', '.': 'E',
        '..-.': 'F', '--.': 'G', '....': 'H', '..': 'I', '.---': 'J',
        '-.-': 'K', '.-..': 'L', '--': 'M', '-.': 'N', '---': 'O',
        '.--.': 'P', '--.-': 'Q', '.-.': 'R', '...': 'S', '-': 'T',
        '..-': 'U', '...-': 'V', '.--': 'W', '-..-': 'X', '-.--': 'Y',
        '--..': 'Z'
    }
    return ' '.join(
        ''.join(table[c] for c in word.split(' '))
        for word in code.split('   ')
    )
```
