---
id: 6a22d77ddf034bc4e35b1d57
title: "도전 과제 350: 문자 거리"
challengeType: 29
dashedName: challenge-350
---

# --description--

길이가 같은 두 문자열이 주어지면 각 문자 쌍 사이의 가장 짧은 거리들의 합을 반환하세요.

- 입력은 소문자만 포함합니다
- 알파벳은 원형으로 취급하므로 `a`와 `z` 사이 거리는 1입니다.

# --hints--

`letter_distance("abc", "bcd")`는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "bcd"), 3)`)
}})
```

`letter_distance("abc", "xyz")`는 `9`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("abc", "xyz"), 9)`)
}})
```

`letter_distance("encrypt", "decrypt")`는 `10`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("encrypt", "decrypt"), 10)`)
}})
```

`letter_distance("algorithm", "codeblock")`는 `43`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("algorithm", "codeblock"), 43)`)
}})
```

`letter_distance("lobster", "penguin")`는 `47`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("lobster", "penguin"), 47)`)
}})
```

`letter_distance("alligator", "crocodile")`는 `55`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(letter_distance("alligator", "crocodile"), 55)`)
}})
```

# --seed--

## --seed-contents--

```py
def letter_distance(str1, str2):

    return str1
```

# --solutions--

```py
def letter_distance(str1, str2):
    total = 0
    for a, b in zip(str1, str2):
        diff = abs(ord(a) - ord(b))
        total += min(diff, 26 - diff)
    return total
```
