---
id: 6a26df3e2988bcdded204895
title: "도전 과제 359: 골프 핸디캡 계산기"
challengeType: 29
dashedName: challenge-359
---

# --description--

골프 점수 배열과 해당하는 코스 파 값 배열이 주어질 때, 다음 방법을 사용해 골퍼의 핸디캡 지수를 반환하세요:

- 각 라운드마다 점수에서 파를 빼서 차이를 계산한 후, 모든 차이의 평균을 소수점 첫째 자리까지 반올림해 반환하세요.

# --hints--

`calculate_handicap([72, 72, 72], [72, 72, 72])`는 `0`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([72, 72, 72], [72, 72, 72]), 0)`)
}})
```

`calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72])`는 `6`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72]), 6)`)
}})
```

`calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36])`는 `8.3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3)`)
}})
```

`calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])`는 `8.8`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8)`)
}})
```

`calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])`는 `11.7`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_handicap(scores, pars):

    return scores
```

# --solutions--

```py
def calculate_handicap(scores, pars):
    differentials = [s - p for s, p in zip(scores, pars)]
    average = sum(differentials) / len(differentials)
    return int(average * 10 + 0.5) / 10
```
