---
id: 6a26df95efa55a2524399743
title: "도전 과제 362: 논노그램 검증기"
challengeType: 29
dashedName: challenge-362
---

# --description--

단서 숫자 배열과 셀 배열이 주어졌을 때, 셀이 논노그램 단서를 만족하는지 판단하세요.

- 단서는 연속된 채워진 셀의 길이를 순서대로 나타내는 숫자 배열입니다. 예를 들어, 단서가 `[3, 2]`라면 3개의 연속된 채워진 셀 다음에 최소 하나의 빈 셀로 구분된 2개의 연속된 채워진 셀이 있어야 합니다.
- 행은 1(채워짐)과 0(빈 셀)로 이루어진 배열입니다.

# --hints--

`is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1])`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1]), True)`)
}})
```

`is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1])`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])`는 `True`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])`는 `False`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_nonogram(clue, cells):

    return clue
```

# --solutions--

```py
def is_valid_nonogram(clue, cells):
    runs = []
    count = 0

    for i in range(len(cells) + 1):
        if i < len(cells) and cells[i] == 1:
            count += 1
        elif count > 0:
            runs.append(count)
            count = 0

    if len(runs) != len(clue):
        return False
    return all(runs[i] == clue[i] for i in range(len(runs)))
```
