---
id: 6a26df95efa55a2524399745
title: "도전 과제 364: 두 양동이 사이"
challengeType: 29
dashedName: challenge-364
---

# --description--

RGB 색상과 가득 찬 정도가 각각 있는 두 양동이의 페인트가 주어지면, 혼합된 RGB 색상을 세 개의 정수로 된 배열로 반환하세요.

- 각 양동이는 `color` 속성(세 개의 정수로 된 배열 `[r, g, b]`)과 `fullness` 속성(0–100)을 가진 객체(JavaScript) 또는 딕셔너리(Python)입니다.
- 혼합된 색상은 두 색상의 각 채널을 가득 찬 정도에 따라 가중 평균한 값이며, 각 채널은 가장 가까운 정수로 반올림합니다.

# --hints--

`mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50})`는 `[125, 125, 125]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 50}, {"color": [0, 0, 0], "fullness": 50}), [125, 125, 125])`)
}})
```

`mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20})`는 `[200, 200, 200]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [250, 250, 250], "fullness": 80}, {"color": [0, 0, 0], "fullness": 20}), [200, 200, 200])`)
}})
```

`mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70})`는 `[100, 150, 200]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [100, 150, 200], "fullness": 30}, {"color": [100, 150, 200], "fullness": 70}), [100, 150, 200])`)
}})
```

`mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90})`는 `[114, 184, 170]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [143, 143, 101], "fullness": 45}, {"color": [100, 204, 204], "fullness": 90}), [114, 184, 170])`)
}})
```

`mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54})`는 `[68, 163, 123]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mix_paint({"color": [15, 134, 249], "fullness": 29}, {"color": [97, 178, 55], "fullness": 54}), [68, 163, 123])`)
}})
```

# --seed--

## --seed-contents--

```py
def mix_paint(bucket1, bucket2):

    return bucket1
```

# --solutions--

```py
def mix_paint(bucket1, bucket2):
    total = bucket1['fullness'] + bucket2['fullness']
    return [
        round((c1 * bucket1['fullness'] + c2 * bucket2['fullness']) / total)
        for c1, c2 in zip(bucket1['color'], bucket2['color'])
    ]
```
