---
id: 6a26df3e2988bcdded204893
title: "도전 과제 357: 먹이 사슬"
challengeType: 29
dashedName: challenge-357
---

# --description--

`[predator, prey]` 쌍으로 이루어진 배열이 주어지면, 최상위 포식자부터 최하위까지 먹이 사슬을 반환하세요.

- 최상위 포식자는 다른 동물의 먹이가 되지 않는 동물입니다.
- 먹이 사슬은 문자열 배열로 반환하세요.

# --hints--

`get_food_chain([["cat", "mouse"]])`는 `["cat", "mouse"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])`는 `["wolf", "deer", "grass"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])`는 `["hawk", "snake", "frog", "fly"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])`는 `["eagle", "fox", "rabbit", "grass"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])`는 `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
