---
id: 6a26df3e2988bcdded204894
title: "도전 과제 358: 이모지 번역기"
challengeType: 29
dashedName: challenge-358
---

# --description--

이모지 문자열이 주어지면, 다음 표를 사용해 구문을 반환하세요:

| Emoji | Word |
|-------|------|
| 👶 | `"baby"` |
| 🐱 | `"cat"` |
| 🐕 | `"dog"` |
| 🐟 | `"fish"` |
| 🥵 | `"hot"` |
| 🧊 | `"ice"` |
| 🪨 | `"rock"` |
| 🦈 | `"shark"` |
| 🍲 | `"soup"` |
| ⭐ | `"star"` |

단어들은 공백으로 구분해 반환하세요.

# --hints--

`get_emoji_phrase("🪨⭐")`는 `"rock star"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🪨⭐"), "rock star")`)
}})
```

`get_emoji_phrase("🥵🐕")`는 `"hot dog"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🥵🐕"), "hot dog")`)
}})
```

`get_emoji_phrase("👶🦈")`는 `"baby shark"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("👶🦈"), "baby shark")`)
}})
```

`get_emoji_phrase("⭐🐟")`는 `"star fish"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("⭐🐟"), "star fish")`)
}})
```

`get_emoji_phrase("🧊🧊👶")`는 `"ice ice baby"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🧊🧊👶"), "ice ice baby")`)
}})
```

`get_emoji_phrase("🐱🐟🍲")`는 `"cat fish soup"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_emoji_phrase("🐱🐟🍲"), "cat fish soup")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_emoji_phrase(s):

    return s
```

# --solutions--

```py
def get_emoji_phrase(s):
    table = {
        '👶': 'baby',
        '🐱': 'cat',
        '🐕': 'dog',
        '🐟': 'fish',
        '🥵': 'hot',
        '🧊': 'ice',
        '🪨': 'rock',
        '🦈': 'shark',
        '🍲': 'soup',
        '⭐': 'star',
    }
    return ' '.join(table[emoji] for emoji in s)
```
