---
id: 6a26df95efa55a2524399744
title: "도전 과제 363: 버킷 채우기 2"
challengeType: 29
dashedName: challenge-363
---

# --description--

한 글자 색 문자열로 이루어진 2D 그리드와 목표 색이 주어질 때, 전체 그리드를 목표 색으로 만들기 위해 필요한 최솟값 플러드 필 "클릭" 수를 반환하세요.

- 각 클릭은 클릭한 셀의 색과 같은 색으로 연결된 전체 영역을 목표 색으로 바꿉니다.
- 셀들은 수평과 수직으로 연결되어 있으며(대각선 연결 아님) 연결됩니다.

# --hints--

`bucket_fill([["R", "R"], ["R", "R"]], "G")`는 `1`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "R"], ["R", "R"]], "G"), 1)`)
}})
```

`bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B")`는 `0`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["B", "B", "B"], ["B", "B", "B"], ["B", "B", "B"]], "B"), 0)`)
}})
```

`bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R")`는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "Y", "Y"], ["G", "Y", "G"], ["Y", "Y", "G"]], "R"), 3)`)
}})
```

`bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P")`는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "P", "Y"], ["O", "P", "P", "P"], ["O", "O", "P", "G"], ["G", "O", "O", "G"]], "P"), 5)`)
}})
```

`bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y")`는 `12`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "C", "C", "O"], ["B", "Y", "B", "Y", "O"], ["B", "J", "O", "J", "B"], ["G", "Y", "Y", "Y", "B"], ["G", "P", "P", "G", "G"]], "Y"), 12)`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, target_color):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, target_color):
    rows = len(grid)
    cols = len(grid[0])
    visited = set()

    def dfs(r, c, color):
        if r < 0 or r >= rows or c < 0 or c >= cols:
            return
        if (r, c) in visited:
            return
        if grid[r][c] != color:
            return
        visited.add((r, c))
        dfs(r + 1, c, color)
        dfs(r - 1, c, color)
        dfs(r, c + 1, color)
        dfs(r, c - 1, color)

    clicks = 0
    for r in range(rows):
        for c in range(cols):
            if (r, c) not in visited and grid[r][c] != target_color:
                dfs(r, c, grid[r][c])
                clicks += 1

    return clicks
```
