---
id: 6a26df3e2988bcdded204896
title: "도전 과제 360: 음성으로 표현한 시간"
challengeType: 29
dashedName: challenge-360
---

# --description--

초 단위 숫자가 주어지면, 음성으로 표현한 시간을 반환하세요.

- 시간을 시간, 분, 초로 나누세요.
- 0인 값은 건너뛰세요.
- 단수 또는 복수를 적절히 사용하세요 (`"1 hour"`, `"2 hours"`).
- 마지막 두 단위가 있으면 `"and"`로 연결하고, 마지막에서 두 번째와 세 번째 단위는 쉼표로 구분하세요 (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`get_spoken_duration(3723)`는 `"1 hour, 2 minutes and 3 seconds"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(3723), "1 hour, 2 minutes and 3 seconds")`)
}})
```

`get_spoken_duration(7295)`는 `"2 hours, 1 minute and 35 seconds"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(7295), "2 hours, 1 minute and 35 seconds")`)
}})
```

`get_spoken_duration(8521)`는 `"2 hours, 22 minutes and 1 second"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(8521), "2 hours, 22 minutes and 1 second")`)
}})
```

`get_spoken_duration(435)`는 `"7 minutes and 15 seconds"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(435), "7 minutes and 15 seconds")`)
}})
```

`get_spoken_duration(14455)`는 `"4 hours and 55 seconds"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(14455), "4 hours and 55 seconds")`)
}})
```

`get_spoken_duration(72000)`는 `"20 hours"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(72000), "20 hours")`)
}})
```

`get_spoken_duration(1)`는 `"1 second"`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(1), "1 second")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_duration(seconds):

    return seconds
```

# --solutions--

```py
def get_spoken_duration(seconds):
    hours = seconds // 3600
    minutes = (seconds % 3600) // 60
    secs = seconds % 60

    parts = []
    if hours:
        parts.append(f"{hours} {'hour' if hours == 1 else 'hours'}")
    if minutes:
        parts.append(f"{minutes} {'minute' if minutes == 1 else 'minutes'}")
    if secs:
        parts.append(f"{secs} {'second' if secs == 1 else 'seconds'}")

    if len(parts) == 1:
        return parts[0]
    if len(parts) == 2:
        return f"{parts[0]} and {parts[1]}"
    return f"{parts[0]}, {parts[1]} and {parts[2]}"
```
