---
id: 6a26df95efa55a2524399746
title: "도전 과제 365: 마지막 도전 과제: 버킷 채우기 3"
challengeType: 29
dashedName: challenge-365
---

# --description--

오늘로 매일 코딩 도전 과제를 시작한 지 1년이 되었습니다. 지금은 마지막 새로운 도전 과제입니다. 행운을 빌어요!

한 글자 색 문자열로 이루어진 2D 격자와 목표 색이 주어질 때, 전체 격자를 그 색으로 만들기 위해 필요한 최소한의 플러드 필 "클릭" 수를 반환하세요.

- 각 클릭은 클릭한 셀의 색과 같은 색으로 연결된 셀 전체 영역(4방향)의 색을 변경합니다.
- 클릭은 목표 색뿐 아니라 중간 단계로 어떤 색도 사용할 수 있습니다.

# --hints--

`bucket_fill([["B", "B"], ["B", "B"]], "R")`는 `1`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["B", "B"], ["B", "B"]], "R"), 1)`)
}})
```

`bucket_fill([["G", "G", "G"], ["G", "G", "G"], ["G", "G", "G"]], "G")`는 `0`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "G"], ["G", "G", "G"], ["G", "G", "G"]], "G"), 0)`)
}})
```

`bucket_fill([["P", "P", "Y"], ["Y", "P", "Y"], ["Y", "P", "P"]], "O")`는 `2`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["P", "P", "Y"], ["Y", "P", "Y"], ["Y", "P", "P"]], "O"), 2)`)
}})
```

`bucket_fill([["G", "Y", "C", "C"], ["Y", "Y", "Y", "B"], ["C", "Y", "B", "B"], ["C", "B", "B", "C"]], "R")`는 `4`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "Y", "C", "C"], ["Y", "Y", "Y", "B"], ["C", "Y", "B", "B"], ["C", "B", "B", "C"]], "R"), 4)`)
}})
```

`bucket_fill([["G", "G", "O", "O"], ["G", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["G", "G", "G", "G"]], "P")`는 `5`를 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["G", "G", "O", "O"], ["G", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["B", "Y", "B", "Y"], ["G", "G", "G", "G"]], "P"), 5)`)
}})
```

`bucket_fill([["R", "G", "R", "G"], ["R", "G", "R", "G"], ["B", "B", "B", "B"], ["B", "B", "B", "B"], ["R", "G", "R", "G"]], "Y")`는 `3`을 반환해야 합니다.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(bucket_fill([["R", "G", "R", "G"], ["R", "G", "R", "G"], ["B", "B", "B", "B"], ["B", "B", "B", "B"], ["R", "G", "R", "G"]], "Y"), 3)`)
}})
```

# --seed--

## --seed-contents--

```py
def bucket_fill(grid, target_color):

    return grid
```

# --solutions--

```py
def bucket_fill(grid, target_color):
    from collections import deque

    rows = len(grid)
    cols = len(grid[0])

    def grid_key(g):
        return tuple(tuple(r) for r in g)

    def get_region(g, row, col):
        color = g[row][col]
        visited = set()
        stack = [(row, col)]
        while stack:
            r, c = stack.pop()
            if (r, c) in visited or not (0 <= r < rows and 0 <= c < cols) or g[r][c] != color:
                continue
            visited.add((r, c))
            stack.extend([(r+1,c),(r-1,c),(r,c+1),(r,c-1)])
        return visited

    def flood_fill(g, row, col, color):
        next_g = [list(r) for r in g]
        for r, c in get_region(g, row, col):
            next_g[r][c] = color
        return next_g

    def is_complete(g):
        return all(cell == target_color for row in g for cell in row)

    initial = [list(r) for r in grid]
    if is_complete(initial):
        return 0

    queue = deque([(initial, 0)])
    seen = {grid_key(initial)}

    while queue:
        current, clicks = queue.popleft()
        all_colors = set(cell for row in current for cell in row) | {target_color}

        visited_all = set()
        region_roots = []
        for r in range(rows):
            for c in range(cols):
                if (r, c) not in visited_all:
                    region = get_region(current, r, c)
                    visited_all |= region
                    region_roots.append((r, c))

        for r, c in region_roots:
            for color in all_colors:
                if color == current[r][c]:
                    continue
                next_g = flood_fill(current, r, c, color)
                if is_complete(next_g):
                    return clicks + 1
                key = grid_key(next_g)
                if key not in seen:
                    seen.add(key)
                    queue.append((next_g, clicks + 1))
```
