---
id: 67116d7584d0b469b14579c0
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Crea otro objeto dentro del array `library` con las siguientes propiedades y valores:

| Propiedad | Valor                                                                   |
| --------- | ----------------------------------------------------------------------- |
| `title`   | `"Atomic Habits"`                                                       |
| `author`  | `"James Clear"`                                                         |
| `about`   | `"A practical book about discarding bad habits and building good ones"` |
| `pages`   | `320`                                                                   |

# --hints--


Debe tener un objeto como el segundo ítem de su array `library`.

```js
assert.isObject(library[1]);
```

Su objeto debe tener una propiedad `title`.

```js
assert.property(library[1], "title")
```

La propiedad `title` debe tener `"Atomic Habits"` como su valor.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Su objeto debe tener una propiedad `author`.

```js
assert.property(library[1], "author")
```

La propiedad `author` debe tener `"James Clear"` como su valor.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Su objeto debe tener una propiedad `about`.

```js
assert.property(library[1], "about")
```

La propiedad `about` debe tener el valor `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Su objeto debe tener una propiedad `pages`.

```js
assert.property(library[1], "pages")
```

El valor de tu propiedad `pages` debe ser un número.

```js
assert.isNumber(library[1]?.pages)
```

La propiedad `pages` debe establecerse en `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--

--fcc-editable-region--
];
```
