---
id: 67116d7584d0b469b14579bf
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Dentro del arreglo `library`, crea un objeto con las siguientes propiedades y valores:

PropiedadValor
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Debes tener un objeto dentro del arreglo `library`.

```js
assert.isObject(library[0])
```

Su objeto debe tener una propiedad `title`.

```js
assert.property(library[0], "title")
```

El valor de tu propiedad `title` debe ser `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Su objeto debe tener una propiedad `author`.

```js
assert.property(library[0], "author")
```

Tu propiedad `author` debe tener el valor `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Su objeto debe tener una propiedad `about`.

```js
assert.property(library[0], "about")
```

Tu propiedad `about` debe tener el valor `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Su objeto debe tener una propiedad `pages`.

```js
assert.property(library[0], "pages")
```

El valor de tu propiedad `pages` debe ser un número.

```js
assert.isNumber(library[0]?.pages)
```

La propiedad `pages` debe establecerse en `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
