---
id: 67116c175a77cb64c5c10d49
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En lecciones anteriores, aprendiste cómo trabajar con funciones de orden superior como `map`, `filter` y `reduce`.

En este taller, construirás una aplicación de gestor de bibliotecas que te dará la oportunidad de practicar trabajando con diferentes funciones de orden superior.

Para comenzar, crea una variable llamada `library` y asígnale un arreglo vacío.

# --hints--

Debes crear un arreglo de `library`.

```js
assert.isArray(library)
```

Tu arreglo `library` debe estar vacío.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
