---
id: 68225ea2f1aae301fa4df0f4
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Mientras el array de resultados está funcionando, el resultado final deseado debería ser una cadena que liste todos los libros con un título y autor.

Aquí es donde entra el método `join()`. En lecciones anteriores, aprendiste que el método `join()` se usa para devolver una nueva cadena con todos los elementos del arreglo concatenados en una sola cadena, con un separador especificado entre cada elemento.

Aquí tienes un repaso:

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

El separador en este caso es el `\n` que representa el carácter de nueva línea. 

Encadena el método `join()` con un `\n` como separador al `map()`. Ahora deberías ver una cadena en la consola en lugar del arreglo de resultados. 

# --hints--

Tu función `getBookInformation` debería devolver una cadena.

```js
assert.isString(getBookInformation(library));
```

Tu función `getBookInformation` debería devolver una sola cadena que represente el resultado final.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`)
}
--fcc-editable-region--

console.log(getBookInformation(library));
```
