---
id: 68ed661e9f463c155bb677a1
title: 1단계
challengeType: 0
dashedName: step-1
---

# --description--

이 워크숍에서는 CSS flexbox를 사용해 색상이 있는 상자들을 디자인하며 연습합니다.

먼저 `body` 요소 안에 `header` 요소를 추가하세요. `h1` 요소 안에 `header` 요소를 중첩하고, 그 안에 `Colored Boxes Layout` 텍스트를 넣으세요.

# --hints--

`header` 요소 안에 `body` 요소가 있어야 합니다.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

`header` 요소 안에 `h1` 요소가 있어야 합니다.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

`h1` 요소 안에 `Colored Boxes Layout` 텍스트가 있어야 합니다.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
