---
id: 68ef0b3237fef991c92d9c19
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이제 `div` 아래에 클래스가 `flex-container`인 `header` 요소를 추가하세요. 그런 다음 방금 만든 `div` 안에 클래스가 `box`인 `div` 요소 여섯 개를 만드세요.

클래스 속성은 스타일 시트의 클래스 이름을 가리키는 데 자주 사용된다는 점을 기억하세요. 이 경우, 모든 `div` 요소는 `.flex-container`와 `.box` 스타일 정의에 따라 동일하게 스타일이 적용됩니다.

# --hints--

`div` 요소 아래에 `header` 요소가 있어야 합니다.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

`div` 요소 아래에 있는 `header` 요소는 클래스 `flex-container`를 가져야 합니다.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

`div` 요소 안에 `.flex-container` 요소 여섯 개가 중첩되어 있어야 합니다.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

`div` 요소 안에 있는 여섯 개의 `.flex-container` 요소 모두 클래스 `box`를 가져야 합니다.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
