---
id: 68ef0d6d156f86a6ac1824c6
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

이제 `align-content` 속성을 사용할 차례입니다. 이 속성은 교차 축을 따라 아이템의 배치를 제어합니다. `align-content`의 기본값은 `stretch`이지만, `align-content`에 사용할 수 있는 가장 관련 있는 옵션은 `space-between`, `start`, `center`, `end`, `space-around` 그리고 `space-evenly`입니다.

`align-content` 속성을 값 `space-between`으로 추가하여 `.flex-container` 요소 안에서 박스들이 어떻게 정렬되는지 확인하세요.

# --hints--

`.flex-container` 선택자에 `align-content` 속성을 값 `space-between`으로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('align-content'), 'space-between');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}

.box {
	max-height: 120px;
	color: #000;
	border: 1px solid #000;
	display: flex;
	flex: 1 1 100px;
	flex-direction: column;
	align-items: center;
	margin: 10px;
	font-weight: bold;
	font-size: 1.125rem;
	border-radius: 5px;
	order: 0; 
}

.flex-container {
	display: flex;
	flex-wrap: wrap;
	width: 70%;
	height: 600px;
	padding: 10px;
	margin: 20px auto;
--fcc-editable-region--
	
--fcc-editable-region--
}
```
