---
id: 68ef1b1c657f1b2396d98592
title: 29단계
challengeType: 0
dashedName: step-29
---

# --description--

`flex` 속성은 플렉스 컨테이너 내부 항목의 크기와 동작을 제어합니다. 이 속성은 `flex-grow`, `flex-shrink`, 그리고 `flex-basis` 세 가지 속성으로 구성됩니다.

`flex-grow` 속성은 컨테이너에 여유 공간이 있을 때 플렉스 항목이 얼마나 추가 공간을 차지할지 제어합니다.

이제 `.box` 선택자로 돌아가서 `flex`의 첫 번째 숫자(`flex-grow` 값에 해당)를 `1`에서 `0`로 변경하세요.

# --hints--

`.box` 선택자에는 `flex` 속성이 `0 1 100px` 값으로 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '0 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}

.box {
	max-height: 120px;
	color: #000;
	border: 1px solid #000;
	display: flex;
--fcc-editable-region--
	flex: 1 1 100px;
--fcc-editable-region--
	flex-direction: column;
	align-items: center;
	margin: 10px;
	font-weight: bold;
	font-size: 1.125rem;
	border-radius: 5px;
	order: 0;
}

.flex-container {
	display: flex;
	flex-wrap: wrap;
	width: 70%;
	height: 600px;
	padding: 10px;
	margin: 20px auto;
	align-content: space-evenly;
}
```
