---
id: 68f46b758eb8ea7043abd73c
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

이제 flexbox 컨테이너 안의 박스들을 정리할 차례입니다. 먼저 `.box1` 선택자를 만들고 `background` 속성을 값 `#f16e79`로 추가하세요.

그다음 `order` 속성을 값 `1`로 추가하세요. 그러면 첫 번째 박스가 끝으로 이동하는 것을 볼 수 있습니다. 이는 `box` 클래스가 이미 `order` 속성에 `0` 값을 가지고 있어서, 더 높은 order 값을 가진 항목이 나중에 나타나기 때문입니다.

또한 `flex-grow` 속성을 값 `1`로 추가하세요. 이렇게 하면 박스가 자신의 줄에서 남는 공간을 채우도록 커집니다.

# --hints--

`.box1` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box1'));
```

`.box1` 선택자에는 `background` 속성이 값 `#f16e79`로 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('background'), 'rgb(241, 110, 121)');
```

`.box1` 선택자에는 `order` 속성이 값 `1`로 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('order'), '1');
```

`.box1` 선택자에는 `flex-grow` 속성이 값 `1`로 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.getPropVal('flex-grow'), '1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--
```
