---
id: 68f47093c85c7da27c7ee5cb
title: 39단계
challengeType: 0
dashedName: step-39
---

# --description--

이제 `.box4` 선택자를 만들고 `background` 속성을 값 `#4caf50`로 추가하세요.

그다음 `order` 선택자에 값이 `3`인 `.box4` 속성을 추가하세요. 그러면 이 상자가 `box3` 클래스를 가진 상자 뒤로 이동하는 것을 볼 수 있습니다.

# --hints--

`.box4` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box4'));
```

`.box4` 선택자에는 값이 `background`인 `#4caf50` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box4')?.getPropVal('background'), 'rgb(76, 175, 80)');
```

`.box4` 선택자에는 값이 `order`인 `3` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box4')?.getPropVal('order'), '3');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box box3">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box box4">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

--fcc-editable-region--

--fcc-editable-region--
```
