---
id: 68f932c803f4d3d9f511b436
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

`flex-wrap` 속성은 플렉스 항목이 한 줄에 강제로 배치되는지 아니면 여러 줄로 감쌀 수 있는지를 정의합니다.

기본적으로 플렉스 항목은 `nowrap` 값을 가지며, 이는 컨테이너를 넘쳐도 한 줄에 머무른다는 뜻입니다. 반면 `wrap` 값은 플렉스 항목이 컨테이너 너비에 맞지 않을 때 새 줄로 감싸져 보다 유연하고 반응형 레이아웃을 만듭니다.

계속하려면 `flex-wrap` 클래스 선택자에 `wrap` 속성을 `.flex-container` 값으로 추가하세요.

# --hints--

`.flex-container` 선택자에 `flex-wrap` 속성을 `wrap` 값으로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('flex-wrap'), 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
