---
id: 68f93a8b58cbf330b4ba45c6
title: 22단계
challengeType: 0
dashedName: step-22
---

# --description--

`flex-container` 클래스를 스타일링하는 작업을 마치려면, 값이 `padding`인 `10px` 속성과 값이 `margin`인 `20px auto` 속성을 추가하세요.

기억하세요: `padding`는 요소 내부(테두리와 내용 사이)에 공간을 추가하고, `margin`는 요소 외부(테두리와 다른 요소 사이)에 공간을 추가합니다.

# --hints--

`.flex-container` 선택자에 `padding` 속성을 값 `10px`로 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('padding'), '10px');
```

`.flex-container` 선택자에 `margin` 속성을 값 `20px auto`로 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('margin'), '20px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
