---
id: 68f91df5776e693049fb8942
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이제 `styles.css` 파일에 CSS 규칙을 추가할 시간입니다. `h1` 요소에 대한 선택자를 먼저 만드세요.

`h1` 요소의 내용을 가운데 정렬하려면 `text-align` 속성을 `center`로 설정하세요. 또한 `margin-bottom` 요소의 아래쪽과 그 아래에 오는 HTML 요소 사이에 여백을 추가하려면 값이 `10px`인 `h1` 속성도 만드세요.

# --hints--

`h1` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` 선택자는 `text-align` 속성을 `center`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

`h1` 선택자는 `margin-bottom` 속성을 `10px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
