---
id: 692f22933ba1cb86cd75a70c
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

이제 박스들의 색상 이름을 만들어야 합니다. 이를 위해 각 `p` 요소 아래에 `h2` 요소를 추가하고, 순서대로 `Red`, `Orange`, `Yellow`, `Green`, `Blue`, `Indigo` 색상 이름을 넣으세요. 박스마다 하나의 색상입니다.

# --hints--

첫 번째 `p` 요소 안의 `Red` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

두 번째 `p` 요소 안의 `Orange` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

세 번째 `p` 요소 안의 `Yellow` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

네 번째 `p` 요소 안의 `Green` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

다섯 번째 `p` 요소 안의 `Blue` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

여섯 번째 `p` 요소 안의 `Indigo` 요소 아래에 텍스트 `h2`를 가진 `.box` 요소를 추가하세요.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
