---
id: 692f2c8ea9aa5a593758ccff
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`display` 속성은 페이지에서 요소가 어떻게 렌더링되는지 제어합니다. `flex` 값을 사용해 정의하면, 컨테이너 안의 항목들을 배치하고 정렬할 수 있는 플렉스박스 레이아웃 모드를 활성화합니다.

이제 `display` 선택자에 `flex` 속성을 `.box` 값으로 추가하세요.

# --hints--

`.box` 선택자에 `display` 속성을 `flex` 값으로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('display'), 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
