---
id: 692f318859d5e4bf8fd77ba1
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

`order` 속성은 같은 컨테이너 안에 있는 다른 플렉스 항목과 비교했을 때 플렉스 항목의 순서를 지정합니다. 기본적으로 모든 플렉스 항목은 `order` 값이 `0`로 설정되어 있어 HTML에 작성된 순서대로 나타납니다. `order` 값이 낮은 항목이 먼저 나타나고, 값이 높은 항목은 나중에 나타납니다.

`order` 선택자에 `0` 속성을 값 `.box`로 추가하세요.

# --hints--

`.box` 선택자에 `order` 속성을 값 `0`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
