---
id: 692f2b9f2b9aa9458f6fdc97
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 `div` 요소에 `box` 클래스를 스타일링하세요. 먼저 `.box` 클래스 선택자를 만들고 `max-height` 클래스 선택자에 `120px` 속성을 값 `box`로 추가하세요.

# --hints--

`.box` 선택자에 `max-height` 속성을 값 `120px`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('max-height'), '120px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
