---
id: 692f30e86027a4b253d0a42d
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

`font-size` 선택자에 `1.125rem` 속성을 값 `.box`으로 추가하세요.

`font-size` 속성은 텍스트 크기를 지정합니다. 픽셀, em, rem, 백분율 등 다양한 단위로 설정할 수 있어 반응형 디자인에 유연하게 사용할 수 있습니다.

이번 경우 값은 루트 `rem` 요소를 기준으로 하는 `html` 단위를 사용합니다. 부모 요소가 아니라 루트 요소를 기준으로 하기 때문에 폰트 크기가 일정하게 유지되어 부모 요소 크기에 영향을 받지 않아 레이아웃 예측이 더 쉽습니다.

# --hints--

`.box` 선택자에 `font-size` 속성을 값 `1.125rem`으로 추가해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('font-size'), '1.125rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
