---
id: 692f314d449329ba8bcd5d2e
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

`border-radius` 속성은 요소 테두리의 모서리를 둥글게 만듭니다. 다음과 같이 지정할 수 있습니다:

- 네 모서리 모두에 적용할 하나의 값.
- 두 개의 값, 첫 번째는 왼쪽 위/오른쪽 아래, 두 번째는 오른쪽 위/왼쪽 아래에 적용.
- 세 개의 값, 각각 왼쪽 위, 오른쪽 위/왼쪽 아래, 오른쪽 아래에 대응.
- 네 개의 값, 시계 방향으로 왼쪽 위 모서리부터 테두리 반경을 설정.

이제 `border-radius` 선택자에 `5px` 속성을 값 `.box`로 추가하세요.

# --hints--

`.box` 선택자에 `border-radius` 속성을 값 `5px`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border-radius'), '5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
