---
id: 692f2d37f34ed866ed42fcd6
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`flex-direction` 속성은 플렉스 항목이 컨테이너 내에서 어떻게 배치되는지 제어합니다. 기본값은 `row`으로, 항목을 왼쪽에서 오른쪽으로 수평으로 배치합니다. 이 속성은 `display: flex` 또는 `display: inline-flex`가 적용된 요소에서만 작동합니다.

`column` 값을 사용하면 항목이 위에서 아래로 수직으로 쌓입니다. 그 밖에 `row`, `row-reverse`, `column-reverse` 값도 가능합니다.

이제 `flex-direction` 선택자에 `column` 속성을 값 `.box`로 추가하세요.

# --hints--

`.box` 선택자에 `flex-direction` 속성을 값 `column`로 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
