---
id: 692f22a2ad44ae887a24c8f9
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

각 `.box` 요소 안에 텍스트 `h2`와 공백, 그리고 `Box`부터 시작하는 순차 번호가 붙은 `1` 요소를 추가하세요.

# --hints--

첫 번째 `h2` 안에 클래스 `Box 1`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

두 번째 `h2` 안에 클래스 `Box 2`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

세 번째 `h2` 안에 클래스 `Box 3`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

네 번째 `h2` 안에 클래스 `Box 4`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

다섯 번째 `h2` 안에 클래스 `Box 5`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

마지막 `h2` 안에 클래스 `Box 6`가 있는 `div` 요소가 텍스트 `box`를 포함하고 있어야 합니다.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
