---
id: 692f2c20cb7f9e5050b935a7
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`border` 선택자에 `1px solid #000` 값을 가진 `.box` 속성을 추가하세요. `border` 속성은 요소 주위에 보이는 윤곽선을 만드는 데 사용됩니다. 이 속성은 `border-width`, `border-style`, `border-color` 측면에서 사용자 지정할 수 있습니다.

이 경우 `1px`은 `border-width`로, 테두리 두께를 정의합니다. `solid`은 `border-style`이며 테두리 모양(실선, 점선, 대시선 등)을 설정합니다. 그리고 `#000`은 `border-color`로, 테두리 색상을 정의합니다.

# --hints--

`.box` 선택자에 `border` 속성을 `1px solid #000` 값으로 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('border'), '1px solid rgb(0, 0, 0)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
