---
id: 694b46a414bb2f5079868ed1
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

`flex` 속성은 유연한 컨테이너 안의 항목 크기와 동작을 제어합니다. `flex-grow`, `flex-shrink`, `flex-basis` 세 가지 속성으로 구성되어 있습니다.

이제 `flex` 선택자에 값이 `1 1 100px`인 `.box` 속성을 추가하세요. 이렇게 하면 각 박스가 `100px`에서 시작하여 사용 가능한 공간을 채우기 위해 확장하거나 필요할 때 축소할 수 있습니다.

# --hints--

`.box` 선택자에 값이 `flex`인 `1 1 100px` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '1 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
