---
id: 6a3e6c0eb5d050ce07781709
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Якщо умова `if` оцінюється як `False` (тобто у вас *є* достатньо яєць), функція пропустить блок `if` і продовжить виконання решти коду.

Тепер потрібно обробити свої дані і повернути нове обчислене значення.

Поза межами вашого блоку `if` (на тому ж рівні відступу, що й ключове слово `if`), використайте виклик `print()` та f-рядок, щоб вивести `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Потім, на наступному рядку, `return` залишкову кількість яєць, віднявши `eggs_to_use` від `available_eggs`.

# --hints--

У вашій функції `use_eggs` має бути інструкція print, яка виводить `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

Ви маєте повернути залишкову кількість яєць, віднявши `eggs_to_use` від `available_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

Ви маєте повернути `available_eggs - eggs_to_use` після інструкції `print`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
