---
id: 6a3e5475c35dbc138d77c6f4
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер, коли ви обчислили свій загальний підсумок, слід відобразити інвентар. Нагадаємо з попередніх уроків, що ви можете використовувати функцію `print()` з **f-рядками**, щоб включати змінні у свої вихідні дані.

Усередині вашої функції `check_kitchen_stock`, під змінною `total_items`, додайте саме ці чотири інструкції виведення:

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

У вашій функції має бути інструкція `print(f'The kitchen has {total_items} total items:')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

Інструкція `print(f'The kitchen has {total_items} total items:')` має бути першим викликом `print()` у вашій функції.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

У вашій функції має бути інструкція `print(f'- {available_eggs} eggs')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

У вашій функції має бути інструкція `print(f'- {available_flour} flour')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

У вашій функції має бути інструкція `print(f'- {available_sugar} sugar')`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
