---
id: 6a3e2b9b7eeca084dd4fb857
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Ласкаво просимо на кухню! Так само, як кухар покладається на рецепти для приготування страви, ви будете покладатися на функції для створення програмного забезпечення. На цьому практичному занятті ви навчитеся писати багаторазові Python-функції та опануєте область видимості, розуміючи точно, які інгредієнти (змінні) доступні у вашому коді в будь-який момент.

Щоб почати, слід визначити базові дані для вашого кухонного інвентарю. Створіть три змінні: `available_eggs`, `available_flour` та `available_sugar` і призначте їм відповідно значення `1`, `2` та `3`.

# --hints--

Визначте змінну з ім’ям `available_eggs` і призначте їй значення `1`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

Визначте змінну з ім’ям `available_flour` і призначте їй значення `2`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

Визначте змінну з ім’ям `available_sugar` і призначте їй значення `3`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
