---
id: 6a3e58891e1ff13d07f4ab48
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Python читає блок `def` і запам’ятовує функцію, але не запускає код всередині неї одразу.

Щоб виконати код і побачити свої вихідні дані, потрібно **викликати** (або ініціювати) функцію. Для цього введіть ім’я функції, а потім круглі дужки.

Наприклад:

```py
my_function()
```

Внизу вашого коду, поза межами функції, викличте вашу функцію `check_kitchen_stock`.

# --hints--

Ви повинні викликати функцію `check_kitchen_stock` в кінці свого коду. Переконайтеся, що вона не відступлена, щоб вона запускалася глобально.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
