---
id: 6a3e5ccb99fd6053af9032a1
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви знаєте, що є у вашій кухні, ви можете використовувати ці елементи для приготування їжі. Для цього створіть нову функцію, яка приймає параметри. Уявіть параметри як заповнювачі для даних, які ви передаєте у функцію під час її виклику.

Ось функція, яка приймає два параметри, `a` та `b`, і виводить їх суму в термінал.

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

Під вашою функцією `check_kitchen_stock` визначте порожню функцію з назвою `use_eggs`, яка приймає два параметри: `available_eggs` та `eggs_to_use` у такому порядку.

Зверніть увагу, що якщо ви дасте параметру точно таку ж назву, як у глобальної змінної (наприклад, `available_eggs`), функція використовуватиме свій локальний параметр `available_eggs` замість глобальної змінної. Це називається затіненням імен, оскільки локальна змінна «затінює» глобальну.

Поки що використовуйте ключове слово `pass` всередині тіла функції.

# --hints--

Ви повинні визначити функцію з назвою `use_eggs`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

Ваша функція `use_eggs` має мати параметри `available_eggs` та `eggs_to_use`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
