---
id: 6a3e71a0a3a44260e675b157
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Щоб розгадати загадку, чому кількість ваших яєць не змінилася, і насправді приготувати сніданок, створіть нову функцію. Уявіть цю функцію як ваш рецепт, який викликає інші функції для виконання роботи.

Під інструкцією `print(available_eggs)` визначте порожню функцію з назвою `make_fried_egg`, яка приймає `available_eggs` як єдиний параметр.

Поки що використайте ключове слово `pass` всередині тіла функції.

# --hints--

Ви повинні визначити функцію з назвою `make_fried_egg`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

Ваша функція `make_fried_egg` має мати `available_eggs` як параметр.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
