---
id: 6a3ea8a55317ba7a08e21d5b
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Тепер ви нарешті можете відповісти на попереднє питання: чому глобальна кількість яєць залишилася `1`? Коли функція повертає значення, вона не перезаписує ваші існуючі змінні автоматично. Ви повинні явно зберегти це нове значення!

Ви використовуєте змінну `available_eggs` всередині функції `make_fried_egg`, але це не та сама змінна, що визначена на початку вашого коду. Та, що всередині функції, є локальною змінною, яка існує лише під час виконання функції. Коли функція завершується, локальна змінна видаляється, а глобальна `available_eggs` залишається без змін.

Щоб оновити фактичний запас у вашій кухні, перезапишіть глобальну змінну `available_eggs` новою кількістю яєць, яку повертає ваша функція `make_fried_egg`.

У самому низу вашого коду викличте функцію `make_fried_egg`, передаючи `available_eggs` як аргумент, і призначте результат назад глобальній змінній `available_eggs`.

# --hints--

Ви повинні переназначити змінну `available_eggs` результатом виклику `make_fried_egg(available_eggs)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
