---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Час створити функцію, щоб перевірити, що у вас є на кухні. У Python ви визначаєте функції, використовуючи ключове слово `def`, за яким іде назва функції, круглі дужки `()`, і двокрапка `:`.

Ось приклад, як визначити просту функцію:

```py
def my_function():
    pass
```

Визначте порожню функцію з іменем `check_kitchen_stock`. Тимчасово використайте ключове слово `pass` всередині тіла функції, щоб Python не видавав `SyntaxError`.

Пам’ятайте, що Python покладається на відступи для визначення блоків коду. Тож обов’язково зробіть відступ у тілі функції.

# --hints--

Ви маєте визначити функцію з назвою `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
