---
id: 6a3e6ad6b9554e31dcc30cd6
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Якщо у вас недостатньо яєць, не варто намагатися готувати! Потрібен спосіб зупинити виконання функції далі.

Для цього можна використати ключове слово `return`. Коли функція зустрічає інструкцію `return`, вона негайно припиняє виконання і повертає значення туди, звідки була викликана.

Наприклад:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

Усередині вашого блоку `if`, під інструкцією виведення, використайте ключове слово `return`, щоб повернути значення `available_eggs`.

# --hints--

Ви повинні повернути `available_eggs` всередині вашої інструкції `if`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

Ви повинні повернути `available_eggs` після інструкції `print` всередині вашого блоку `if`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
