---
id: 6a3e6fd08125dfa615cb431b
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер, коли ваша `use_eggs` функція завершена, ви можете її протестувати. Коли ви викликаєте функцію, яка потребує **параметрів**, ви передаєте конкретні значення всередині круглих дужок. Ці значення називаються **аргументами**.

Ось приклад того, як викликати функцію з аргументами:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

У цьому прикладі перші дві лінії — це визначення функції, де `a` та `b` є параметрами. Остання лінія, `calculate_sum(3, 1)`, — це виклик функції. Значення `3` та `1` є аргументами. Вони співставляються з параметрами за позицією, тому `3` призначається `a`, а `1` призначається `b`.

Внизу вашого коду викличте вашу `use_eggs` функцію і передайте їй два аргументи: глобальну змінну `available_eggs` та число `1` (що представляє кількість яєць, які ви хочете використати).

# --hints--

Ви повинні викликати функцію `use_eggs` і передати `available_eggs` та `1` як аргументи.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
