---
id: 6a3e524fc3394ef52ed5b797
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Функції мають свою локальну область видимості. Це означає, що змінні, створені всередині функції, схожі на інгредієнти, які зберігаються в коморі тієї кухні: до них можна отримати доступ лише всередині цієї конкретної функції.

Однак функції не є повністю ізольованими. Вони можуть читати глобальні змінні, визначені поза межами їхньої області видимості. Вони можуть використовувати ці змінні для виконання обчислень, прийняття рішень або створення нових локальних змінних.

Всередині вашої функції `check_kitchen_stock` замініть ключове слово `pass` на нову змінну з назвою `total_items` і призначте їй суму `available_eggs`, `available_flour` та `available_sugar`.

# --hints--

Ви повинні визначити змінну з назвою `total_items` всередині вашої функції `check_kitchen_stock`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

Ваша змінна `total_items` має бути сумою `available_eggs`, `available_flour` та `available_sugar`.

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
