---
id: 6a3e70b94f67d9ee3ecd8486
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Ви щойно викликали `use_eggs(available_eggs, 1)`, щоб використати одне яйце. Це означає, що у вашій кухні тепер має залишитися `0` яєць, так?

Ви дізнаєтеся! Внизу свого коду використайте функцію `print()`, щоб вивести значення змінної `available_eggs`.

Коли ви запустите код, помітите щось дивне: вихідні дані все ще `1`! Чому кількість яєць не змінилася? Ви дослідите цю загадку в наступних кількох кроках.

# --hints--

Ви маєте вивести змінну `available_eggs` в кінці свого коду.

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
