---
id: 6a3e73eabddd614f91cc235a
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Тепер ви розв’яжете загадку з раніше! Коли ваша `use_eggs` функція повертала залишок яєць, ви не зберегли це нове значення ніде. Воно просто зникло в нікуди!

Щоб справді оновити свій інвентар, потрібно захопити це повернене значення, призначивши його змінній.

Всередині вашої `make_fried_egg` функції використайте інструкцію `if/else`, щоб перевірити змінну `has_enough_eggs`.

- Якщо `has_enough_eggs` дорівнює `True`, переназначте змінну `available_eggs` результатом виклику `use_eggs(available_eggs, 1)` і виведіть повідомлення `Made a fried egg. Yummy!`.

- Інакше, у блоці `else`, виведіть повідомлення `Could not make a fried egg. Not enough eggs!`.

# --hints--

Ви повинні мати інструкцію `if` і використовувати `has_enough_eggs` як її умову.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

Всередині блоку `if` ви повинні переназначити `available_eggs` результатом виклику `use_eggs(available_eggs, 1)`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

Ви повинні вивести повідомлення `Made a fried egg. Yummy!`, коли `has_enough_eggs` дорівнює `True`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

Ви повинні мати блок `else`, який виводить повідомлення `Could not make a fried egg. Not enough eggs!`, коли `has_enough_eggs` дорівнює `False`.

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
