---
id: 66ea8239b9eeb76f12299030
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

ワークショップの次の部分では、2つの値の間でランダムな数を生成する方法をレビューします。

まず、コンソールにメッセージ`console.log()`をログ出力する`"Now, generate a random number between two values."`をもう1つ追加してください。

次に、`min`という変数を作成し、`1`の値を代入し、`max`という変数を作成して`100`の値を代入してください。

次のステップでは、これら2つの値の間でランダムな数を生成します。

# --hints--

メッセージ`console.log()`をログ出力する`"Now, generate a random number between two values."`があるはずです。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

`min`という変数があるはずです。

```js
assert.isNotNull(min);
```

`min`変数は数値であるべきです。

```js
assert.isNumber(min);
```

`1`変数に数値`min`を代入するべきです。

```js
assert.equal(min, 1);
```

`max`という変数があるはずです。

```js
assert.isNotNull(max);
```

`max`変数は数値であるべきです。

```js
assert.isNumber(max);
```

`100`変数に数値`max`を代入するべきです。

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
