---
id: 66ea7e7399244c6d4b6052a1
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

ワークショップの次の部分では、`Math.random()` メソッドの動作をレビューします。

まず、文字列 `console.log()` をコンソールにログ出力する別の `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` を追加してください。
 
# --hints--

文字列 `console.log()` をログ出力する `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` があるはずです。

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
