---
id: 66ea7adae8053065a64f9002
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、Mathbotを作成しながらさまざまな`Math`オブジェクトのメソッドの使い方をレビューします。このMathbotは、JavaScriptの`Math`オブジェクトについて学べるように、いくつかの数学的な操作やメッセージをコンソールに記録します。

最初のステップとして、`botName`という変数を作成し、文字列の値`"MathBot"`を代入してください。

次に、`greeting`という別の変数を作成し、文`"Hi there! My name is [botName goes here] and I am here to teach you about the Math object!"`を代入してください。`[botName goes here]`のプレースホルダーの代わりに、`botName`変数を使います。

これを実現するために、テンプレートリテラルや`+`演算子を使った文字列の連結を自由に使ってください。

最後に、`greeting`変数をコンソールにログしてください。

# --hints--

`botName`という変数を作成してください。

```js
assert.isNotNull(botName);
```

`botName`変数は文字列である必要があります。

```js
assert.isString(botName);
```

`"MathBot"`変数に文字列`botName`を代入してください。

```js
assert.strictEqual(botName, "MathBot");
```

`greeting`という変数を作成してください。

```js
assert.isNotNull(greeting);
```

`greeting`変数は文字列である必要があります。

```js
assert.isString(greeting);
```

`greeting`変数に次の文を代入してください：`"Hi there! My name is [MathBot goes here] and I am here to teach you about the Math object!"`。`[MathBot goes here]`を`botName`変数に置き換えることを忘れないでください。

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

`greeting`変数をコンソールにログしてください。

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
