---
id: 66ea7fc048a6e26e33c598fe
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

次は、`Math.random()` メソッドを使ってランダムな数値を生成しましょう。

`randomNum` という変数を作成し、`Math.random()` メソッドを呼び出した結果の値を代入してください。

その後、`randomNum` 変数をコンソールに出力しましょう。

コードにスペースを追加してボットを再実行し、生成される異なるランダムな数値を確認してみてください。

# --hints--

`randomNum` という変数があるはずです。

```js
assert.isNotNull(randomNum);
```

`Math.random()` メソッドを呼び出した結果を `randomNum` 変数に代入してください。

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

`randomNum` 変数にランダムな数値をハードコードしてはいけません。ランダムな数値を生成するには `Math.random()` メソッドを使うことを確認してください。

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

`randomNum` 変数をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
