---
id: 66ea845f54195c7020cc521b
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

2つの値の間でランダムな数を生成するための式は次のとおりです：

```js
Math.random() * (maximum - minimum) + minimum;
```

これは、2つの値の間の浮動小数点数を結果として生成します。

`randomNum2` という変数を作成し、`min` と `max` の値の間で値を生成した結果を代入してください。

その後、結果を確認するために `randomNum2` 変数をコンソールに表示してください。ページを更新すると異なる結果が表示されます。

# --hints--

`randomNum2` という変数を用意してください。

```js
assert.isNotNull(randomNum2);
```

`randomNum2` 変数は数値である必要があります。

```js
assert.isNumber(randomNum2);
```

`min` と `max` の値の間で値を生成した結果を代入してください。値をハードコードしないようにしてください。上記の式を使う必要があります。

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

`randomNum2` 変数をコンソールに表示してください。

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
