---
id: 66eb3ec3b07dfe8f898646c0
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

前のレッスンで学んだように、`Math.floor()` メソッドは値を最も近い整数に切り捨てます。

```js
const price = 10.99;
Math.floor(price); // 10
```

`numRoundedDown` の浮動小数点数を最も近い整数に切り捨てた結果を `6.7` という変数に代入してください。

その後、結果を確認するために `numRoundedDown` 変数をコンソールに出力してください。

# --hints--

`numRoundedDown` という変数を用意してください。

```js
assert.isNotNull(numRoundedDown);
```

`numRoundedDown` 変数は数値である必要があります。

```js
assert.isNumber(numRoundedDown);
```

`numRoundedDown` 変数は `6.7` の浮動小数点数を最も近い整数に切り捨てた結果と等しくなるはずです。

```js
assert.equal(numRoundedDown, 6);
```

`6` の値を `numRoundedDown` 変数に直接書き込まないでください。

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

`numRoundedDown` 変数をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
