---
id: 66eb488c0129e192542f0865
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

前のレッスンでは、`Math.ceil()` メソッドの使い方を学びました。

```js
const price = 10.01;
Math.ceil(price); // 11
```

`Math.ceil()` メソッドは値を最も近い整数に切り上げます。

`numRoundedUp` という変数を作成し、浮動小数点数 `3.2` を最も近い整数に切り上げた結果を代入してください。

その後、結果を確認するために `numRoundedUp` 変数をコンソールに出力してください。

# --hints--

`numRoundedUp` という変数があるはずです。

```js
assert.isNotNull(numRoundedUp);
```

`numRoundedUp` 変数は数値であるべきです。

```js
assert.isNumber(numRoundedUp);
```

`numRoundedUp` 変数は浮動小数点数 `3.2` を最も近い整数に切り捨てた結果と等しいはずです。

```js
assert.equal(numRoundedUp, 4);
```

`4` の値を `numRoundedUp` 変数にハードコードしてはいけません。

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

`numRoundedUp` 変数をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
