---
id: 66ece147323c91f654bf8dd9
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

前のレッスンでは、次のように数値の範囲から最大値と最小値を取得する方法を学びました。

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

`maxNum` という変数を作成し、`3, 125, 55, 24` の数値の中から最大の数を見つけた結果を代入してください。その後、`maxNum` の値をコンソールに出力してください。

その下に、`minNum` という変数を作成し、`6, 90, 14, 90, 2` の数値の中から最小の数を見つけた結果を代入してください。その後、`minNum` の値をコンソールに出力してください。

# --hints--

`maxNum` という変数を作成しているはずです。

```js
assert.isNotNull(maxNum);
```

`maxNum` 変数には、`3, 125, 55, 24` の数値の中から最大の数を見つけた結果を代入してください。

```js
assert.equal(maxNum, 125);
```

`125` の値を `maxNum` 変数に直接書き込まないでください。代わりに `Math.max()` メソッドを使うことを確認してください。

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

`maxNum` の値をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

`minNum` という変数を作成しているはずです。

```js
assert.isNotNull(minNum);
```

`minNum` 変数には、`6, 90, 14, 90, 2` の数値の中から最小の数を見つけた結果を代入してください。

```js
assert.equal(minNum, 2);
```

`2` の値を `minNum` 変数に直接書き込まないでください。代わりに `Math.min()` メソッドを使うことを確認してください。

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

`minNum` の値をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
