---
id: 69218e46f4490b33f2481a1a
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

2つの値の間でランダムな整数を生成するための式は次の通りです。

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

これは、2つの値の間の整数を生成します。

`randomInt` という変数を作成し、`min` と `max` の値の間で値を生成した結果を代入してください。

その後、結果を確認するために `randomInt` 変数をコンソールに出力してください。ページを更新すると異なる結果が表示されます。

# --hints--

`randomInt` という変数を作成しているはずです。

```js
assert.isNotNull(randomInt);
```

`randomInt` 変数は数値である必要があります。

```js
assert.isNumber(randomInt);
```

`min` と `max` の値の間で値を生成した結果を代入してください。値をハードコードしないようにしてください。上記の式を使う必要があります。

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

`randomInt` 変数をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
