---
id: 69218a15fe536ee1be38aefb
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

ワークショップの次の部分では、2つの値の間でランダムな整数を生成する方法をレビューします。

まず、`console.log()`というメッセージをコンソールにログ出力する別の`"Now, generate a random integer between two values."`を追加してください。

次のステップでは、すでに作成した`min`と`max`変数の値の間でランダムな整数を生成します。

# --hints--

`console.log()`というメッセージをログ出力する`"Now, generate a random integer between two values."`があるはずです。

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
