---
id: 5f33294a6af5e9188dbdb8f3
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este taller, practicarás los conceptos básicos de CSS (Cascading Style Sheets) construyendo un menú de café.

Let's start by adding some menu content. Añade un elemento `main` dentro del elemento `body` existente. Contendrá información sobre los precios de los cafés y postres ofrecidos en la cafetería.

# --hints--

Tu código debe tener una etiqueta de apertura `<main>`.

```js
assert.match(code, /<main>/i);
```

Tu código debe tener una etiqueta de cierre `</main>`.

```js
assert.match(code, /<\/main>/i);
```

No debes cambiar el elemento `body`. Asegúrate de que no has borrado accidentalmente la etiqueta de cierre.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

La etiqueta `main` debe estar dentro de la etiqueta `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
