---
id: 5f344fad8bf01691e71a30eb
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Hasta ahora, has tenido un control limitado sobre la presentaci칩n y apariencia de tu contenido. Para cambiar eso, a침ade un elemento `style` dentro del elemento `head`.

# --hints--

Tu c칩digo debe tener una etiqueta de apertura `<style>`.

```js
assert.match(code, /<style\s*>/i);
```

Tu c칩digo debe tener una etiqueta de cierre `</style>`.

```js
assert.match(code, /<\/style\s*>/);
```

Tu elemento `style` debe estar anidado dentro del elemento `head`.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
