---
id: 5f344fbc22624a2976425065
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Crea un elemento `h2` en el elemento `section` y dale el texto `Coffee`.

# --hints--

Debes tener una etiqueta de apertura `<h2>`.

```js
assert.match(code, /<h2\s*>/i);
```

Debes tener una etiqueta de cierre `</h2>`.

```js
assert.match(code, /<\/h2\s*>/i);
```

No debes modificar el elemento `section` existente. Asegúrate de que no borras la etiqueta de cierre.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Tu elemento `h2` debe estar anidado dentro de tu elemento `section`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

Tu elemento `h2` debe tener el texto `Coffee`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
