---
id: 5f3477aefa51bfc29327200b
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Has dado estilo a tres elementos escribiendo CSS dentro de la etiqueta `style`. Esto funciona, pero dado que habrá muchos más estilos, es mejor poner todos los estilos dentro de un archivo separado y enlazarlo.

Se ha creado un archivo `styles.css` separado para tí. Puedes cambiar entre archivos con las pestañas en la parte superior del editor.

Comienza reescribiendo, los estilos que ya has creado, en el archivo `styles.css`. Asegurate de excluir las etiquetas `style` de apertura y cierre.

# --hints--

Tu archivo `styles.css` debe tener los selectores de tipo `h1, h2, p`.

```js
assert.match(editableContents.replace(/[\s\n]*/g, ""), /(h1|h2|p),(h1|h2|p),(h1|h2|p){/);
```

Tu selector debe establecer la propiedad `text-align` al valor `center`.

```js
assert.match(editableContents, /text-align:\s*center;?/);
```

Tu código no debe contener las etiquetas `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
