---
id: 5f356ed60785e1f3e9850b6e
title: Paso 20
challengeType: 0
dashedName: step-20
---

# --description--

Ahora es fácil ver que el texto está centrado dentro del elemento `#menu`. Por ahora, el ancho (width) del elemento `#menu` está en píxeles (`px`). 

Cambia el valor de la propiedad `width` a `80%`, para que sea el `80%` del ancho de su elemento padre (`body`).

# --hints--

La propiedad `width` debe tener el valor `80%`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert.isTrue(hasWidth);
```

No debes tener ni una propiedad `width` con el valor `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isFalse(hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
--fcc-editable-region--
  width: 300px;
--fcc-editable-region--
  background-color: burlywood;
}
```

