---
id: 5f3ef6e0a81099d9a697b550
title: Paso 61
challengeType: 0
dashedName: step-61
---

# --description--

Dentro de `address`, añade un elemento `p`. Después, anida un elemento anchor (`a`) en el elemento `p` que enlace a `https://www.freecodecamp.org` y tenga el texto `Visit our website`.

Asegúrate de que el enlace se abre en una nueva pestaña agregando un atributo `target` con el valor `_blank`.

# --hints--

No debes modificar el elemento `footer` existente.

```js
assert.lengthOf(document.querySelectorAll("footer"), 1);
```

No debes modificar el elemento `address` existente.

```js
assert.lengthOf(document.querySelectorAll("address"), 1);
```

Tu nuevo elemento `p` debería estar anidado dentro de tu elemento `address`. Debes tener exactamente un elemento `p`.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p"), 1);
assert.lengthOf(document.querySelectorAll("footer p"), 1);
```

Tu nuevo elemento `a` debe ser anidado dentro de tu nuevo elemento `p`. Debes tener exactamente un elemento `a`.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p > a"), 1);
assert.lengthOf(document.querySelectorAll("footer a"), 1);
```

Tu nuevo elemento `a` debe tener el texto `Visit our website`.

```js
assert.equal(document.querySelector("footer > address > p > a")?.innerText.trim(), "Visit our website");
```

Tu nuevo elemento `a` te debe enlazar a `https://www.freecodecamp.org`. Recuerda que los elementos `a` utilizan el atributo `href` para crear un enlace.

```js
assert.equal(document.querySelector("footer > address > p > a")?.href, "https://www.freecodecamp.org/");
```

Tu nuevo elemento `a` debe tener el atributo `target` establecido en `_blank`.

```js
assert.equal(document.querySelector("footer > address > p > a")?.target, "_blank");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <footer>
        <address>
--fcc-editable-region--
          
--fcc-editable-region--
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

