---
id: 61969aa6acef5b12200f672e
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

لإنشاء بعض المناظر في الخلفية، ستضيف جبلين.

فوق العنصر `.penguin`، أضف `div` مع `class` اسمه `left-mountain`.

# --hints--

يجب إضافة `div` جديد داخل `body`. كان من المتوقع رؤية عناصر `div` --fcc-expected--، لكن وُجد --fcc-actual--.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 3);
```

يجب إعطاء `div` الخاصية `class` بقيمة `left-mountain`.

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

يجب وضع `.left-mountain` قبل `.penguin`.

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
--fcc-editable-region--
    
--fcc-editable-region--
    <div class="penguin"></div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}

```
