---
id: 61967e74a8e3690ab6292daa
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

استهدف عنصر `body` لتعيين `background` إلى تدرج لوني خطي بزاوية 45 درجة باتجاه عقارب الساعة، يبدأ من `rgb(118, 201, 255)` وينتهي عند `rgb(247, 255, 222)`.

# --hints--

يجب أن تستخدم محدِّد العنصر `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب استخدام خاصية `background` في محدِّد `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

يجب تعيين `background` إلى `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
