---
id: 619665c9abd72906f3ad30f9
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ستبني بطريق Flappy سعيد، وستتمرن أكثر على تحويلات CSS والرسوم المتحركة في هذه العملية.

ابدأ بربط ورقة الأنماط بالصفحة.

# --hints--

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert.exists(document.querySelector('link'));
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
