---
id: 6196ce0415498d2463989e84
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

لإعطاء تأثير سلسلة جبال، أضف جبلًا آخر بإنشاء `div` جديد مباشرة بعد `.left-mountain`، وأعطِ الـ `div` الجديد الـ `class` الخاص بـ `back-mountain`.

# --hints--

يجب إضافة `div` جديد داخل `body`.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 4);
```

يجب إعطاء الـ `div` الـ `class` الخاص بـ `back-mountain`.

```js
assert.include(document.querySelector('div:not(.left-mountain, .ground, .penguin)')?.className, 'back-mountain');
```

يجب وضع `.back-mountain` بعد `.left-mountain`.

```js
assert.strictEqual(document.querySelector('.left-mountain')?.nextElementSibling, document.querySelector('.back-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
--fcc-editable-region--
    
--fcc-editable-region--
    <div class="penguin"></div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}

```
