---
id: 587d8259367417b2b2512c85
title: Tekeleza kupanga kwa kuchagua
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

Hapa tuta-tekeleza kupanga kwa kuchagua. Kupanga kwa kuchagua hufanya kazi kwa kuchagua thamani ya kiwango cha chini katika orodha na kubadilisha nafasi yake na thamani ya kwanza katika orodha. Kisha huanza kwenye nafasi ya pili, huchagua thamani ndogo zaidi katika orodha iliyobaki, na kubadilisha nafasi yake na kipengele cha pili. Huu huendelea kurudia kupitia orodha na kubadilisha vipengele hadi ifikie mwisho wa orodha. Sasa orodha imepangwa. Kupanga kwa kuchagua kuna ugumu wa wakati wa mduara wa mraba katika kesi zote.

**Maelekezo**: Andika kitendakazi `selectionSort` kinachopokea safu ya nambari kamili kama ingizo na kurudisha safu ya nambari hizi zikiwa zimepangwa kutoka kiwango cha chini hadi kikubwa zaidi.

# --hints--

`selectionSort` inapaswa kuwa kitendakazi.

```js
assert.isFunction(selectionSort);
```

`selectionSort` inapaswa kurudisha safu iliyopangwa (kutoka kiwango cha chini hadi kikubwa zaidi).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` inapaswa kurudisha safu ambayo haijabadilika isipokuwa kwa mpangilio.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` haipaswi kutumia njia iliyo-jengwa ndani ya `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
