---
id: 61abc7ebf3029b56226de5b6
title: Tekeleza utafutaji wa binary
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

Utafutaji wa binary ni algoriti yenye ufanisi wa **O(log(n))** kwa utafuta katika safu ya data iliyopangwa ili kupata kipengele. Inafanya kazi kwa kutumia hatua zifuatazo:

1. Tafuta `value` ya katikati ya safu ya data iliyopangwa. Ikiwa `value == target` rudisha `true` (Thamani imepatikana na utafuta umekamilika).
1. Ikiwa katikati `value < target`, tafuta nusu ya kulia ya safu katika kulinganisha lijalo.
1. Ikiwa katikati `value > target`, tafuta nusu ya kushoto ya safu katika kulinganisha lijalo.
1. Ikiwa baada ya kutafuta katika safu yote thamani haipo, rudisha `false` (Safu imefutwa na thamani haipo katika safu).

Kama unavyoona, unagawanya safu ya data mara kwa mara, jambo linalokupa ufanisi wa log(n). Kwa changamoto hii, tunataka uonyeshe kazi yako - jinsi ulivyofikia thamani lengwa... njia uliyotumia!

# --instructions--

Andika kitendakazi `binarySearch` kinachotekeleza algoriti ya utafutaji wa binary kwenye safu ya data, kikirejesha njia uliyotumia (kila kulinganisha thamani ya katikati) kupata lengwa katika safu.

Kitendakazi kinapokea safu ya data iliyopangwa ya nambari kamili na thamani lengwa kama ingizo. Kinarejesha safu ya data yenye (kwa mpangilio) thamani ya katikati uliyopata kila ugawaji wa safu ya awali hadi upate thamani lengwa. Thamani lengwa inapaswa kuwa kipengele cha mwisho cha safu inayorejeshwa. Ikiwa thamani haipatikani, rudisha mfuatano wa herufi `Value Not Found`.

Kwa mfano, `binarySearch([1,2,3,4,5,6,7], 5)` itarejesha `[4,6,5]`.

Kwa changamoto hii, wakati wa kugawanya, LAZIMA utumie `Math.floor()` wakati wa kugawanya: `Math.floor(x/2)`. Hii itatoa njia thabiti inayoweza kujaribiwa.

**Kumbuka:** Safu ifuatayo itatumika katika mitihani:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch` inapaswa kuwa kitendakazi.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` inapaswa kurudisha `[13, 5, 2, 0]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` inapaswa kurudisha `[13, 5, 2, 0, 1]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` inapaswa kurudisha `[13, 5, 2]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` inapaswa kurudisha mfuatano wa herufi `Value Not Found`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` inapaswa kurudisha `[13, 5, 10, 11]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` inapaswa kurudisha `[13]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` inapaswa kurudisha `[13, 19, 22, 49, 70]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
