---
id: a3f503de51cfab748ff001aa
title: Wapare wa vipande viwili
challengeType: 1
forumTopicId: 301617
dashedName: pairwise
---

# --description--

Ikiwa umepewa safu ya data `arr`, tafuta wapare wa vipande viwili ambavyo jumla yao ni sawa na hoja ya pili `arg` na rudisha jumla ya vielezo vyao.

Unaweza kutumia wapare wengi wenye vipande vya nambari sawa lakini vikiwa na vielezo tofauti. Kila pande inapaswa kutumia vielezo vidogo vinavyopatikana. Mara kipande kimechaguliwa hakiwezi kutumika tena kuunganishwa na kipande kingine. Kwa mfano, `pairwise([1, 1, 2], 3)` huunda pande `[2, 1]` kwa kutumia 1 iliyo kwenye kielezo 0 badala ya 1 iliyo kwenye kielezo 1, kwa sababu 0+2 &lt; 1+2.

Kwa mfano `pairwise([7, 9, 11, 13, 15], 20)` hurudisha `6`. Wapare ambao jumla yao ni 20 ni `[7, 13]` na `[9, 11]`. Tunaweza kisha kuandika safu ya data na vielezo na thamani zao.

<div style='margin-left: 2em;'>

| Kielezo | 0 | 1 | 2  | 3  | 4  |
| ------- | - | - | -- | -- | -- |
| Thamani | 7 | 9 | 11 | 13 | 15 |

</div>

Hapa chini tutachukua vielezo vinavyolingana na kuviunganisha.

<div style='margin-left: 2em;'>

7 + 13 = 20 → Vielezo 0 + 3 = 3  
9 + 11 = 20 → Vielezo 1 + 2 = 3  
3 + 3 = 6 → Rudisha `6`

</div>

# --hints--

`pairwise([1, 4, 2, 3, 0, 5], 7)` inapaswa kurudisha 11.

```js
assert.deepEqual(pairwise([1, 4, 2, 3, 0, 5], 7), 11);
```

`pairwise([1, 3, 2, 4], 4)` inapaswa kurudisha 1.

```js
assert.deepEqual(pairwise([1, 3, 2, 4], 4), 1);
```

`pairwise([1, 1, 1], 2)` inapaswa kurudisha 1.

```js
assert.deepEqual(pairwise([1, 1, 1], 2), 1);
```

`pairwise([0, 0, 0, 0, 1, 1], 1)` inapaswa kurudisha 10.

```js
assert.deepEqual(pairwise([0, 0, 0, 0, 1, 1], 1), 10);
```

`pairwise([], 100)` inapaswa kurudisha 0.

```js
assert.deepEqual(pairwise([], 100), 0);
```

# --seed--

## --seed-contents--

```js
function pairwise(arr, arg) {
  return arg;
}

pairwise([1,4,2,3,0,5], 7);
```

# --solutions--

```js
function pairwise(arr, arg) {
  var sum = 0;
  arr.forEach(function(e, i, a) {
    if (e != null) {
      var diff = arg-e;
      a[i] = null;
      var dix = a.indexOf(diff);
      if (dix !== -1) {
        sum += dix;
        sum += i;
        a[dix] = null;
      }
    }
  });
  return sum;
}

pairwise([1,4,2,3,0,5], 7);
```
