---
id: a7bf700cd123b9a54eef01d5
title: Tafadhali usirudie
challengeType: 1
forumTopicId: 16037
dashedName: no-repeats-please
---

# --description--

Rudisha idadi ya mabadiliko yote ya mfuatano wa herufi uliotolewa ambayo hayana herufi zinazojirudia mfululizo. Kubali kuwa herufi zote katika mfuatano uliotolewa ni za kipekee kila moja.

Kwa mfano, `aab` inapaswa kurudisha 2 kwa sababu ina mabadiliko 6 yote (`aab`, `aab`, `aba`, `aba`, `baa`, `baa`), lakini ni 2 tu kati yao (`aba` na `aba`) ambazo hazina herufi ile ile (katika kesi hii `a`) inayojirudia.

# --hints--

`permAlone("aab")` inapaswa kurudisha nambari.

```js
assert.isNumber(permAlone('aab'));
```

`permAlone("aab")` inapaswa kurudisha 2.

```js
assert.strictEqual(permAlone('aab'), 2);
```

`permAlone("aaa")` inapaswa kurudisha 0.

```js
assert.strictEqual(permAlone('aaa'), 0);
```

`permAlone("aabb")` inapaswa kurudisha 8.

```js
assert.strictEqual(permAlone('aabb'), 8);
```

`permAlone("abcdefa")` inapaswa kurudisha 3600.

```js
assert.strictEqual(permAlone('abcdefa'), 3600);
```

`permAlone("abfdefa")` inapaswa kurudisha 2640.

```js
assert.strictEqual(permAlone('abfdefa'), 2640);
```

`permAlone("zzzzzzzz")` inapaswa kurudisha 0.

```js
assert.strictEqual(permAlone('zzzzzzzz'), 0);
```

`permAlone("a")` inapaswa kurudisha 1.

```js
assert.strictEqual(permAlone('a'), 1);
```

`permAlone("aaab")` inapaswa kurudisha 0.

```js
assert.strictEqual(permAlone('aaab'), 0);
```

`permAlone("aaabb")` inapaswa kurudisha 12.

```js
assert.strictEqual(permAlone('aaabb'), 12);
```

# --seed--

## --seed-contents--

```js
function permAlone(str) {
  return str;
}

permAlone('aab');
```

# --solutions--

```js
function permAlone(str) {
  return permuter(str).filter(function(perm) {
    return !perm.match(/(.)\1/g);
  }).length;
}

function permuter(str) {
  // http://staff.roguecc.edu/JMiller/JavaScript/permute.html
  //permArr: Global array which holds the list of permutations
  //usedChars: Global utility array which holds a list of "currently-in-use" characters
  var permArr = [], usedChars = [];
  function permute(input) {
    //convert input into a char array (one element for each character)
    var i, ch, chars = input.split("");
    for (i = 0; i < chars.length; i++) {
      //get and remove character at index "i" from char array
      ch = chars.splice(i, 1);
      //add removed character to the end of used characters
      usedChars.push(ch);
      //when there are no more characters left in char array to add, add used chars to list of permutations
      if (chars.length === 0) permArr[permArr.length] = usedChars.join("");
      //send characters (minus the removed one from above) from char array to be permuted
      permute(chars.join(""));
      //add removed character back into char array in original position
      chars.splice(i, 0, ch);
      //remove the last character used off the end of used characters array
      usedChars.pop();
    }
  }
  permute(str);
  return permArr;
}

permAlone('aab');
```
