---
id: 66bb6a9c2dd58b73cd759034
title: "Crea un'app per il registro dei voti"
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `getAverage` che prende un array di punteggi dei test come parametro e restituisce la media dei punteggi.
1. Dovresti avere una funzione chiamata `getGrade` che prende un punteggio di uno studente come parametro e restituisce una stringa che rappresenta un voto in lettere basato sul punteggio. Ecco i punteggi e i corrispondenti voti in lettere:

    | Intervallo punteggi | Voto   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. Dovresti avere una funzione chiamata `hasPassingGrade` che prende un punteggio come parametro e restituisce `true` o `false` a seconda che il punteggio corrisponda a un voto sufficiente.
1. La funzione `hasPassingGrade` dovrebbe usare la funzione `getGrade` per ottenere il voto in lettere e usarlo per determinare se il voto è sufficiente. Un voto sufficiente è qualsiasi voto diverso da `"F"`.
1. Dovresti avere una funzione chiamata `studentMsg` che prende un array di punteggi e un punteggio di uno studente come parametri. La funzione dovrebbe restituire una stringa con il formato:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, se lo studente ha superato il corso.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, se lo studente ha fallito il corso.

    Sostituisci `average-goes-here` con la media totale dei punteggi e `grade-goes-here` con il voto dello studente.
    Usa `getAverage` per ottenere la media dei punteggi e `getGrade` per ottenere il voto dello studente.

# --hints--

Dovresti avere una funzione chiamata `getAverage`.

```js
assert.isFunction(getAverage);
```

La tua funzione `getAverage` dovrebbe restituire un numero.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` dovrebbe restituire `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` dovrebbe restituire `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` dovrebbe restituire `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` dovrebbe restituire `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` dovrebbe restituire `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

La tua funzione `getAverage` dovrebbe restituire la media dei punteggi dei test.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Dovresti avere una funzione chiamata `getGrade`.

```js
assert.isFunction(getGrade);
```

La tua funzione `getGrade` dovrebbe restituire `"A+"` se il punteggio è `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

La tua funzione `getGrade` dovrebbe restituire `"A"` se il punteggio è tra `90` e `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

La tua funzione `getGrade` dovrebbe restituire `"B"` se il punteggio è tra `80` e `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

La tua funzione `getGrade` dovrebbe restituire `"C"` se il punteggio è tra `70` e `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

La tua funzione `getGrade` dovrebbe restituire `"D"` se il punteggio è tra `60` e `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

La tua funzione `getGrade` dovrebbe restituire `"F"` se il punteggio è tra `0` e `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Dovresti avere una funzione chiamata `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

La tua funzione `hasPassingGrade` dovrebbe restituire un valore booleano.

```js
assert.isBoolean(hasPassingGrade(100));
```

La tua funzione `hasPassingGrade` dovrebbe restituire `true` se il voto è un `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

La tua funzione `hasPassingGrade` dovrebbe restituire `false` se il voto è un `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

La tua funzione `hasPassingGrade` dovrebbe restituire `true` per tutti i voti sufficienti.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Dovresti avere una funzione chiamata `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` dovrebbe restituire il seguente messaggio: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` dovrebbe restituire il seguente messaggio: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` dovrebbe restituire il seguente messaggio: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` dovrebbe restituire il seguente messaggio: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

La tua funzione `studentMsg` dovrebbe restituire il messaggio corretto basato sul punteggio dello studente e sulla media della classe.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
