---
id: 68ce56d54d3cad1149a1c4cf
title: 기능 선택 페이지 디자인하기
challengeType: 25
dashedName: design-a-feature-selection-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 페이지에 `h1` 요소가 있고 텍스트는 `Feature Selection`이어야 합니다.
2. 페이지에 `div` 요소가 있고 `feature-card-container` 클래스를 가져야 합니다.
3. `label` 안에 `feature-card` 클래스를 가진 `div` 요소가 최소 두 개 있어야 합니다.
4. 각 `label` 요소에는 레이블 텍스트와 `input` 타입의 `checkbox` 요소가 있어야 합니다.
5. 체크박스를 대상으로 하는 선택자를 만들고 다음 스타일을 적용하세요:
    - 모든 체크박스 요소는 CSS에서 `appearance: none;`로 설정되어야 합니다.
    - 모든 체크박스 요소는 원하는 테두리 너비, 색상, 스타일을 가져야 합니다.
6. 체크박스가 선택되면 체크박스 중앙에 체크 표시 `✓`가 보여야 합니다.
7. 체크박스가 선택되면 체크박스 배경색이 원하는 색으로 바뀌어야 합니다.
8. 체크박스가 선택되면 체크박스 테두리 색상이 원하는 색으로 바뀌어야 합니다.

**참고:** CSS 파일을 HTML에 반드시 연결하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1`에는 `Feature Selection` 텍스트가 있어야 합니다.

```js
assert.equal(document.querySelector('h1').innerText.trim(), `Feature Selection`);
```

`div` 요소가 있고 `feature-card-container` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector('div.feature-card-container'));
```

`label` 안에 최소 두 개의 `.feature-card-container` 요소가 있어야 합니다.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isAtLeast(labels.length, 2);
```

`label` 안의 각 `.feature-card-container` 요소는 `feature-card` 클래스를 가져야 합니다.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isTrue(label.classList.contains('feature-card')));
```

각 `label` 요소 안에는 레이블 텍스트가 있어야 합니다.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.isNotEmpty(label.innerText.trim()));
```

각 레이블 요소 안에는 `input` 타입의 `checkbox` 요소가 하나 있어야 합니다.

```js
const labels = document.querySelectorAll('.feature-card-container label');
assert.isNotEmpty(labels);

labels.forEach(label => assert.exists(label.querySelector('input[type="checkbox"]')));
```

체크박스를 위한 선택자가 있어야 합니다.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

체크박스는 `appearance: none`로 설정되어야 합니다.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.equal(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('appearance'), 'none');
```

체크박스는 원하는 테두리 너비를 가져야 합니다.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-width'));
```

체크박스는 원하는 테두리 색상을 가져야 합니다.

```js
const selectors = [
  '.feature-card input[type="checkbox"]',
  '.feature-card [type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input',
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-color'));
```

체크박스는 원하는 테두리 스타일을 가져야 합니다.

```js
const selectors = [  
  '.feature-card [type="checkbox"]',
  '.feature-card input[type="checkbox"]',
  'input[type="checkbox"]',
  '[type="checkbox"]',
  '.feature-card input', // there aren't other input elements, so this is a valid approach
  'input'
]
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyleAny(selectors).getPropVal('border-style'));
```

체크박스가 선택되면 체크박스 중앙에 체크 표시 `✓`가 있어야 합니다.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;

    checkbox.checked = true;
    const after = window.getComputedStyle(checkbox, ':after');
    assert.equal(after.content, '"✓"');
});
```

체크박스가 선택되면 배경색이 원하는 색으로 바뀌어야 합니다.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBG = window.getComputedStyle(checkbox).backgroundColor;
    checkbox.checked = true;
    const checkedBG = window.getComputedStyle(checkbox).backgroundColor;

    assert.notEqual(checkedBG, uncheckedBG);
});
```

체크박스가 선택되면 테두리 색상이 원하는 색으로 바뀌어야 합니다.

```js
const checkboxes = document.querySelectorAll('.feature-card input[type="checkbox"]');
assert.isNotEmpty(checkboxes);

checkboxes.forEach((checkbox) => { 
    checkbox.checked = false;
    const uncheckedBC = window.getComputedStyle(checkbox).borderColor;
    checkbox.checked = true;
    const checkedBC = window.getComputedStyle(checkbox).borderColor;

    assert.notEqual(checkedBC, uncheckedBC);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Selection Feature Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Feature Selection</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Feature Selection</h1>
    <div class="feature-card-container">
      <label class="feature-card">
        <input type="checkbox" />
        <h4>Cloud Storage</h4>
        <p>100 Gigabyte secure storage</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Dedicated Support</h4>
        <p>24/7 customer help</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Advanced Analytics</h4>
        <p>Insights & reports</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Custom User Themes</h4>
        <p>Personalized dashboard design</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>Multi-User Collab</h4>
        <p>Team access and sharing</p>
      </label>

      <label class="feature-card">
        <input type="checkbox" />
        <h4>API Access</h4>
        <p>Integrate with your custom built tools</p>
      </label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  flex-direction: column;
  height: 100vh;
}

.feature-card-container {
  display: flex;
  flex-wrap: wrap;
  gap: 20px;
  max-width: 600px;
  justify-content: center;
}

.feature-card-container > * {
  flex: 1 1 180px;
}

.feature-card {
  border: 2px solid #ccc;
  border-radius: 12px;
  padding: 20px;
  text-align: center;
  position: relative;
  cursor: pointer;
}

.feature-card:hover {
  border-color: #f1be32;
  box-shadow: 0 4px 12px rgba(255, 201, 53, 0.6);
}

.feature-card input[type='checkbox'] {
  position: absolute;
  top: 5px;
  right: 5px;
  width: 24px;
  height: 24px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #ccc;
  border-radius: 4px;
  background-color: white;
}

.feature-card input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #f1be32;
}

.feature-card input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}

.feature-card h4 {
  margin: 10px 0 5px;
  font-size: 1.2rem;
}

@media (max-width: 600px) {
  body {
    height: auto;
  }
}
```
