---
id: 6564a767a7241362f7d8d664
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

독립된 단일 밑줄은 신경 쓰지 않거나 코드에서 사용하지 않을 값을 나타내는 데 사용됩니다.
반복 변수는 실제로 사용되지 않습니다.

`i` 변수를 단일 밑줄로 수정하세요.

# --hints--

`i` 변수를 `_`로 수정해야 합니다.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/for\s+_\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for i in range(length):
        password += secrets.choice(all_characters)
    
--fcc-editable-region--
```
