---
id: 65649122c7f77f519aaf0975
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

마지막으로, 변수 `new_password`를 선언하고 `generate_password`을 호출한 결과를 할당하세요. `8` 호출에 `generate_password`를 인수로 전달하세요.

# --hints--

`generate_password`을 호출할 때 `8`를 인수로 전달해야 합니다.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

변수 `generate_password(8)`에 `new_password`를 할당해야 합니다.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
