---
id: 6564cf2e8642517abdf9d6e2
title: 37단계
challengeType: 20
dashedName: step-37
---

# --description--

이제 제약 튜플의 패턴을 수정하여 대괄호를 사용해 모든 숫자의 범위를 나타내세요.

# --hints--

`(nums, '[0-9]')` 목록에 `constraints` 튜플이 있어야 합니다.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\[0-9\]\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '[0123456789]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = '[a-z]t'
quote = 'Not all those who wander are lost.'
#print(re.findall(pattern, quote))

```
