---
id: 6564d27c8e4c197d367c3b1c
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

`constraints` 리스트에 새 튜플을 추가하세요. 첫 번째 항목으로 `lowercase`을 사용하고, 두 번째 항목으로 소문자 한 글자와 일치하는 정규식 패턴을 사용하세요.

# --hints--

`constraints` 리스트에 두 번째 튜플을 추가해야 합니다. 첫 번째 항목으로 `lowercase` 매개변수를 사용하고, 두 번째 항목으로 `[a-z]`를 사용하세요.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\[0-9\]\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*("|')\[a-z\]\2\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '[0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
