---
id: 6564683821b2ee3174e7250e
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

가져온 모듈 안에 정의된 유틸리티는 점 표기법을 통해 접근할 수 있습니다. 점 표기법은 모듈 이름 뒤에 점과 유틸리티 이름을 붙이는 방식으로 작동합니다. 예를 들어, `ascii_lowercase` 상수에 접근하는 방법은 다음과 같습니다:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

이 프로젝트에서는 `string` 모듈의 다양한 상수를 사용할 것입니다.

새 변수 `letters`를 선언하고 `string.ascii_letters`를 이 변수에 할당하세요.

# --hints--

`letters`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

`string.ascii_letters`를 `letters` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
