---
id: 6564d75a923d21815caaa445
title: 43단계
challengeType: 20
dashedName: step-43
---

# --description--

패턴에서 `.` 또는 `+`처럼 특별한 의미를 가진 문자를 일치시켜야 할 경우, 역슬래시 문자 `\`를 앞에 붙여서 이스케이프할 수 있습니다. 예를 들어, 이것은 문자 그대로의 더하기 기호를 일치시킵니다: `\+`.

`pattern`를 수정해서 문자 그대로의 점 하나와 일치하도록 하세요.

# --hints--

`pattern` 변수를 `\.`로 수정해야 합니다.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
