---
id: 6564f32b18480893cf7799fd
title: 52단계
challengeType: 20
dashedName: step-52
---

# --description--

이제 원시 문자열과 f-문자열을 결합하고 `symbols` 변수를 문자 클래스 안에 보간하세요. f-문자열 내에서 변수를 중괄호 `{ }`를 사용해 보간할 수 있다는 점을 기억하세요.

# --hints--

네 번째 제약 조건 튜플의 두 번째 항목은 문자열 `fr'[{symbols}]'`여야 합니다.

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
