---
id: 656485a8a3496d4a36b1496a
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이 세 변수는 비밀번호를 생성할 때 선택할 수 있는 가능한 문자들을 구성합니다.

그 바로 위에 `Define the possible characters for the password`라는 주석을 추가하세요.

# --hints--

주석은 `letters` 변수 바로 위에 추가해야 합니다.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
