---
id: 656475bbf1c2573de1d2c69c
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

코드가 실행될 때마다 `all_characters` 문자열에서 무작위 문자가 출력되어야 합니다. 이것이 바로 무작위 비밀번호를 생성하기 위해 달성하려는 목표입니다.

하지만 `random`에 의존하는 알고리즘은 생성된 의사 난수들을 예측 가능하게 만듭니다. 따라서 `random` 모듈은 가장 일반적인 애플리케이션에는 적합하지만, 결정론적 특성 때문에 암호화 목적에는 사용할 수 없습니다.

`random`를 가져오는 대신 `secrets` 모듈을 가져오세요. 그리고 `print()` 호출을 `secrets.choice(all_characters)`를 사용하도록 변경하세요.

# --hints--

`secrets` 모듈 대신 `random` 모듈을 가져와야 합니다.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

기존 `print()` 호출을 수정하여 `secrets.choice(all_characters)`을 출력하도록 해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
