---
id: 6564d68c34027a8072a704f4
title: 42단계
challengeType: 20
dashedName: step-42
---

# --description--

점 문자(dot)는 문자열에서 모든 문자와 일치하는 와일드카드입니다 — 기본적으로 줄바꿈 문자는 제외합니다. 현재 패턴을 `pattern` 뒤에 `.` 수량자를 붙인 것으로 바꿔 `+`가 전체 문자열과 일치하도록 수정하세요.

# --hints--

`pattern` 변수를 전체 `quote` 문자열과 일치하도록 수정하세요.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), ".+") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[^a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
