---
id: 6564dd65c3c2fa873a83d213
title: 47단계
challengeType: 20
dashedName: step-47
---

# --description--

문자 클래스에서는 대괄호 안에 여러 범위를 추가 문자 없이 연달아 작성하여 여러 범위를 결합할 수 있습니다. 예를 들어 `[a-d3-6]`는 `[a-d]`과 `[3-6]`의 결합입니다.

이제 마지막 정규식 패턴을 수정하여 영숫자가 아닌 모든 문자를 매치하세요. `a-z`, `A-Z`, `0-9` 범위를 하나의 문자 클래스로 결합하고, 패턴을 부정하기 위해 `^`를 첫 문자로 추가하세요.

# --hints--

네 번째 패턴을 완성하여 영숫자가 아닌 모든 문자를 매치하세요.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
