---
id: 6564b8c9349bd76dc037967b
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

`search()` 모듈의 `re` 함수는 인수로 전달된 문자열을 분석하여 정규식 패턴이 문자열과 처음으로 일치하는 위치를 찾습니다.

`quote`라는 변수를 선언하고 문자열 `'Not all those who wander are lost.'`을 이 변수에 할당하세요. 그런 다음 `pattern.search(quote)`의 결과를 출력하세요.

# --hints--

`quote` 변수가 있어야 합니다.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

제공된 문자열을 새 `quote` 변수에 할당해야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

`pattern.search(quote)`을 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
