---
id: 65648f4b2281ba50051ae39c
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

다음으로, `for` 반복문을 `i`을 반복 변수로 사용해 작성하세요. `range()` 함수를 사용해 `length` 값까지 반복하세요.

반복문 안에서는 덧셈 대입 연산자를 사용해 `all_characters`에서 무작위 문자를 현재 `password` 값에 더하세요. 이를 위해 `choice()` 모듈의 `secrets` 함수를 사용하세요.

# --hints--

`for`를 반복하는 `range(length)` 반복문을 작성해야 합니다.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

덧셈 대입 연산자 `+=`를 사용해 `all_characters`에서 무작위 문자를 현재 `password` 값에 더해야 합니다.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
