---
id: 65647c71a5d2bd431596f629
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

`generate_password` 함수에 몇 개의 매개변수가 필요합니다. `length` 매개변수를 추가하는 것부터 시작하세요.

# --hints--

`generate_password` 함수는 `length` 매개변수를 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == "(length)"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import secrets
import string


def generate_password():    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

--fcc-editable-region--
```
