---
id: 65646ffeaed2d238c562a014
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

코드 맨 위에 `import` 문을 두는 것이 일반적인 관례입니다. 그리고 `import` 문이 여러 개일 경우, 가독성을 높이기 위해 알파벳 순서로 정렬하세요.

코드 맨 위에 `random` 모듈을 가져오세요.

# --hints--

`random` 모듈을 가져와야 합니다.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
