---
id: 656491fa4e69005287eb5a9a
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

새 변수를 출력해서 결과를 확인하세요.

# --hints--

`new_password`를 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)

--fcc-editable-region--
```
