---
id: 656470d517833a39bb8b5608
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`random` 모듈에는 의사 난수 생성기가 포함되어 있습니다. 대부분의 기능은 `random()` 함수에 의존하며, 이 함수는 `0.0` (포함)와 `1.0` (미포함) 사이의 부동 소수점 숫자를 반환합니다.

`random()` 모듈에서 `random` 함수를 호출하고 결과를 출력하세요.

# --hints--

`random.random()` 함수를 호출한 결과를 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
