---
id: 6564c67db26c417561ab510d
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

출력에서 볼 수 있듯이, 이제 정규식이 문자열 안의 첫 번째 `l`와 일치합니다.

패턴에서 문자 뒤에 수량자를 추가해 그 문자가 몇 번 반복되어야 하는지 지정할 수 있습니다. 예를 들어, `+` 수량자는 한 번 이상 반복되어야 함을 의미합니다.

패턴에 `+` 수량자를 추가하세요.

# --hints--

`pattern` 변수를 `re.compile('l+')`로 수정해야 합니다.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
