---
id: 65646ab6bced52360f44670f
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이제 `all_characters`라는 변수를 선언하고 기존 변수들을 연결한 결과를 할당하세요.

# --hints--

`all_characters`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

`letters`, `digits`, `symbols`를 연결하고 그 결과를 `all_characters` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
