---
id: 6564de10a0887f882b0012d3
title: 48단계
challengeType: 20
dashedName: step-48
---

# --description--

`[0-9]` 클래스가 `\d`과 같은 것처럼, `[^0-9]` 클래스도 `\D`과 같습니다. 영숫자 문자는 `\w`로 매칭할 수 있고, 영숫자가 아닌 문자는 `\W`로 매칭할 수 있습니다.

`[^a-zA-Z0-9]` 문자 클래스를 `\W`로 바꾸세요.

# --hints--

`[^a-zA-Z0-9]` 문자 클래스를 `\W`로 바꿔야 합니다.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\\W\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'[^a-zA-Z0-9]')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
