---
id: 6564efa70114b591b74d5679
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

공백 문자와 마지막 마침표가 일치하는데, 이는 문자열에서 유일한 영숫자가 아닌 문자들이기 때문입니다.

이제 `quote` 문자열을 단일 밑줄 문자로 바꾸세요.

# --hints--

`quote` 변수는 `'_'`여야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "_") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
