---
id: 656494269ccce754411a2112
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

`re` 모듈을 사용하면 코드에서 *정규 표현식*을 사용할 수 있습니다. 정규 표현식에 대해서는 곧 더 자세히 배우게 됩니다.

지금은 코드 맨 위에 `import` 문장을 추가하여 `re` 모듈을 가져오세요.

# --hints--

`re` 모듈을 가져와야 합니다.

```js
({ test: () => assert.match(code, /^import\s+re/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

import secrets
import string
--fcc-editable-region--

def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
