---
id: 656489e0a5d0954c1a64fc93
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 `all_characters` 변수를 출력해서 어떻게 생겼는지 확인하세요.

# --hints--

`all_characters` 변수를 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*all_characters\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
