---
id: 6564cbb6311a40783b3f5de6
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

문자 클래스는 일치시킬 문자 범위를 지정할 수도 있습니다. 시작 문자와 끝 문자를 하이픈으로 구분해 지정해야 합니다, `-`. 문자는 유니코드 순서를 따릅니다.

`pattern` 변수를 수정해 `t` 변수에 있는 소문자 앞에 오는 `quote`의 모든 문자와 일치하도록 만드세요. 이를 위해 `a`부터 `z`까지의 문자 범위를 사용하세요.

# --hints--

`pattern` 변수를 수정해 `t`에 있는 소문자 앞에 오는 `quote`의 모든 문자와 일치하도록 `[a-z]` 클래스를 사용하세요.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
