---
id: 6564f0279e23ce924eedd1b2
title: 51단계
challengeType: 20
dashedName: step-51
---

# --description--

밑줄 문자는 변수 이름에 유효한 문자이므로, 변수 이름과 편리하게 매치할 수 있는 `\w` 문자 클래스(`[a-zA-Z0-9_]`와 동일)에 포함되어 있습니다.

따라서 `\W` 문자 클래스는 밑줄 문자를 제외한 `[^a-zA-Z0-9_]`와 동일합니다. 이 때문에 모든 특수 문자를 매치하는 데 사용할 수 없습니다.

코드에서 마지막 세 줄을 삭제하세요.

# --hints--

코드에서 마지막 세 줄을 삭제해야 합니다.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
