---
id: 65649a631440b757f9f7b428
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

정규 표현식, 또는 정규식은 문자열 안에서 특정 문자 조합과 일치시키기 위해 사용하는 패턴입니다. 문자열 안에서 유효성 검사, 문자 교체 등 패턴을 찾거나 일치시켜야 할 때 `if`/`else` 조건문 대신 사용할 수 있는 유효한 대안입니다.

`compile()` 모듈의 `re` 함수는 인수로 전달된 문자열을 정규식 객체로 컴파일하며, 이 객체는 다른 `re` 메서드에서 사용할 수 있습니다.

새로운 `pattern` 변수를 선언하고 `re.compile('i')` 값을 이 변수에 할당하세요.

# --hints--

`pattern` 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

`re.compile('i')` 변수에 `pattern` 값을 할당해야 합니다.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
