---
id: 656469e62377ab34a5bcb8d1
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`digits`와 `symbols`이라는 두 개의 새 변수를 선언하고 각각 `string.digits`와 `string.punctuation`을 할당하세요.

# --hints--

`digits`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

`string.digits`를 `digits` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

`symbols`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

`string.punctuation`를 `symbols` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
