---
id: 656474314bf37d3c83a85143
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`choice()` 모듈의 `random` 함수는 수열을 받아 수열에서 무작위 항목을 반환합니다.

`print()` 호출을 수정하여 `choice()` 함수를 사용하고 `all_characters`를 인수로 전달하세요.

# --hints--

기존 `print()` 호출을 수정하여 `random.choice(all_characters)`를 출력하세요.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
