---
id: 65648e4c5b316c4ec5e4fddc
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

함수 맨 아래에 `password` 변수를 선언하고 이 변수에 빈 문자열을 할당하세요.

# --hints--

함수 안에 빈 문자열 값을 가진 `password` 변수가 있어야 합니다. 들여쓰기에 주의하세요.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
