---
id: 6564883669b5af4b69f794cc
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`all_characters` 변수는 모든 소문자와 대문자, 10개의 숫자, 그리고 여러 특수 문자로 이루어진 문자열입니다.

그 바로 위에 `Combine all characters`라는 주석을 추가하세요.

# --hints--

주석은 `all_characters` 변수 바로 위에 추가해야 합니다.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
