---
id: 6564fd78ffff909b1531d3c3
title: 55단계
challengeType: 20
dashedName: step-55
---

# --description--

`for` 반복문 안에서 `findall()` 함수를 호출하며 `pattern`와 `password`을 인수로 전달하세요.

# --hints--

`re.findall()`를 호출하며 `pattern`와 `password`을 인수로 전달해야 합니다.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1\s{4}re\.findall\s*\(\s*pattern\s*,\s*password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        for constraint, pattern in constraints:
            pass
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
