---
id: 656479aa5f298441c190bf8f
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`generate_password` 함수를 선언하고 `import` 줄을 제외한 모든 코드를 함수 본문 안에 작성하세요.

# --hints--

`generate_password`라는 이름의 함수를 선언해야 합니다.

```js
({
  test: () => assert(runPython(`
  _Node(_code).has_function('generate_password')
  `))
})
```

가져오기 문은 여전히 함수 밖에 있어야 합니다.

```js
({
  test: () => assert(runPython(`
  imports_list = _Node(_code).find_imports()
  any([imp.is_equivalent('import secrets') for imp in imports_list]) and any([imp.is_equivalent('import string') for imp in imports_list])
  `))
})
```

네 개의 변수 선언은 함수 안으로 옮겨야 합니다.

```js
({
  test: () => assert(runPython(`
  func = _Node(_code).find_function('generate_password')
  vars = ['letters', 'digits', 'symbols', 'all_characters']
  func.find_variable('letters').is_equivalent('letters = string.ascii_letters') and \\
  func.find_variable('digits').is_equivalent('digits = string.digits') and \\
  func.find_variable('symbols').is_equivalent('symbols = string.punctuation') and \\
  func.find_variable('all_characters').is_equivalent('all_characters = letters + digits + symbols') and \\
  all(not _Node(_code).has_variable(var) for var in vars)
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
