---
id: 656477845006313fbfea0ad1
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

효과가 `random.choice()`와 같아 보일 수 있지만, `secrets`는 운영 체제가 제공할 수 있는 가장 안전한 무작위성 원천을 보장합니다.

이제 두 개의 `print()` 호출을 삭제하세요.

# --hints--

두 개의 `print()` 호출을 삭제해야 합니다.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
