---
id: 6564642ce403fe29cc23fee7
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

파이썬 모듈은 특정 작업을 수행하도록 설계된 코드를 포함하는 파일입니다. 파이썬 표준 라이브러리에는 가져와서 코드에서 사용할 수 있는 많은 모듈이 들어 있습니다. `import` 문 다음에 모듈 이름을 써서 이 작업을 할 수 있습니다.

이 프로젝트는 `string` 모듈을 가져오는 것으로 시작하세요.

# --hints--

`string` 모듈을 가져와야 합니다.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
