---
id: 6564d096a55d707bd77ab67b
title: 38단계
challengeType: 20
dashedName: step-38
---

# --description--

캐럿 `^`을 문자 클래스 맨 앞에 두면, 클래스에 지정된 문자를 제외한 모든 문자와 일치합니다.

문자 클래스 안에 `^`을 첫 번째 문자로 추가하고 어떤 일이 일어나는지 확인하세요.

# --hints--

문자 클래스 안에 `^`을 첫 번째 문자로 추가해야 합니다.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[^a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
