---
id: 6564d97a5ef99783216a6229
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

Python은 *raw* 문자열이라는 특별한 유형의 문자열을 제공합니다. Raw 문자열은 `r` 접두사가 붙습니다. 일반 문자열과의 주요 차이점은 역슬래시 문자를 처리하는 방식에 있습니다: raw 문자열에서는 역슬래시가 이스케이프 문자가 아니라 문자 그대로 취급됩니다. 정규 표현식을 작성할 때는 보통 많은 `\` 문자를 포함할 수 있으므로 raw 문자열을 사용하는 것이 좋은 습관입니다.

`pattern` 문자열을 `r` 접두사를 붙여 raw 문자열로 바꾸세요.

# --hints--

`pattern` 변수를 `r'\.'`로 수정하세요.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
