---
id: 6565bb128adfcd5ec362382d
title: 57단계
challengeType: 20
dashedName: step-57
---

# --description--

`for` 반복문 안에서 `constraint`과 `findall()`가 반환하는 리스트의 길이를 비교하세요. 이때 `<=` 연산자를 사용하세요.

# --hints--

`constraint` 반복문 안에서 `findall`와 `<=`가 반환하는 리스트의 길이를 `for` 연산자를 사용해 비교해야 합니다.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1\s{4}constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        for constraint, pattern in constraints:
            len(re.findall(pattern, password))
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
