---
id: 6565c13fdb798865c161d8f8
title: 65단계
challengeType: 20
dashedName: step-65
---

# --description--

이제 함수를 테스트할 차례입니다. 코드에서 마지막 두 줄의 주석을 해제하고 함수 호출에 5개의 인수를 전달하도록 수정하세요. 길이에는 `8`를, 나머지 네 개 제약 조건에는 `1`을 사용하세요.

# --hints--

제공된 인수로 `generate_password`를 호출해야 합니다.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*,\s*1\s*,\s*1\s*,\s*1\s*,\s*1\s*\)/m) })
```

`new_password` 변수를 출력해야 합니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
        
        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break

    return password
    
--fcc-editable-region--
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
```
