---
id: 656ef5bc5c0cd464be1df675
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

다음으로, 생성된 비밀번호에 제약 조건으로 작용할 매개변수를 함수에 더 추가할 것입니다.

기존 `nums` 매개변수 뒤에 `special_chars`, `uppercase`, `lowercase`, `length`를 이 순서대로 추가하여 함수 선언을 수정하세요.

# --hints--

함수는 `length`, `nums`, `special_chars`, `uppercase`, `lowercase`를 매개변수로 받아야 합니다. 순서가 중요합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
