---
id: 656ef89b4b486b7a16077864
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

튜플은 Python의 또 다른 내장 자료 구조입니다. 튜플은 리스트와 매우 비슷하지만, 대괄호 대신 소괄호 `()`로 정의합니다. 또한, 튜플은 리스트와 달리 변경할 수 없습니다.

```py
my_tuple = ('larch', 1, True)
```

`constraints` 리스트는 튜플을 저장할 것입니다. 각 튜플의 첫 번째 항목은 제약 매개변수가 됩니다.

`constraints` 리스트 할당을 수정하여 리스트에 튜플을 추가하세요. 첫 번째 항목은 `nums`를 사용하고 두 번째 항목은 빈 문자열로 하세요.

# --hints--

`constraints` 리스트에 튜플을 추가해야 합니다.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

`(nums, '')`를 `constraints` 리스트에 추가해야 합니다.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
