---
id: 656ef54070c72161e6feb90f
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

모든 것이 괜찮아 보이지만, 생성된 비밀번호가 특정 기능을 충족하는지 확인하는 방법이 있으면 좋겠습니다. 예를 들어, 특수 문자, 숫자, 대문자/소문자의 최솟값 같은 조건입니다. 곧 그 부분을 처리할 것입니다.

지금은 코드의 마지막 두 줄을 주석 처리하세요.

# --hints--

코드의 마지막 두 줄을 주석 처리해야 합니다.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
