---
id: 656eec2f40d18056cc58b229
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

파싱된 문자열 안에 `None`이 없기 때문에 `'i'` 값이 반환됩니다.

이제 `re.compile()`에 전달된 문자열을 `'l'`으로 수정하고 결과를 확인하세요.

# --hints--

`pattern` 변수를 `re.compile('l')`으로 수정해야 합니다.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
