---
id: 656ef783bba6976de014eaa8
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

`password` 변수 선언과 그 다음 `for` 반복문을 `while` 반복문 안에 넣으세요. 새 반복문의 조건으로 `True`을 사용하세요.

# --hints--

기존 `while True` 선언과 `password` 반복문을 감싸는 `for` 반복문을 만드세요.

```js
({ test: () => {
    
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /while\s+True\s*/);
    assert(block_body.match(/^\s+password\s*=\s*("|')\1\s*^\s+for\s+_\s+in\s+range\s*\(\s*length\s*\s*\)\s*:\s*^\s+password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*$/m));    
    const {block_body: for_body} = __helpers.python.getBlock(commentless_code, /for\s+_\s+in\s+range\s*\(\s*length\s*\s*\)\s*/);
    assert(for_body.match(/^\s+password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*$/m));  
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
--fcc-editable-region--        
    return password

# new_password = generate_password(8)
# print(new_password)
```
