---
id: 656f0d6bc60b58b9777a36c9
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

문자 클래스는 대괄호로 표시하며, 괄호 안에 지정된 문자 중 하나와 일치합니다. 예를 들어, `ma[dnt]`는 `mad`, `man`, 또는 `mat` 중 하나와 일치합니다.

패턴을 수정하여 `w` 다음에 `h` 또는 `a` 중 하나가 오는 경우와 일치하도록 하세요.

# --hints--

패턴은 문자 클래스를 사용하여 `w` 다음에 `h` 또는 `a` 중 하나가 오는 경우와 일치해야 합니다.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
