---
id: 656f10458134d4c4e283a2f1
title: 35단계
challengeType: 20
dashedName: step-35
---

# --description--

이제 제약 튜플의 빈 문자열을 단일 숫자와 일치하는 정규식 패턴으로 바꾸세요. `0`부터 `9`까지 각 숫자를 지정하는 문자 클래스를 사용하세요.

# --hints--

`(nums, '[0123456789]')` 목록에 `constraints` 튜플이 있어야 합니다.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\s*\(\s*nums\s*,\s*("|')\[\d{10}\]\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, '')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
