---
id: 656f04b20397d9a574cc7eb2
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

`compile()` 함수를 사용하지 않고도 같은 결과를 얻을 수 있습니다. `pattern` 변수를 리터럴 문자열 `'l+'`로 수정하세요. 그런 다음 `print()` 호출을 `re.search(pattern, quote)`를 출력하도록 변경하세요.

# --hints--

`pattern` 변수를 리터럴 문자열 `'l+'`로 수정하세요.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "l+") })
```

`re.search(pattern, quote)`를 출력하세요.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.search\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l+')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
