---
id: 656f0bc259c1f6b57486ed68
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

생성된 비밀번호가 요구하는 기능을 충족하는지 확인하려면 `findall()` 모듈의 `re` 함수를 사용하세요. 이 함수는 `search`와 비슷하지만, 일치하는 패턴의 모든 발생을 포함하는 리스트를 반환합니다.

`search()` 호출을 `findall()`로 바꾸고 출력 결과를 확인하세요.

# --hints--

기존 `print()` 호출에서 `search()`를 `findall()`로 바꾸면 됩니다.

```js
({ test: () => assert.match(code, /^print\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.search(pattern, quote))
--fcc-editable-region--
```
