---
id: 6577319039f4f7de9251b822
title: 61단계
challengeType: 20
dashedName: step-61
---

# --description--

루프와 카운터 변수를 사용하는 대신, 다음 몇 단계에서 구현할 다른 방법으로 같은 결과를 얻을 수 있습니다.

`all()`는 주어진 반복 가능한 객체 안의 모든 요소가 `True`로 평가되면 `True`을 반환하는 내장 Python 함수입니다. 그렇지 않으면 `False`을 반환합니다.

기존의 `for` 루프와 두 개의 `if` 문을 하나의 `if` 문으로 바꾸세요. `if` 조건에는 `all()` 함수를 호출하고 빈 리스트를 인수로 전달하세요.

# --hints--

기존의 `for` 루프와 두 개의 `if` 문을 하나의 `if` 문으로 바꾸세요.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

새로운 `if` 조건은 `all([])`이어야 합니다.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

새로운 `break` 본문 안에는 `if`가 있어야 합니다.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
