---
id: 6577333feab1e8e927014f03
title: 63단계
challengeType: 20
dashedName: step-63
---

# --description--

`all([expression for i in iterable])`가 있다는 것은 `expression` 안의 각 `i`에 대해 `iterable`을 평가하여 새 리스트가 만들어진다는 뜻입니다. `all()` 함수가 새로 만들어진 리스트를 반복한 후에는 더 이상 필요하지 않으므로 리스트가 자동으로 삭제됩니다.

제너레이터 표현식을 사용하면 메모리를 절약할 수 있습니다. 제너레이터 표현식은 리스트 내포 구문을 따르지만 대괄호 대신 소괄호를 사용합니다.

대괄호를 제거하여 리스트 내포를 제너레이터 표현식으로 바꾸세요.

# --hints--

리스트 내포를 제너레이터 표현식으로 바꾸려면 대괄호를 제거하세요.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        if all(
            [
                constraint <= len(re.findall(pattern, password))
                for constraint, pattern in constraints
            ]
        ):
--fcc-editable-region--
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
