---
id: 657732654845d8e2fb1217e6
title: 62단계
challengeType: 20
dashedName: step-62
---

# --description--

지금 `all()`는 빈 리스트를 인수로 받고 있습니다. 리스트 내포 구문을 사용해 빈 리스트를 채우고, `constraint <= len(re.findall(pattern, password))` 리스트에 있는 각 `constraint`-`pattern` 튜플에 대해 `constraints` 표현식을 평가한 결과를 저장하세요.

이렇게 하면 모든 요구 사항이 충족된 후에만 `while` 반복문을 빠져나가게 됩니다.

# --hints--

`[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]`를 `all()` 함수에 전달해야 합니다.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break
    
    return password

# new_password = generate_password(8)
# print(new_password)
```
