---
id: 657db4cb77190e33a20e852a
title: 69단계
challengeType: 20
dashedName: step-69
---

# --description--

기본 인수와 키워드 인수를 결합하면 함수가 요구하는 인수보다 명시적으로 적은 인수를 전달할 수 있습니다. 함수 호출에 명시적으로 전달되지 않은 인수는 기본값을 갖게 됩니다.

`generate_password()` 호출을 `length=8`만 받도록 수정하세요.

# --hints--

`length=8` 호출에 `generate_password()`만 전달해야 합니다.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region--
new_password = generate_password(nums=1, length=8, special_chars=1, uppercase=1, lowercase=1)
print(new_password)
--fcc-editable-region--
```
