---
id: 657dadf4d8b93c1704f3a57c
title: 66단계
challengeType: 20
dashedName: step-66
---

# --description--

작동은 하지만, 여전히 개선할 수 있는 몇 가지가 있습니다. 우선, 5개의 인수를 가진 함수를 호출하면 어떤 값이 어떤 매개변수에 할당될지 혼란스러울 수 있습니다.

키워드 인수를 사용해 함수를 호출할 수 있습니다. 즉, 매개변수 이름을 명시하고 대입 연산자와 값을 뒤따르게 작성하는 방식입니다. 예를 들어:

```py
def add(x, y):
    return x + y

add(x=1, y=7) # 8
```

함수 호출을 키워드 인수를 사용하도록 수정하세요.

# --hints--

`new_password` 변수는 `generate_password(length=8, nums=1, special_chars=1, uppercase=1, lowercase=1)` 값을 가져야 합니다.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*,\s*nums\s*=\s*1\s*,\s*special_chars\s*=\s*1\s*,\s*uppercase\s*=\s*1\s*,\s*lowercase\s*=\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
   
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region-- 
new_password = generate_password(8, 1, 1, 1, 1)
print(new_password)
--fcc-editable-region--
```
