---
id: 6823f9df49cc206af5471a30
title: Camperbot 프로필 페이지 디버그하기
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot이 프로필 페이지를 만들고 있습니다. 친구에게 코드를 봐달라고 했는데 몇 가지 오류가 있다고 합니다.

Camperbot의 모든 오류를 고쳐서 프로필 페이지를 계속 만들 수 있게 하세요. 아래 사용자 스토리의 항목을 완료하고 "Run the Tests"를 눌러 모든 오류가 고쳐졌는지 확인하세요.

**사용자 스토리:**

1. Camperbot이 `heading2` 요소를 사용하려고 하는데, 그 요소는 존재하지 않습니다. 올바른 2단계 제목 요소를 사용하도록 태그를 고치세요.
2. Camperbot이 `pp` 두 개를 추가하려고 하는데, 그것들도 존재하지 않습니다. 올바른 단락 태그를 사용하도록 고치세요.
3. Camperbot이 `h3` 요소를 `Background and Interests` 부제목에 사용했는데 구문 오류가 있습니다. 문제를 찾아 해결하세요.

# --hints--

페이지에 `h1` 요소가 정확히 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

페이지에서 `<h1>Hello from Camperbot!</h1>` 요소를 제거하면 안 됩니다.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

페이지에 `h2` 요소가 정확히 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

텍스트가 `h2`인 `About` 요소가 있어야 합니다. 예시는 다음과 같습니다: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

텍스트가 `My name is Camperbot and I love learning new things.`인 단락 요소가 있어야 합니다. 예시는 다음과 같습니다: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

텍스트가 `I enjoy solving puzzles.`인 단락 요소가 있어야 합니다. 예시는 다음과 같습니다: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

페이지에 단락 요소가 정확히 두 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

페이지에 `h3` 요소가 정확히 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

텍스트가 `h3`인 `Background and Interests` 요소가 있어야 합니다. 예시는 다음과 같습니다: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

페이지에 총 다섯 개의 요소가 정확히 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
