---
id: a97fd23d9b809dac9921074f
title: 构建一个选择性参数求和函数
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

在本实验中，你将构建一个函数，该函数最多接受两个参数并将它们相加，但如果只有一个参数，则返回一个函数，该函数等待第二个数字以进行相加。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 创建一个函数 `addTogether`。
2. 如果函数接收两个参数，函数应返回这两个参数的和。
3. 如果只提供一个参数，`addTogether` 应该返回一个函数。
   - 当返回的函数以单个参数调用时，它应返回和。

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. 如果任一参数不是有效数字，返回 `undefined`。

# --hints--

你应该有 `addTogether` 函数。

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` 应返回 5。

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` 应返回 53.4。

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` 应返回 `undefined`。

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` 应返回 `undefined`。

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` 应返回 `undefined`。

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` 应返回一个函数。

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` 应返回 12。

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` 应返回 `undefined`。

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` 应返回 `undefined`。

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
