---
id: 6733b5e6f210c1bd1afdf8b3
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

剩余的植物对象已为你添加。现在是创建你的目录的时候了。

像普通的 JavaScript 对象一样，`Map` 对象包含一个键值对集合。要创建一个 `Map` 对象，你可以使用 `Map` 构造函数：

```js
const myMap = new Map();
```

声明一个 `catalog` 变量并赋值一个 `Map` 对象。

# --hints--

你应该有一个名为 `catalog` 的变量。

```js
assert.exists(catalog);
```

你应该将一个 `Map` 对象赋值给你的 `catalog` 变量。

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
