---
id: 673339d96c56d36e963e888e
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将构建一个植物苗圃目录，同时练习 JavaScript 的 Map 和 Set 对象。

为了表现你目录中的每个植物，你将使用对象，每个对象有三个属性：

- `commonName`，植物的通用名称
- `scientificName`，植物的学名（或拉丁名）
- `cultivar`，植物的品种

首先为名为 `Ballerina` 的品种的 `Spanish lavender`（学名 `Lavandula stoechas`）创建一个具有上述三个属性的对象，并将其赋值给名为 `ballerina` 的变量。

# --hints--

你应该有一个名为 `ballerina` 的变量。

```js
assert.exists(ballerina);
```

你应该为你的 `ballerina` 变量赋值一个对象。

```js
assert.isObject(ballerina);
```

你的 `ballerina` 对象应具有设置为 `"Spanish lavender"` 的 `commonName` 属性。

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

你的 `ballerina` 对象应具有设置为 `"Lavandula stoechas"` 的 `scientificName` 属性。

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

你的 `ballerina` 对象应具有设置为 `"Ballerina"` 的 `cultivar` 属性。

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

你的 `ballerina` 对象应有三个密钥。

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
