---
id: 6733b5632fb854bc6c58dc0d
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

现在创建另一个对象，具有相同的三个属性，用于表现另一种名为 `Pretty Polly` 的西班牙薰衣草品种，并将其赋值给名为 `prettyPolly` 的变量。

# --hints--

你应该有一个名为 `prettyPolly` 的变量。

```js
assert.exists(prettyPolly);
```

你应该为你的 `prettyPolly` 变量赋值一个对象。

```js
assert.isObject(prettyPolly);
```

你的 `prettyPolly` 对象应具有设置为 `"Spanish lavender"` 的 `commonName` 属性。

```js
assert.strictEqual(prettyPolly.commonName, "Spanish lavender");
```

你的 `prettyPolly` 对象应具有设置为 `"Lavandula stoechas"` 的 `scientificName` 属性。

```js
assert.strictEqual(prettyPolly.scientificName, "Lavandula stoechas");
```

你的 `prettyPolly` 对象应具有设置为 `"Pretty Polly"` 的 `cultivar` 属性。

```js
assert.strictEqual(prettyPolly.cultivar, "Pretty Polly");
```

你的 `prettyPolly` 对象应有三个密钥。

```js
assert.lengthOf(Object.keys(prettyPolly), 3);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}
--fcc-editable-region--

--fcc-editable-region--
```
