---
id: 6734a2385c52c54d916a8b86
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

为你的目录添加另一个键值点对。使用你的 `prettyPolly` 变量作为密钥，并使值成为一个对象，该对象具有 `small`、`medium` 和 `large` 属性，其值再次为整数。

# --hints--

你应该将你的 `prettyPolly` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(prettyPolly));
```

`catalog` 映射中 `prettyPolly` 密钥的值应为具有 `small`、`medium` 和 `large` 属性的对象。

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog` 映射中与 `prettyPolly` 密钥关联的对象的 `small` 属性应为正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog` 映射中与 `prettyPolly` 密钥关联的对象的 `medium` 属性应为正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog` 映射中与 `prettyPolly` 密钥关联的对象的 `large` 属性应为正整数。

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你仍然应该在你的 `catalog` 映射中拥有 `ballerina` 密钥。

```js
assert.isTrue(catalog.has(ballerina));
```

你的 `catalog` 地图应包含两个元素。

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
