---
id: 6734a606e5968a5a19b77171
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

继续像之前的步骤那样将剩余的植物添加到你的目录中。

# --hints--

你应该将你的 `willowVale` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(willowVale));
```

`catalog` 映射中 `willowVale` 密钥的值应为一个具有属性 `small`、`medium` 和 `large` 的对象，每个属性的值都是正整数。

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的 `hidcote` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(hidcote));
```

`catalog` 映射中 `hidcote` 密钥的值应为一个具有属性 `small`、`medium` 和 `large` 的对象，每个属性的值都是正整数。

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的 `imperialGem` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(imperialGem));
```

`catalog` 映射中 `imperialGem` 密钥的值应为一个具有属性 `small`、`medium` 和 `large` 的对象，每个属性的值都是正整数。

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你应该将你的 `royalCrown` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(royalCrown));
```

`catalog` 映射中 `royalCrown` 密钥的值应为一个具有属性 `small`、`medium` 和 `large` 的对象，每个属性的值都是正整数。

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你仍然应该在你的 `catalog` 映射中拥有 `ballerina` 密钥。

```js
assert.isTrue(catalog.has(ballerina));
```

你仍然应该在你的 `catalog` 映射中拥有 `prettyPolly` 密钥。

```js
assert.isTrue(catalog.has(prettyPolly));
```

你的 `catalog` 地图应包含六个元素。

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
