---
id: 6734a13a29acdf47f2b8e161
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

与普通对象不同，普通对象的密钥必须始终是字串，`map` 对象允许你使用任何数据类型作为密钥。在你的目录中，你将使用对象-对象点对，将你的每个植物对象与包含不同尺寸可用花盆数量的对象关联起来。

修改你的 `set` 调用，将你的 `ballerina` 对象用作密钥。对于值，使用包含三个属性的对象，`small`、`medium` 和 `large`，每个属性的值都是正整数。 

# --hints--

你应该将你的 `ballerina` 对象设置为你的 `catalog` 映射的密钥。

```js
assert.isTrue(catalog.has(ballerina));
```

`catalog` 映射中 `ballerina` 密钥的值应为具有属性 `small`、`medium` 和 `large` 的对象。

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

`catalog` 映射中与 `ballerina` 密钥关联的对象的 `small` 属性应为正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

`catalog` 映射中与 `ballerina` 密钥关联的对象的 `medium` 属性应为正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

`catalog` 映射中与 `ballerina` 密钥关联的对象的 `large` 属性应为正整数。

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

你的 `catalog` 地图应包含单个密钥-值点对。

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
