---
id: 67349ba230cff237c0f2bdf2
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

映射对象会记住密钥插入映射中的顺序。要在你的映射对象中设置一个密钥-值点对，你可以使用 `set` 方法。以下代码将密钥 `"name"` 设置为值 `"Jerry"`，在 `myMap` 映射对象中。

```js
myMap.set("name", "Jerry");
```

将 `"ballerina"` 密钥设置为你 `catalog` 映射中 `ballerina` 对象的值。然后，将你的 catalog 日志输出到控制台。

请注意，要正确查看你的地图对象，你必须打开浏览器控制台。

# --hints--

你应该将 `"ballerina"` 密钥设置为你 `catalog` 映射中 `ballerina` 对象的值。

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

你的 `catalog` 地图应包含单个密钥-值点对。

```js
assert.lengthOf(catalog, 1);
```

你应该将你的 `catalog` 映射记录到控制台日志。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
