---
id: 673521668be7905059b2d555
title: 步骤 22
challengeType: 1
dashedName: step-22
---

# --description--

最后，你可以使用 `forEach` 方法在你的映射中的每个密钥-值点对上运行一个函数。

```js
myMap.forEach((value, key) => {console.log(key, value)});
```

请注意，传递给 `forEach` 的回调通过其第一个参数访问值，通过其第二个参数访问每个条目的密钥。 

修改你的函数以返回包含你整个目录的字串，并将结果记录到控制台。使用以下结构为 `catalog` 中的每个条目构建你的字串：

```md
<scientific-name> '<cultivar>': <small-pots> S, <medium-pots> M, <large-pots> L

```

请注意，每个条目的表现应以换行字符终止。

# --hints--

你应该将调用 `displayCatalog` 的结果记录到控制台日志中。

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*displayCatalog\s*\(\s*\)\s*\)/)
```

你的 `displayCatalog` 函数应在 `catalog` 上使用 `forEach` 方法。

```js
let flag = false;
const temp = catalog.forEach;
try {
    catalog.forEach = obj => {flag = true};
    displayCatalog();    
    assert.isTrue(flag);    
} finally {
    catalog.forEach = temp;
}
```

`displayCatalog` 应返回 `--fcc-expected--`，但发现了 `--fcc-actual--`。

```js
const displayCatalogTest = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`   })
    return catalogString
}
assert.equal(displayCatalog(), displayCatalogTest())
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const removePlant = plant => catalog.delete(plant);

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}
--fcc-editable-region--
const displayCatalog = () => {
    for (const entry of catalog.entries()) {
        console.log(entry);
    }
}

displayCatalog()
--fcc-editable-region--
```
