---
id: 64b163c20e59cbd4a64940b0
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

创建一个将两个参数相加的函数。

# --hints--

3 和 5 相加应返回 8。

```js
({
  test: () => assert.equal(__userGlobals.get("add")(3,5), 8)
})
```

测试 `__name__` 的值

```js
({
  test: () => assert(runPython(`__name__ == '__main__'`))
})
```

Test cleaner syntax

```js
({
  test: () => assert(runPython(`add(4,5) == 9`))
})
```

Testing getDef

```js
({
  test: () => {
    const add = __helpers.python.getDef(code, "add");
    assert.deepEqual(add, {
        def: "def add(a, b):\n    return a + b\n",
        function_parameters: "a, b",
        function_body: "    return a + b\n",
        function_indentation: 0
    });
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py

def add(a, b):
    return a + b

```
