---
id: 64b171849f925b0773aa434c
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

创建一个从 1 到 3 的数字列表，并将其分配到变量 `xs`。

# --hints--

`xs` 应该是列表 `[1,2,3]`。

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
