---
id: 68da604b8b6012a33e2344fb
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、幅優先探索（BFS）を使って丸括弧のすべての有効な組み合わせを生成する関数を実装します。たとえば、2組の丸括弧の有効な組み合わせは `(())` と `()()` です。

まず、`gen_parentheses` という単一のパラメータを持つ `pairs` という名前の関数を作成してください。今は関数から空のリストを返してください。

# --hints--

`gen_parentheses` という名前の関数を定義してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

`gen_parentheses` 関数は `pairs` という単一のパラメータを持つべきです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

関数は空のリストを返すべきです。

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
