---
id: 68da6138d70960c197ec5eee
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

これで結果を保存するデータ構造を設定します。`result` という名前の変数を作成し、空のリストとして初期化してください。このリストに、生成するすべての有効な丸括弧の組み合わせを保存します。

return文を更新して、空のリストの代わりに `result` を返すようにしてください。

# --hints--

`result` 関数内で `gen_parentheses` という名前の変数を宣言してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

`result` を空のリストとして初期化してください。

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

関数の最後で `result` を返してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
