---
id: 68da6138d70960c197ec5eef
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

幅優先探索の方法では、キューを使って丸括弧の組み合わせを作成する過程で異なる状態を追跡します。各状態は3つの要素を含むタプルで表されます：

- 現在作成中の文字列
- これまでに使った開き丸括弧の数
- これまでに使った閉じ丸括弧の数

`queue` という名前の変数を作成し、空の文字列と丸括弧がゼロの状態を表す1つのタプル `('', 0, 0)` を含むリストで初期化してください。

# --hints--

`queue` 関数内で `gen_parentheses` という名前の変数を宣言してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

`queue` をタプル `('', 0, 0)` を含むリストで初期化してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
