---
id: 68da6138d70960c197ec5eec
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

コアアルゴリズムを実装する前に、入力を検証する必要があります。`pairs` パラメータは整数である必要があります。これは生成する丸括弧のペア数を表しているためです。

関数の先頭に `if` 文を追加して、`pairs` が整数でないかどうかをチェックしてください。そのために `isinstance()` 関数を使います。

条件が真の場合、文字列 `The number of pairs should be an integer` を返してください。

# --hints--

`if` 関数内に `gen_parentheses` 文があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

`if` 文は `pairs` を使って `not isinstance(pairs, int)` が整数でないかどうかをチェックする必要があります。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

入力が整数でない場合は文字列 `The number of pairs should be an integer` を返してください。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
