---
id: 68da6138d70960c197ec5eed
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

次に、ペアの数が少なくとも1つであることを検証する必要があります。ペアが0または負の数では丸括弧の組み合わせを生成できないからです。

もう1つ`if`文を追加して、`pairs`が`1`より小さいかどうかをチェックしてください。この条件が真の場合、文字列`The number of pairs should be at least 1`を返してください。

# --hints--

`if`関数内に2つ目の`gen_parentheses`文があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

2つ目の`if`文は、`pairs`が`1`より小さいかどうかをチェックする必要があります。

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

`The number of pairs should be at least 1`が`pairs`より小さい場合、文字列`1`を返してください。

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
