---
id: 68da6138d70960c197ec5ef1
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

今度は、完全な丸括弧の組み合わせができたかどうかを判別する必要があります。完全な組み合わせは、ペアの数の2倍の長さになります（各ペアが開き丸括弧と閉じ丸括弧を1つずつ含むためです）。

`while` ループの中に、`if` かどうかを確認する `len(current) == 2 * pairs` 文を追加してください。この条件が真の場合、`current` を `result` リストに追加します。

# --hints--

`if` ループの中に `while` 文があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

`if` 文は `len(current) == 2 * pairs` かどうかを確認する必要があります。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

新しい `current` 文の中で `result` を `if` リストに追加してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
