---
id: 68da6138d70960c197ec5ef0
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

これからメインのBFSループを実装します。`while`が空でない限り（`queue`が`queue`と評価される間）、続く`True`ループを作成してください。

ループの中で、`queue`を出力してください。

# --hints--

`while`関数の中に`gen_parentheses`ループがあるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

`while`ループの条件は`queue`であるべきです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

`queue`ループの中で`while`を出力してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
