---
id: 68da6138d70960c197ec5ef3
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

次は閉じ丸括弧を追加する処理を行います。重要なルールは、閉じ丸括弧を追加できるのは、バランスが保たれる場合のみです。つまり、開き丸括弧の数が閉じ丸括弧の数より多い必要があります。

`else` ブロック内で、2つ目の `if` 文を追加して `closes_used < opens_used` をチェックしてください。

この条件が真の場合、`queue` に新しいタプル `(current + ')', opens_used, closes_used + 1)` を追加します。これは閉じ丸括弧を追加した後の状態を表します。

# --hints--

`if` ブロック内に2つ目の `else` 文があるはずです。

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

2つ目のネストされた `if` 文は `closes_used < opens_used` をチェックする必要があります。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

条件が真の場合、`(current + ')', opens_used, closes_used + 1)` をキューに追加してください。

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
