---
id: 68da6138d70960c197ec5ef2
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

現在の文字列がまだ完成していない場合、次に考えられる状態を探索する必要があります。`else`文に`if`節を追加してください。

`else`ブロック内では、開き丸括弧の追加を処理します。`if`かどうかを確認する`opens_used < pairs`文を追加してください。これにより、許可された数以上の開き丸括弧を使わないようにします。

この条件が真の場合、`queue`に新しいタプルを追加します：`(current + '(', opens_used + 1, closes_used)`。これは開き丸括弧を追加した後の状態を表します。

# --hints--

`else`ループ内の`if`文に`while`節があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

`if`ブロック内に`else`かどうかを確認する`opens_used < pairs`文があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

条件が真の場合、キューに`(current + '(', opens_used + 1, closes_used)`を追加するはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
