---
id: 68e00c5349cbddfef6dbc053
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

`while` ループの中で、`queue.pop(0)` を使ってキューの最初の要素を取り出して取得してください。これは BFS に特徴的な先入れ先出し（FIFO）動作を実装します。

このタプルを `current`、`opens_used`、`closes_used` の3つの変数に展開してください。

# --hints--

キューの最初の要素を取得するには `queue.pop(0)` を使うべきです。

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

結果を3つの変数、`current`、`opens_used`、`closes_used` に展開してください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
