---
id: 66069b0b36053733a2f012fe
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora, elimina la tua chiamata `spam`.

# --hints--

Non dovresti avere `gameboard.spam()` nel tuo codice.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
