---
id: 6606c14182435d4bab0de2ee
title: Passaggio 30
challengeType: 20
dashedName: step-30
---

# --description--

Adesso creerai un metodo che verifica se un numero può essere inserito in una colonna specifica della griglia di sudoku controllando che il numero non sia già presente in quella colonna.

All'interno della classe `Board`, crea un metodo chiamato `valid_in_col` e dagli tre parametri: `self`, `col` e `num`.
Dove `col` e `num` sono rispettivamente l'indice della colonna e il numero da controllare.

# --hints--

Dovresti creare un nuovo metodo chiamato `valid_in_col` all'interno della classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

Il tuo metodo `valid_in_col` dovrebbe avere tre parametri: `self`, `col` e `num`, in questo ordine.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
