---
id: 6606b224a69a293f98f8db8f
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Il metodo `.index()` genera un'eccezione `ValueError` quando il valore non viene trovato. Per evitare che il programma si blocchi, devi annidare questa riga di codice all'interno di un blocco `try`. L'istruzione `try` serve a racchiudere il codice che potrebbe generare un'eccezione. La clausola `except`, invece, offre un codice alternativo da eseguire se si verifica un'eccezione:

```py
try:
    <code>
except:
    <code>
```

Metti l'assegnazione di `col` all'interno di un blocco `try`. Poi, crea una clausola `except` e riempi il suo corpo con `pass`.

# --hints--

Dovresti mettere l'assegnazione di `col` all'interno di un blocco `try` e creare una clausola `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
