---
id: 6606cd69f56e27516583b0cc
title: Passaggio 37
challengeType: 20
dashedName: step-37
---

# --description--

Ora devi calcolare l'indice della riga di partenza per il quadrato 3x3 all'interno della griglia della board e assicurarti che l'indice della riga di partenza per ogni quadrato 3x3 sia un multiplo di 3.

Questo si può ottenere prendendo il risultato della divisione intera `row // 3` moltiplicato per `3`. Sostituisci `pass` con una variabile `row_start` e assegnale `(row // 3) * 3`.


# --hints--

Dovresti eliminare `pass` e dichiarare una variabile `row_start` con il valore di `(row // 3) * 3`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").find_body().is_equivalent("row_start = (row//3)*3")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    def valid_in_square(self, row, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
