---
id: 6606bc4e5535c0484990ccd5
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

Adesso lavorerai su un metodo che verifica se un numero dato può essere inserito in una riga specifica della griglia del sudoku.

All'interno della classe `Board`, crea un metodo chiamato `valid_in_row` e dagli tre parametri: `self`, `row` e `num`. Dove `self` rappresenta l'istanza della classe, e `row` e `num` sono rispettivamente l'indice della riga e il numero da controllare.


# --hints--

Dovresti creare un nuovo metodo chiamato `valid_in_row` all'interno della classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Il tuo metodo `valid_in_row` dovrebbe avere tre parametri: `self`, `row` e `num`, in questo ordine.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
