---
id: 6606a4641ec48b3a9fe8c2fc
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

La funzione integrata `enumerate` prende un iterabile come argomento e restituisce un oggetto enumerate su cui puoi iterare. Fornisce il conteggio (che per impostazione predefinita parte da zero) e il valore dall'iterabile.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

Il loop dell'esempio sopra produrrebbe in output le tuple `0, a`, `1, b` e `2, c`.

All'interno del metodo `find_empty_cell`, sostituisci `pass` con un loop `for` che usa la funzione `enumerate()` per iterare su ogni riga della griglia sudoku. Usa `row` come indice della riga corrente e `contents` per gli elementi della riga corrente.


# --hints--

Dovresti sostituire `pass` con un loop `for` che itera su `enumerate(self.board)`. Usa `row` e `contents` come variabili del loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
