---
id: 66069b992c1c5e3451f3deb0
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

L'istanza crea un oggetto vuoto. Il metodo `__init__` è un metodo speciale che ti permette di istanziare un oggetto in uno stato personalizzato. Quando una classe implementa un metodo `__init__`, `__init__` viene chiamato automaticamente all'istanza.

All'interno della tua classe `Board`, elimina il metodo `spam` e sostituiscilo con un metodo `__init__` che includa un parametro `self`.

# --hints--

Non dovresti avere un metodo `spam` nella tua classe `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Dovresti definire un metodo `__init__` nella tua classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Il tuo metodo `__init__` dovrebbe avere un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
