---
id: 66069f86f58f85371d47123e
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Un attributo è una variabile associata a un oggetto, usata per archiviare dati come variabili normali.

All’interno del metodo `__init__`, assegna il parametro `board` (che viene passato quando si crea un’istanza della classe `Board`) a un attributo dell’istanza `board` usando `self.board`.

`self.board` si riferisce all’attributo `board` dell’istanza della classe. È una variabile che appartiene all’oggetto creato dalla classe `Board`.

# --hints--

Dovresti eliminare `pass` e assegnare il parametro `board` a `self.board` all’interno del metodo `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
