---
id: 66069167b3307b2f4067b22b
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Una nuova istanza di una classe si crea usando la notazione a funzione, che consiste nell'aggiungere una coppia di parentesi tonde al nome della classe.

All'esterno della definizione della classe, crea un'istanza della classe `Board` e assegnala a una variabile chiamata `gameboard`.

# --hints--

Dovresti dichiarare una variabile `gameboard` all'esterno della classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

La tua variabile `gameboard` dovrebbe avere il valore di `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
