---
id: 6606927d010be4300a4e5330
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

L'istanza crea un oggetto vuoto. Ma le classi possono avere metodi, che sono come funzioni locali per ogni istanza. All'interno di una classe, i metodi si dichiarano nel modo seguente:

```py
class ClassName:
    def method_name():
        pass
```

All'interno della classe `Board`, sostituisci `pass` con un metodo vuoto `spam`.

# --hints--

Dovresti definire un metodo chiamato `spam` all'interno della classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
