---
id: 660699aabc59c532f2d556e5
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Per chiamare un metodo di un'istanza, devi usare la notazione a punto:

```py
instance_name.method_name()
```

Dove `instance_name` è l'istanza o l'oggetto, e `method_name` è il metodo che vuoi chiamare.

Chiama il metodo `spam` dell'oggetto `gameboard`.

# --hints--

Dovresti chiamare il metodo `spam` dell'oggetto `gameboard` con `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
