---
id: 6606bf4561f8794a0d345919
title: Passaggio 27
challengeType: 20
dashedName: step-27
---

# --description--

È il momento di testare il metodo `valid_in_row`. Chiama `valid_in_row` su `gameboard`. Passagli `0` e `8` come argomenti e stampa il risultato.

Ancora una volta, nota come il metodo sia definito con tre parametri, ma viene chiamato con solo due argomenti perché `self` viene passato automaticamente come oggetto su cui il metodo viene chiamato.

# --hints--

Dovresti stampare il risultato della chiamata a `valid_in_row(0, 8)` su `gameboard`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
