---
id: 6606cc754a8834509cd0afb6
title: Passaggio 36
challengeType: 20
dashedName: step-36
---

# --description--

Un'altra cosa da verificare è se un numero può essere inserito in una griglia 3x3.

All'interno della classe `Board`, crea un metodo chiamato `valid_in_square` con quattro parametri: `self`, `row`, `col` e `num`. Dove `row`, `col` e `num` rappresentano rispettivamente l'indice della riga, l'indice della colonna e il numero da controllare.

# --hints--

Dovresti creare un nuovo metodo chiamato `valid_in_square` all'interno della classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

Il tuo metodo `valid_in_square` dovrebbe avere quattro parametri: `self`, `row`, `col` e `num`, in questo ordine.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
