---
id: 66068fb0bfddba2b7977eb60
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto imparerai a conoscere le classi e gli oggetti costruendo un risolutore di sudoku.

In Python, una classe è un modello per creare oggetti. Gli oggetti creati da una classe sono istanze di quella classe. Puoi creare una classe usando questa sintassi:

```py
class ClassName:
    pass
```

Dove `class` è la parola chiave necessaria per definire la classe e `ClassName` è il nome della classe, scritto per convenzione in *PascalCase*.

Inizia creando una classe `Board`.

# --hints--

Dovresti creare una classe chiamata `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
