---
id: 6606933d6813a8308c962dd1
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Per essere un metodo di istanza, un metodo richiede un parametro speciale, chiamato `self` per convenzione. Questo parametro è un riferimento all’istanza della classe e deve essere sempre il primo parametro.

Aggiungi un parametro `self` al tuo metodo `spam`.

# --hints--

Il tuo metodo `spam` dovrebbe avere un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
