---
id: 660699119472f332798860ad
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora, sostituisci `pass` con una chiamata a `print` e passagli la stringa `'Spam!'`.

# --hints--

Dovresti eliminare `pass` e stampare la stringa `'Spam!'` all'interno del metodo `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
