---
id: 66069e5759b800364707988e
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Tornando al metodo `__init__`, richiede un parametro aggiuntivo che rappresenta il puzzle da risolvere.

Aggiungi un secondo parametro chiamato `board` al metodo `__init__` e correggi l'istanza di `gameboard` passando la lista `puzzle` come faresti con un argomento in una chiamata di funzione.

# --hints--

Il tuo metodo `__init__` dovrebbe avere due parametri nell'ordine: `self` e `board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

La tua variabile `gameboard` dovrebbe avere il valore di `Board(puzzle)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
