---
id: 66c051d13a6a20255a963695
title: Створіть стилізований список справ
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

У цій лабораторній роботі ви попрактикуєтесь у використанні різних стилів, які можна застосовувати до посилань при наведенні, фокусуванні, натисканні та переході за ними.

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати невпорядкований список з класом `todo-list`.
2. Невпорядкований список має містити чотири елементи списку.
3. Всередині кожного елемента списку мають бути:

   - елемент `input` з типом `checkbox` та `id` з унікальним значенням;
   - елемент `label` з атрибутом `for`, значення якого відповідає `id` відповідного елемента `input`;
   - невпорядкований список з класом `sub-item`;
   - елемент списку, який містить елемент-посилання. Цей елемент-посилання повинен мати клас `sub-item-link`, дійсне значення `href` і значення `target`, завдяки якому посилання відкривається у новій вкладці.

4. Your `a` elements should not have any text decorations.
5. You should set the text color of unvisited links to a color of your choice.
6. Якщо посилання відвідано, його колір має змінитись на будь-який інший.
7. Якщо посилання наведено, його колір має змінитись на будь-який інший.
8. Якщо посилання сфокусовано, його контур має бути кольоровим.
9. Якщо посилання натиснуто, його колір має змінитись на будь-який інший.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Ви повинні мати невпорядкований список з класом `todo-list`.

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

У невпорядкованому списку мають бути чотири елементи списку.

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

Елементи `li` в межах `ul` з класом `todo-list` повинні мати елемент `input` з `type` зі значенням `checkbox`.

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

Елементи `li` в межах `ul` з класом `todo-list` повинні мати елемент `label`.

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

Всі елементи `input` повинні мати `id`.

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

Всі елементи `label` повинні мати атрибут `for`.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

Всі елементи `label` повинні мати текст.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

Атрибути `id` та `for` парних елементів `input` й `label` повинні мати відповідні значення.

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

Після елементів `label` має бути невпорядкований список з класом `sub-item`.

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

Елементи `li` в межах `ul` з класом `sub-item` повинні мати елемент-посилання з класом `sub-item-link`.

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

Всі елементи-посилання повинні мати дійсне значення атрибута `href`.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

Всі елементи-посилання повинні мати текст.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

Всі елементи `a` повинні мати атрибут `target` зі значенням `_blank`.

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

Your `a` elements should not have any text decorations.

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

Unvisited links should have a text color set to a color of your choice.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

Посилання повинні змінювати колір при наведенні курсору.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

Посилання повинні змінювати колір при натисканні на них.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

Посилання повинні мати контур при фокусуванні.

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

Посилання повинні змінювати колір після переходу на них.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
