---
id: 6859a484a55387c6543ab2c0
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

各惑星の事実を追加する時間です。最初の `p` 要素の中に、次のテキストを追加してください：
`Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.`

2番目の `p` 要素の中に、次のテキストを追加してください：
`Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.`

3番目の `p` 要素の中に、次のテキストを追加してください：
`Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.`

# --hints--

最初の `p` 要素には地球に関する事実を含めてください。テキストが提供された説明と一致していることを確認してください。

```js
assert.equal(
  document.querySelector('div[role="tabpanel"] p').textContent.trim(),
  'Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It\'s the only planet in the solar system known to harbor life.'
);
```

2番目の `p` 要素には土星に関する事実を含めてください。テキストが提供された説明と一致していることを確認してください。

```js
assert.equal(
  document.querySelectorAll('div[role="tabpanel"] p')[1].textContent.trim(),
  'Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It\'s a gas giant with dozens of moons orbiting it.'
);
```

3番目の `p` 要素には火星に関する事実を含めてください。テキストが提供された説明と一致していることを確認してください。

```js
assert.equal(
  document.querySelectorAll('div[role="tabpanel"] p')[2].textContent.trim(),
  'Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It\'s a key focus for exploration in the search for past or present life.'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--
      <div role="tabpanel">
        <p></p>
      </div>
      <div role="tabpanel">
        <p></p>
      </div>
      <div role="tabpanel">
        <p></p>
      </div>
      --fcc-editable-region--
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
