---
id: 68599633407374c0d574587b
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

次のステップは、異なる惑星とその事実を切り替える役割を持つボタンを追加することです。

`tablist` 要素の中に、`button` 属性が `role` に設定された `"tab"` 要素を3つ追加してください。最初の `button` 要素にはテキスト `🌍 Earth` を入れてください。2番目の `button` にはテキスト `🪐 Saturn` を入れてください。3番目の `button` 要素にはテキスト `🔴 Mars` を入れてください。  

# --hints--

`button` 要素の中に `tablist` 要素が3つあるはずです。

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

それぞれの `button` 要素には `role` 属性が `"tab"` に設定されている必要があります。

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

最初の `button` 要素にはテキスト `🌍 Earth` を入れてください。

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

2番目の `button` 要素にはテキスト `🪐 Saturn` を入れてください。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

3番目の `button` 要素にはテキスト `🔴 Mars` を入れてください。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
