---
id: 68599af28ead38c22590a89d
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

ボタンとそれが表すコンテンツをつなげる方法が必要です。これは各ボタンに`aria-controls`属性を使うことで行います。後で、コンテンツと対応する`id`属性をコンテンツ要素に追加します。

まず、各`aria-controls`要素に`button`属性を追加してください。最初のボタンの`aria-controls`属性の値は`"panel-earth"`に設定します。2番目のボタンの`aria-controls`属性は`"panel-saturn"`に設定します。3番目のボタンの`aria-controls`属性は`"panel-mars"`に設定してください。

# --hints--

各`button`要素には`aria-controls`属性が必要です。

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

最初の`button`要素には`aria-controls`属性を`"panel-earth"`に設定してください。

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

2番目の`button`要素には`aria-controls`属性を`"panel-saturn"`に設定してください。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

3番目の`button`要素には`aria-controls`属性を`"panel-mars"`に設定してください。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
