---
id: 6859a1270ac255c571352189
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

各惑星の事実を表示するパネルのマークアップを作成する時間です。

まず、`div` 要素を3つ追加し、それぞれに `role` 属性を `"tabpanel"` に設定してください。各 `div` 要素の中には、`p` 要素が1つ含まれている必要があります。

# --hints--

`div` 要素を3つ用意し、それぞれに `role` 属性を `"tabpanel"` に設定してください。

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"]'), 3);
```

各 `tabpanel` 要素の中には `p` 要素が含まれている必要があります。

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"] p'), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
