---
id: 685976a92ffab8b2dc744b72
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`.tabs` 要素の中に、`div` 属性が `role` に設定され、`"tablist"` 属性が `aria-labelledby` に設定された `"tabs-title"` 要素を追加してください。

ここで `aria-labelledby` 属性は、タブリストを目に見えるラベルと関連付けるために必要で、スクリーンリーダーのような支援技術がタブリストの目的をユーザーに伝えるのに役立ちます。

# --hints--

`div` 要素の中に `.tabs` 要素があるはずです。

```js
assert.exists(document.querySelector(".tabs div"));
```

`div` 要素には `role` 属性が `"tablist"` に設定されているはずです。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

`div` 要素には `aria-labelledby` 属性が `"tabs-title"` に設定されているはずです。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
