---
id: 685c49c6e1db1c74a9762eae
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

HTMLを見ると、最初の`tab`はデフォルトで選択されており、`aria-selected`属性が`"true"`に設定されています。ユーザーが別のタブをクリックしたとき、すべてのタブはまず`aria-selected="false"`に設定される必要があります。

これを実現するには、`setAttribute`メソッドを使って各タブの`aria-selected`属性を変更できます。

以下は、`setAttribute`要素のリストに対して`div`メソッドを使う例です。

```js
const divs = document.querySelectorAll("div");
divs.forEach(div => {
  div.setAttribute("class", "new-class");
});
```

`forEach`コレクションに`tabs`を使って各タブをループし、その`aria-selected`属性を`"false"`に設定してください。コールバック関数のパラメータ名は任意ですが、外側の`tab`ループで既に使われている`forEach`は避けるのが良いです。

# --hints--

ユーザーがタブをクリックしたとき、すべてのタブの`aria-selected`属性を`"false"`に設定してください。

```js
const allTabs = document.querySelectorAll('[role="tab"]');
allTabs.forEach(tab => {
  tab.click();
  assert.equal(tab.getAttribute("aria-selected"), "false");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
const tabs = document.querySelectorAll('[role="tab"]');
const panels = document.querySelectorAll('[role="tabpanel"]');

--fcc-editable-region--
tabs.forEach(tab => {
  tab.addEventListener("click", () => {

  });
});
--fcc-editable-region--
```
