---
id: 685db0bd1280154ae85bcf9f
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、HTML、CSS、JavaScriptを使ってアクセシブルな惑星のタブリストを作成します。HTMLのボイラープレートとCSSは用意されています。

最初の数ステップでは、タブリストの残りのHTMLを作成します。

まず、クラスが`div`に設定された`"tabs"`要素を作成してください。その`div`要素の中に、`h2`属性が`id`に設定され、テキスト内容が`"tabs-title"`の`Planets`要素を追加してください。

# --hints--

`div`要素があるはずです。

```js
assert.exists(document.querySelector("div"));
```

`div`要素のクラスは`"tabs"`に設定してください。

```js
assert.exists(document.querySelector("div.tabs"));
```

`div`要素の中に`h2`要素があるはずです。

```js
assert.exists(document.querySelector(".tabs h2"));
```

`h2`要素の`id`属性は`"tabs-title"`に設定してください。

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

`h2`要素のテキスト内容は`Planets`にしてください。

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
