---
id: ab306dbdcc907c7ddfc30830
title: 深い配列のフラット化ツールを作成する
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

このラボでは、配列をフラット化するアルゴリズムを実装します。

配列のフラット化とは、任意の深さでネストされた配列を、1次元の単一配列に変換することです。この処理では、配列だけを展開し、他の型はそのままにします。

例えば：

| 元の配列 | フラット化後 |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**目的：** 以下のストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ストーリー：**

1. `steamrollArray` という名前の関数を用意してください。
1. `steamrollArray` 関数は1つの引数を受け取る必要があります：ネストされた配列です。
1. 関数は、ネストの深さが異なる配列をフラット化する必要があります。
1. 解答では `Array.prototype.flat()` や `Array.prototype.flatMap()` メソッドを使わないでください。
1. グローバル変数は使わないでください。

# --hints--

`steamrollArray` 関数を用意してください。

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` は `["a", "b"]` を返す必要があります。

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` は `[1, 2, 3, 4]` を返す必要があります。

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` は `[1, 3, 4]` を返す必要があります。

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` は `[1, {}, 3, 4]` を返す必要があります。

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

解答では `Array.prototype.flat()` や `Array.prototype.flatMap()` メソッドを使わないでください。

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

グローバル変数は使わないでください。

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
