---
id: acda2fb1324d9b0fa741e6b5
title: 문자열 끝 확인 도구 만들기
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

이 실습에서는 문자열이 주어진 대상 문자열로 끝나는지 확인하는 함수를 구현할 것입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 확인할 문자열과 비교할 문자열 두 개의 매개변수를 받는 `confirmEnding`라는 이름의 함수를 만들어야 합니다.
2. 첫 번째 문자열이 두 번째 문자열로 끝나면 `true`을 반환하고, 그렇지 않으면 `false`를 반환해야 합니다.
3. `.endsWith()` 메서드를 사용하지 말고, 대신 JavaScript의 문자열 자르기 메서드 중 하나를 사용하세요.

# --hints--

`confirmEnding`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding`는 2개의 매개변수를 받아야 합니다.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")`는 `true`을 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")`는 `false`을 반환해야 합니다.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")`는 `true`를 반환해야 합니다.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

이 실습을 해결할 때 내장 메서드 `.endsWith()`를 사용하지 않아야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
