---
id: 684cca09aac0fc035db4099f
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`p` 요소 안에 `Error! Something went wrong. Please try again.`가 `Error!` 요소로 감싸져야 하는 곳에 `strong` 텍스트를 추가하세요.

# --hints--

`strong` 요소 안에 `p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("p strong"));
```

`strong` 요소에는 `Error!` 텍스트가 있어야 합니다.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

`p` 요소에는 `Something went wrong. Please try again.`로 끝나는 올바른 텍스트가 있어야 합니다.

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
