---
id: 684ccbdc1ef89e042b744e09
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

워크숍의 다음 부분은 닫기 버튼을 추가하는 것입니다.

`p` 요소 아래에 `button` 요소를 추가하고 `aria-label`를 `"Dismiss"`로 설정하세요. 그 `button` 요소 안에 `span` 요소를 `aria-hidden` 속성과 함께 추가하세요. `span` 요소의 텍스트는 `&times;`이어야 합니다.

*참고*: 이 워크숍은 Tailwind CSS 스타일 작업에만 집중하므로 `button`에는 어떤 종류의 상호작용도 추가하지 않습니다.

# --hints--

`button` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("button"));
```
 
`button` 요소에 `aria-label` 속성이 있어야 합니다.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

`button` 요소에 `aria-label` 속성이 `"Dismiss"`로 설정되어야 합니다.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

`button` 요소 안에 `span` 요소가 중첩되어 있어야 합니다.

```js
assert.exists(document.querySelector("button span"));
```

`span` 요소에 `aria-hidden` 속성이 있어야 합니다.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

`span` 요소에 `&times;` 텍스트가 있어야 합니다.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
