---
id: 684cff0bf0a4ca0ad7e42990
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

다음 단계는 테두리의 모서리를 둥글게 만드는 것입니다. Tailwind CSS를 사용해 `border-radius` 속성을 설정하려면, `rounded-size`가 반경을 나타내는 다음 패턴의 유틸리티 클래스를 사용할 수 있습니다: `size`.

다음은 요소에 매우 큰 테두리 반경을 적용하는 예시입니다:

```html
<p class="rounded-xl">This is border radius.</p>
```

열리는 `div` 태그 안에 `md` 크기의 테두리 반경을 적용하세요.

# --hints--

열리는 `div` 태그에 `rounded-md` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
