---
id: 684cd027744c2a0520ba4d58
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS는 마크업에 직접 적용하는 유틸리티 클래스를 사용합니다. 이 클래스들은 폰트, 색상, Flexbox 레이아웃, Grid 등 다양한 스타일링을 포함합니다.

Tailwind CSS에서 배경색을 적용하려면 `bg` 접두사 다음에 색상 이름과 숫자를 사용하세요. 숫자는 색상의 명도를 나타냅니다.

다음은 단락 요소에 어두운 파란색 배경색을 적용하는 예시입니다.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

열리는 `div` 태그에 연한 빨간색 배경색을 나타내는 유틸리티 클래스 `bg-red-100`를 포함한 클래스를 추가하세요.

# --hints--

열리는 `div` 태그에는 클래스 속성이 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

클래스 속성에는 유틸리티 클래스 `bg-red-100`가 포함되어야 합니다.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
