---
id: 684cc59eed2706fa50e8933c
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 오류 메시지 컴포넌트를 만들어 Tailwind CSS 기본을 시작하는 방법을 배우세요.

처음 몇 단계에서는 초기 HTML 마크업을 만드는 데 집중하세요.

`body` 요소 안에 `div` 요소를 추가하세요. 그런 다음 그 `div` 요소 안에 `p` 요소를 추가하세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("div"));
```

`p` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("p"));
```

`p` 요소는 `div` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
