---
id: 684cd59ec94b5e063bf776f0
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

연한 빨간 배경 주위에 테두리가 있으면 좋겠습니다.

Tailwind CSS에서 `border-width`를 설정하려면 `border-number` 유틸리티 클래스를 사용할 수 있으며, 여기서 `number`는 테두리의 밝기나 두께를 나타냅니다.

모든 면에 `border-width`의 `4px`를 적용하는 예는 다음과 같습니다.

```html
<p class="border-4">This is how to set the border-width.</p>
```

열린 `div` 태그 안에 모든 면에 `border-width`로 설정된 `2px`에 맞는 유틸리티 클래스를 추가하세요.

# --hints--

열린 `div` 태그에는 `border-2` 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
