---
id: 684cd8bf5445de089af9c837
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

이제 테두리 색상을 적용할 시간입니다. Tailwind CSS를 사용해 `border-color` 속성을 설정하려면, `border-color-number`가 테두리 색상의 음영을 나타내는 다음 패턴의 유틸리티 클래스를 사용할 수 있습니다: `number`.

다음은 요소에 보라색 테두리 색상을 적용하는 예시입니다:

```html
<p class="border-purple-400">This is border-color.</p>
```

열리는 `div` 태그 안에 `300` 음영의 빨간 테두리 색상을 적용하세요.

# --hints--

열리는 `div` 태그에는 `border-red-300` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
