---
id: 684d03df2a0ddf0c735865f0
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

이제 요소 너비가 `body`의 절반 크기인 것을 볼 수 있습니다. 이는 큰 기기에서는 잘 작동하지만 작은 기기에서는 그렇지 않습니다.

Tailwind CSS는 모바일 우선 접근 방식을 따릅니다. 따라서 다른 기기에 맞는 반응형 디자인을 만들고 싶다면 `md`와 `lg` 같은 Tailwind의 반응형 변형을 이렇게 사용할 수 있습니다:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

다음 예시는 `66%` 이상의 기기 크기에 대해 부모 요소 너비의 `1024px`로 너비를 설정합니다.

열린 `div` 태그 안에서 `w-1/2`를 `768px` 이상의 기기 크기에만 적용되도록 수정하세요. 이 클래스를 `md` 변형으로 접두사할 수 있습니다.

# --hints--

열린 `div` 태그에는 `md:w-1/2` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

열린 `w-1/2` 태그에는 더 이상 `div` 클래스가 없어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
