---
id: 684d1942a1904a14a3709c6c
title: 20단계
challengeType: 0
dashedName: step-20
---

# --description--

에러 메시지 컴포넌트의 스타일링이 거의 완성되었지만, 고쳐야 할 부분이 하나 남아 있습니다. 단락과 버튼 텍스트에 서로 다른 폰트 크기를 설정했기 때문에, 플렉스 항목들 간의 수직 정렬이 맞지 않습니다.

Tailwind CSS를 사용해 컨테이너의 교차 축을 따라 플렉스 항목을 배치하려면, 다음 패턴을 사용할 수 있습니다: `items-value`. `value`는 `start`, `end`, `center`, `baseline` 등 무엇이든 될 수 있습니다.

열리는 `div` 태그 안에 `align-items` 속성을 `center`로 설정하는 올바른 유틸리티 클래스를 사용하세요.

이 마지막 변경으로 에러 메시지 컴포넌트가 완성됩니다!

# --hints--

열리는 `div` 태그에 `items-center` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
