---
id: 684d17a3b350bb13d3c19c0f
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

이제 `button` 요소에 스타일을 적용할 차례입니다.

먼저 `button` 요소에 클래스 속성을 추가하세요. 다음으로, 빨간색 텍스트 색상을 설정하는 올바른 유틸리티 클래스를 `700` 음영으로 적용하세요.

그런 다음, 여는 `button` 태그 안에 `font-size`를 `2xl`로 설정하는 올바른 유틸리티 클래스를 사용하세요. 이는 `1.5rem` 또는 (`24px`)을 나타냅니다.

# --hints--

`button` 요소에는 클래스 속성이 있어야 합니다.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

여는 `button` 태그에는 `text-red-700` 유틸리티 클래스가 있어야 합니다.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

여는 `button` 태그에는 `text-2xl` 유틸리티 클래스가 있어야 합니다.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
