---
id: 684d0d369e2b3b0f288d2c20
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`div` 요소를 수평으로 가운데 정렬하면 좋겠습니다. `margin`을 양쪽에 동시에 적용하려면 다음 패턴을 사용할 수 있습니다: `mx-number` 또는 `my-number`.

`mx`는 좌우 여백을 설정하고, `my`는 상하 여백을 설정합니다.

다음은 요소에 상하 여백을 적용하는 예시입니다:

```html
<p class="my-8">This is margin.</p>
```

열린 `div` 태그 안에 좌우 여백을 `auto`로 설정하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열린 `div` 태그에는 `mx-auto` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
