---
id: 684d0e11094c3e0fe1fb98f5
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

현재 단락 요소가 dismiss `button` 요소 위에 쌓여 있습니다. 하지만 `button` 요소가 단락과 같은 행에 있으면 좋겠습니다. 이럴 때 Flexbox를 사용합니다.

Tailwind CSS에서 `div` 요소를 Flexbox 레이아웃으로 설정하려면 `flex` 유틸리티 클래스를 사용하세요.

# --hints--

열리는 `div` 태그에 `flex` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
