---
id: 684d11ffd15a531225d081e9
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

이제 `div` 요소에 스타일이 적용되었으니 `p` 요소에도 스타일을 적용할 차례입니다.

현재 단락 요소의 텍스트 색상은 검정색입니다. 하지만 색상이 진한 빨간색이면 좋겠습니다.

이전 단계에서 이렇게 배경색을 다루는 방법을 배웠습니다: `bg-red-100`. 텍스트 색상도 이와 비슷한 방식으로 다음 패턴을 따라 적용할 수 있습니다: `text-color-shade`.

다음은 요소에 파란색 텍스트 색상을 적용하는 예시입니다:

```html
<p class="text-blue-500">This has blue text.</p>
```

먼저 `p` 요소에 클래스 속성을 추가하세요. 그런 다음 빨간색 텍스트 색상을 설정하는 올바른 유틸리티 클래스를 적용하되, 색상 음영은 `700`로 설정하세요.

# --hints--

`p` 요소에는 클래스 속성이 있어야 합니다.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

열리는 `p` 태그에는 `text-red-700` 유틸리티 클래스가 있어야 합니다.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
