---
id: 684d09cad4ec900e41db2e74
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

현재 `div` 요소가 미리보기 화면의 왼쪽 상단 모서리에 배치되어 있습니다. 하지만 요소 상단에 여백을 추가하면 좋겠습니다. Tailwind CSS를 사용해 요소에 `margin`을 적용하려면 다음 패턴을 사용할 수 있습니다: `m-number`, 여기서 `number`은 네 면 모두에 적용할 고정 픽셀 수를 나타냅니다.

한쪽 면에만 여백을 적용하려면 다음 패턴 중 하나를 사용할 수 있습니다: `mt-number` (상단 여백), `mb-number` (하단 여백), `mr-number` (오른쪽 여백), `ml-number` (왼쪽 여백).

요소에 `margin-bottom`을 적용하는 예시는 다음과 같습니다:

```html
<p class="mb-8">This is margin.</p>
```

열린 `div` 태그 안에 `margin-top`의 `4px`을 적용하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열린 `div` 태그에는 `mt-4` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
