---
id: 684d105e488dbd1166a717ae
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

미리보기 창 크기를 키우면 모든 항목이 여전히 컨테이너 안에서 왼쪽 정렬되어 있는 것을 볼 수 있습니다. 플렉스 항목이 가운데 정렬되면 좋겠습니다.

Tailwind CSS로 컨테이너의 주축을 따라 플렉스 항목을 배치하려면 다음 패턴을 사용할 수 있습니다: `justify-value`. `value`은 `start`, `end`, `center`, `between` 등일 수 있습니다. `justify-content`와 함께 바닐라 CSS에서 사용하는 모든 유효한 값은 Tailwind CSS에서도 사용할 수 있습니다.

열리는 `div` 태그 안에 `justify-content` 속성을 `center`로 설정하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열리는 `div` 태그에 `justify-center` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
