---
id: 684d160a2370b01310ae0fda
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

새로운 빨간색이 멋지지만, 텍스트가 더 크면 좋겠습니다. Tailwind CSS를 사용해 요소의 `font-size`를 설정하려면 다음 패턴을 사용할 수 있습니다: `text-value`.

요소의 `font-size`를 `0.875rem` 또는 (`14px`)로 설정하는 예시는 다음과 같습니다:

```html
<p class="text-sm">This is for font-size.</p>
```

열리는 `p` 요소 안에, `font-size` 또는 (`xl`)를 나타내는 `1.25rem`로 `20px`를 설정하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열리는 `p` 태그에 `text-xl` 유틸리티 클래스가 있어야 합니다.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
