---
id: 684d009994a3840b98c9aa3a
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

현재 `div` 요소의 너비가 페이지 전체 너비를 차지하고 있습니다. Tailwind CSS에서 너비를 적용하려면 고정 너비에 `w-number` 패턴을 사용하거나 백분율에 `w-fraction` 패턴을 사용하는 등 몇 가지 방법이 있습니다.

다음은 부모 컨테이너 너비의 `66%`로 너비를 설정하는 예시입니다:

```html
<p class="w-2/3">This is a percentage width.</p>
```

열린 `div` 태그 안에 부모 컨테이너 너비의 `50%`로 너비를 설정하세요.

# --hints--

열린 `div` 태그에 `w-1/2` 유틸리티 클래스를 추가하세요.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
