---
id: 684d0eeb9c81e5109e52fe8b
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`div` 요소가 점점 더 좋아지고 있습니다. 하지만 이제 dismiss `button`가 단락 요소와 너무 가까워졌습니다.

플렉스 항목 사이에 공간을 만들려면 다음 패턴을 사용할 수 있습니다: `gap-number`. `number`는 플렉스 레이아웃에서 행과 열 사이의 고정된 공간을 나타냅니다.

요소에 `gap`를 적용한 예시는 다음과 같습니다:

```html
<p class="flex gap-8">This is gap.</p>
```

열린 `div` 태그 안에 `gap`를 `4`로 설정하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열린 `div` 태그에는 `gap-4` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
