---
id: 684d07cbfa43a20d6a002ed7
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`div` 요소의 스타일이 점점 완성되고 있습니다. 하지만 텍스트 주변에 약간의 여백이 있으면 좋겠습니다.

Tailwind CSS를 사용해 요소에 `padding`을 적용하려면, `p-number`이 네 면 모두에 적용할 픽셀 수를 나타내는 `number` 패턴을 사용할 수 있습니다.

다음은 요소에 `padding`의 `8px`를 적용하는 예시입니다:

```html
<p class="p-8">This is padding.</p>
```

열린 `div` 태그 안에 `padding`의 `4px`를 적용하는 올바른 유틸리티 클래스를 사용하세요.

# --hints--

열린 `div` 태그에는 `p-4` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
