---
id: 68501d32b69f9d7756334ffc
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

미리보기 창을 확대하면 `div` 요소가 페이지 가로 중앙에 배치된 것을 볼 수 있습니다. 하지만 작은 기기에서는 `div` 요소의 너비가 페이지 전체 너비를 차지합니다.

`mx-auto` 클래스를 `md` 기기 이상에만 적용되도록 업데이트하세요.

# --hints--

열리는 `div` 태그에 `md:mx-auto` 유틸리티 클래스가 있어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

열리는 `mx-auto` 태그에 더 이상 `div` 클래스가 없어야 합니다.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
