---
id: 657a0ea50da0c8d9d6d7950a
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 프로젝트에서는 주사위 게임을 만들어 알고리즘적 사고를 배우게 됩니다. 총 6라운드가 있으며, 각 라운드마다 플레이어는 주사위를 최대 3번 굴리고 점수를 모을 수 있습니다.

HTML과 CSS는 이미 제공되어 있습니다. 자유롭게 살펴보세요. 준비가 되면 HTML 변수를 설정해야 합니다. 모든 `.die` 요소를 가져와 `listOfAllDice` 변수에 할당하세요. 점수 입력란(`input` 안의 `#score-options` 요소)과 점수 스팬을 가져와 각각 `scoreInputs`와 `scoreSpans`에 할당하세요. `#current-round` 요소를 `roundElement`에, `#current-round-rolls` 요소를 `rollsElement`에 할당한 다음 `#total-score`과 `#score-history` 요소도 같은 방식으로 처리하세요. `#roll-dice-btn`, `#keep-score-btn`, `#rules-btn`, `.rules-container`도 적절한 이름의 변수에 할당하세요.

사용자가 `Show rules` 버튼을 클릭하면 게임 규칙을 보여주거나 숨길 수 있어야 합니다. 게임 규칙 상태를 추적할 `isModalShowing` 변수를 만드세요.

사용자가 주사위를 굴릴 때마다 모든 주사위 값을 추적해야 합니다. 이를 위해 `diceValuesArr` 변수를 만드세요.

게임 내내 현재 점수, 총 점수, 굴린 횟수, 플레이어가 있는 라운드를 추적해야 합니다. 이를 위해 `rolls`, `score`, `round` 변수를 선언하세요.

이 변수들의 시작 값이 무엇이어야 할지 생각해 보세요. 이 변수들은 모두 재할당할 수 있어야 합니다.

# --hints--

`.die` 요소들을 `listOfAllDice`에 할당해야 합니다.

```js
assert.deepEqual(listOfAllDice, document.querySelectorAll(".die"));
```

`#score-options input` 요소들을 `scoreInputs`에 할당해야 합니다.

```js
assert.deepEqual(scoreInputs, document.querySelectorAll("#score-options input"));
```

`#score-options span` 요소들을 `scoreSpans`에 할당해야 합니다.

```js
assert.deepEqual(scoreSpans, document.querySelectorAll("#score-options span"));
```

`#current-round` 요소를 `roundElement`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("current-round"), document.querySelector("#current-round")], roundElement);
```

`#current-round-rolls` 요소를 `rollsElement`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("current-round-rolls"), document.querySelector("#current-round-rolls")], rollsElement);
```

`#total-score` 요소를 `totalScoreElement`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("total-score"), document.querySelector("#total-score")], totalScoreElement);
```

`#score-history` 요소를 `scoreHistory`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("score-history"), document.querySelector("#score-history")], scoreHistory);
```

`#roll-dice-btn` 요소를 `rollDiceBtn`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("roll-dice-btn"), document.querySelector("#roll-dice-btn")], rollDiceBtn);
```

`#keep-score-btn` 요소를 `keepScoreBtn`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("keep-score-btn"), document.querySelector("#keep-score-btn")], keepScoreBtn);
```

`#rules-btn` 요소를 `rulesBtn`에 할당해야 합니다.

```js
assert.deepInclude([document.getElementById("rules-btn"), document.querySelector("#rules-btn")], rulesBtn);
```

`.rules-container` 요소를 `rulesContainer`에 할당해야 합니다.

```js
assert.deepEqual(rulesContainer, document.querySelector(".rules-container"));
```

`isModalShowing` 변수는 `false` 값을 가져야 합니다.

```js
assert.isFalse(isModalShowing);
```

`diceValuesArr` 변수는 빈 배열을 값으로 가져야 합니다.

```js
assert.isArray(diceValuesArr);
assert.isEmpty(diceValuesArr);
```

`rolls` 변수는 `0` 값을 가져야 합니다.

```js
assert.strictEqual(rolls, 0);
```

`score` 변수는 `0` 값을 가져야 합니다.

```js
assert.strictEqual(score, 0);
```

`round` 변수는 `1` 값을 가져야 합니다.

```js
assert.strictEqual(round, 1);
```

모든 새 변수는 `let`로 선언해야 합니다.

```js
assert.match(code, /let isModalShowing/);
assert.match(code, /let diceValuesArr/);
assert.match(code, /let rolls/);
assert.match(code, /let score/);
assert.match(code, /let round/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Advanced Dice Game</title>
  <link rel="stylesheet" href="styles.css" />
</head>

<body>
  <header>
    <h1>Advanced Dice Game</h1>
    <button class="btn" id="rules-btn" type="button">Show rules</button>
    <div class="rules-container">
      <h2>Rules</h2>
      <ul>
        <li>There are total of six rounds</li>
        <li>You can only roll the dice three times per round</li>
        <li>To start the game, roll the dice</li>
        <li>
          Then, choose from one of the selected scores or roll the dice again
        </li>
        <li>
          If you choose a selected score, then you will move to the next round
        </li>
        <li>
          If you decline to choose a selected score, then you can roll the
          dice again two more times
        </li>
      </ul>
      <h2 class="points">Points</h2>
      <ul>
        <li>Three of a kind: Sum of all five dice</li>
        <li>Four of a kind: Sum of all five dice</li>
        <li>Full house: Three of a kind and a pair - 25 points</li>
        <li>
          Small straight: Four of the dice have consecutive values - 30 points
        </li>
        <li>
          Large straight: All five dice have consecutive values - 40 points
        </li>
      </ul>
    </div>
  </header>

  <main>
    <form id="game">
      <div id="dice">
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
        <div class="die"></div>
      </div>

      <p class="rounds-text">
        <strong>Rolls:</strong> <span id="current-round-rolls">0</span> |
        <strong>Round:</strong> <span id="current-round">1</span>
      </p>

      <div id="score-options">
        <div>
          <input type="radio" name="score-options" id="three-of-a-kind" value="three-of-a-kind" disabled />
          <label for="three-of-a-kind">Three of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="four-of-a-kind" value="four-of-a-kind" disabled />
          <label for="four-of-a-kind">Four of a kind<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="full-house" value="full-house" disabled />
          <label for="full-house">Full house<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="small-straight" value="small-straight" disabled />
          <label for="small-straight">Small straight<span></span></label>
        </div>
        <div>
          <input type="radio" name="score-options" id="large-straight" value="large-straight" disabled />
          <label for="large-straight">Large straight<span></span></label>
        </div>

        <div>
          <input type="radio" name="score-options" id="none" value="none" disabled />
          <label for="none">None of the above<span></span></label>
        </div>
      </div>

      <button class="btn" id="keep-score-btn" type="button">
        Keep the above selected score
      </button>
      <button class="btn" id="roll-dice-btn" type="button">
        Roll the dice
      </button>
    </form>

    <div id="scores">
      <h3>Score history (Total score: <span id="total-score">0</span>)</h3>
      <ol id="score-history"></ol>
    </div>
  </main>
  <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

header {
  color: var(--light-grey);
  text-align: center;
}

h1 {
  font-size: 2.5rem;
  margin: 25px 0;
}

.rules-container {
  display: none;
  background-color: var(--light-grey);
  color: var(--black);
  width: 50%;
  margin: 20px auto;
  height: 300px;
  border-radius: 10px;
  overflow-y: scroll;
}

.rules-container ul {
  list-style-type: none;
}

.points {
  margin-top: 15px;
}

main {
  background-color: var(--light-grey);
  padding: 20px;
  display: grid;
  grid-template-columns: repeat(auto-fit, minmax(500px, 1fr));
  gap: 1rem;
  margin: auto;
  justify-items: center;
  width: 50%;
  border-radius: 10px;
}

#dice {
  display: flex;
  justify-content: space-around;
  margin-bottom: 15px;
}

.die {
  width: 40px;
  height: 40px;
  text-align: center;
  margin-right: 15px;
  border: 4px solid var(--black);
  padding: 10px;
}

.rounds-text {
  text-align: center;
}

input[type="radio"]:disabled + label {
  color: var(--grey);
}

#score-history {
  margin-top: 15px;
  text-align: center;
  list-style-position: inside;
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px 0 10px 0.5rem;
  color: var(--black);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover:enabled {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

@media (max-width: 992px) {
  main {
    width: 100%;
  }
}

@media (max-width: 500px) {
  .btn {
    width: 120px;
  }
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
