---
id: 69b279bae0f20eda63163cf2
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

接下來，你需要創建一個迴圈，只要 `leftPart` 和 `rightPart` 中還有元素就持續執行。

為此，創建一個 `while` 迴圈，包含兩個條件：一個檢查 `leftArrayIndex` 是否小於 `leftPart` 的長度，另一個檢查 `rightArrayIndex` 是否小於 `rightPart` 的長度。

# --hints--

你應該創建一個 `while` 迴圈，檢查 `leftArrayIndex` 是否小於 `leftPart` 的長度，以及 `rightArrayIndex` 是否小於 `rightPart` 的長度。

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
