---
id: 69b2762b17d23fbee40a6c00
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

合併排序演算法執行三個動作：

- 將未排序的項目序列分割成子零件
- 對子零件中的項目進行排序
- 合併已排序的子零件

上述過程會遞迴發生，直到子零件合併成完整排序的序列。從將序列分割開始。

首先，在 `mergeSort` 函式（程式）內，使用 `middlePoint` 宣告一個變數 `const` 並指定 `Math.floor(array.length / 2)` 給它，這樣你就可以在遞迴之前將序列分成兩半。

# --hints--

你應該宣告一個 `middlePoint` 變數並指定 `Math.floor(array.length / 2)` 給它。

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
