---
id: 69b279bae0f20eda63163cf0
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

在你的函式（程式）主體底部，再次呼叫 `mergeSort()` 函式（程式）。這次，將 `rightPart` 作為引數傳入函式（程式）呼叫中。

# --hints--

你應該在函式（程式）主體底部使用引數 `mergeSort()` 呼叫 `rightPart`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
