---
id: 69b2713ff6c21fed2d3ad88b
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

合併排序是一種使用分而治之原則來排序資料群集的排序演算法。也就是說，它會將一個群集「分割」成較小的子部分，並透過獨立排序這些子部分來「征服」它們，最後再將排序好的子部分合併。

首先建立一個名為 `mergeSort` 的函式（程式），並帶有參數 `array`。

# --hints--

你應該宣告一個名為 `mergeSort` 的函式（程式），並帶有一個 `array` 參數。

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
