---
id: 69b278d82f0e514f32acb724
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

在切出左半邊後，你仍然需要右半邊來完成合併步驟。

使用 slice 語法來擷取 `array` 的右半邊，並將其指定給名為 `rightPart` 的變數。

# --hints--

你應該在 `rightPart` 函式內有一個 `mergeSort` 變數。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

你應該使用 slice 語法來擷取陣列的右半邊，並將其指定給 `rightPart` 變數。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
