---
id: 69b279bae0f20eda63163cf8
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

你在前面步驟創建的 `if` 和 `else` 述語會將元素指定到 sorted 陣列。

每個指定到 sorted 陣列的元素都會佔據列表中的一個索引。因此，每次指派後，你必須移動到 sorted 陣列的下一個索引。

在 `if`/`else` 區段下方，但仍在 `while` 迴圈內，將 `sortedIndex` 增加 `1`。

# --hints--

你應該在 `sortedIndex` 語句之後，將 `1` 的當前值增加 `else`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
