---
id: 69b279bae0f20eda63163cef
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

現在你已經將 `array` 列表分成兩個獨立的列表，接下來你要持續將每個列表再分割，直到每個元素都獨立成為一個只有單一元素的列表。只有一個數字的列表一定是已排序的。

為此，請在你的函式（程式）中遞迴呼叫 `mergeSort`，並將 `leftPart` 作為引數傳入該呼叫。

# --hints--

你應該在同一個 `mergeSort(leftPart)` 函式（程式）中呼叫 `mergeSort`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
