---
id: 69b279bae0f20eda63163cf9
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

你創建的 `while` 迴圈會比較 `leftPart` 中的一個元素與 `rightPart` 中的另一個元素，然後將較小的元素添加到主要的 `array` 列表中。

它會持續執行這個操作直到沒有元素可比較為止。但 `leftPart` 可能仍有元素剩下，而 `rightPart` 則沒有，反之亦然。

創建另一個 `while` 迴圈，當 `leftArrayIndex` 小於 `leftPart` 的長度時執行。在接下來的步驟中，你會用它將 `leftPart` 中剩下的元素複製到 `array` 列表中。

# --hints--

你應該創建一個 `while` 迴圈，條件是 `leftArrayIndex < leftPart.length`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
