---
id: 69b279bae0f20eda63163cf3
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在迴圈內，開始比較每一半中目前的元素，這樣你才能先合併較小的值。

新增一個 `if` 區塊，檢查 `leftPart` 中目前的元素是否小於 `rightPart` 中目前的元素。 
# --hints--

`if` 條件應該檢查 `leftPart[leftArrayIndex]` 是否小於 `rightPart[rightArrayIndex]`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
