---
id: 69b279bae0f20eda63163cf1
title: 第 7 步驟
challengeType: 1
dashedName: step-7
---

# --description--

現在是時候將列表 `leftPart` 和 `rightPart` 排序並合併回原本的 `array`。

你可以透過比較兩個列表中的元素，將較小的元素合併到主要列表中。你將對 `leftPart` 和 `rightPart` 中的所有索引進行這樣的比較。

創建三個變數：`leftArrayIndex`、`rightArrayIndex` 和 `sortedIndex`，並將它們的值設為 `0`。這些變數將幫助你在排序過程中追蹤每個索引。

# --hints--

你應該宣告一個 `leftArrayIndex` 變數，並初始化為 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

你應該宣告一個 `rightArrayIndex` 變數，並初始化為 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

你應該宣告一個 `sortedIndex` 變數，並初始化為 `0`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
