---
id: 69b2762b17d23fbee40a6c00
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

合併排序演算法執行三個動作：

- 將未排序的項目序列分割成子部分
- 對子部分中的項目進行排序
- 合併已排序的子部分

上述過程會遞迴進行，直到子部分合併成完整的排序序列。讓我們先從分割序列開始。

首先，在 `mergeSort` 函式（程式）內，使用 `middlePoint` 宣告一個變數 `const` 並指定 `Math.floor(array.length / 2)` 給它，這樣你就可以在遞迴之前將序列分成兩半。

# --hints--

你應該宣告一個 `middlePoint` 變數並指定 `Math.floor(array.length / 2)` 給它。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
