---
id: 69b278519ef46b2fd8bf1546
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

你已經有一個 `middlePoint` 值。使用它將陣列分成兩半，並準備好左半部進行排序。

使用 slice 語法擷取 `array` 的左半部，並指定給名為 `leftPart` 的變數。

# --hints--

你應該在 `leftPart` 函式內宣告一個 `mergeSort` 變數。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

你應該將 `array.slice(0, middlePoint)` 指定給 `leftPart`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
