---
id: 69b279bae0f20eda63163cf4
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

當 `if` 條件被評估為 `true` 時，代表 `leftPart` 列表中的元素比它在 `rightPart` 列表中被比較的元素還要小。

在這種情況下，你可以將在 `leftArrayIndex` 找到的元素放入已排序陣列中的 `leftPart`。

在 `if` 區塊內，將 `leftPart[leftArrayIndex]` 指定給 `array[sortedIndex]`。

# --hints--

你應該在 `leftPart[leftArrayIndex]` 內將 `array[sortedIndex]` 指定給 `if`。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
