---
id: 69b278519ef46b2fd8bf1546
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

你已經有一個 `middlePoint` 值。使用它將陣列分成兩半，並準備好左半部進行排序。

使用 slice 語法擷取 `array` 的左半部，並指定給名為 `leftPart` 的變數。

# --hints--

你應該在 `mergeSort` 函式內宣告一個 `leftPart` 變數。

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
```

你應該將 `array.slice(0, middlePoint)` 指定給 `leftPart`。

```js
const explorer = await __helpers.Explorer(code);
const { leftPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftPart,
  'You should declare a leftPart variable inside mergeSort'
);
assert.isTrue(
  leftPart.value.matches('array.slice(0, middlePoint)'),
  'You should assign array.slice(0, middlePoint) to leftPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
