---
id: 69b2713ff6c21fed2d3ad88b
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

合併排序是一種使用分而治之原則來排序資料群集的排序演算法。也就是說，它將一個群集「分割」成較小的子零件，並透過獨立排序這些子零件來「征服」它們，然後合併已排序的子零件。

首先建立一個名為 `mergeSort` 的函式（程式），並帶有參數 `array`。

# --hints--

你應該宣告一個名為 `mergeSort` 的函式（程式），並帶有一個 `array` 參數。

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
