---
id: 69b279bae0f20eda63163cfe
title: 第 20 步
challengeType: 1
dashedName: step-20
---

# --description--

在測試 `mergeSort()` 函式（程式）之前，你需要創建一個基底情況，當 `array` 的長度小於或等於 `1` 時，停止函式（程式）執行。

這個基底情況會停止遞迴呼叫。若沒有它，合併排序操作將會持續執行，即使列表已經排序完成或是沒有元素。

在函式（程式）宣告式之後，創建一個 `if` 述語，檢查 `array.length` 是否小於或等於 1。在 `if` 述語的主體中，加入一個 `return` 述語來停止函式（程式）的執行。

# --hints--

你應該創建一個帶有條件 `if` 的 `array.length <= 1` 述語。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

你應該在 `return` 述語的主體中有一個 `if` 述語。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
