---
id: 69b279bae0f20eda63163cf6
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在前面的步驟中，你將 `leftPart` 中的元素指定給 `array` 列表，因為它比較小。但情況不一定總是如此。在某些比較情況下，右邊的元素可能比較小。

創建一個 `else` 語句，當 `leftPart` 中的元素不小於 `rightPart` 中的元素時執行。

在 `else` 區塊內，將 `rightPart[rightArrayIndex]` 指定給 `array[sortedIndex]`。

# --hints--

你應該創建一個包含 `else` 的 `array[sortedIndex] = rightPart[rightArrayIndex]` 區塊。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
