---
id: 594810f028c0303b75339ad2
title: 向量叉積
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

向量被定義為具有三個維度，表述為由三個數字組成的有序群集：(X、Y、Z)。

# --instructions--

撰寫一個函式（程式），該函式接受兩個向量（陣列）作為輸入並計算它們的叉積。當輸入無效，例如向量長度不同時，你的函式應傳回 `null`。

# --hints--

`crossProduct` 應該是一個函式（程式）。

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()` 應該傳回 null。

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])` 應該傳回 `[-3, 6, -3]`。

```js
assert.deepEqual(res12, exp12);
```

# --seed--

## --after-user-code--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
