---
id: 594810f028c0303b75339ad1
title: 快樂數
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

快樂數是由以下處理程序定義的：

從任何正整數開始，將該數字取代為其各位數字平方和，並重複此處理程序，直到該數字相等於 `1`（此時會保持不變），或是無限迴圈在不包含 `1` 的循環中。那些此處理程序最終以 `1` 結束的數字稱為快樂數字，而那些不以 `1` 結束的則稱為不快樂數字。

# --instructions--

實作一個函式，如果數字是快樂數則傳回 true，否則傳回 false。

# --hints--

`happy` 應該是一個函式（程式）。

```js
assert(typeof happy === 'function');
```

`happy(1)` 應該傳回一個布林值。

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` 應該傳回 `true`。

```js
assert(happy(1));
```

`happy(2)` 應該傳回 `false`。

```js
assert(!happy(2));
```

`happy(7)` 應該傳回 `true`。

```js
assert(happy(7));
```

`happy(10)` 應該傳回 `true`。

```js
assert(happy(10));
```

`happy(13)` 應該傳回 `true`。

```js
assert(happy(13));
```

`happy(19)` 應該傳回 `true`。

```js
assert(happy(19));
```

`happy(23)` 應該傳回 `true`。

```js
assert(happy(23));
```

`happy(28)` 應該傳回 `true`。

```js
assert(happy(28));
```

`happy(31)` 應該傳回 `true`。

```js
assert(happy(31));
```

`happy(32)` 應該傳回 `true`。

```js
assert(happy(32));
```

`happy(33)` 應該傳回 `false`。

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
