---
id: 594810f028c0303b75339ad8
title: 之字形矩陣
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

「之」字形陣列是將前 $N^2$ 個整數以方陣排列，數字沿著陣列的反對角線以之字形順序遞增。

例如，對於輸入 `5`，應該產生以下結果：

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

撰寫一個函式（程式），該函式（程式）接受之字形矩陣的大小，並傳回對應的矩陣作為二維陣列。

# --hints--

ZigZagMatrix 應該是一個函式（程式）。

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

ZigZagMatrix 應該傳回陣列。

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

ZigZagMatrix 應該傳回一個巢狀陣列的陣列。

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

ZigZagMatrix(1) 應該傳回 \[[0]]。

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

ZigZagMatrix(2) 應該傳回 \[[0, 1], [2, 3]]。

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

ZigZagMatrix(5) 應該傳回指定的矩陣。

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --after-user-code--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
