---
id: 594810f028c0303b75339acd
title: '豐數、虧數和完全數分類'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

這些定義了基於其真因數的三種正整數類別。

設 $P(n)$ 為 `n` 的適當因數之和，其中適當因數是指所有正整數 `n` 除了 `n` 本身以外。

如果 `P(n) < n`，則 `n` 被歸類為 `deficient`

如果 `P(n) === n`，則 `n` 被歸類為 `perfect`

如果 `P(n) > n`，則 `n` 被歸類為 `abundant`

**範例**：`6` 的適當因數有 `1`、`2` 和 `3`。`1 + 2 + 3 = 6`，所以 `6` 被歸類為完美數。

# --instructions--

實作一個函式（程式），計算從 `1` 到 `num`（包含）中有多少整數屬於這三個類別。將結果以陣列形式輸出，格式如下 `[deficient, perfect, abundant]`。

# --hints--

`getDPA` 應該是一個函式（程式）。

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` 應該傳回一個陣列。

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` 傳回的陣列長度應該為 `3`。

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` 應該傳回 `[3758, 3, 1239]`。

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` 應該傳回 `[7508, 4, 2488]`。

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` 應該傳回 `[15043, 4, 4953]`。

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
