---
id: 594810f028c0303b75339acc
title: ABC 問題
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

你有一組 ABC 區塊群集（例如，童年字母區塊）。共有 20 個區塊，每個區塊有兩個字母。所有區塊的各面保證包含完整字母表。區塊群集範例：

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

實作一個函式（程式），該函式接收一個字串（字、單字），並判斷該字是否可以用給定的區塊群集拼出來。

需要記住的一些規則：

<ul>
  <li>Once a letter on a block is used, that block cannot be used again.</li>
  <li>The function should be case-insensitive.</li>
</ul>

# --hints--

`canMakeWord` 應該是一個函式（程式）。

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` 應該傳回一個布林值。

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` 應該傳回 true。

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` 應該傳回 false。

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` 應該傳回 true。

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` 應該傳回 false。

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` 應該傳回 true。

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` 應該傳回 true。

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --after-user-code--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
