---
id: 594810f028c0303b75339ad3
title: 向量點積
challengeType: 1
forumTopicId: 302343
dashedName: vector-dot-product
---

# --description--

一個向量可以有一個或多個由有序群集表述的值。範例可以是 (x)、(x, y) 或 (x, y, z)。

# --instructions--

撰寫一個函式，該函式接受兩個向量（表述為一維陣列）作為輸入並計算它們的內積。若輸入無效，例如向量長度不同或傳入的不是兩個向量，則你的函式應傳回 `null`。

# --hints--

`dotProduct` 應該是一個函式（程式）。

```js
assert.equal(typeof dotProduct, 'function');
```

`dotProduct()` 應該傳回 `null`。

```js
assert.equal(dotProduct(), null);
```

`dotProduct([1], [1])` 應該傳回 `1`。

```js
assert.equal(dotProduct([1], [1]), 1);
```

`dotProduct([1], [1, 2])` 應該傳回 `null`。

```js
assert.equal(dotProduct([1], [1, 2]), null);
```

`dotProduct([1, 3, -5], [4, -2, -1])` 應該傳回 `3`。

```js
assert.equal(dotProduct([1, 3, -5], [4, -2, -1]), 3);
```

`dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1])` 應該傳回 `null`。

```js
assert.equal(dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1]), null);
```

`dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ])` 應該傳回 `360`。

```js
assert.equal(dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ]), 360);
```

# --seed--

## --seed-contents--

```js
function dotProduct(...vectors) {

}
```

# --solutions--

```js
function dotProduct(...vectors) {
  if (!vectors || !vectors.length || vectors.length > 2 || vectors[0].length !== vectors[1].length) {
    return null;
  }
  const vectorLen = vectors[0].length;

  let prod = 0;
  let sum = 0;
  let j = vectorLen;
  let i = 2;
  // Sum terms
  while (j--) {
    i = 2;
    prod = 1;

    while (i--) {
      prod *= vectors[i][j];
    }
    sum += prod;
  }
  return sum;
}
```
