---
id: 594810f028c0303b75339ace
title: 累加器工廠
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Paul Graham 提出的一個問題是創建一個函式，該函式接受單一（數值型）引數，並傳回另一個函式，該函式是一個累加器。傳回的累加器函式同樣接受單一數值型引數，並傳回到目前為止傳入該累加器的所有數值（包括建立累加器時傳入的初始值）之和。

# --instructions--

創建一個函式（程式），該函式接收一個數字 $n$ 並產生累加器函式（程式），這些函式（程式）會傳回所有傳入它們的數字的總和。

**規則：**

請勿使用全域變數。

**提示：**

閉包會存檔外部狀態。

# --hints--

`accumulator` 應該是一個函式（程式）。

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)` 應該傳回一個函式（程式）。

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)` 應該傳回一個數字。

```js
assert(typeof accumulator(0)(2) === 'number');
```

傳入值 3、-4、1.5 和 5 應該會傳回 5.5。

```js
assert(testFn(5) === 5.5);
```

# --seed--

## --after-user-code--

```js
const testFn = typeof accumulator(3) === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
