---
id: 5949b579404977fbaefcd736
title: 9 十億個神的名字 整數
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

這個任務是 Arthur C. Clarke 短篇小說的變體。

解決方案應該意識到完成此任務的後果。

詳細來說，要指定「name」的含義：

<ul>
  <li>The integer 1 has 1 name "1".</li>
  <li>The integer 2 has 2 names "1+1" and "2".</li>
  <li>The integer 3 has 3 names "1+1+1", "2+1",  and "3".</li>
  <li>The integer 4 has 5 names "1+1+1+1", "2+1+1", "2+2", "3+1", "4".</li>
  <li>The integer 5 has 7 names "1+1+1+1+1", "2+1+1+1", "2+2+1", "3+1+1", "3+2", "4+1", "5".</li>
</ul>

這可以在以下表單中視覺化：

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

其中列 $n$ 對應整數 $n$，而第 $m$ 列中從左到右的每個行 $C$ 對應以 $C$ 開頭的名稱數量。

選用地註記第 $n$ 列的總和 $P(n)$ 是整數分割函數。

# --instructions--

實作一個函式（程式），傳回第 $n$ 列的總和。

# --hints--

`numberOfNames` 應該是函式（程式）。

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` 應該相等 7。

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` 應該相等 77。

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` 應該相等 385。

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` 應該相等 1255。

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` 應該相等 53174。

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` 應該相等 2552338241。

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
