---
id: 594da033de4190850b893874
title: 平均值／均方根
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

計算包含 1 到 10 的數字的均方根（RMS）。

<abbr title="平方根平均值">RMS</abbr> 是透過取數字平方平均值的平方根計算得出，其公式為：

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

`rms` 應該是一個函式（程式）。

```js
assert(typeof rms === 'function');
```

`rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` 應該相等於 `6.2048368229954285`。

```js
assert.equal(rms(arr1), answer1);
```

# --seed--

## --after-user-code--

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
```

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
