---
id: 594dc6c729e5700999302b45
title: 平衡括號
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

判斷一個產生的括號字串是否平衡；也就是說，該字串是否完全由開啟/關閉括號的對組（依此順序）組成，且沒有任何錯誤巢狀。

**範例：**
| Input                     | Output |
| ------------------------- | ------ |
| <code>\[]</code>          | true   |
| <code>]\[</code>          | false  |
| <code>[][]</code>         | true   |
| <code>]\[]</code>         | false  |
| <code>\[]]\[\[]</code>    | false  |
| <code>\[\[\[\[]]]]</code> | true   |

# --hints--

`isBalanced` 應該是一個函式（程式）。

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` 應該傳回 true。

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` 應該傳回 false。

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` 應該傳回 true。

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` 應該傳回 false。

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` 應該傳回 false。

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` 應該傳回 false。

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` 應該傳回 false。

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` 應該傳回 false。

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` 應該傳回 false。

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` 應該傳回 false。

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` 應該傳回 false。

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` 應該傳回 true。

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` 應該傳回 true。

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` 應該傳回 false。

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` 應該傳回 false。

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` 應該傳回 false。

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` 應該傳回 false。

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` 應該傳回 true。

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --after-user-code--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
