---
id: 594d8d0ab97724821379b1e6
title: 平均值／眾數
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

撰寫一個函式（程式）`mode`，用來尋找在陣列中出現次數最多的值。

可以忽略群集為空的情況。必須小心處理模式為非唯一的情況。

如果不適合或無法支援一般群集，則在可能的情況下使用向量（陣列）。如果不適合或無法支援未指定的值型別，則使用整數。

# --hints--

`mode` 應該是一個函式（程式）。

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])` 應該相等於 `[6]`

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])` 應該相等於 `[1, 4]`。

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --after-user-code--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
