---
id: 594ecc0d9a8cf816e3340187
title: 計程車數字
challengeType: 1
forumTopicId: 302337
dashedName: taxicab-numbers
---

# --description--

計程車數（此處使用的定義）是可以用兩種以上方式表示為兩個正立方數之和的正整數。

第一個計程車數是 `1729`，它是：

1<sup>3</sup> + 12<sup>3</sup> 和

9<sup>3</sup> + 10<sup>3</sup>。

計程車數字也稱為：

<ul>
  <li>taxi numbers</li>
  <li>taxi-cab numbers</li>
  <li>taxi cab numbers</li>
  <li>Hardy-Ramanujan numbers</li>
</ul>

# --instructions--

撰寫一個函式（程式），傳回最低的 `n` 個 taxicab 數。對於每個 taxicab 數，顯示該數字以及其組成的立方數。

# --hints--

`taxicabNumbers` 應該是一個函式（程式）。

```js
assert(typeof taxicabNumbers === 'function');
```

`taxicabNumbers` 應該傳回一個陣列。

```js
assert(typeof taxicabNumbers(2) === 'object');
```

`taxicabNumbers` 應該傳回一個數字陣列。

```js
assert(typeof taxicabNumbers(100)[0] === 'number');
```

`taxicabNumbers(4)` 應該傳回 [1729, 4104, 13832, 20683]。

```js
assert.deepEqual(taxicabNumbers(4), res4);
```

`taxicabNumbers(25)` 應該傳回 [1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656, 110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763, 373464, 402597]

```js
assert.deepEqual(taxicabNumbers(25), res25);
```

`taxicabNumbers(39)` 產生的 20 - 29 之間的數字應該是 [314496,320264,327763,373464,402597,439101,443889,513000,513856]。

```js
assert.deepEqual(taxicabNumbers(39).slice(20, 29), res39From20To29);
```

# --seed--

## --after-user-code--

```js
const res4 = [1729, 4104, 13832, 20683];
const res25 = [
  1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656,
  110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763,
  373464, 402597
];

const res39From20To29 = [314496, 320264, 327763, 373464, 402597, 439101, 443889, 513000, 513856];
```

## --seed-contents--

```js
function taxicabNumbers(n) {

  return true;
}
```

# --solutions--

```js
function taxicabNumbers(nNumbers) {
  const cubeN = [];
  const s3s = {};

  const e = 100;
  for (let n = 1; n < e; n += 1) {
    cubeN[n] = n * n * n;
  }

  for (let a = 1; a < e - 1; a += 1) {
    const a3 = cubeN[a];
    for (let b = a; b < e; b += 1) {
      const b3 = cubeN[b];
      const s3 = a3 + b3;

      let abs = s3s[s3];
      if (!abs) {
        s3s[s3] = abs = [];
      }
      abs.push([a, b]);
    }
  }

  let i = 0;
  const res = [];
  Object.keys(s3s).forEach(s3 => {
    const abs = s3s[s3];
    if (abs.length >= 2) { // No two cube pairs found
      i += 1;
      if (i <= nNumbers) {
        res.push(s3);
      }
    }
  });
  return res.map(item => parseInt(item, 10));
}
```
