---
id: 5951ed8945deab770972ae56
title: 河內塔
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

解決河內塔問題。物件的數量將作為第一個參數給出，接著是用來識別每個物件堆疊的字串。創建一個巢狀陣列，包含移動列表，`["source", "destination"]`。
 
例如，參數 `(4, 'A', 'B', 'C')`，將產生巢狀的移動陣列 `[['A', 'C'], ['B', 'A']]`，表示第 1 次移動是將一個物件從堆疊 `A` 移動到 `C`，第 2 次移動是將一個物件從堆疊 `B` 移動到 `A`。

# --instructions--

撰寫一個函式（程式），傳回將物件堆疊在巢狀陣列中的移動步驟。

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi` 應該是一個函式（程式）。

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)` 應該傳回 7 次移動。

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')` 應該傳回 `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`。

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

`towerOfHanoi(5, "X", "Y", "Z")` 第 10 步應該是 Y -> X。

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')` 前十個移動應該是 `[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
