---
id: 595668ca4cfe1af2fb9818d4
title: Harshad 或 Niven 數列
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

Harshad 或 Niven 數是大於或等於 1 的正整數，且可被其各位數字之和整除。

例如，`42` 是一個 Harshad 數，因為 `42` 可以被 `(4 + 2)` 整除且沒有餘數。

假設該數列定義為遞增順序的數字。

# --instructions--

實作一個函式（程式）來產生 Harshad 數列的連續成員。

用它來傳回一個包含十個成員的陣列，該序列從大於 `n` 的第一個 Harshad 數開始。

# --hints--

`isHarshadOrNiven` 應該是一個函式（程式）。

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)` 應該傳回 `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)` 應該傳回 `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)` 應該傳回 `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
