---
id: 59637c4d89f6786115efd814
title: 霍夫斯塔特 Q 序列
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

Hofstadter Q 序列定義為：

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\\big(n-Q(n-1)\\big)+Q\\big(n-Q(n-2)), \\quad n>2.$

它的定義類似費波那契數列，但費波那契數列中的下一項是前兩項的和，而在 Q 數列中，前兩項告訴你要往回多少項去找兩個數字，將它們相加以產生數列的下一項。

# --instructions--

將 Hofstadter Q Sequence 方程式實作為一個函式。該函式應接受數字 `n`，並傳回一個整數。

# --hints--

`hofstadterQ` 應該是一個函式（程式）。

```js
assert(typeof hofstadterQ === 'function');
```

`hofstadterQ()` 應該傳回 `integer`

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

`hofstadterQ(1000)` 應該傳回 `502`

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

`hofstadterQ(1500)` 應該傳回 `755`

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

`hofstadterQ(2000)` 應該傳回 `1005`

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

`hofstadterQ(2500)` 應該傳回 `1261`

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --after-user-code--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
