---
id: 59669d08d75b60482359409f
title: 日期格式
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

傳回一個包含兩個當前日期字串的陣列，符合以下規格：

- 第一個字串的日期順序應該是以年份數字、月份數字和日期數字，並以破折號（`-`）分隔。
- 第一個字串的年份應為四位數長度。
- 第一個字串的月份和日期不應包含任何前導零。
- 第二個字串的星期幾和月份名稱不應縮寫。
- 第二個字串的日期不應包含任何前導零。

範例輸出：

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` 應該是一個函式（程式）。

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` 應該傳回一個物件。

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` 應該傳回一個有 2 個元素的陣列。

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` 應該傳回正確格式的正確日期

```js
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --after-user-code--

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
```

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
