---
id: 5966f99c45e8976909a85575
title: 星期幾
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

一家公司決定，只要 Xmas 落在星期日，他們就會給員工所有額外的有薪假期，這樣員工連同任何公共假期在內，就不必在接下來的那一週（12 月 25 日到 1 月 1 日之間）工作。

# --instructions--

撰寫一個函式（程式），該函式（程式）接受一個開始年份和一個結束年份，並傳回一個包含所有 12 月 25 日為星期日的年份的陣列。

# --hints--

`findXmasSunday` 應該是一個函式（程式）。

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)` 應該傳回一個陣列。

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)` 應該傳回 `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)` 應該傳回 `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --after-user-code--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
