---
id: 5966c21cf732a95f1b67dd28
title: 日期操作
challengeType: 1
forumTopicId: 302244
dashedName: date-manipulation
---

# --description--

給定一個 EST 的日期字串，輸出將時間加上 12 小時後的日期字串。時區應保持不變。

2009 年 3 月 6 日 7:30pm EST

2009 年 3 月 7 日 上午 7:30 EST

# --hints--

`add12Hours` 應該是一個函式（程式）。

```js
assert(typeof add12Hours === 'function');
```

`add12Hours(dateString)` 應該傳回一個字串。

```js
assert(typeof add12Hours('January 17 2017 11:43am EST') === 'string');
```

`add12Hours("January 17 2017 11:43am EST")` 應該傳回 `"January 17 2017 11:43pm EST"`

```js
assert(
  add12Hours('January 17 2017 11:43am EST') === 'January 17 2017 11:43pm EST'
);
```

應該處理日期變更。`add12Hours("March 6 2009 7:30pm EST")` 應該傳回 `"March 7 2009 7:30am EST"`

```js
assert(add12Hours('March 6 2009 7:30pm EST') === 'March 7 2009 7:30am EST');
```

應該處理閏年的月份變更。`add12Hours("February 29 2004 9:15pm EST")` 應該傳回 `"March 1 2004 9:15am EST"`

```js
assert(add12Hours('February 29 2004 9:15pm EST') === 'March 1 2004 9:15am EST');
```

應該處理平年中的月份變更。`add12Hours("February 28 1999 3:15pm EST")` 應該傳回 `"March 1 1999 3:15am EST"`

```js
assert(add12Hours('February 28 1999 3:15pm EST') === 'March 1 1999 3:15am EST');
```

應該處理年份變更。`add12Hours("December 31 2020 1:45pm EST")` 應該傳回 `"January 1 2021 1:45am EST"`

```js
assert(
  add12Hours('December 31 2020 1:45pm EST') === 'January 1 2021 1:45am EST'
);
```

# --seed--

## --seed-contents--

```js
function add12Hours(dateString) {

  return true;
}
```

# --solutions--

```js
function add12Hours(dateString) {
  const months = ['January', 'February', 'March', 'April', 'May', 'June',
    'July', 'August', 'September', 'October', 'November', 'December'];
  // Get the parts of the string
  const parts = dateString.split(' ');
  const month = months.indexOf(parts[0]);
  const day = parseInt(parts[1], 10);
  const year = parseInt(parts[2], 10);
  const time = parts[3].split(':');
  let hours = parseInt(time[0], 10);
  if (time[1].slice(-2) === 'pm') {
    hours += 12;
  }
  const minutes = parseInt(time[1].slice(0, -2), 10);

  // Create a date with given parts, and updated hours
  const date = new Date();
  date.setFullYear(year, month, day);
  date.setHours(hours + 12);  // Add 12 hours
  date.setMinutes(minutes);

  let hoursOutput = date.getHours();
  let abbreviation = 'am';
  if (hoursOutput > 12) {
    hoursOutput -= 12;
    abbreviation = 'pm';
  }

  return `${months[date.getMonth()]} ${date.getDate()} ${date.getFullYear()} ${hoursOutput}:${date.getMinutes()}${abbreviation} EST`;
}
```
