---
id: 596a8888ab7c01048de257d5
title: 深拷貝
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

撰寫一個函式（程式），傳回給定物件的深層複本。該複本不得與所給的物件相同。

此任務不會測試：

<ul>
  <li>Objects with properties that are functions</li>
  <li>Date objects or object with properties that are Date objects</li>
  <li>RegEx or object with properties that are RegEx objects</li>
  <li>Prototype copying</li>
</ul>

# --hints--

`deepcopy` 應該是一個函式（程式）。

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` 應該傳回一個物件。

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` 不應該傳回提供的相同物件。

```js
assert(deepcopy(obj2) != obj2);
```

當傳入一個包含陣列的物件時，`deepcopy` 應該傳回該物件的深層複本。

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

當傳入一個包含另一個物件的物件時，`deepcopy` 應該傳回該物件的深層複本。

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --after-user-code--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
