---
id: 596e414344c3b2872167f0fe
title: 逗號吹毛求疵
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">Comma quibbling</a> 是 Eric Lippert 在他的部落格中最初設定的任務。

# --instructions--

撰寫一個函式（程式），用來產生一個字串輸出，該輸出是從列表／序列中輸入字的串接，條件如下：

<ol>
  <li>An input of no words produces the output string of just the two brace characters (<code>"{}"</code>)</li>
  <li>An input of just one word, e.g. <code>["ABC"]</code>, produces the output string of the word inside the two braces, e.g. <code>"{ABC}"</code></li>
  <li>An input of two words, e.g. <code>["ABC", "DEF"]</code>, produces the output string of the two words inside the two braces with the words separated by the string <code>" and "</code>, e.g. <code>"{ABC and DEF}"</code></li>
  <li>An input of three or more words, e.g. <code>["ABC", "DEF", "G", "H"]</code>, produces the output string of all but the last word separated by <code>", "</code> with the last word separated by <code>" and "</code> and all within braces; e.g. <code>"{ABC, DEF, G and H}"</code></li>
</ol>

使用以下一系列輸入測試你的函式（程式）、並在此頁面上顯示你的輸出：

<ul>
  <li>[] # (No input words).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**注意：** 假設字是非空的大寫字元字串。

# --hints--

`quibble` 應該是一個函式（程式）。

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` 應該傳回一個字串。

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` 應該傳回 "{}"。

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` 應該傳回 `"{ABC}"`。

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` 應該傳回 `"{ABC and DEF}"`。

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` 應該傳回 `"{ABC, DEF, G and H}"`。

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --after-user-code--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
