---
id: 596e457071c35c882915b3e4
title: 比較字串列表
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

列表是一組有序的值，可能包含重複項。以下是一個範例：

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

給定一個任意多字串的列表，為以下每個條件實作一個函式。

<ul>
  <li>test if they are all lexically equal</li>
  <li>test if every string is lexically less than the one after it  (i.e. whether the list is in strict ascending order)</li>
</ul>

# --hints--

`allEqual` 應該是一個函式（程式）。

```js
assert(typeof allEqual === 'function');
```

`azSorted` 應該是一個函式（程式）。

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` 應該傳回 true。

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` 應該傳回 false。

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` 應該傳回 false。

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` 應該傳回 true。

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` 應該傳回 true。

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` 應該傳回 true。

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` 應該傳回 true。

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` 應該傳回 true。

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` 應該傳回 false。

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` 應該傳回 false。

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --after-user-code--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
