---
id: 596fda99c69f779975a1b67d
title: 計算子字串的出現次數
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

創建一個函式（程式），或顯示一個內建函式（程式），用來計算字串中不重疊子字串出現的次數。

該函式（程式）應該接受兩個引數：

<ul>
  <li>the first argument being the string to search, and</li>
  <li>the second a substring to be searched for.</li>
</ul>

它應該傳回一個整數計數。

匹配應該產生最高數量的不重疊匹配。

一般來說，這本質上表示從左到右或從右到左進行匹配。

# --hints--

`countSubstring` 應該是一個函式（程式）。

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` 應該傳回 `3`。

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` 應該傳回 `2`。

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` 應該傳回 `2`。

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --after-user-code--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
