---
id: 597089c87eec450c68aa1643
title: 定義原始資料型別
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

定義一個型別，其行為類似整數，但最低有效值為 1，最高有效值為 10。

錯誤處理：

<ul>
  <li>If you try to instantiate a <code>Num</code> with a value outside of 1 - 10, it should throw a <code>TypeError</code> with an error message of <code>'Out of range'</code>.</li>
  <li>If you try to instantiate a <code>Num</code> with a value that is not a number, it should throw a <code>TypeError</code> with an error message of <code>'Not a Number'</code>.</li>
</ul>

# --hints--

`Num` 應該是一個函式（程式）。

```js
assert(typeof Num === 'function');
```

`new Num(4)` 應該傳回一個物件。

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` 應該丟出一個帶有訊息 'Not a Number' 的 TypeError。

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` 應該丟出帶有訊息 'Out of range' 的 TypeError。

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` 應該丟出帶有訊息 'Out of range' 的 TypeError。

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` 應該丟出帶有訊息 'Out of range' 的 TypeError。

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` 應該丟出帶有訊息 'Out of range' 的 TypeError。

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` 應該相等 7。

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` 應該相等於 -1。

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` 應該相等 12。

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` 應該相等 0.75。

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` 應該為 true。

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` 應該是 false。

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` 應該傳回 '5'

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
