---
id: 59713bd26bdeb8a594fb9413
title: 數硬幣
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

美國貨幣中有四種常見的硬幣型別、類型：

<ul>
  <li>quarters (25 cents)</li>
  <li>dimes (10 cents)</li>
  <li>nickels (5 cents), and</li>
  <li>pennies (1 cent)</li>
</ul>

<p>There are six ways to make change for 15 cents:</p>

<ul>
  <li>A dime and a nickel</li>
  <li>A dime and 5 pennies</li>
  <li>3 nickels</li>
  <li>2 nickels and 5 pennies</li>
  <li>A nickel and 10 pennies</li>
  <li>15 pennies</li>
</ul>

# --instructions--

實作一個函式（程式），用來判斷有多少種方法可以用這些常見硬幣，對代表美國分幣金額的輸入 `cents` 找零。

# --hints--

`countCoins` 應該是一個函式（程式）。

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)` 應該傳回 `6`。

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)` 應該傳回 `163`。

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)` 應該傳回 `242`。

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
