---
id: 597f1e7fbc206f0e9ba95dc4
title: 整數的因數
challengeType: 1
forumTopicId: 302265
dashedName: factors-of-an-integer
---

# --description--

撰寫一個函式，傳回正整數的因數陣列。

這些因數是能將被分解的數字整除並產生正整數結果的正整數。

# --hints--

`factors` 應該是一個函式（程式）。

```js
assert(typeof factors === 'function');
```

`factors(45)` 應該傳回 `[1,3,5,9,15,45]`。

```js
assert.deepEqual(factors(45), ans[0]);
```

`factors(53)` 應該傳回 `[1,53]`。

```js
assert.deepEqual(factors(53), ans[1]);
```

`factors(64)` 應該傳回 `[1,2,4,8,16,32,64]`。

```js
assert.deepEqual(factors(64), ans[2]);
```

# --seed--

## --after-user-code--

```js
const ans=[[1,3,5,9,15,45],[1,53],[1,2,4,8,16,32,64]];
```

## --seed-contents--

```js
function factors(num) {

}
```

# --solutions--

```js
function factors(num)
{
 let n_factors = [], i, sqr=Math.floor(Math.sqrt(num));

 for (i = 1; i <=sqr ; i += 1)
  if (num % i === 0)
  {
   n_factors.push(i);
   if (num / i !== i)
    n_factors.push(num / i);
  }
 n_factors.sort(function(a, b){return a - b;});
 return n_factors;
}
```
