---
id: 59880443fb36441083c6c20e
title: 歐拉方法
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

Euler 方法數值近似一階常微分方程（ODE）在給定初始值下的解。它是一種用於解決初始值問題（IVPs）的明確方法，如<a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">這篇文章</a>所述。

ODE 必須以以下表單提供：

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

帶有初始值

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

為了獲得數值解法，我們用有限差分近似取代左側的導數：

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

然後求解 $y(t+h)$：

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

這與…相同

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

反覆的解決方案規則如下：

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

其中 $h$ 是步長，是影響解法準確度最相關的參數。較小的步長會提高準確度，但也會增加計算成本，因此必須根據手頭的問題手動選擇。

**範例：牛頓冷卻定律**

牛頓冷卻定律描述初始溫度為 $T(t_0) = T_0$ 的物件如何在環境溫度 $T_R$ 中冷卻：

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

或

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

它表示物件的冷卻速率 $\\frac{dT(t)}{dt}$ 與當前溫差 $\\Delta T = (T(t) - T_R)$ 與周圍環境成正比。

我們將與數值近似比較的解析解是

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

實作一個 Euler 的常式，然後使用它來解決給定的牛頓冷卻定律範例，步長為三種不同的大小：

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code> and</li>
  <li><code>10 s</code></li>
</ul>

並與解析解法比較。

**初始值：**

<ul>
  <li>initial temperature <big>$T_0$</big> shall be <code>100 °C</code></li>
  <li>room temperature <big>$T_R$</big> shall be <code>20 °C</code></li>
  <li>cooling constant <big>$k$</big> shall be <code>0.07</code></li>
  <li>time interval to calculate shall be from <code>0 s</code> to <code>100 s</code></li>
</ul>  

函式的第一個參數是初始時間，第二個參數是初始溫度，第三個參數是經過時間，第四個參數是步長。

# --hints--

`eulersMethod` 應該是一個函式（程式）。

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` 應該傳回一個數字。

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` 應該傳回 20.0424631833732。

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` 應該傳回 20.01449963666907。

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` 應該傳回 20.000472392。

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
