---
id: 598eef80ba501f1268170e1e
title: Fibonacci n 步數列
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

這些數列是常規 Fibonacci 數列的擴展，其中：

<ol>
  <li>當 $n = 2$ 時，我們得到 Fibonacci 數列；初始值為 $[1, 1]$，且 $F_k^2 = F_{k-1}^2 + F_{k-2}^2$</li>
  <li>當 $n = 3$ 時，我們得到 tribonacci 數列；初始值為 $[1, 1, 2]$，且 $F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$</li>
  <li>當 $n = 4$ 時，我們得到 tetranacci 數列；初始值為 $[1, 1, 2, 4]$，且 $F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$…</li>
  <li>對於一般 $n>2$，我們有 Fibonacci $n$ 步數列 - $F_k^n$；初始值為第 $(n-1)$ 個 Fibonacci $n$ 步數列 $F_k^{n-1}$ 的前 $n$ 個值；此第 $n$ 個數列的第 $k$ 個值為 $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$</li>
</ol>

對於較小的 $n$ 值，有時會使用希臘數字前置式來個別命名每個數列。

Fibonacci $n$ 步數列：

| $n$ | 數列名稱 | 數值                                                    |
| --- | -------- | ------------------------------------------------------- |
| 2   | fibonacci   | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...         |
| 3   | tribonacci  | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...    |
| 4   | tetranacci  | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...  |
| 5   | pentanacci  | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...  |
| 6   | hexanacci   | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...  |
| 7   | heptanacci  | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ... |
| 8   | octonacci   | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ... |
| 9   | nonanacci   | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ... |
| 10  | decanacci   | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ... |

也可以產生相關數列，改變初始值：Lucas 數列像 $n=2$ 的 fibonacci 數列一樣將前兩個值相加，但初始值為 $\[2, 1]$。

# --instructions--

撰寫一個函式來產生 Fibonacci $n$ 步數列和 Lucas 數列。第一個參數是 $n$。第二個參數是要傳回的元素數量。第三個參數指定要輸出 Fibonacci 數列還是 Lucas 數列。如果參數是 `"f"`，則傳回 Fibonacci 數列；如果是 `"l"`，則傳回 Lucas 數列。數列必須以陣列形式傳回。

# --hints--

`fib_luc` 應該是一個函式。

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` 應該傳回 `[1,1,2,3,5,8,13,21,34,55]`。

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` 應該傳回 `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]`。

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` 應該傳回 `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]`。

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` 應該傳回 `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]`。

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` 應該傳回 `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]`。

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` 應該傳回 `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]`。

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` 應該傳回 `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]`。

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
