---
id: 5992e222d397f00d21122931
title: 費波那契字串
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

費波那契字串序列的創建方式類似於費波那契數列，但它著重於迭代串接。

<pre>定義  F_Word<sub>1</sub>  為  <strong>1</strong>
定義  F_Word<sub>2</sub>  為  <strong>0</strong>
形成   F_Word<sub>3</sub>  為  F_Word<sub>2</sub>   串接 F_Word<sub>1</sub> ，即：  <strong>01</strong>
形成   F_Word<sub>n</sub>  為  F_Word<sub>n-1</sub>  串接 F_Word<sub>n-2</sub>
</pre>

此挑戰需要計算熵，<a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">詳見此 Rosetta Code 挑戰</a>

# --instructions--

撰寫一個函式以傳回前 `n` 個費波那契字串。`n` 的數量會作為參數傳入函式。函式應傳回一個物件陣列。物件的形式應為 `{ N: 1, Length: 1, Entropy: 0, Word: '1' }`。`Entropy` 是針對字串 `Word` 計算的熵，並四捨五入至小數點後 8 位。請注意此序列的索引從 `1` 開始。

# --hints--

`fibWord` 應該是一個函式。

```js
assert(typeof fibWord === 'function');
```

`fibWord(5)` 應該傳回一個陣列。

```js
assert(Array.isArray(fibWord(5)));
```

`fibWord(5)` 應該傳回 `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }]`。

```js
assert.deepEqual(fibWord(5), globalThis.words5);
```

`fibWord(7)` 應該傳回 `[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`。

```js
assert.deepEqual(fibWord(7), globalThis.words7);
```

# --before-all--

```js
globalThis.words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

globalThis.words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

# --seed--

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
