---
id: 599d15309e88c813a40baf58
title: 熵
challengeType: 1
forumTopicId: 302254
dashedName: entropy
---

# --description--

計算給定輸入字串的 Shannon 熵 H。

給定離散隨機變數 $X$，它是一個由 $n$ 種不同字元（n=2 代表二進位）組成的 $N$ 個「符號」（總字元數）字串，$X$ 的香農熵以位元／符號為單位為：

$H_2(X) = -\\sum\_{i=1}^n \\frac{count_i}{N} \\log_2 \\left(\\frac{count_i}{N}\\right)$

其中 $count_i$ 是字元 $n_i$ 的計數。

# --hints--

`entropy` 應該是一個函式（程式）。

```js
assert(typeof entropy === 'function');
```

`entropy("0")` 應該傳回 `0`

```js
assert.equal(entropy('0'), 0);
```

`entropy("01")` 應該傳回 `1`

```js
assert.equal(entropy('01'), 1);
```

`entropy("0123")` 應該傳回 `2`

```js
assert.equal(entropy('0123'), 2);
```

`entropy("01234567")` 應該傳回 `3`

```js
assert.equal(entropy('01234567'), 3);
```

`entropy("0123456789abcdef")` 應該傳回 `4`

```js
assert.equal(entropy('0123456789abcdef'), 4);
```

`entropy("1223334444")` 應該傳回 `1.8464393446710154`

```js
assert.equal(entropy('1223334444'), 1.8464393446710154);
```

# --seed--

## --seed-contents--

```js
function entropy(s) {

}
```

# --solutions--

```js
function entropy(s) {
    // Create a dictionary of character frequencies and iterate over it.
  function process(s, evaluator) {
    let h = Object.create(null),
      k;
    s.split('').forEach(c => {
      h[c] && h[c]++ || (h[c] = 1); });
    if (evaluator) for (k in h) evaluator(k, h[k]);
    return h;
  }
    // Measure the entropy of a string in bits per symbol.

  let sum = 0,
    len = s.length;
  process(s, (k, f) => {
    const p = f / len;
    sum -= p * Math.log(p) / Math.log(2);
  });
  return sum;
}
```
