---
id: 599d1566a02b571412643b84
title: 衣索比亞乘法
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

埃塞俄比亞乘法是一種僅使用加法、加倍和除以二來乘法整數的方法。

**方法：**

<ol>
  <li>Take two numbers to be multiplied and write them down at the top of two columns</li>
  <li>In the left-hand column repeatedly halve the last number, discarding any remainders, and write the result below the last in the same column, until you write a value of <code>1</code></li>
  <li>In the right-hand column repeatedly double the last number and write the result below. stop when you add a result in the same row as where the left hand column shows <code>1</code></li>
  <li>Examine the table produced and discard any row where the value in the left column is even</li>
  <li>Sum the values in the right-hand column that remain to produce the result of multiplying the original two numbers together</li>
</ol>

**例如：** `17 × 34`

<pre>17   34
</pre>

將第一行減半：

<pre>17   34
8
4
2
1
</pre>

將第二行加倍：

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

刪除第一個儲存格為偶數的列：

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

將右側行中剩餘的數字加總：

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

所以以埃塞俄比亞方法，`17` 乘以 `34` 是 `578`。

# --instructions--

任務是定義三個具名的函式（程式）／方法／程序／常式：

<ol>
  <li>one to halve an integer,</li>
  <li>one to double an integer, and</li>
  <li>one to state if an integer is even</li>
</ol>

使用這些函式（程式）來創建一個執行埃塞俄比亞乘法的函式（程式）。

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` 應該是一個函式（程式）。

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` 應該傳回 `578`。

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` 應該傳回 `1058`。

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` 應該傳回 `324`。

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` 應該傳回 `5488`。

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` 應該傳回 `4662`。

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
