---
id: 599d1566a02b571412643b84
title: 衣索比亞乘法
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

埃塞俄比亞乘法是一種僅使用加法、加倍和除以二來乘法整數的方法。

**方法：**

<ol>
  <li>取兩個要相乘的數字，並將它們寫在兩個行的頂端</li>
  <li>在左邊的行中，反覆將最後一個數字除以二並捨去餘數，將結果寫在同一行的上一個數字下方，直到寫出值為 <code>1</code></li>
  <li>在右邊的行中，反覆將最後一個數字加倍並寫在下方。當你在左邊行顯示 <code>1</code> 的同一列加入結果時停止</li>
  <li>檢查產生的表格，捨棄左邊行中值為偶數的任何列</li>
  <li>將右邊行中剩下的值相加，得到原本兩個數字相乘的結果</li>
</ol>

**例如：** `17 × 34`

<pre>17   34
</pre>

將第一行減半：

<pre>17   34
8
4
2
1
</pre>

將第二行加倍：

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

刪除第一個儲存格為偶數的列：

<pre>17   34
8    <s>68</s>
4   <s>136</s>
2   <s>272</s>
1   544
</pre>

將右側行中剩餘的數字加總：

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

所以以埃塞俄比亞方法，`17` 乘以 `34` 是 `578`。

# --instructions--

任務是定義三個具名的函式（程式）／方法／程序／常式：

<ol>
  <li>一個用來將整數除以二的函式，</li>
  <li>一個用來將整數加倍的函式，和</li>
  <li>一個用來判斷整數是否為偶數的函式</li>
</ol>

使用這些函式（程式）來創建一個執行埃塞俄比亞乘法的函式（程式）。

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` 應該是一個函式（程式）。

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` 應該傳回 `578`。

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` 應該傳回 `1058`。

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` 應該傳回 `324`。

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` 應該傳回 `5488`。

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` 應該傳回 `4662`。

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
