---
id: 599d0ba974141b0f508b37d5
title: 反向質數
challengeType: 1
forumTopicId: 302253
dashedName: emirp-primes
---

# --description--

一個 emirp（**質數** 反過來拼寫）是指當其反轉（在十進位表述中）後仍為不同的質數。

# --instructions--

撰寫一個函式（程式），該函式：

<ul>
  <li>Shows the first <code>n</code> emirp numbers.</li>
  <li>Shows the emirp numbers in a range.</li>
  <li>Shows the number of emirps in a range.</li>
  <li>Shows the <code>n<sup>th</sup></code> emirp number.</li>
</ul>

該函式應接受兩個參數。第一個將接收 `n` 或作為陣列的範圍。第二個將接收一個布林值，用以指定函式是否傳回陣列形式的 emirps 或單一數字（範圍內的質數數量或第 <code>n<sup>th</sup></code> 個質數）。根據參數，函式應傳回陣列或數字。

# --hints--

`emirps` 應該是一個函式（程式）。

```js
assert(typeof emirps === 'function');
```

`emirps(20,true)` 應該傳回 `[13,17,31,37,71,73,79,97,107,113,149,157,167,179,199,311,337,347,359,389]`

```js
assert.deepEqual(emirps(20, true), [
  13,
  17,
  31,
  37,
  71,
  73,
  79,
  97,
  107,
  113,
  149,
  157,
  167,
  179,
  199,
  311,
  337,
  347,
  359,
  389
]);
```

`emirps(1000)` 應該傳回 `70529`

```js
assert.deepEqual(emirps(1000), 70529);
```

`emirps([7700,8000],true)` 應該傳回 `[7717,7757,7817,7841,7867,7879,7901,7927,7949,7951,7963]`

```js
assert.deepEqual(emirps([7700, 8000], true), [
  7717,
  7757,
  7817,
  7841,
  7867,
  7879,
  7901,
  7927,
  7949,
  7951,
  7963
]);
```

`emirps([7700,8000],false)` 應該傳回 `11`

```js
assert.deepEqual(emirps([7700, 8000], false), 11);
```

# --seed--

## --seed-contents--

```js
function emirps(n) {

}
```

# --solutions--

```js
function emirps(num, showEmirps)
{
  const is_prime = function(n)
    {
    if (!(n % 2) || !(n % 3)) return false;
    let p = 1;
    while (p * p < n)
                    { if (n % (p += 4) == 0 || n % (p += 2) == 0)
                            { return false; } }
    return true;
  };
  const is_emirp = function(n) {
    const r = parseInt(n.toString().split('').reverse().join(''));
    return r != n && is_prime(n) && is_prime(r);
  };

  let i,
    arr = [];
  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (is_emirp(i)) arr.push(i);
    // first x emirps
    if (showEmirps) return arr;
    // xth emirp
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (is_emirp(i)) arr.push(i);
    // emirps between x .. y
    if (showEmirps) return arr;
    // number of emirps between x .. y
    return arr.length;
  }
}
```
