---
id: 59d9c6bc214c613ba73ff012
title: SEDOL 編碼
challengeType: 1
forumTopicId: 302305
dashedName: sedols
---

# --description--

<abbr title="Stock Exchange Daily Official List">SEDOL</abbr> 是由倫敦證券交易所發行的證券識別號碼列表。

對於每個由 6 位數 <abbr title="Stock Exchange Daily Official List">SEDOL</abbr> 組成的數字列表，計算並附加校驗碼數字。也就是說，給定左側的輸入字串，你的函式應該傳回右側對應的字串：

<pre>
710889 => 7108899
B0YBKJ => B0YBKJ7
406566 => 4065663
B0YBLH => B0YBLH2
228276 => 2282765
B0YBKL => B0YBKL9
557910 => 5579107
B0YBKR => B0YBKR5
585284 => 5852842
B0YBKT => B0YBKT7
B00030 => B000300
</pre>

檢查每個輸入是否正確形成，特別是關於 SEDOL 字串中允許的有效字元。你的函式應該在輸入無效時傳回 `null`。

# --hints--

`sedol` 應該是一個函式（程式）。

```js
assert(typeof sedol === 'function');
```

`sedol('a')` 應該傳回 null。

```js
assert(sedol('a') === null);
```

`sedol('710889')` 應該傳回 '7108899'。

```js
assert(sedol('710889') === '7108899');
```

`sedol('BOATER')` 應該傳回 null。

```js
assert(sedol('BOATER') === null);
```

`sedol('228276')` 應該傳回 '2282765'。

```js
assert(sedol('228276') === '2282765');
```

# --seed--

## --seed-contents--

```js
function sedol(input) {
  const checkSum = 0
  
  return checkSum;
}
```

# --solutions--

```js
function sedol(input) {
  const checkDigit = sedolCheckDigit(input);
  if (checkDigit !== null) {
    return input + checkDigit;
  }
  return null;
}

const weight = [1, 3, 1, 7, 3, 9, 1];
function sedolCheckDigit(char6) {
  if (char6.search(/^[0-9BCDFGHJKLMNPQRSTVWXYZ]{6}$/) === -1) {
    return null;
  }

  let sum = 0;
  for (let i = 0; i < char6.length; i++) {
    sum += weight[i] * parseInt(char6.charAt(i), 36);
  }
  const checkSum = (10 - (sum % 10)) % 10;
  return checkSum.toString();
}
```
