---
id: 59da22823d04c95919d46269
title: '水手、椰子和猴子問題'
challengeType: 1
forumTopicId: 302304
dashedName: sailors-coconuts-and-a-monkey-problem
---

# --description--

五名水手在島上遇難，並在白天收集了一大堆椰子。 

那天晚上第一個水手醒來，決定提前拿他的第一份，於是嘗試將椰子堆平均分成五堆，但發現剩下一顆椰子，於是他把那顆椰子丟給一隻猴子，然後藏起「他」的五堆中其中一堆大小相等的椰子，並將另外四堆推在一起，再次形成一堆可見的椰子，然後去睡覺。 

長話短說，每位水手輪流在夜間起身一次，執行相同的動作：將椰子堆分成五堆，發現剩下一顆椰子，並將那顆剩餘的椰子給猴子，然後將其他四堆合併成一堆。 

在早晨（在五名水手各自偷偷且分開行動過夜之後），剩下的椰子被分成五堆相等的堆，每名水手各得一堆，結果發現椰子堆能被水手們平均分配且沒有餘數。（早晨猴子沒有得到任何東西。）

# --instructions--

創建一個函式（程式），傳回在一天中為 `N` 個水手收集的初始椰子堆的最小可能大小。**注意：**當然這個故事發生在一個世界裡，當天收集任意數量的椰子以及多次分堆等情況都能在符合故事情節的時間內發生，因此不影響數學計算。

# --hints--

`splitCoconuts` 應該是一個函式（程式）。

```js
assert(typeof splitCoconuts === 'function');
```

`splitCoconuts(5)` 應該傳回 3121。

```js
assert(splitCoconuts(5) === 3121);
```

`splitCoconuts(6)` 應該傳回 233275。

```js
assert(splitCoconuts(6) === 233275);
```

`splitCoconuts(7)` 應該傳回 823537。

```js
assert(splitCoconuts(7) === 823537);
```

# --seed--

## --seed-contents--

```js
function splitCoconuts(intSailors) {

  return true;
}
```

# --solutions--

```js
function splitCoconuts(intSailors) {
  let intNuts = intSailors;
  let result = splitCoconutsHelper(intNuts, intSailors);
  while (!result) {
    intNuts += 1;
    result = splitCoconutsHelper(intNuts, intSailors);
  }

  return intNuts;
}

function splitCoconutsHelper(intNuts, intSailors, intDepth) {
  const nDepth = intDepth !== undefined ? intDepth : intSailors;
  const portion = Math.floor(intNuts / intSailors);
  const remain = intNuts % intSailors;

  if (portion <= 0 || remain !== (nDepth ? 1 : 0)) {
    return null;
  }

  if (nDepth) {
    return splitCoconutsHelper(
      intNuts - portion - remain, intSailors, nDepth - 1
    );
  }

  return intNuts;
}
```
