---
id: 5a23c84252665b21eecc7ee0
title: 左階乘
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**左階乘**、$ !n $，可能指的是 *錯排階乘* 或 *階乘和*。相同的符號可能會令人混淆地用於這兩種不同的定義。有時，*錯排階乘*（也稱為 *錯置排列*）可能會使用以下任一符號：

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

（這可能看起來不明顯，但最後一個範例使用了顛倒的驚嘆號。）此任務將使用此公式計算**左階乘**：

$ !n = \sum_{k=0}^{n-1} k! $

當 $!0 = 0$

# --instructions--

撰寫一個函式（程式）來計算給定數字的左階乘。

# --hints--

`leftFactorial` 應該是一個函式（程式）。

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)` 應該傳回一個數字。

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)` 應該傳回 `0`。

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)` 應該傳回 `1`。

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)` 應該傳回 `2`。

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)` 應該傳回 `4`。

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)` 應該傳回 `409114`。

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)` 應該傳回 `22324392524314`。

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)` 應該傳回 `6780385526348314`。

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
