---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Soundex 演算法** 處理字的*意圖*。它為發音相似的字產生一個表述。 

它用於搜尋<em>名稱</em>和<em>位址</em>。這表示填寫<em>名稱</em>的人可以專注於它的發音，而不是更正<em>名稱</em>的拼寫。

例如： 

如果你是第一次聽到 `Quenci` 這個名稱，且拼錯了，你將會得到 **Soundex** 程式碼 `Q520`。 

當你下次正確拼寫 `Quincy` 這個名稱時，你仍然會得到相同的程式碼 `Q520`，這表示你可以將多個名稱發音連結到同一個 <em>person</em>，而不需要為每個拼寫添加。 

規則如下： 

<ul>
  <li>If a vowel (A, E, I, O, U) separates two consonants that have the same soundex code, the consonant to the right of the vowel is coded. Tymczak is coded as T-522 (T, 5 for the M, 2 for the C, Z ignored (see "Side-by-Side" rule above), 2 for the K). Since the vowel "A" separates the Z and K, the K is coded.</li>
  <li>If "H" or "W" separate two consonants that have the same soundex code, the consonant to the right of the vowel is not coded. Example: Ashcraft is coded A-261 (A, 2 for the S, C ignored, 6 for the R, 1 for the F). It is not coded A-226.</li>
</ul>

# --instructions--

撰寫一個函式（程式），該函式（程式）以字串作為參數並傳回編碼後的字串。

# --hints--

`soundex` 應該是一個函式（程式）。

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")` 應該傳回一個字串。

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")` 應該傳回 `"S532"`。

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")` 應該傳回 `"E251"`。

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")` 應該傳回 `"S532"`。

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")` 應該傳回 `"E251"`。

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")` 應該傳回 `"E460"`。

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")` 應該傳回 `"G200"`。

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")` 應該傳回 `"H416"`。

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")` 應該傳回 `"K530"`。

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")` 應該傳回 `"L300"`。

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")` 應該傳回 `"L222"`。

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
