---
id: 5a23c84252665b21eecc8014
title: 排序穩定性
challengeType: 1
forumTopicId: 302308
dashedName: sort-stability
---

# --description--

當依特定行或欄位對資料表中的記錄進行排序時，<a href="https://www.freecodecamp.org/news/stability-in-sorting-algorithms-a-treatment-of-equality-fa3140a5a539/" target="_blank" rel="noopener noreferrer nofollow">穩定排序</a>將始終保留具有相同鍵的記錄的相對順序。

例如，在這個國家和城市的表格中，對**第二**行，也就是城市，進行穩定排序會保持美國 Birmingham 在英國 Birmingham 之上。（雖然不穩定排序*可能*在此情況下將美國 Birmingham 排在英國 Birmingham 之上，但穩定排序常式會*保證*如此）。

<pre>UK  London
US  New York
US  Birmingham
UK  Birmingham
</pre>

同樣地，僅對第一行進行穩定排序會產生「UK London」作為第一項目，以及「US Birmingham」作為最後一項目（因為具有相同第一字 – 「UK」或「US」– 的元素順序會被維護）。

# --instructions--

撰寫一個函式（程式），該函式（程式）以二維陣列作為參數。每個元素都有 2 個元素，類似上述範例。該函式（程式）應該依先前提到的方式對陣列進行排序，並傳回排序後的陣列。

# --hints--

`stableSort` 應該是一個函式（程式）。

```js
assert(typeof stableSort == 'function');
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` 應該傳回一個陣列。

```js
assert(
  Array.isArray(
    stableSort([
      ['UK', 'London'],
      ['US', 'New York'],
      ['US', 'Birmingham'],
      ['UK', 'Birmingham']
    ])
  )
);
```

`stableSort([["UK", "London"], ["US", "New York"], ["US", "Birmingham"], ["UK", "Birmingham"]])` 應該傳回 `[["US", "Birmingham"], ["UK", "Birmingham"], ["UK", "London"], ["US", "New York"]]`。

```js
assert.deepEqual(
  stableSort([
    ['UK', 'London'],
    ['US', 'New York'],
    ['US', 'Birmingham'],
    ['UK', 'Birmingham']
  ]),
  [
    ['US', 'Birmingham'],
    ['UK', 'Birmingham'],
    ['UK', 'London'],
    ['US', 'New York']
  ]
);
```

`stableSort([[2, 2], [1, 2], [1, 4], [1, 5]])` 應該傳回 `[[2, 2], [1, 2], [1, 4], [1, 5]]`。

```js
assert.deepEqual(
  stableSort([
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]),
  [
    [2, 2],
    [1, 2],
    [1, 4],
    [1, 5]
  ]
);
```

`stableSort([[11, 55], [12, 45], [11, 45], [32, 45]])` 應該傳回 `[[12, 45], [11, 45], [32, 45], [11, 55]]`。

```js
assert.deepEqual(
  stableSort([
    [11, 55],
    [12, 45],
    [11, 45],
    [32, 45]
  ]),
  [
    [12, 45],
    [11, 45],
    [32, 45],
    [11, 55]
  ]
);
```

`stableSort([[10, 22], [1, 2], [1, 4], [1, 5], [10, 9]])` 應該傳回 `[[1, 2], [1, 4], [1, 5], [10, 9], [10, 22]]`。

```js
assert.deepEqual(
  stableSort([
    [10, 22],
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9]
  ]),
  [
    [1, 2],
    [1, 4],
    [1, 5],
    [10, 9],
    [10, 22]
  ]
);
```

`stableSort([[55, 54], [12, 22], [31, 43], [31, 54], [10, 49]])` 應該傳回 `[[12, 22], [31, 43], [10, 49], [55, 54], [31, 54]]`。

```js
assert.deepEqual(
  stableSort([
    [55, 54],
    [12, 22],
    [31, 43],
    [31, 54],
    [10, 49]
  ]),
  [
    [12, 22],
    [31, 43],
    [10, 49],
    [55, 54],
    [31, 54]
  ]
);
```

# --seed--

## --seed-contents--

```js
function stableSort(arr) {

}
```

# --solutions--

```js
function stableSort(arr) {
  arr.sort(function(a, b) {
    return a[1] < b[1] ? -1 : a[1] > b[1] ? 1 : 0;
  });
  return arr;
}
```
