---
id: 5a23c84252665b21eecc8013
title: 排序演算法/串列排序
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

**Strand sort** 會創建已排序的子集合，然後合併以產生最終結果。 

考慮一個 `unsortedArray = [3, 1, 4, 2]`。選取第一個條款 `3` 並將它複製到一個獨立的陣列中。搜尋此條款後面是否有更大的條款。當你找到一個較大的條款時，在此例中為 `4`，將它複製到獨立陣列中，即 `[3, 4]`，並將後續條款與這個新值 `4` 進行比較。

當你到達陣列的結尾後，移除你複製的項目 `[3, 4]`，然後從 `unsortedArray` 中剩下的第一個項目開始重新開始，在這個例子中是 `1`。

依照此處理程序會產生兩個已排序的陣列，`[3, 4]` 和 `[1, 2]`。將這兩個陣列合併以創建 `strandSortedArray`。

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

撰寫一個函式（程式）來使用 **Strand sort** 對陣列進行排序。該函式（程式）應傳回排序後的陣列。


# --hints--

`strandSort` 應該是一個函式（程式）。

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])` 應該傳回一個陣列。

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])` 應該傳回 `[7, 12, 20, 25, 32]`。

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])` 應該傳回 `[8, 13, 35, 38, 45]`。

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])` 應該傳回 `[20, 24, 34, 36, 43]`。

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])` 應該傳回 `[1, 12, 16, 18, 26, 33, 38]`。

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])` 應該傳回 `[1, 3, 4, 16, 29, 39, 48]`。

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
