---
id: 5a23c84252665b21eecc7ec2
title: Jaro 距離
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

Jaro 距離是衡量兩個字串相似度的指標。兩個字串的 Jaro 距離越高，字串越相似。分數經過正規化，`0` 表示完全不相似，`1` 則表示完全相符。

**定義**

兩個給定字串 \\(s_1\\) 和 \\(s_2\\) 的 Jaro 距離 \\( d_j \\) 是

\\begin{align}d_j = \\begin{cases}0& & \\text{如果 }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{否則}\\end{cases}\\end{align}

其中：

<ul>
  <li>\(m\) is the number of <i>matching characters</i>;</li>
  <li> \(t\) is half the number of <i>transpositions</i>.</li>
</ul>

分別來自 \\(s_1\\) 和 \\(s_2\\) 的兩個字元，只有在它們相同且距離不超過 \\(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\) 時，才被視為*匹配*。

\\(s_1\\) 的每個字元都會與 \\(s_2\\) 中所有相符的字元進行比較。相符（但序列順序不同）的字元數除以 2 即定義為 *transpositions* 的數量。

**範例**

給定字串 \\(s_1\\) *DWAYNE* 和 \\(s_2\\) *DUANE* 我們發現：

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

我們得到的 Jaro 分數為：\\(d_j = \\frac{1}{3}\\left(\\frac{4}{6} + \\frac{4}{5} + \\frac{4-0}{4}\\right) = 0.822\\)。

# --instructions--

撰寫一個函式（程式）`a`，該函式接受兩個字串作為參數並傳回相關的 Jaro 距離。

# --hints--

`jaro` 應該是一個函式（程式）。

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")` 應該傳回一個數字。

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")` 應該傳回 `0.9444444444444445`。

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")` 應該傳回 `0.7666666666666666`。

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")` 應該傳回 `0.8962962962962964`。

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")` 應該傳回 `0.888888888888889`。

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")` 應該傳回 `0.8333333333333334`。

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
