---
id: 5a23c84252665b21eecc7ede
title: 閏年
challengeType: 1
forumTopicId: 302300
dashedName: leap-year
---

# --description--

判斷給定的年份是否為格里曆的閏年。

# --hints--

`isLeapYear` 應該是一個函式（程式）。

```js
assert(typeof isLeapYear == 'function');
```

`isLeapYear()` 應該傳回一個布林值。

```js
assert(typeof isLeapYear(2018) == 'boolean');
```

`isLeapYear(2018)` 應該傳回 `false`。

```js
assert.equal(isLeapYear(2018), false);
```

`isLeapYear(2016)` 應該傳回 `true`。

```js
assert.equal(isLeapYear(2016), true);
```

`isLeapYear(2000)` 應該傳回 `true`。

```js
assert.equal(isLeapYear(2000), true);
```

`isLeapYear(1900)` 應該傳回 `false`。

```js
assert.equal(isLeapYear(1900), false);
```

`isLeapYear(1996)` 應該傳回 `true`。

```js
assert.equal(isLeapYear(1996), true);
```

`isLeapYear(1800)` 應該傳回 `false`。

```js
assert.equal(isLeapYear(1800), false);
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

}
```

# --solutions--

```js
function isLeapYear(year) {
  return year % 100 === 0 ? year % 400 === 0 : year % 4 === 0;
}
```
