---
id: 5a23c84252665b21eecc8040
title: 求 3 和 5 的倍數總和
challengeType: 1
forumTopicId: 302332
dashedName: sum-multiples-of-3-and-5
---

# --description--

目標是撰寫一個函式（程式），用來尋找所有小於 *n* 的 3 或 5 的正倍數的總和。

# --hints--

`sumMults` 應該是一個函式（程式）。

```js
assert(typeof sumMults == 'function');
```

`sumMults(10)` 應該傳回一個數字。

```js
assert(typeof sumMults(10) == 'number');
```

`sumMults(10)` 應該傳回 `23`。

```js
assert.equal(sumMults(10), 23);
```

`sumMults(100)` 應該傳回 `2318`。

```js
assert.equal(sumMults(100), 2318);
```

`sumMults(1000)` 應該傳回 `233168`。

```js
assert.equal(sumMults(1000), 233168);
```

`sumMults(10000)` 應該傳回 `23331668`。

```js
assert.equal(sumMults(10000), 23331668);
```

`sumMults(100000)` 應該傳回 `2333316668`。

```js
assert.equal(sumMults(100000), 2333316668);
```

# --seed--

## --seed-contents--

```js
function sumMults(n) {

}
```

# --solutions--

```js
function sumMults(n) {
  var sum = 0;
  for (var i = 1; i < n; i++) {
    if (i % 3 == 0 || i % 5 == 0) sum += i;
  }
  return sum;
}
```
