---
id: 5a23c84252665b21eecc7ec1
title: 迭代數字平方
challengeType: 1
forumTopicId: 302291
dashedName: iterated-digits-squaring
---

# --description--

如果你將一個自然數（大於零的整數）各位數字的平方相加，你最終總是會得到 1 或 89：

<pre>15 -> 26 -> 40 -> 16 -> 37 -> 58 -> 89
7 -> 49 -> 97 -> 130 -> 10 -> 1
</pre>

# --instructions--

撰寫一個函式（程式），該函式接受一個數字作為參數，並在執行所述處理程序後傳回 1 或 89。

# --hints--

`iteratedSquare` 應該是一個函式（程式）。

```js
assert(typeof iteratedSquare == 'function');
```

`iteratedSquare(4)` 應該傳回一個數字。

```js
assert(typeof iteratedSquare(4) == 'number');
```

`iteratedSquare(4)` 應該傳回 `89`。

```js
assert.equal(iteratedSquare(4), 89);
```

`iteratedSquare(7)` 應該傳回 `1`。

```js
assert.equal(iteratedSquare(7), 1);
```

`iteratedSquare(15)` 應該傳回 `89`。

```js
assert.equal(iteratedSquare(15), 89);
```

`iteratedSquare(20)` 應該傳回 `89`。

```js
assert.equal(iteratedSquare(20), 89);
```

`iteratedSquare(70)` 應該傳回 `1`。

```js
assert.equal(iteratedSquare(70), 1);
```

`iteratedSquare(100)` 應該傳回 `1`。

```js
assert.equal(iteratedSquare(100), 1);
```

# --seed--

## --seed-contents--

```js
function iteratedSquare(n) {

}
```

# --solutions--

```js
function iteratedSquare(n) {
    var total;
    while (n != 89 && n != 1) {
        total = 0;
        while (n > 0) {
            total += Math.pow(n % 10, 2);
            n = Math.floor(n/10);
        }
        n = total;
    }
    return n;
}
```
