---
id: 5a23c84252665b21eecc7ec5
title: 約瑟夫問題
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

約瑟夫問題是一個數學謎題，描述相當嚴峻：$n$ 名囚犯站成一個圓圈，依序編號從 $0$ 到 $n-1$。

一名刽子手沿著圓圈行走，從囚犯 $0$ 開始，每隔 $k$ 名囚犯就將其移除並處決。

隨著處理程序的進行，圓圈變得越來越小，直到只剩下一名囚犯，然後該囚犯被釋放。

例如，如果有 $n=5$ 名囚犯且 $k=2$，囚犯被殺的順序（我們稱之為「殺戮序列」）將是 1、3、0 和 4，倖存者將是 #2。

給定任意 $n, k > 0$，找出最後存活的囚犯是誰。

在其中一個事件中，有 41 名囚犯，每第 3<sup>rd</sup> 名囚犯被處決（$k=3$）。

其中有一位聰明的傢伙名叫 Josephus，他解出了這個問題，站在存活的位置上，並活著講述這個故事。

他是第幾號？

# --instructions--

撰寫一個函式，該函式以初始囚犯數量和 `k` 作為參數，並傳回存活囚犯的編號。

# --hints--

`josephus` 應該是一個函式（程式）。

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)` 應該傳回一個數字。

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)` 應該傳回 `28`。

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)` 應該傳回 `2`。

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)` 應該傳回 `8`。

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)` 應該傳回 `1`。

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)` 應該傳回 `1`。

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
