---
id: 5a23c84252665b21eecc8046
title: 對稱差異
challengeType: 1
forumTopicId: 16086
dashedName: symmetric-difference
---

# --description--

給定兩個集合 *A* 和 *B*，計算 $(A \setminus B) \cup (B \setminus A)$。也就是列舉在 *A* 或 *B* 中但不在兩者皆有的項目。這個集合稱為 *A* 和 *B* 的對稱差集。換句話說：$(A \cup B) \setminus (A \cap B)$（至少在 *A* 或 *B* 其中一個集合中的項目集合，減去同時在 *A* 和 *B* 中的項目集合）。

範例：

對於集合 `A = [1, 2, 3]` 和 `B = [1, 3, 4]`，*A* 和 *B* 的對稱差集是 `[2, 4]`。

# --instructions--

撰寫一個函式（程式），該函式（程式）接受兩個陣列作為參數，並傳回對稱差集。在傳回結果陣列之前先將其排序。

# --hints--

`symmetricDifference` 應該是一個函式（程式）。

```js
assert(typeof symmetricDifference == 'function');
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` 應該傳回一個陣列。

```js
assert(
  Array.isArray(
    symmetricDifference(
      ['John', 'Bob', 'Mary', 'Serena'],
      ['Jim', 'Mary', 'John', 'Bob']
    )
  )
);
```

`symmetricDifference(["John", "Bob", "Mary", "Serena"], ["Jim", "Mary", "John", "Bob"])` 應該傳回 `["Jim", "Serena"]`。

```js
assert.deepEqual(
  symmetricDifference(
    ['John', 'Bob', 'Mary', 'Serena'],
    ['Jim', 'Mary', 'John', 'Bob']
  ),
  ['Jim', 'Serena']
);
```

`symmetricDifference([1, 2, 3], [3, 4])` 應該傳回 `[1, 2, 4]`。

```js
assert.deepEqual(symmetricDifference([1, 2, 3], [3, 4]), [1, 2, 4]);
```

`symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7])` 應該傳回 `[1, 2, 5, 7, 8]`。

```js
assert.deepEqual(symmetricDifference([1, 2, 3, 4, 5], [3, 4, 8, 7]), [
  1,
  2,
  5,
  7,
  8
]);
```

`symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9])` 應該傳回 `[2, 4, 9]`。

```js
assert.deepEqual(
  symmetricDifference([1, 2, 3, 4, 5, 6, 7, 8], [1, 3, 5, 6, 7, 8, 9]),
  [2, 4, 9]
);
```

`symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9])` 應該傳回 `[1, 3, 4, 8]`。

```js
assert.deepEqual(symmetricDifference([1, 2, 4, 7, 9], [2, 3, 7, 8, 9]), [
  1,
  3,
  4,
  8
]);
```

# --seed--

## --seed-contents--

```js
function symmetricDifference(A, B) {

}
```

# --solutions--

```js
function symmetricDifference(A, B) {
  function relative_complement(A, B) {
    return A.filter(function(elem) {
      return B.indexOf(elem) == -1;
    });
  }

  function unique(ary) {
    var u = ary.concat().sort();
    for (var i = 1; i < u.length; ) {
      if (u[i - 1] === u[i]) u.splice(i, 1);
      else i++;
    }
    return u;
  }

  return unique(
    relative_complement(A, B).concat(relative_complement(B, A))
  ).sort();
}
```
