---
id: 5a23c84252665b21eecc8002
title: 排序演算法/Bogosort
challengeType: 1
forumTopicId: 302311
dashedName: sorting-algorithmsbogosort
---

# --description--

對一個數字 `list` 進行 Bogosort。

Bogosort 只是隨機洗牌一個群集直到它被排序。

「Bogosort」是一種極度低效的演算法，僅用作內部笑話。

它的平均執行時間是 O(n!)，因為任何給定集合的隨機排列以排序順序結束的機率約為 *n* 階乘分之一，而最壞情況是無限大，因為無法保證隨機排列會產生排序序列。

最佳情況是 O(n)，因為可能只需對元素進行一次遍歷即可排序。

偽程式碼：

<pre><b>while not</b> InOrder(list) <b>do</b>
  Shuffle(list)
<b>done</b>
</pre>

# --hints--

`bogosort` 應該是一個函式（程式）。

```js
assert(typeof bogosort == 'function');
```

`bogosort([25, 32, 12, 7, 20])` 應該傳回一個陣列。

```js
assert(Array.isArray(bogosort([25, 32, 12, 7, 20])));
```

`bogosort([25, 32, 12, 7, 20])` 應該傳回 `[7, 12, 20, 25, 32]`。

```js
assert.deepEqual(bogosort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`bogosort([38, 45, 35, 8, 13])` 應該傳回 `[8, 13, 35, 38, 45]`。

```js
assert.deepEqual(bogosort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`bogosort([43, 36, 20, 34, 24])` 應該傳回 `[20, 24, 34, 36, 43]`。

```js
assert.deepEqual(bogosort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`bogosort([12, 33, 26, 18, 1, 16, 38])` 應該傳回 `[1, 12, 16, 18, 26, 33, 38]`。

```js
assert.deepEqual(bogosort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`bogosort([3, 39, 48, 16, 1, 4, 29])` 應該傳回 `[1, 3, 4, 16, 29, 39, 48]`。

```js
assert.deepEqual(bogosort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function bogosort(v) {

}
```

# --solutions--

```js
function bogosort(v) {
  function shuffle(v) {
    for (
      var j, x, i = v.length;
      i;
      j = Math.floor(Math.random() * i), x = v[--i], v[i] = v[j], v[j] = x
    );
    return v;
  }

  function isSorted(v) {
    for (var i = 1; i < v.length; i++) {
      if (v[i - 1] > v[i]) {
        return false;
      }
    }
    return true;
  }
  var sorted = false;
  while (sorted == false) {
    v = shuffle(v);
    sorted = isSorted(v);
  }
  return v;
}
```
