---
id: 5a23c84252665b21eecc7edc
title: 每個月的最後一個星期五
challengeType: 1
forumTopicId: 302299
dashedName: last-friday-of-each-month
---

# --description--

撰寫一個函式（程式），傳回指定年份中指定月份的最後一個星期五的日期。

# --hints--

`lastFriday` 應該是一個函式（程式）。

```js
assert(typeof lastFriday == 'function');
```

`lastFriday(2018, 1)` 應該傳回一個數字。

```js
assert(typeof lastFriday(2018, 1) == 'number');
```

`lastFriday(2018, 1)` 應該傳回 `26`。

```js
assert.equal(lastFriday(2018, 1), 26);
```

`lastFriday(2017, 2)` 應該傳回 `24`。

```js
assert.equal(lastFriday(2017, 2), 24);
```

`lastFriday(2012, 3)` 應該傳回 `30`。

```js
assert.equal(lastFriday(2012, 3), 30);
```

`lastFriday(1900, 4)` 應該傳回 `27`。

```js
assert.equal(lastFriday(1900, 4), 27);
```

`lastFriday(2000, 5)` 應該傳回 `26`。

```js
assert.equal(lastFriday(2000, 5), 26);
```

`lastFriday(2006, 6)` 應該傳回 `30`。

```js
assert.equal(lastFriday(2006, 6), 30);
```

`lastFriday(2010, 7)` 應該傳回 `30`。

```js
assert.equal(lastFriday(2010, 7), 30);
```

`lastFriday(2005, 8)` 應該傳回 `26`。

```js
assert.equal(lastFriday(2005, 8), 26);
```

# --seed--

## --seed-contents--

```js
function lastFriday(year, month) {

}
```

# --solutions--

```js
function lastFriday(year, month) {
  var i, last_day;
  i = 0;
  while (true) {
    last_day = new Date(year, month, i);
    if (last_day.getDay() === 5) {
      return last_day.getDate();
    }
    i -= 1;
  }
}
```
