---
id: 5a23c84252665b21eecc7e7b
title: 產生器／指數
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

產生器是一種可執行的物體（例如函式或程序），其中包含會一次產生一個值序列的程式碼，因此每次你呼叫該產生器時，都會提供序列中的下一個值。

產生器通常建立在協同程序或物件之上，以便物件的內部狀態能「自然地」被控制代碼。

產生器通常用於序列可能是無限的情況，且只需最小狀態即可構建序列的下一個值。

# --instructions--

撰寫一個使用產生器來產生平方和立方的函式（程式）。創建一個新的產生器，從平方的產生器中篩選出所有立方。

該函式應該傳回篩選後產生器的第 \\( n^{th} \\) 個值。

例如對於 \\(n=7\\)，該函式應該傳回 81，因為序列將會是 4、9、16、25、36、49、81。這裡 64 被過濾掉，因為它是立方數。

# --hints--

`exponentialGenerator` 應該是一個函式（程式）。

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()` 應該傳回一個數字。

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)` 應該傳回 `144`。

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)` 應該傳回 `196`。

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)` 應該傳回 `256`。

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)` 應該傳回 `484`。

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)` 應該傳回 `784`。

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
