---
id: 5a23c84252665b21eecc7eb1
title: 單位矩陣
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

*單位矩陣* 是一個大小為 \\( n \\times n \\) 的方陣，其中對角線上的元素全都是 `1`（一），其他所有元素全都是 `0`（零）。

<ul>
  <li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\)</li>
</ul>

# --instructions--

撰寫一個函式（程式），該函式接受一個數字 `n` 作為參數，並傳回階數為 \\( n \times n \\) 的單位矩陣。

# --hints--

`idMatrix` 應該是一個函式（程式）。

```js
assert(typeof idMatrix == 'function');
```

`idMatrix(1)` 應該傳回一個陣列。

```js
assert(Array.isArray(idMatrix(1)));
```

`idMatrix(1)` 應該傳回 `[ [ 1 ] ]`。

```js
assert.deepEqual(idMatrix(1), results[0]);
```

`idMatrix(2)` 應該傳回 `[ [ 1, 0 ], [ 0, 1 ] ]`。

```js
assert.deepEqual(idMatrix(2), results[1]);
```

`idMatrix(3)` 應該傳回 `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]`。

```js
assert.deepEqual(idMatrix(3), results[2]);
```

`idMatrix(4)` 應該傳回 `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]`。

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --after-user-code--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
