---
id: 5a23c84252665b21eecc8036
title: 從字串中剝除控制碼和擴展字元
challengeType: 1
forumTopicId: 302327
dashedName: strip-control-codes-and-extended-characters-from-a-string
---

# --description--

此任務是從字串中剝除控制碼和擴展字元。解法應示範如何達成以下每個結果：剝除控制碼和擴展字元的字串。在 ASCII 中，控制碼的十進位代碼為 0 到 31 及 127。在基於 ASCII 的系統中，如果剝除控制碼，結果字串的所有字元將位於 ASCII 表中十進位 32 到 126 的範圍內。在非基於 ASCII 的系統中，我們將在此任務中視為擴展字元的字元定義為那些在 ASCII 表（十進位 32 到 126 範圍內）沒有對應字形的字元。

# --hints--

`strip` 應該是一個函式（程式）。

```js
assert(typeof strip == 'function');
```

`strip("abc")` 應該傳回一個字串。

```js
assert(typeof strip('abc') == 'string');
```

`strip("\ba\x00b\n\rc\fd\xc3")` 應該傳回 `"abcd"`。

```js
assert.equal(strip('\ba\x00b\n\rc\fd\xc3'), 'abcd');
```

`strip("\u0000\n abc\u00E9def\u007F")` 應該傳回 `" abcdef"`。

```js
assert.equal(strip('\u0000\n abc\u00E9def\u007F'), ' abcdef');
```

`strip("a\n\tb\u2102d\u2147f")` 應該傳回 `"abdf"`。

```js
assert.equal(strip('a\n\tb\u2102d\u2147f'), 'abdf');
```

`strip("Français.")` 應該傳回 `"Franais."`。

```js
assert.equal(strip('Français.'), 'Franais.');
```

`strip("123\tabc\u0007DEF\u007F+-*/€æŧðłþ")` 應該傳回 `"123abcDEF+-*/"`。

```js
assert.equal(strip('123\tabc\u0007DEF\u007F+-*/€æŧðłþ'), '123abcDEF+-*/');
```

# --seed--

## --seed-contents--

```js
function strip(s) {

}
```

# --solutions--

```js
function strip(s) {
  return s
    .split('')
    .filter(function(x) {
      var n = x.charCodeAt(0);

      return 31 < n && 127 > n;
    })
    .join('');
}
```
