---
id: 5a23c84252665b21eecc8041
title: 級數的總和
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

計算一個<em>級數</em>的第 **n**<sup>項</sup>，即對應<em>數列</em>的前 **n** 項的總和。非正式地，這個值或當 **n** 趨近無限大時的極限，也稱為*級數的和*，因此這個任務的標題即為此。對於這個任務，使用：$S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$。

# --instructions--

撰寫一個函式（程式），該函式（程式）以 $a$ 和 $b$ 作為參數，並傳回序列中第 $a$ 個到第 $b$ 個成員的總和。

# --hints--

`sum` 應該是一個函式（程式）。

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` 應該傳回一個數字。

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` 應該傳回 `1.6349839001848923`。

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` 應該傳回 `0.009262256361481223`。

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` 應該傳回 `0.044086990748706555`。

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` 應該傳回 `0.08619778593108679`。

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` 應該傳回 `1.5497677311665408`。

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
