---
id: 5a23c84252665b21eecc8010
title: 排序演算法/希爾排序
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

撰寫一個函式，使用 Shell sort 演算法（一種遞減增量排序法）來排序一個元素陣列。該函式應傳回排序後的陣列。

Shell 排序（也稱為 Shellsort 或 Shell 的方法）以其發明者 Donald Shell 命名，他於 1959 年發表了該演算法。

Shell sort 是一系列基於遞增序列的交錯插入排序。每次排序後遞增大小會減少，直到遞增大小為 1。

當增量大小為 1 時，排序是基本的插入排序，但此時資料已保證幾乎排序完成，這是插入排序的「最佳情況」。

只要序列以 1 結尾，任何序列都會對資料進行排序，但有些序列的效果比其他序列更好。

實證研究顯示，公比約為 2.2 的等比遞增序列在實務中效果良好。

# --hints--

`shellSort` 應該是一個函式（程式）。

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` 應該傳回一個陣列。

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` 應該傳回 `[7, 12, 20, 25, 32]`。

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` 應該傳回 `[8, 13, 35, 38, 45]`。

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` 應該傳回 `[20, 24, 34, 36, 43]`。

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` 應該傳回 `[1, 12, 16, 18, 26, 33, 38]`。

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` 應該傳回 `[1, 3, 4, 16, 29, 39, 48]`。

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
