---
id: 5a23c84252665b21eecc803f
title: 計算整數的各位數字總和
challengeType: 1
forumTopicId: 302331
dashedName: sum-digits-of-an-integer
---

# --description--

撰寫一個函式，該函式接受一個字串作為參數。此字串表述一個可以是任意進位（小於 37）的數字，並傳回其各位數字的總和。

<ul>
  <li><b>1</b><sub>10</sub> sums to <b>1</b></li>
  <li><b>1234</b><sub>10</sub> sums to <b>10</b></li>
  <li><b>fe</b><sub>16</sub> sums to <b>29</b></li>
  <li><b>f0e</b><sub>16</sub> sums to <b>29</b></li>
</ul>

# --hints--

`sumDigits` 應該是一個函式（程式）。

```js
assert(typeof sumDigits == 'function');
```

`sumDigits("1")` 應該傳回一個數字。

```js
assert(typeof sumDigits('1') == 'number');
```

`sumDigits("1")` 應該傳回 `1`。

```js
assert.equal(sumDigits('1'), 1);
```

`sumDigits("12345")` 應該傳回 `15`。

```js
assert.equal(sumDigits('12345'), 15);
```

`sumDigits("254")` 應該傳回 `11`。

```js
assert.equal(sumDigits('254'), 11);
```

`sumDigits("fe")` 應該傳回 `29`。

```js
assert.equal(sumDigits('fe'), 29);
```

`sumDigits("f0e")` 應該傳回 `29`。

```js
assert.equal(sumDigits('f0e'), 29);
```

`sumDigits("999ABCXYZ")` 應該傳回 `162`。

```js
assert.equal(sumDigits('999ABCXYZ'), 162);
```

# --seed--

## --seed-contents--

```js
function sumDigits(n) {

}
```

# --solutions--

```js
function sumDigits(n) {
  n += '';
  for (var s = 0, i = 0, e = n.length; i < e; i += 1)
    s += parseInt(n.charAt(i), 36);
  return s;
}
```
