---
id: 5a23c84252665b21eecc7eca
title: Kaprekar 數
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

如果一個正整數是 Kaprekar 數，則：

<ul>
  <li>It is 1, or,</li>
  <li>The decimal representation of its square may be split once into two parts consisting of positive integers which sum to the original number. </li>
</ul>

請注意，分割後若產生的零件完全由 0 組成，則無效，因為 0 不被視為正數。

範例 Kaprekar 數字：

<ul>
  <li><code>2223</code> is a Kaprekar number, as <code>2223 * 2223 = 4941729</code>, <code>4941729</code> may be split to <code>494</code> and <code>1729</code>, and <code>494 + 1729 = 2223</code></li>
  <li>The series of Kaprekar numbers is known as A006886, and begins as <code>1, 9, 45, 55, ...</code></li>
</ul>

# --instructions--

撰寫一個函式，該函式接受一個數字 $n$、一個基底 $bs$，並在該數字為給定基底的 Kaprekar 數時傳回 true。否則，該函式傳回 false。

# --hints--

`isKaprekar` 應該是一個函式（程式）。

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` 應該傳回一個布林值。

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` 應該傳回 `true`。

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` 應該傳回 `true`。

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` 應該傳回 `true`。

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` 應該傳回 `false`。

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` 應該傳回 `false`。

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` 應該傳回 `true`。

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` 應該傳回 `false`。

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
