---
id: 5a23c84252665b21eecc7ec4
title: JortSort
challengeType: 1
forumTopicId: 302293
dashedName: jortsort
---

# --description--

jortSort 是一個排序工具組，讓使用者自己完成工作並保證效率，因為你再也不需要排序。它最初由 Jenn "Moneydollars" Schiffer 在著名的 JSConf2014 發表。

jortSort 應該是一個接受單一可比較物件陣列作為其引數的函式。它接著將陣列以遞增順序排序，並將排序後的陣列與原本提供的陣列進行比較。如果陣列相符（即原始陣列已經排序），該函式傳回 true。如果陣列不相符（即原始陣列未排序），該函式傳回 false。

# --hints--

`jortsort` 應該是一個函式（程式）。

```js
assert(typeof jortsort == 'function');
```

`jortsort([1,2,3,4,5])` 應該傳回一個布林值。

```js
assert(typeof jortsort([1, 2, 3, 4, 5]) == 'boolean');
```

`jortsort([1,2,3,4,5])` 應該傳回 `true`。

```js
assert.equal(jortsort([1, 2, 3, 4, 5]), true);
```

`jortsort([1,2,13,4,5])` 應該傳回 `false`。

```js
assert.equal(jortsort([1, 2, 13, 4, 5]), false);
```

`jortsort([12,4,51,2,4])` 應該傳回 `false`。

```js
assert.equal(jortsort([12, 4, 51, 2, 4]), false);
```

`jortsort([1,2])` 應該傳回 `true`。

```js
assert.equal(jortsort([1, 2]), true);
```

`jortsort([5,4,3,2,1])` 應該傳回 `false`。

```js
assert.equal(jortsort([5, 4, 3, 2, 1]), false);
```

`jortsort([1,1,1,1,1])` 應該傳回 `true`。

```js
assert.equal(jortsort([1, 1, 1, 1, 1]), true);
```

# --seed--

## --seed-contents--

```js
function jortsort(array) {

}
```

# --solutions--

```js
function jortsort(array) {
  // sort the array
  var originalArray = array.slice(0);
  array.sort( function(a,b){return a - b} );

  // compare to see if it was originally sorted
  for (var i = 0; i < originalArray.length; ++i) {
    if (originalArray[i] !== array[i]) return false;
  }

  return true;
};
```
