---
id: 5e4ce2a1ac708cc68c1df25d
title: 長乘法
challengeType: 1
forumTopicId: 385269
dashedName: long-multiplication
---

# --description--

明確地實作長乘法。

這是一種可能的任意精度整數代數方法。

# --instructions--

撰寫一個函式（程式），該函式（程式）接受兩個大型數字的字串作為參數。你的函式（程式）應該傳回這兩個大型數字相乘的結果，並以字串形式表示。

**注意：** 在 JavaScript 中，對大數字進行算術操作不精確，所以你必須自行實作精確乘法。

# --hints--

`mult` 應該是一個函式（程式）。

```js
assert(typeof mult == 'function');
```

`mult("18446744073709551616", "18446744073709551616")` 應該傳回一個字串。

```js
assert(typeof mult('18446744073709551616', '18446744073709551616') == 'string');
```

`mult("18446744073709551616", "18446744073709551616")` 應該傳回 `"340282366920938463463374607431768211456"`。

```js
assert.equal(
  mult('18446744073709551616', '18446744073709551616'),
  '340282366920938463463374607431768211456'
);
```

`mult("31844674073709551616", "1844674407309551616")` 應該傳回 `"58743055272886011737990786529368211456"`。

```js
assert.equal(
  mult('31844674073709551616', '1844674407309551616'),
  '58743055272886011737990786529368211456'
);
```

`mult("1846744073709551616", "44844644073709551616")` 應該傳回 `"82816580680737279241781007431768211456"`。

```js
assert.equal(
  mult('1846744073709551616', '44844644073709551616'),
  '82816580680737279241781007431768211456'
);
```

`mult("1844674407370951616", "1844674407709551616")` 應該傳回 `"3402823669833978308014392742590611456"`。

```js
assert.equal(
  mult('1844674407370951616', '1844674407709551616'),
  '3402823669833978308014392742590611456'
);
```

`mult("2844674407370951616", "1844674407370955616")` 應該傳回 `"5247498076580334548376218009219475456"`。

```js
assert.equal(
  mult('2844674407370951616', '1844674407370955616'),
  '5247498076580334548376218009219475456'
);
```

# --seed--

## --seed-contents--

```js
function mult(strNum1, strNum2) {

}
```

# --solutions--

```js
function mult(strNum1, strNum2) {
    var a1 = strNum1.split("").reverse();
    var a2 = strNum2.toString().split("").reverse();
    var aResult = new Array;

    for ( var iterNum1 = 0; iterNum1 < a1.length; iterNum1++ ) {
        for ( var iterNum2 = 0; iterNum2 < a2.length; iterNum2++ ) {
            var idxIter = iterNum1 + iterNum2;    // Get the current array position.
            aResult[idxIter] = a1[iterNum1] * a2[iterNum2] + ( idxIter >= aResult.length ? 0 : aResult[idxIter] );

            if ( aResult[idxIter] > 9 ) {    // Carrying
                aResult[idxIter + 1] = Math.floor( aResult[idxIter] / 10 ) + ( idxIter + 1 >= aResult.length ? 0 : aResult[idxIter + 1] );
                aResult[idxIter] %= 10;
            }
        }
    }
    return aResult.reverse().join("");
}
```
