---
id: 5e4ce2eaac708cc68c1df260
title: Levenshtein 距離
challengeType: 1
forumTopicId: 385264
dashedName: levenshtein-distance
---

# --description--

在資訊理論和計算機科學中，**Levenshtein distance** 是用來衡量兩個序列之間差異量的度量（即編輯距離）。兩個字串之間的 Levenshtein distance 定義為將一個字串轉換成另一個字串所需的最少編輯次數，允許的編輯操作為插入、刪除或替換單一字元。

範例：

「**kitten**」和「**sitting**」之間的 Levenshtein 距離是 3，因為以下三個編輯將其中一個變成另一個，且沒有比三個編輯更少的方式做到這點：

<ul>
  <li><strong>k</strong>itten   <strong>s</strong>itten    (substitution of 'k' with 's')</li>
  <li>sitt<strong>e</strong>n   sitt<strong>i</strong>n    (substitution of 'e' with 'i')</li>
  <li>sittin   sittin<strong>g</strong>    (insert 'g' at the end).</li>
</ul>

*「**rosettacode**」與「**raisethysword**」之間的 Levenshtein 距離是 **8**。*

*兩個字串之間的距離與兩個字串皆反轉後的距離相同。*

# --instructions--

撰寫一個函式，傳回兩個作為參數給定的字串之間的 Levenshtein 距離。

# --hints--

`levenshtein` 應該是一個函式（程式）。

```js
assert(typeof levenshtein == 'function');
```

`levenshtein("mist", "dist")` 應該傳回一個數字。

```js
assert(typeof levenshtein('mist', 'dist') == 'number');
```

`levenshtein("mist", "dist")` 應該傳回 `1`。

```js
assert.equal(levenshtein('mist', 'dist'), 1);
```

`levenshtein("tier", "tor")` 應該傳回 `2`。

```js
assert.equal(levenshtein('tier', 'tor'), 2);
```

`levenshtein("kitten", "sitting")` 應該傳回 `3`。

```js
assert.equal(levenshtein('kitten', 'sitting'), 3);
```

`levenshtein("stop", "tops")` 應該傳回 `2`。

```js
assert.equal(levenshtein('stop', 'tops'), 2);
```

`levenshtein("rosettacode", "raisethysword")` 應該傳回 `8`。

```js
assert.equal(levenshtein('rosettacode', 'raisethysword'), 8);
```

`levenshtein("mississippi", "swiss miss")` 應該傳回 `8`。

```js
assert.equal(levenshtein('mississippi', 'swiss miss'), 8);
```

# --seed--

## --seed-contents--

```js
function levenshtein(a, b) {

}
```

# --solutions--

```js
function levenshtein(a, b) {
  var t = [], u, i, j, m = a.length, n = b.length;
  if (!m) { return n; }
  if (!n) { return m; }
  for (j = 0; j <= n; j++) { t[j] = j; }
  for (i = 1; i <= m; i++) {
    for (u = [i], j = 1; j <= n; j++) {
      u[j] = a[i - 1] === b[j - 1] ? t[j - 1] : Math.min(t[j - 1], t[j], u[j - 1]) + 1;
    } t = u;
  } return u[n];
}
```
