---
id: 5e4ce2f5ac708cc68c1df261
title: 線性同餘產生器
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

線性同餘產生器（LCG）是一種<em>演算法</em>，它透過不連續分段線性方程式計算出一連串偽隨機數字。所有線性同餘產生器都使用此公式：

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

其中：

<ul>
<li>$ r_0 $ is a seed.</li>
<li>$r_1$, $r_2$, $r_3$, ..., are the random numbers.</li>
<li>$a$, $c$, $m$ are constants.</li>
</ul>

如果仔細選擇 $a$、$c$ 和 $m$ 的值，則產生器會產生從 $0$ 到 $m - 1$ 的均勻整數分佈。

<abbr title="linear congruential generator">LCG</abbr> 數字的品質很差。$r_n$ 和 $r\_{n + 1}$ 並不像真正的隨機數那樣獨立。任何知道 $r_n$ 的人都可以預測 $r\_{n + 1}$，因此 <abbr title="linear congruential generator">LCG</abbr> 並不具備密碼學安全性。<abbr title="linear congruential generator">LCG</abbr> 對於像 Miller-Rabin 質數測試或 FreeCell 發牌這類簡單任務仍然足夠。<abbr title="linear congruential generator">LCG</abbr> 的利益之一是可以輕易從相同的 $r_0$ 重現一串數字。也可以用不同的程式語言重現這樣的序列，因為公式非常簡單。

# --instructions--

撰寫一個函式（程式），該函式（程式）以 $r_0,a,c,m,n$ 作為參數並傳回 $r_n$。

# --hints--

`linearCongGenerator` 應該是一個函式（程式）。

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` 應該傳回一個數字。

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` 應該傳回 `855`。

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` 應該傳回 `1110`。

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` 應該傳回 `62217`。

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` 應該傳回 `12345`。

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` 應該傳回 `1406932606`。

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
