---
id: 5e6dee7749a0b85a3f1fc7d5
title: Lucas-Lehmer 測試
challengeType: 1
forumTopicId: 385281
dashedName: lucas-lehmer-test
---

# --description--

Lucas-Lehmer 測試：對於奇質數 $p$，若且唯若 $2^p-1$ 整除 $S(p-1)$，則梅森數 $2^p-1$ 為質數，其中 $S(n+1)=(S(n))^2-2$，且 $S(1)=4$。

# --instructions--

撰寫一個函式（程式），用來判斷給定的 Mersenne 數是否為質數。

# --hints--

`lucasLehmer` 應該是一個函式（程式）。

```js
assert(typeof lucasLehmer == 'function');
```

`lucasLehmer(11)` 應該傳回一個布林值。

```js
assert(typeof lucasLehmer(11) == 'boolean');
```

`lucasLehmer(11)` 應該傳回 `false`。

```js
assert.equal(lucasLehmer(11), false);
```

`lucasLehmer(15)` 應該傳回 `false`。

```js
assert.equal(lucasLehmer(15), false);
```

`lucasLehmer(13)` 應該傳回 `true`。

```js
assert.equal(lucasLehmer(13), true);
```

`lucasLehmer(17)` 應該傳回 `true`。

```js
assert.equal(lucasLehmer(17), true);
```

`lucasLehmer(19)` 應該傳回 `true`。

```js
assert.equal(lucasLehmer(19), true);
```

`lucasLehmer(21)` 應該傳回 `false`。

```js
assert.equal(lucasLehmer(21), false);
```

# --seed--

## --seed-contents--

```js
function lucasLehmer(p) {

}
```

# --solutions--

```js
function lucasLehmer(p) {
    function isPrime(p) {
        if (p == 2)
            return true;
        else if (p <= 1 || p % 2 == 0)
            return false;
        else {
            var to = Math.sqrt(p);
            for (var i = 3; i <= to; i += 2)
                if (p % i == 0)
                    return false;
            return true;
        }
    }

    function isMersennePrime(p) {
        if (p == 2)
            return true;
        else {
            var m_p = Math.pow(2, p) - 1
            var s = 4;
            for (var i = 3; i <= p; i++)
                s = (s * s - 2) % m_p
            return s == 0;
        }
    }

    return isPrime(p) && isMersennePrime(p)
}
```
