---
id: 5e6dd15004c88cf00d2a78b3
title: 同時對多個陣列進行迴圈
challengeType: 1
forumTopicId: 385279
dashedName: loop-over-multiple-arrays-simultaneously
---

# --description--

對多個陣列進行迴圈，並創建一個新陣列，其第 $i^{th}$ 個元素是每個給定陣列第 $i^{th}$ 個元素的串接。

在這個範例中，如果你得到這個陣列的陣列：

```js
[ ["a", "b", "c"], ["A", "B", "C"], [1, 2, 3] ]
```

輸出應該是：

```js
["aA1","bB2","cC3"]
```

# --instructions--

撰寫一個函式（程式），該函式（程式）以陣列的陣列作為參數，並傳回符合給定描述的字串陣列。

# --hints--

`loopSimult` 應該是一個函式（程式）。

```js
assert(typeof loopSimult == 'function');
```

`loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])` 應該傳回一個陣列。

```js
assert(
  Array.isArray(
    loopSimult([
      ['a', 'b', 'c'],
      ['A', 'B', 'C'],
      [1, 2, 3]
    ])
  )
);
```

`loopSimult([["a", "b", "c"], ["A", "B", "C"], [1, 2, 3]])` 應該傳回 `["aA1", "bB2", "cC3"]`。

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c'],
    ['A', 'B', 'C'],
    [1, 2, 3]
  ]),
  ['aA1', 'bB2', 'cC3']
);
```

`loopSimult([["c", "b", "c"], ["4", "5", "C"], [7, 7, 3]])` 應該傳回 `["c47", "b57", "cC3"]`。

```js
assert.deepEqual(
  loopSimult([
    ['c', 'b', 'c'],
    ['4', '5', 'C'],
    [7, 7, 3]
  ]),
  ['c47', 'b57', 'cC3']
);
```

`loopSimult([["a", "b", "c", "d"], ["A", "B", "C", "d"], [1, 2, 3, 4]])` 應該傳回 `["aA1", "bB2", "cC3", "dd4"]`。

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b', 'c', 'd'],
    ['A', 'B', 'C', 'd'],
    [1, 2, 3, 4]
  ]),
  ['aA1', 'bB2', 'cC3', 'dd4']
);
```

`loopSimult([["a", "b"], ["A", "B"], [1, 2]])` 應該傳回 `["aA1", "bB2"]`。

```js
assert.deepEqual(
  loopSimult([
    ['a', 'b'],
    ['A', 'B'],
    [1, 2]
  ]),
  ['aA1', 'bB2']
);
```

`loopSimult([["b", "c"], ["B", "C"], [2, 3]])` 應該傳回 `["bB2", "cC3"]`。

```js
assert.deepEqual(
  loopSimult([
    ['b', 'c'],
    ['B', 'C'],
    [2, 3]
  ]),
  ['bB2', 'cC3']
);
```

# --seed--

## --seed-contents--

```js
function loopSimult(A) {

}
```

# --solutions--

```js
function loopSimult(A) {
    var res = [], output;
    for (var i = 0; i < A[0].length; i += 1) {
        output = "";
        for (var j = 0; j < A.length; j++)
            output += A[j][i];
        res.push(output);
    }
    return res;
}
```
