---
id: 5e6dd1278e6ca105cde40ea9
title: 最長共同子序列
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**最長共同子序列**（或 **LCS**）是群組 A 和 B 中，兩個群組共有且在各自群組中順序相同的最長元素群組。 例如，序列 `1234` 和 `1224533324` 的 LCS 是 `1234`：

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

以字串範例來說，考慮序列 `thisisatest` 和 `testing123testing`。一個 LCS 將會是 `tsitest`：

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing。

你的程式碼只需要處理字串。

# --instructions--

撰寫一個區分大小寫的函式（程式），傳回兩個字串的 LCS。你不需要顯示多個 LCS。

# --hints--

`lcs` 應該是一個函式（程式）。

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")` 應該傳回一個字串。

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")` 應該傳回 `"tsitest"`。

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")` 應該傳回 `"ADH"`。

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")` 應該傳回 `"GTAB"`。

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")` 應該傳回 `"BDCB"`。

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")` 應該傳回 `"ABAD"`。

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
