---
id: 5e94a54cc7b022105bf0fd2c
title: 字頻率
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

給定一個文字字串和一個整數 n，傳回檔案中出現頻率最高的 n 個字（以及它們出現的次數），並依頻率由高到低排列。

# --instructions--

撰寫一個函式（程式）來計算每個字出現的次數，並傳回出現頻率由高到低的 n 個最常見字及其出現次數。

該函式（程式）應傳回一個二維陣列，陣列中的每個元素皆為以下表單：`[word, freq]`。`word` 應為該字的全小寫版本，`freq` 則為表示計數的數字。

如果沒有提供字串，該函式應傳回空陣列。

該函式（程式）應該不區分大小寫，例如，字串 "Hello" 和 "hello" 應該被視為相同。

你可以將包含底線、破折號、撇號、逗號等特殊字元的字、單字視為不同的字、單字。

例如，給定字串 "Hello hello goodbye"，你的函式（程式）應該傳回 `[['hello', 2], ['goodbye', 1]]`。

# --hints--

`wordFrequency` 應該是一個函式（程式）。

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` 應該傳回一個陣列。

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` 應該傳回 `[['hello', 2], ['world', 1]]`

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` 應該傳回 `[['the', 2]]`

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` 應該傳回 `[['opensource', 2]]`

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` 應該傳回 `[['app', 2], ['apple', 2], ['apply', 1]]` 或 `[['apple', 2], ['app', 2], ['apply', 1]]`

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` 應該傳回 `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]`

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` 應該傳回 `[]`

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
