---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

撰寫一個程式，產生一個從 1 到 100（含）整數的陣列。但：

<ul>
    <li>for multiples of 3, add <code>"Fizz"</code> to the array instead of the number</li>
    <li>for multiples of 5, add <code>"Buzz"</code> to the array instead of the number</li>
    <li>for multiples of 3 and 5, add <code>"FizzBuzz"</code> to the array instead of the number</li>
</ul>

# --instructions--

你的程式應該傳回一個包含根據上述規則結果的陣列。

# --hints--

`fizzBuzz` 應該是一個函式（程式）。

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` 應該傳回一個陣列。

```js
assert(Array.isArray(fizzBuzz()) == true);
```

只有能被 3 整除的數字應該傳回 `"Fizz"`。

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

只有能被 5 整除的數字應該傳回 `"Buzz"`。

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

同時可被 3 和 5 整除的數字應該傳回 `"FizzBuzz"`。

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

不能被 3 或 5 整除的數字應該傳回該數字本身。

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
