---
id: 5eaf48389ee512d4d103684b
title: 自我描述數字
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

有幾個所謂的「自我描述」整數。

如果一個整數具有這樣的屬性：當數字位置標記為 0 到 N-1 時，每個位置的數字相等於該數字在整數中出現的次數，則該整數稱為「自我描述」。

例如，**2020** 是一個四位數的自我描述數字：

<ul>
    <li> position 0 has value 2 and there are two 0s in the number; </li>
    <li> position 1 has value 0 and there are no 1s in the number; </li>
    <li> position 2 has value 2 and there are two 2s; </li>
    <li> position 3 has value 0 and there are zero 3s; </li>
</ul>

自我描述數字且小於 100,000,000 的有：1210、2020、21200、3211000、42101000。

# --instructions--

撰寫一個函式，該函式以正整數作為參數。如果它是自我描述的，則傳回 true。否則，傳回 false。

# --hints--

`isSelfDescribing` 應該是一個函式（程式）。

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` 應該傳回一個布林值。

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` 應該傳回 `true`。

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` 應該傳回 `false`。

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` 應該傳回 `true`。

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
