---
id: 5eb3e4a21f462f409d656c73
title: 自我取用序列
challengeType: 1
forumTopicId: 385317
dashedName: self-referential-sequence
---

# --description--

有幾種方法可以產生自我取用序列。一個非常常見的方法（<a href="https://rosettacode.org/wiki/Look-and-say_sequence" target="_blank" rel="noopener noreferrer nofollow">Look-and-say sequence</a>）是從一個正整數開始，然後透過串接相鄰相同數字的列舉群組來產生下一個項目：

<pre>0, 10, 1110, 3110, 132110, 1113122110, 311311222110 ...</pre>

產生的詞彙長度呈幾何級數增長，且永遠不會收斂。

另一種產生自我參考序列的方法是總結前一項。

計算每個相同數字的數量，然後將每個排序後的列舉數字的總和與數字串接起來。請注意，前五項與前一個序列相同。

<pre>0, 10, 1110, 3110, 132110, 13123110, 23124110 ...</pre>

將數字從大到小排序。不要包含前一項中未出現的數字計數。

根據種子值，以此方式產生的序列總是會收斂到一個穩定值或一個短週期樣式。（就我們的目的而言，收斂表示一個元素與先前出現過的元素相符。）所示序列以種子值 0 為例，在 11 次迭代後收斂到穩定值 1433223110。收斂最快的種子值是 22。它在第一個元素後即變得穩定。（下一個元素是 22，這是之前見過的。）

# --instructions--

撰寫一個函式（程式），該函式以種子值作為參數，產生一個自我參考序列直到它收斂，並將其傳回為一個陣列。

# --hints--

`selfReferential` 應該是一個函式（程式）。

```js
assert(typeof selfReferential === 'function');
```

`selfReferential(40)` 應該傳回一個陣列。

```js
assert(Array.isArray(selfReferential(40)));
```

`selfReferential(40)` 應該傳回 `["40", "1410", "142110", "14123110", "1413124110", "2413125110", "151413224110", "152413225110", "251413324110", "152423224110", "152413423110"]`。

```js
assert.deepEqual(selfReferential(40), [
  '40',
  '1410',
  '142110',
  '14123110',
  '1413124110',
  '2413125110',
  '151413224110',
  '152413225110',
  '251413324110',
  '152423224110',
  '152413423110'
]);
```

`selfReferential(132110)` 應該傳回 `["132110", "13123110", "23124110", "1413223110", "1423224110", "2413323110", "1433223110"]`。

```js
assert.deepEqual(selfReferential(132110), [
  '132110',
  '13123110',
  '23124110',
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(132211)` 應該傳回 `["132211", "132231", "232221", "134211", "14131231", "14231241", "24132231", "14233221"]`。

```js
assert.deepEqual(selfReferential(132211), [
  '132211',
  '132231',
  '232221',
  '134211',
  '14131231',
  '14231241',
  '24132231',
  '14233221'
]);
```

`selfReferential(1413223110)` 應該傳回 `["1413223110", "1423224110", "2413323110", "1433223110"]`。

```js
assert.deepEqual(selfReferential(1413223110), [
  '1413223110',
  '1423224110',
  '2413323110',
  '1433223110'
]);
```

`selfReferential(251413126110)` 應該傳回 `["251413126110", "16151413225110", "16251413226110", "26151413325110", "16251423225110", "16251413424110", "16153413225110"]`。

```js
assert.deepEqual(selfReferential(251413126110), [
  '251413126110',
  '16151413225110',
  '16251413226110',
  '26151413325110',
  '16251423225110',
  '16251413424110',
  '16153413225110'
]);
```

# --seed--

## --seed-contents--

```js
function selfReferential(n) {

}
```

# --solutions--

```js
function selfReferential(n) {
  var descending,
    i,
    incr,
    j,
    max_i,
    max_len,
    max_seq,
    seq,
    sequence,
    indexOf =
      [].indexOf ||
      function(item) {
        for (var i = 0, l = this.length; i < l; i++) {
          if (i in this && this[i] === item) return i;
        }
        return -1;
      };

  sequence = function(n) {
    var c, cnt, cnts, d, digit, i, j, l, len, new_cnts, ref, s, seq;
    cnts = {};
    ref = n.toString();
    for (j = 0, len = ref.length; j < len; j++) {
      c = ref[j];
      d = parseInt(c);
      incr(cnts, d);
    }
    seq = [ref];
    while (true) {
      s = '';
      for (i = l = 9; l >= 0; i = --l) {
        if (cnts[i]) {
          s += '' + cnts[i] + i;
        }
      }
      if (indexOf.call(seq, s) >= 0) {
        break;
      }
      seq.push(s);
      new_cnts = {};
      for (digit in cnts) {
        cnt = cnts[digit];
        incr(new_cnts, cnt);
        incr(new_cnts, digit);
      }
      cnts = new_cnts;
    }
    return seq;
  };

  incr = function(h, k) {
    if (h[k] == null) {
      h[k] = 0;
    }
    return (h[k] += 1);
  };

  descending = function(n) {
    var tens;
    if (n < 10) {
      return true;
    }
    tens = n / 10;
    if (n % 10 > tens % 10) {
      return false;
    }
    return descending(tens);
  };

  return sequence(n);
}
```
