---
id: 5eb3e497b8d6d7f63c5517ea
title: 搜尋 `list` 中的 `record`
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

一筆記錄由描述一個物體的屬性組成。每個屬性都有一個名稱和一個值。例如，一個人可以有一個屬性 `age`，其值為 25。對一個記錄列表來說，一個重要的操作是尋找具有特定屬性值的記錄。

# --instructions--

撰寫一個函式（程式），該函式（程式）以字串作為參數。該函式（程式）應傳回 `list` 中 `name` 屬性值與給定字串相符的項目的索引。

# --hints--

`searchCity` 應該是一個函式（程式）。

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")` 應該傳回一個數字。

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")` 應該傳回 `6`。

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")` 應該傳回 `9`。

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")` 應該傳回 `1`。

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")` 應該傳回 `4`。

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")` 應該傳回 `0`。

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
