---
id: 680fc849a6f2be0a8597c593
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

الآن، أنشئ متغيرات الحالة والمُعدِّلات لـ `powerSource` و `powers`. يجب أن تكون القيمة الابتدائية لـ `powerSource` سلسلة نصية فارغة، ولـ `powers` مصفوفة فارغة.

# --hints--

يجب أن تستخدم التفكيك للمصفوفة لتعيين متغير حالة `powerSource` ومُعدِّل `setPowerSource`.

```js
assert.match(code, /(const|let|var)\s+\[\s*powerSource\s*,\s*setPowerSource\s*\]/);
```

يجب أن يستخدم كل من `powerSource` و `setPowerSource` الخطاف `useState`.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  assert.isAtLeast(abuseState.calls.length, 3);
```

يجب أن يحتوي الخطاف `useState` الخاص بـ `powerSource` على قيمة ابتدائية سلسلة نصية فارغة.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  const newStateCalls = abuseState.calls.slice(2);
  assert.isTrue(newStateCalls.some(call => call[0] === ""));
```

يجب أن تستخدم التفكيك للمصفوفة لتعيين متغير حالة `powers` ومُعدِّل `setPowers`.

```js
assert.match(code, /(const|let|var)\s+\[\s*powers\s*,\s*setPowers\s*\]/);
```

يجب أن يستخدم كل من `powers` و `setPowers` الخطاف `useState`.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  assert.isAtLeast(abuseState.calls.length, 4);
```

يجب أن يحتوي الخطاف `useState` الخاص بـ `powers` على قيمة ابتدائية مصفوفة فارغة.

```js
  const abuseState = __helpers.spyOn(React, "useState");
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.SuperheroForm);

  const newStateCalls = abuseState.calls.slice(2);
  assert.isTrue(newStateCalls.some(call => Array.isArray(call[0]) && call[0].length === 0));
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {
  
  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  --fcc-editable-region--
  
  --fcc-editable-region--
  
  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
    </div>
  )
};
```
