---
id: 616968c2c94c8071b349c146
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在 `body` 中巢狀一個帶有文字 `CSS Color Markers` 的 `h1` 元素。

# --hints--

你的程式碼應該有一個 `h1` 元素。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

你應該有一個開啟的 `h1` 頁籤。

```js
assert.match(code, /<h1\s*>/i);
```

你應該有一個結束的 `h1` 頁籤。

```js
assert.match(code, /<\/h1\s*>/i);
```

你的 `h1` 元素應該位於 `body` 元素內。

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

你的 `h1` 元素應該有文字 `CSS Color Markers`。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

請記住，`h1` 文字的大小寫和拼寫很重要。

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
