---
id: 616d525007b8c5d8b3308b1c
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

如果你的標記置中於頁面上會看起來更好。一個簡單的方法是使用 `margin` <dfn>shorthand property</dfn>。

你可以使用像 `margin-top` 和 `margin-left` 這樣的屬性分別設定元素的外邊距區域，`margin` 簡寫屬性則讓你輕鬆同時設定多個外邊距區域。

若要將你的標記置中於頁面，請將其 `margin` 屬性設為 `auto`。這會將 `margin-top`、`margin-right`、`margin-bottom` 和 `margin-left` 全部設為 `auto`。

# --hints--

你的 `.marker` CSS 規則應該有一個設定為 `auto` 的 `margin` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
