---
id: 616d4a84b756d9c4b8255093
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

是時候為標記添加一些顏色了。請記住，為元素添加顏色的一種方法是使用像 `black`、`cyan` 或 `yellow` 這樣的<dfn>顏色關鍵字</dfn>。

提醒你，以下是如何標的 `freecodecamp` 類別：

```css
.freecodecamp {
  
}
```

創建一個新的 CSS 規則，標的為類別 `marker`，並將其 `background-color` 屬性設為 `red`。

**注意：** 為你的 CSS 添加後，你不會看到任何變化。

# --hints--

你應該創建一個類別選擇器來標的 `marker` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

你的 `.marker` CSS 規則應該有一個 `background-color` 屬性設定為 `red`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
