---
id: 616d3a67ccf800ad94ec89ae
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

現在你將添加一些元素，最終你會將它們設計成彩色標記。

首先，在 `body` 元素內，為你的 `div` 元素添加一個 `class` 屬性並設為 `container`。確保 `div` 元素位於 `h1` 元素下方。

# --hints--

你應該有一個開啟的 `div` 頁籤。

```js
assert.match(code, /<div.*>/i);
```

你應該有一個結束的 `div` 頁籤。

```js
assert.match(code, /<\/div\s*>/i);
```

你的 `div` 元素應該位於你的 `body` 元素內。

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

你的 `div` 元素應該在 `h1` 元素之後。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素應該有一個設定為 `container` 的 `class` 屬性。

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
