---
id: 6176519636a76b810ab1219a
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

創建一個 CSS 規則，標的為類別 `two`，並將其 `background-color` 屬性設定為 `green`。 

此外，創建一個獨立的 CSS 規則，標的為類別 `three`，並將其 `background-color` 設定為 `blue`。

# --hints--

你應該使用類別選擇器來標的 `two` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設定為 `green`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

你應該使用類別選擇器來標的 `three` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

你的 `.three` CSS 規則應該有一個 `background-color` 屬性設定為 `blue`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
