---
id: 61764f7e5240eb7ccc7f6a0a
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

然後，創建一個新的 CSS 規則，標的為類別 `one`，並將其 `background-color` 屬性設為 `red`。

# --hints--

你應該使用類別選擇器來標的 `one` 類別。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.one'));
```

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設為 `red`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
