---
id: 61764c602bee6974e7790f35
title: 步驟11
challengeType: 0
dashedName: step-11
---

# --description--

爲了讓馬克筆有不同的顏色，需要給每支筆設置一個不同的類。 通過在 `class` 屬性中列出多個類，並用空格隔開，你可以爲一個元素添加多個類。 比如，下面的代碼給 `div` 元素添加了 `animal` 和 `dog` 類：

```html
<div class="animal dog">
```

如果你給一個 HTML 元素添加多個類，你在樣式表中列出的前幾個類的樣式可能會被之後的類覆蓋。

首先，給第一個馬克筆 `div` 元素添加類 `one`。


# --hints--

你應該給第一個馬克筆 `div` 元素添加類 `one`。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

你的第一個馬克筆 `div` 應該有類 `marker` 和 `one`。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
