---
id: 61764c602bee6974e7790f35
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

為了讓標記有不同的顏色，你需要為每個標記添加一個獨特的類別。你可以透過在 `class` 屬性中列出多個類別並以空格分隔，將多個類別添加到一個元素中。例如，以下程式碼將 `animal` 和 `dog` 兩個類別添加到一個 `div` 元素：

```html
<div class="animal dog">
```

如果你將多個類別添加到 HTML 元素，樣式表中你列出的第一個類別的樣式可能會被後面的類別覆蓋。

首先，為第一個標記 `div` 元素添加類別 `one`。


# --hints--

你應該為第一個標記 `div` 元素添加類別 `one`。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

你的第一個標記 `div` 應該有 `marker` 和 `one` 這兩個類別。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
