---
id: 61764fdda535587e1fefb3aa
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

為你的第二個標記 `div` 添加類別 `two`，為你的第三個標記 `div` 添加類別 `three`。

# --hints--

你應該為 `.container` 容器中的第二個標記 `div` 元素添加類別 `two`。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('two'));
```

你的第二個標記 `div` 應該有 `marker` 和 `two` 這兩個類別。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild.classList?.contains('two'));
```

你應該為 `.container` 容器中的第三個標記 `div` 元素添加類別 `three`。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

你的第三個標記 `div` 應該有 `marker` 和 `three` 這兩個類別。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
