---
id: 617bbb6b97a83f6d1f7d6e38
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

現在你已經練習過次級色，讓我們回顧如何創建<dfn>三級色</dfn>。三級色是透過將一個原色與相鄰的次級色組合而生成的。

為了創建第三色橙色，請在 `.one` CSS 規則中更新 `rgb` 函式（程式），將紅色設為最大值（`255`）且綠色設為 `127`。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(255, 127, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 255, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
