---
id: 617bc3386dc7d07d6469bf20
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

請注意，為了創建橙色，你增加了 `rgb` 值中紅色的強度並降低了綠色的強度。這是因為橙色是由紅色與黃色混合產生的。

要創建第三色春綠色，將青色與綠色混合。在 `.two` CSS 規則中更新 `rgb` 函式（程式），將綠色設為最大值，藍色設為 `127`。

# --hints--

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(0, 255, 127)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
