---
id: 617b65579ce424bf5f02ca73
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

如你在課程中所學，有兩種主要的色彩模型：用於電子裝置的 <dfn>additive</dfn> RGB（紅、綠、藍）模型，以及用於列印的 <dfn>subtractive</dfn> CMYK（青、品紅、黃、黑）模型。

在這個專案中，你將使用 RGB 模型。這表示顏色從黑色開始，隨著不同層次的紅色、綠色和藍色被介紹而改變。一個簡單的方法是使用 CSS 的 `rgb` 函式（程式）。

創建一個新的 CSS 規則，標的為類別 `container`，並將其 `background-color` 設為黑色，使用 `rgb(0, 0, 0)`。

# --hints--

你應該使用類別選擇器來標的 `container` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

你的 `.container` CSS 規則應該有一個 `background-color` 屬性設為 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
