---
id: 617b92b9de349513466f6156
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

雖然紅色和藍色標記看起來相同，但綠色標記比之前要淺得多。這是因為 `green` 顏色關鍵字實際上是較深的色調，大約介於黑色和綠色最大值的中間。

在 `.two` CSS 規則中，將 `rgb` 函式（程式）中的綠色值設為 `127` 以降低其強度。

# --hints--

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(0, 127, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
