---
id: 617b8b38f32bf2080a140675
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

正如你在之前的課程中所學，函式（程式）是一段可以接受輸入並執行特定動作的程式碼。CSS 的 `rgb` 函式接受紅色、綠色和藍色的值，或稱為<dfn>引數</dfn>，並產生一種顏色：

```css
rgb(red, green, blue);
```

每個紅色、綠色和藍色的數值都是從 `0` 到 `255` 的數字。`0` 表示該顏色的比例為 0％，且為黑色。`255` 表示該顏色的比例為 100％。

在 `.one` CSS 規則中，將顏色關鍵字 `red` 替換為 `rgb` 函式（程式）。對於 `rgb` 函式（程式），將紅色的值設為 `255`，綠色的值設為 `0`，藍色的值設為 `0`。

# --hints--

你的 `.one` CSS 規則不應該使用 `red` 顏色關鍵字來設定 `background-color` 屬性。

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(255, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
