---
id: 617bb5624a75e86463b7e638
title: 第 23 步
challengeType: 0
dashedName: step-23
---

# --description--

<dfn>次要色彩</dfn> 是你將主要色彩混合後得到的色彩。當你改變紅色、綠色和藍色的數值時，可能已經注意到上一個步驟中的一些次要色彩。

為了創建第一個次要顏色黃色，請更新 `.one` CSS 規則中的 `rgb` 函式（程式），以結合純紅色和純綠色。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(255, 255, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
