---
id: 619b7fd56aa2253778bcf5f7
title: 第 43 步
challengeType: 0
dashedName: step-43
---

# --description--

<dfn>HSL</dfn> 顏色模型或色相、飽和度和亮度是表示顏色的另一種方式。

CSS `hsl` 函數接受 3 個值：0 到 360 的數字表示色相，0 到 100 的百分比表示飽和度，0 到 100 的百分比表示亮度。

如果你想象一個色輪，色相紅色是 0 度，綠色是 120 度，藍色是 240 度。

飽和度指顏色強度從 0% 或灰色到 100% 的純色。 你必須在飽和度和亮度值中添加百分號 `%` 。

亮度是指顏色的明亮程度，從 0% 或全黑到 100% 或全白，50% 爲中性色。

在 `.blue` CSS 規則中，使用 `hsl` 函數將 `background-color` 屬性更改爲純藍色。 將色相設置爲 `240`，將飽和度設置爲 `100%`，將亮度設置爲 `50%`。

# --hints--

你的 `.blue` CSS 規則應該有一個 `background-color` 屬性，設置爲 `hsl(240, 100%, 50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
