---
id: 619b74fa777a2b2473c94f82
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

更新 CSS 類別選擇器 `.two`，使其標的為新的 `green` 類別。並更新 `.three` 類別選擇器，使其標的為新的 `blue` 類別。

# --hints--

你的程式碼中不應再有 `.two` 類別選擇器。

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.two'));
```

你應該使用類別選擇器來標的 `green` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.green'));
```

你的 `.green` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 0, 0)');
```

你的程式碼中不應再有 `.three` 類別選擇器。

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.three'));
```

你應該使用類別選擇器來標的 `blue` 類別。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.blue'));
```

你的 `.blue` CSS 規則應該有一個 `background-color` 屬性設定為 `rgb(0, 0, 0)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
