---
id: 619b7fd56aa2253778bcf5f7
title: 步驟 43
challengeType: 0
dashedName: step-43
---

# --description--

<dfn>HSL</dfn> 色彩模型，或稱色相、飽和度與明度，是另一種表述顏色的方式。

CSS 的 `hsl` 函式接受 3 個值：色相為 0 到 360 的數字、飽和度為 0 到 100 的百分比，以及亮度為 0 到 100 的百分比。

如果你想像一個色輪，色相紅位於 0 度，綠位於 120 度，藍位於 240 度。

飽和度是顏色的強度，從 0％（灰色）到 100％（純色）。你必須為飽和度和亮度值添加百分比符號 `%`。

明度是顏色看起來的亮度，從 0％（完全黑色）到 100％（完全白色），50％ 為中性。

在 `.blue` CSS 規則中，使用 `hsl` 函式（程式）將 `background-color` 屬性更改為純藍色。將色相設為 `240`，飽和度設為 `100%`，亮度設為 `50%`。

# --hints--

你的 `.blue` CSS 規則應該有一個 `background-color` 屬性設定為 `hsl(240, 100%, 50%)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
