---
id: 61a493ead935c148d2b76312
title: 步驟 45
challengeType: 0
dashedName: step-45
---

# --description--

`linear-gradient` 函式非常靈活 — 這裡是你將在本教學中檢視的基本語法：

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` 是用於過渡的線條方向。`color1` 和 `color2` 是顏色引數，這些顏色將用於過渡本身。這些可以是任何型別的顏色，包括顏色關鍵字、十六進位、`rgb` 或 `hsl`。

現在你將沿著 90 度線對第一個標記應用紅到綠的漸層。

首先，在 `.red` CSS 規則中，將 `background` 屬性設為 `linear-gradient()`，並將值 `90deg` 作為 `gradientDirection` 傳入。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性，值為 `linear-gradient(90deg)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
