---
id: 61a4ae5f29eb5584187201c3
title: 步驟 56
challengeType: 0
dashedName: step-56
---

# --description--

對於此標記，你將使用十六進位色碼作為你的漸層。

使用 `linear-gradient` 函式（程式）並將 `gradientDirection` 設定為 `180deg`。而對於第一個顏色引數，使用紅色為 `55`、綠色為 `68`、藍色為 `0D` 的十六進位色碼。

# --hints--

你的 `.green` CSS 規則應該有一個 `background` 屬性，值為 `linear-gradient(180deg, #55680D)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:linear-gradient\(180deg,#55680D\)/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

.green {
--fcc-editable-region--
  background: #007F00;
--fcc-editable-region--
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
