---
id: 61a5c906ab73313316e342f0
title: 步驟 60
challengeType: 0
dashedName: step-60
---

# --description--

如果沒有為 `linear-gradient` 函式（程式）提供 `gradientDirection` 引數，則預設會將顏色從上到下排列，或沿著 180 度線排列。

透過從兩個 `linear-gradient` 函式中移除 `gradientDirection` 引數，進一步清理你的程式碼。

# --hints--

你應該從你的 `.red` 和 `.green` CSS 規則中的 `linear-gradient` 函式（程式）移除 `gradientDirection` 引數。

```js
assert.notMatch(code, /linear-gradient\s*\(\s*180deg/gi);
```

你的 `.red` CSS 規則應該有一個 `background` 屬性，設定為 `linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

你的 `.green` CSS 規則應該有一個 `background` 屬性，設定為 `linear-gradient(#55680D, #71F53E, #116C31)`。

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
