---
id: 61b093179e7fc020b43b1bb3
title: 步驟 68
challengeType: 0
dashedName: step-68
---

# --description--

<dfn>Opacity</dfn> 描述某物的不透明度或非透通性。例如，實心牆是不透明的，光線無法穿透。但飲用玻璃則較為透通，你可以透過玻璃看到另一側。

使用 CSS 的 `opacity` 屬性，你可以控制元素的不透明度或透明度。當值為 `0` 或 0％ 時，元素將完全透明；當值為 `1.0` 或 100％ 時，元素將完全不透明，就像預設一樣。

在 `.sleeve` CSS 規則中，將 `opacity` 屬性設為 `0.5`。

# --hints--

你的 `.sleeve` CSS 規則應該有一個 `opacity` 屬性設定為 `0.5`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity, '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
--fcc-editable-region--
  
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
