---
id: 61b093329e7fc020b43b1bb5
title: 步驟 70
challengeType: 0
dashedName: step-70
---

# --description--

你正在使用 `rgb` 函式（程式）來設定顏色。若要為 `rgb` 顏色添加 alpha 通道，請改用 `rgba` 函式（程式）。

`rgba` 函式（程式）運作方式與 `rgb` 函式（程式）相同，但多接受一個從 `0` 到 `1.0` 的數字作為 alpha 通道：

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

你也可以對 `hsl` 和 `hex` 顏色使用透明通道。你很快就會看到如何操作。

在 `.sleeve` 規則中，使用 `rgba` 函式（程式）將 `background-color` 屬性設為純白色且不透明度為 50%。

# --hints--

你的 `.sleeve` CSS 規則應該有一個 `background-color` 屬性設為 `rgba(255, 255, 255, 0.5)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
--fcc-editable-region--
  background-color: white;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
