---
id: 61b0a1b2af494934b7ec1a72
title: 步驟 80
challengeType: 0
dashedName: step-80
---

# --description--

太棒了。你的紅色標記筆看起來很不錯。現在你只需要為你的其他標記筆添加筆蓋和筆套。

為你的綠色和藍色標記添加帽子和套筒。你可以直接複製紅色標記中的 `div` 元素，然後貼到另外兩個標記中。

# --hints--

你的綠色標記 `div` 應該包含兩個 `div` 元素。

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
greenMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'))
assert.lengthOf(greenMarkerChildren, 2);
```

你的綠色標記筆蓋的 `div` 元素應該在套筒的 `div` 元素之前。

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
const greenMarkerCap = document.querySelector('.green .cap');
const greenMarkerSleeve = document.querySelector('.green .sleeve');
assert.isBelow(greenMarkerChildren.indexOf(greenMarkerCap), greenMarkerChildren.indexOf(greenMarkerSleeve));
```

你的藍色標記 `div` 應該包含兩個 `div` 元素。

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
blueMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(blueMarkerChildren, 2);
```

你的藍色標記筆的蓋子 `div` 元素應該在套筒 `div` 元素之前。

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
const blueMarkerCap = document.querySelector('.blue .cap');
const blueMarkerSleeve = document.querySelector('.blue .sleeve');
assert.isBelow(blueMarkerChildren.indexOf(blueMarkerCap), blueMarkerChildren.indexOf(blueMarkerSleeve));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      --fcc-editable-region--
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
