---
id: 68507cadbf36aa089a84ce2d
title: تنفيذ خوارزمية Luhn
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

خوارزمية Luhn، المعروفة أيضًا باسم خوارزمية "معامل باقي القسمة 10" أو "mod 10"، هي صيغة تحقق بسيطة تُستخدم للتحقق من صحة مجموعة متنوعة من أرقام التعريف، مثل أرقام بطاقات الائتمان. هذه هي الخطوات للتحقق من رقم باستخدام خوارزمية Luhn:

- ابدأ من اليمين، واستبعد الرقم الأيمن الأخير (رقم التحقق)، ثم ضاعف قيمة كل رقم ثاني.
- إذا كانت نتيجة مضاعفة الرقم أكبر من `9`، اجمع أرقامه للحصول على رقم واحد. بدلاً من ذلك، يمكنك طرح `9` من النتيجة.
- اجمع كل الأرقام بما في ذلك رقم التحقق.
- إذا كان مجموع كل الأرقام من مضاعفات `10`، فالرقم صحيح؛ وإلا فهو غير صحيح.

على سبيل المثال، اعتبر الرقم `453914881`. خطوات التحقق منه باستخدام خوارزمية Luhn هي:

 ```md 
 Account number      4   5   3   9   1   4   8   8   1 
 Double every other  4  10   3  18   1   8   8  16   1 
 Sum 2-char digits   4   1   3   9   1   8   8   7   1 
 ``` 

- ثم اجمع كل الأرقام، `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- بما أن `42` ليست من مضاعفات `10`، فالرقم غير صحيح.

في هذا المختبر، ستبني مدقق صحة لبطاقة الائتمان باستخدام خوارزمية Luhn.

**الهدف:** نفذ قصص المستخدم أدناه واجعل كل الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `verify_card_number` تأخذ سلسلة نصية من الأرقام (تمثل رقم بطاقة) وتتحقق مما إذا كانت صحيحة وفقًا لخوارزمية Luhn.

2. داخل دالة `verify_card_number`:

- يجب أن تتعامل مع أي شرطات أو فراغات قد تكون موجودة في رقم البطاقة الممرر إليها.
   - أرجع `VALID!` إذا كان رقم البطاقة صحيحًا؛ وإلا أرجع `INVALID!`.

عند إكمال المشروع، يجب أن ترى الرسائل التالية حسب المُدخل:

| رقم البطاقة         | الرسالة  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

يجب أن يكون لديك دالة باسم `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` يجب أن تُرجع `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` يجب أن تُرجع `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` يجب أن تُرجع `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` يجب أن تُرجع `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

يجب أن تُرجع `verify_card_number` القيمة `VALID!` عند استدعائها برقم بطاقة ائتمان صحيح.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

يجب أن تُرجع `verify_card_number` القيمة `INVALID!` عند استدعائها برقم بطاقة ائتمان غير صحيح.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
