---
id: 68c33e502cd97f2fd52cb63b
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتعلم كيفية جعل عناصر الجدول قابلة للوصول من خلال بناء جدول زمني لمؤتمر تقني.

ابدأ بإضافة عنصر `h1` يحتوي على النص `Tech Conference 2025 Schedule`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

يجب أن يحتوي عنصر `h1` على النص `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
