---
id: 68c34c608cf9e53f5d36b1d5
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

بعد ذلك، أضف صفًا ثالثًا إلى الجدول. ابدأ بإضافة عنصر `tr` آخر. داخل هذا العنصر `tr`، أضف عنصر `th` مع خاصية `scope` مضبوطة على `"row"` ومحتوى نصي `11:00 AM`. ثم أسفل هذا العنصر `th`، أضف عنصر `td` بمحتوى نصي `Break`.

# --hints--

يجب أن يكون لديك عنصر `tr` ثالث داخل عنصر `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

يجب أن يكون لديك عنصر `th` داخل عنصر `tr` الثالث.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

يجب أن يحتوي عنصر `th` على المحتوى النصي `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

يجب أن يحتوي عنصر `th` على خاصية `scope` مضبوطة على `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

يجب أن يكون لديك عنصر `td` داخل عنصر `tr` الثالث.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

يجب أن يحتوي عنصر `td` على المحتوى النصي `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
