---
id: 68c34f4aeb2b1a4170506751
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

بعد ذلك، أضف الصف الخامس إلى الجدول. ابدأ بإضافة عنصر `tr` آخر. داخل هذا العنصر `tr`، أضف عنصر `th` مع خاصية `scope` مضبوطة على القيمة `"row"` ومحتوى نصي هو `12:30 PM`. ثم أسفل هذا العنصر `th`، أضف عنصر `td` بمحتوى نصي هو `Lunch Break`. يجب أن يحتوي عنصر `td` أيضًا على خاصية `colspan` مضبوطة على القيمة `3` ليغطي جميع المسارات الثلاثة.

# --hints--

يجب أن يكون لديك عنصر `tr` خامس داخل عنصر `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

يجب أن يكون لديك عنصر `th` داخل عنصر `tr` الخامس.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

يجب أن يحتوي عنصر `th` على المحتوى النصي `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

يجب أن تحتوي عنصر `th` على خاصية `scope` مضبوطة على القيمة `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

يجب أن يكون لديك عنصر `td` داخل عنصر `tr` الخامس.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

يجب أن يحتوي عنصر `td` على المحتوى النصي `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

يجب أن يحتوي عنصر `td` على خاصية `colspan` مضبوطة على القيمة `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
