---
id: 68c34329d430583a6cd5f7f1
title: خطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل عنصر `tr` الخاص بك، أضف أربعة عناصر `th`. يجب أن يحتوي أول عنصر `th` على النص `Time`، والثاني على النص `Track A`، والثالث على النص `Track B`، والرابع على النص `Track C`.

# --hints--

يجب أن يحتوي عنصر `tr` الخاص بك على أربعة عناصر `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

يجب أن يحتوي أول عنصر `th` على النص `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

يجب أن يحتوي ثاني عنصر `th` على النص `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

يجب أن يحتوي ثالث عنصر `th` على النص `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

يجب أن يحتوي رابع عنصر `th` على النص `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
