---
id: 68c34d57ed874e4004f7b5d4
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

في الوقت الحالي، عنصر `td` الذي يحتوي على النص `Break` يمتد عبر عمود واحد فقط. لكن سيكون من الجيد لو امتد عبر الأعمدة الثلاثة كلها.

كما تتذكر من ورش العمل والدروس السابقة، يمكنك استخدام الخاصية `colspan` لجعل خلية الجدول تمتد عبر عدة أعمدة.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

أضف خاصية `colspan` إلى عنصر `td` وحدد قيمتها بـ `3`.

# --hints--

يجب أن يحتوي عنصر `td` الذي يحتوي على النص `Break` على خاصية `colspan` بقيمة `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
