---
id: 68c34ace730ca43e5fa7b8d6
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

حان الوقت الآن لإضافة صف آخر إلى الجدول.

ابدأ بإضافة عنصر `tr` آخر. داخل هذا العنصر `tr`، أضف عنصر `th` مع خاصية `scope` مضبوطة على القيمة `"row"` ومحتوى نصي هو `10:00 AM`.

ثم أضف ثلاثة عناصر `td` بالمحتوى النصي التالي:

- `Accessibility Deep Dive`
- `CSS for Beginners`
- `Inclusive Design Principles`

# --hints--

يجب أن يكون لديك عنصر `tr` ثاني داخل عنصر `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 2);
```

يجب أن يكون لديك عنصر `th` داخل عنصر `tr` الثاني.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[1].querySelector("th"));
```

يجب أن يحتوي عنصر `th` على المحتوى النصي `10:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.textContent.trim(), "10:00 AM");
```

يجب أن يحتوي عنصر `th` على خاصية `scope` مضبوطة على القيمة `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelector("th")?.getAttribute("scope"), "row");
```

يجب أن يحتوي عنصر `tr` الثاني على ثلاثة عناصر `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[1].querySelectorAll("td"), 3);
```

يجب أن يحتوي عنصر `td` الأول على المحتوى النصي `Accessibility Deep Dive`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[0]?.textContent.trim(), "Accessibility Deep Dive");
```

يجب أن يحتوي عنصر `td` الثاني على المحتوى النصي `CSS for Beginners`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[1]?.textContent.trim(), "CSS for Beginners");
```   

يجب أن يحتوي عنصر `td` الثالث على المحتوى النصي `Inclusive Design Principles`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[1].querySelectorAll("td")[2]?.textContent.trim(), "Inclusive Design Principles");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
