---
id: 68c3420cb0176439a95af1b8
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

الخطوة التالية هي إضافة رأس الجدول.

ابدأ بإضافة عنصر `thead` أسفل عنصر `caption`. داخل `thead`، أضف عنصر `tr`.

# --hints--

يجب أن يكون لديك عنصر `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

يجب أن يكون لديك عنصر `tr` داخل `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
