---
id: 68c343692285523aef0fc6bf
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

تُستخدم خاصية `scope` لتحديد ما إذا كانت خلية القسم العلوي هي قسم علوي لصف أو عمود أو مجموعة من الصفوف أو الأعمدة. إليك مثالًا:

```html
<th scope="col">Example Header</th>
```

هذا يساعد قارئ الشاشة على فهم العلاقة بين خلايا القسم العلوي وخلايا البيانات.

لكل عنصر `th`، أضف خاصية `scope` بقيمة `col`.

# --hints--

يجب أن تحتوي كل عنصر `th` على خاصية `scope` بقيمة `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
