---
id: 68c3482278a14e3cf225e897
title: خطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

داخل عنصر `tr`، أضف عنصر `th` يحتوي على النص `9:00 AM`. ثم أسفل ذلك العنصر `th`، أضف ثلاثة عناصر `td` تحتوي على النصوص `Keynote: Tech Future` و`Intro to Web Dev` و`UX for All`.

# --hints--

يجب أن يكون لديك عنصر `th` داخل عنصر `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

يجب أن يحتوي عنصر `th` على النص `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

يجب أن يكون لديك ثلاثة عناصر `td` أسفل عنصر `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

يجب أن يحتوي أول عنصر `td` على النص `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

يجب أن يحتوي ثاني عنصر `td` على النص `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

يجب أن يحتوي ثالث عنصر `td` على النص `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
