---
id: 68c34e201b347c40acda8111
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

حان الوقت الآن لإضافة صف آخر إلى الجدول.

ابدأ بإضافة عنصر `tr` آخر. داخل هذا العنصر `tr`، أضف عنصر `th` مع خاصية `scope` مضبوطة على القيمة `"row"` ومحتوى نصي هو `11:30 AM`.

ثم أضف ثلاثة عناصر `td` بالمحتوى النصي التالي:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

يجب أن يكون لديك عنصر `tr` رابع داخل عنصر `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

يجب أن يكون لديك عنصر `th` داخل عنصر `tr` الرابع.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

يجب أن يحتوي عنصر `th` على المحتوى النصي `11:30 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

يجب أن تحتوي عنصر `th` على خاصية `scope` مضبوطة على القيمة `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

يجب أن يحتوي عنصر `tr` الرابع على ثلاثة عناصر `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

يجب أن يحتوي أول عنصر `td` على المحتوى النصي `AR/VR in Education`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

يجب أن يحتوي ثاني عنصر `td` على المحتوى النصي `JavaScript Fundamentals`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

يجب أن يحتوي ثالث عنصر `td` على المحتوى النصي `Design Systems at Scale`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
