---
id: 68c341332d98fe38fae733af
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

حان الوقت الآن لبدء بناء الجدول.

ابدأ بإضافة عنصر `table` أسفل عنصر `h1`. داخل عنصر `table`، أضف عنصر `caption` بالنص `Schedule by Track and Time`.

# --hints--

يجب أن يكون لديك عنصر `table`.

```js
assert.exists(document.querySelector("table"));
```

يجب أن يحتوي عنصر `table` على عنصر `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

يجب أن يحتوي عنصر `caption` على النص `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
