---
id: 68dbc6b7e8f794a2ae69bc74
title: Implementa el Algoritmo de las N-Reinas
challengeType: 27
dashedName: step-1
---

# --description--

El problema de las N-Reinas te pide colocar N reinas en un tablero de ajedrez de N×N de manera que ninguna de las dos reinas se ataque entre sí (ninguna comparta fila, columna o diagonal). 

Por ejemplo, si hay un tablero 4x4, una estructura válida es:

```md
[1, 3, 0, 2]
```

Eso significa que en la fila 0, la reina está colocada en la columna 1; en la fila 1, la reina está colocada en la columna 3; en la fila 2, la reina está colocada en la columna 0; y en la fila 3, la reina está colocada en la columna 2.

Visualmente, esta estructura se ve así:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Donde `Q` representa una reina y `.` representa una casilla vacía.

En este laboratorio, implementarás el solucionador del problema de las N-Reinas usando el enfoque de búsqueda en profundidad.


**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `dfs_n_queens`.
2. La función debe aceptar exactamente un argumento: un entero `n`.
3. Si `n` es menor que `1`, la función debe devolver una lista vacía (`[]`).
4. La función debe devolver una lista de soluciones; cada solución es a su vez una lista de longitud `n`, donde el elemento en el índice `i` es el índice de columna (basado en 0) de la reina en la fila `i`.

# --hints--

Debes tener una función llamada `dfs_n_queens` que tome un argumento.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

Si `n` es menor que `1`, la función debe devolver una lista vacía.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

La función debe devolver una lista de soluciones, donde cada solución es una lista de longitud `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` debería devolver `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`dfs_n_queens(2)` debería devolver `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`dfs_n_queens(3)` debería devolver `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`dfs_n_queens(4)` debería devolver `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` debería devolver `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
  assert dfs_n_queens(3) == []
`) })
```

`len(dfs_n_queens(5))` debería ser `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
  assert len(dfs_n_queens(6)) == 4
`) })
```

`len(dfs_n_queens(8))` debería ser `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
  assert len(dfs_n_queens(5)) == 10

`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
