---
id: 69008059d43d6751385883c1
title: Depurar erros do operador de incremento e decremento em um aplicativo com bugs
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Seu amigo pediu para você ajudá-lo a depurar o código dele. Ele está praticando os operadores de incremento e decremento mas está obtendo resultados inesperados no console.

Sua tarefa é ajudá-los a identificar os erros e atualizar o código para que as saídas do console correspondam aos resultados esperados.

Você pode consultar as lições anteriores sobre operadores de incremento e decremento assim como notação prefixa e postfix.

**Histórias de Usuário:**

1. Você deve ter uma variável chamada `updatedCounter` que produza corretamente o valor `11` usando o operador de incremento. Atualize a expressão para usar a notação prefixa ou posfixa correta em vez de codificar o valor diretamente.
2. Você deve ter uma variável chamada `finalScore` que produza corretamente o valor `8` usando o operador de incremento. Atualize a expressão para usar a notação prefixa ou posfixa correta em vez de codificar o valor diretamente.
3. Você deve ter uma variável chamada `updatedCoins` que produza corretamente o valor `2` usando o operador de decremento. Atualize a expressão para usar a notação prefixa ou posfixa correta em vez de codificar o valor diretamente.
4. Você deve ter uma variável chamada `newHealth` que produza corretamente o valor `7` usando o operador de decremento. Atualize a expressão para usar a notação prefixa ou postfixa correta em vez de codificar o valor diretamente.

# --hints--

Você deve ter uma variável chamada `updatedCounter` que usa a notação de incremento prefixo (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` deve ser igual a `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Você deve ter uma variável chamada `finalScore` que usa a notação de incremento pós-fixada (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` deve ser igual a `8`.

```js
assert.strictEqual(finalScore, 8);
```

Você deve ter uma variável chamada `updatedCoins` que usa a notação de decremento prefixo (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` deve ser igual a `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Você deve ter uma variável chamada `newHealth` que usa a notação de decremento pós-fixado (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` deve ser igual a `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
