---
id: 612e89d254fe5d3df7d6693d
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sasa umefanya upya mfano wa kisanduku cha html, unahitaji pia kuutumia huu kwa vipengele vilivyo ndani yake. Ili kufanya hivyo, utahitaji kuweka sifa ya `box-sizing` kwa vipengele vyote vingine kuwa vinarithiwa, ambayo inahakikisha vipengele vinavyolengwa vinachukua thamani sawa na kipengele cha mzazi wake.

Pia utahitaji kulenga vipengele vya bandia, ambavyo ni maneno muhimu maalum yanayofuata kichaguzi. Vipengele viwili vya bandia utakavyotumia ni vipengele vya bandia vya `::before` na `::after`.

Kichaguzi cha `::before` huunda kipengele cha bandia ambacho ni mtoto wa kwanza wa kipengele kilichochaguliwa, wakati kichaguzi cha `::after` huunda kipengele cha bandia ambacho ni mtoto wa mwisho wa kipengele kilichochaguliwa. Vipengele hivi vya bandia mara nyingi hutumika kuunda maudhui ya mapambo, ambayo utaona baadaye katika mradi huu.

Kwa sasa, tengeneza kichaguzi cha CSS kulenga vipengele vyote kwa kutumia `*`, na jumuisha vipengele vya bandia `::before` na `::after`. Weka sifa ya `box-sizing` kuwa `inherit`.

# --hints--

Unapaswa kuwa na kichaguzi cha `html``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Kichaguzi chako cha `html` kinapaswa kuwa na sifa ya `box-sizing` ikiwa imewekwa kuwa `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

Unapaswa kuwa na kichaguzi cha `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Kichaguzi chako cha `*, ::before, ::after` kinapaswa kuwa na sifa ya `box-sizing` ikiwa imewekwa kuwa `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
