---
id: 612e78af05201622d4bab8aa
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Unda kipengele cha `div` ndani ya kipengele chako cha `body` ukiwa na `id` imewekwa kuwa `piano`.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` yako inapaswa kuwa ndani ya kipengele chako cha `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

`div` yako inapaswa kuwa na `id` imewekwa kuwa `piano`.

```js
const div = document.querySelector('div');
assert.equal(div?.getAttribute('id'), 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
