---
id: 612ea4c4993aba52ab4aa32e
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Sasa ni wakati wa kutumia vichaguzi vya bandia ulivyoweka awali. Ili kuunda funguo za rangi nyeusi, ongeza kichaguzi kipya cha `.key.black--key::after`. Hiki kinawalenga vipengele vyenye darasa la vitu la `key` na `black--key` na huchagua kipengele cha bandia kilichoundwa baada ya vipengele hivi katika HTML.

Katika kichaguzi kipya, weka `background-color` kuwa `#1d1e22`. Pia weka sifa ya `content` kuwa `""`. Hii itafanya vipengele vya bandia kuwa tupu.

Sifa ya `content` hutumika kuweka au ghairi maudhui ya kipengele. Kwa chaguo-msingi, vipengele vya bandia vilivyoundwa na vichaguzi vya `::before` na `::after` ni tupu, ambayo ina maana havionyeshwi kwenye ukurasa. Kwa kuweka sifa ya `content` kuwa mfuatano wa herufi tupu `""`, unahakikisha vipengele vya bandia vinaonyeshwa, huku bado vikiwa tupu.

Ikiwa ungependa kujaribu, jaribu kuondoa sifa ya `background-color` na kuweka thamani tofauti kwa sifa ya `content`, kama `"♥"`. Kumbuka kurudisha mabadiliko haya unapoisha ili vipimo vipite.

# --hints--

Unapaswa kuwa na kichaguzi cha `.key.black--key::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Kichaguzi chako cha `.key.black--key::after` kinapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `#1d1e22`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Kichaguzi chako cha `.key.black--key::after` kinapaswa kuwa na sifa ya `content` iliyowekwa kuwa `""`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
