---
id: 612e813b3ba67633222cbe54
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kumbuka, sifa ya kipengele `class` inaweza kuwa na maadili mengi. Ili kutofautisha kati ya funguo zako za rangi nyeupe na nyeusi, ongeza thamani ya pili ya `class` ya `black--key` kwa vipengele vyako vya pili, vya tatu, vya tano, vya sita, na vya saba vya `.key`.

# --hints--

Kipengele chako cha pili cha `.key` kinapaswa pia kuwa na `class` ya `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

Kipengele chako cha tatu cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

Kipengele chako cha tano cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

Kipengele chako cha sita cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

Kipengele chako cha saba cha `.key` kinapaswa kuwa na `black--key` katika `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

Unapaswa kuwa na vipengele vitano vya `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

Unapaswa kuwa na vipengele saba vya `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
--fcc-editable-region--
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
