---
id: 612e89562043183c86df287c
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Vivinjari mara nyingi huweka thamani za kiasili za ukingo na nafasi ya ndani kwa vipengele maalum. Ili kuhakikisha mpangilio wako wa piano unaonyeshwa kwa usahihi, unahitaji kusahihisha mfano wa kisanduku.

Ongeza kichaguzi cha `html` kwenye faili lako la CSS, na weka sifa ya `box-sizing` kuwa `border-box`.

# --hints--

Unapaswa kuwa na kichaguzi cha `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Kichaguzi chako cha `html` kinapaswa kuwa na sifa ya `box-sizing` ikiwa imewekwa kuwa `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
