---
id: 612e804c54d5e7308d7ebe56
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Ndani ya kipengele chako cha `.keys`, ongeza vipengele saba vya `div`, kila kikiwa na darasa la vitu la `key`.

# --hints--

Unapaswa kuunda vipengele saba vipya vya `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

Vipengele vyako saba vipya vya `div` vinapaswa kuwa ndani ya kipengele chako cha `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

Vipengele vyako saba vipya vya `div` vinapaswa vyote kuwa na seti ya `class` kuwa `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
