---
id: 612e8279827a28352ce83a72
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sasa, nakili seti ya vipengele saba vya `.key` na ubandike seti mbili za ziada ndani ya div ya `.keys`.

# --hints--

Unapaswa kuwa na vipengele 21 vya `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

Unapaswa kuwa na vipengele 15 vya `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Vipengele vyote vya `.key` vinapaswa kuwa ndani ya kipengele chako cha `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Vipengele vyote vya `.black--key` vinapaswa kuwa ndani ya kipengele chako cha `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css

```
