---
id: 612e9d142affc44a453655db
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Ni wakati wa kupamba funguo zenyewe. Unda kichaguzi cha `class` kwa vipengele vya `.key`. Weka seti ya `background-color` kwa thamani ya `#ffffff`, sifa ya `position` kwa `relative`, sifa ya `width` kwa `41px`, na sifa ya `height` kwa `175px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.key`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key'));
```

Kichaguzi chako cha `.key` kinapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `#ffffff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor, 'rgb(255, 255, 255)');
```

Kichaguzi chako cha `.key` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `relative`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.position,'relative');
```

Kichaguzi chako cha `.key` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `41px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.width, '41px');
```

Kichaguzi chako cha `.key` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `175px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.height, '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
