---
id: 66bb6a9c2dd58b73cd759034
title: Jenga programu ya daftari la alama
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `getAverage` kinachopokea safu ya data ya alama za mtihani kama kigezo na kurudisha alama ya wastani.
1. Unapaswa kuwa na kitendakazi kinachoitwa `getGrade` kinachopokea alama ya mwanafunzi kama kigezo na kurudisha mfuatano wa herufi unaowakilisha daraja la herufi kulingana na alama. Hapa ni alama na daraja zao za herufi zinazolingana:

    | Eneo la alama | Daraja   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. Unapaswa kuwa na kitendakazi kinachoitwa `hasPassingGrade` kinachopokea alama kama kigezo na kurudisha `true` au `false` kulingana na kama alama inawakilisha daraja la kupita.
1. Kitendakazi cha `hasPassingGrade` kinapaswa kutumia kitendakazi cha `getGrade` kupata daraja la herufi, na kukitumia kuamua kama daraja ni la kupita. Daraja la kupita ni lolote tofauti na `"F"`.
1. Unapaswa kuwa na kitendakazi kinachoitwa `studentMsg` kinachopokea safu ya alama na alama ya mwanafunzi kama vigezo. Kitendakazi kinapaswa kurudisha mfuatano wa herufi wenye muundo:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, ikiwa mwanafunzi alipita kozi.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, ikiwa mwanafunzi alipata daraja la kushindwa kozi.

    Badilisha `average-goes-here` na wastani wa alama zote na `grade-goes-here` na daraja la mwanafunzi.
    Tumia `getAverage` kupata alama ya wastani na `getGrade` kupata daraja la mwanafunzi.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `getAverage`.

```js
assert.isFunction(getAverage);
```

Kitendakazi chako cha `getAverage` kinapaswa kurudisha nambari.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` inapaswa kurudisha `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` inapaswa kurudisha `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` inapaswa kurudisha `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` inapaswa kurudisha `47.25`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` inapaswa kurudisha `47.125`.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Kitendakazi chako cha `getAverage` kinapaswa kurudisha alama ya wastani ya alama za mtihani.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Unapaswa kuwa na kitendakazi kinachoitwa `getGrade`.

```js
assert.isFunction(getGrade);
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"A+"` ikiwa alama ni `100`.

```js
assert.strictEqual(getGrade(100), "A+");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"A"` ikiwa alama iko kati ya `90` na `99`.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"B"` ikiwa alama iko kati ya `80` na `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"C"` ikiwa alama iko kati ya `70` na `79`.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"D"` ikiwa alama iko kati ya `60` na `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Kitendakazi chako cha `getGrade` kinapaswa kurudisha `"F"` ikiwa alama iko kati ya `0` na `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Unapaswa kuwa na kitendakazi kinachoitwa `hasPassingGrade`.

```js
assert.isFunction(hasPassingGrade);
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha thamani ya kweli au siyo kweli.

```js
assert.isBoolean(hasPassingGrade(100));
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `true` ikiwa daraja ni `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `false` ikiwa daraja ni `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Kitendakazi chako cha `hasPassingGrade` kinapaswa kurudisha `true` kwa daraja zote za kupita.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Unapaswa kuwa na kitendakazi kinachoitwa `studentMsg`.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` inapaswa kurudisha ujumbe ufuatao: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Kitendakazi chako cha `studentMsg` kinapaswa kurudisha ujumbe sahihi kulingana na alama ya mwanafunzi na wastani wa darasa.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
