---
id: 587d78af367417b2b2512b00
title: 使用 `align-self` 屬性
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

最後一個用於彈性項目的屬性是 `align-self`。這個屬性允許你單獨調整每個項目的對齊方式，而不是一次設定全部。這很有用，因為其他常見的調整技術，如使用 CSS 屬性 `float`、`clear` 和 `vertical-align`，在彈性項目上無法生效。

`align-self` 接受與 `align-items` 相同的值，並且會改寫由 `align-items` 屬性設定的任何值。

# --instructions--

為你的 `#box-1` 和 `#box-2` 添加 CSS 屬性 `align-self`。將 `#box-1` 的值設為 `center`，將 `#box-2` 的值設為 `flex-end`。

# --hints--

`#box-1` 元素的 `align-self` 屬性應設為 `center`。

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

`#box-2` 元素的 `align-self` 屬性應設定為 `flex-end`。

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
