---
id: 587d78ae367417b2b2512afd
title: 使用 flex-basis 屬性設定項目的初始大小
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

`flex-basis` 屬性指定在 CSS 使用 `flex-shrink` 或 `flex-grow` 調整之前，項目的初始大小。

`flex-basis` 屬性使用的單位與其他尺寸屬性相同（`px`、`em`、`%` 等）。值 `auto` 根據內容調整項目的大小。

# --instructions--

使用 `flex-basis` 設定盒子的初始大小。為 `#box-1` 和 `#box-2` 都添加 CSS 屬性 `flex-basis`。給 `#box-1` 設定值為 `10em`，給 `#box-2` 設定值為 `20em`。

# --hints--

`#box-1` 元素應該具有 `flex-basis` 屬性。

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-1` 元素的 `flex-basis` 值應為 `10em`。

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

`#box-2` 元素應該有 `flex-basis` 屬性。

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

`#box-2` 元素的 `flex-basis` 值應為 `20em`。

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
