---
id: 587d78ae367417b2b2512aff
title: 使用 order 屬性重新排列項目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

`order` 屬性用來告訴 CSS 彈性容器中彈性項目的顯示順序。預設情況下，項目會以它們在原始碼 HTML 中出現的相同順序顯示。該屬性接受數字作為值，且可以使用負數。

# --instructions--

為你的 `#box-1` 和 `#box-2` 添加 CSS 屬性 `order`。將 `#box-1` 的值設為 `2`，將 `#box-2` 的值設為 `1`。

# --hints--

`#box-1` 元素的 `order` 屬性應設定為值 `2`。

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

`#box-2` 元素的 `order` 屬性應設定為值 `1`。

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
