---
id: 587d78ab367417b2b2512af2
title: 使用 `flex-direction` 屬性來建立一個列
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

為元素添加 `display: flex` 會將其轉變為彈性容器。這使得可以將該元素的任何子項目排列成列或行。你可以透過為父項目添加 `flex-direction` 屬性並將其設為 row 或 column 來達成。建立列會使子項目水平排列，建立行則會使子項目垂直排列。

`flex-direction` 的其他選項是 `row-reverse` 和 `column-reverse`。

**注意：** `flex-direction` 屬性的預設值是 `row`。

# --instructions--

為你的 `#box-container` 元素添加 CSS 屬性 `flex-direction`，並將其值設為 `row-reverse`。

# --hints--

`#box-container` 元素應該設定 `flex-direction` 屬性為 `row-reverse`。

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
