---
id: 587d78ad367417b2b2512af8
title: "使用 align-items 屬性對齊元素"
challengeType: 0
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

`align-items` 屬性類似於 `justify-content`。回想一下，`justify-content` 屬性沿主軸對齊彈性項目。對於列來說，主軸是水平線，對於行來說則是垂直線。

Flex 容器也有一個 **交叉軸**，它與主軸相反。對於列來說，交叉軸是垂直的；對於行來說，交叉軸是水平的。

CSS 提供 `align-items` 屬性來沿著交叉軸對齊彈性項目。對於列，它告訴 CSS 如何將整列的項目在容器內向上或向下推動。對於行，它則告訴 CSS 如何將所有項目在容器內向左或向右推動。

`align-items` 可用的不同值包括：

<ul><li><code>flex-start</code>：將項目對齊到彈性容器的起始位置。對於列，這會將項目對齊到容器的頂部。對於行，這會將項目對齊到容器的左側。</li><li><code>flex-end</code>：將項目對齊到彈性容器的結束位置。對於列，這會將項目對齊到容器的底部。對於行，這會將項目對齊到容器的右側。</li><li><code>center</code>：將項目對齊到中心。對於列，這會垂直對齊項目（項目上下空間相等）。對於行，這會水平對齊項目（項目左右空間相等）。</li><li><code>stretch</code>：將項目拉伸以填滿彈性容器。例如，列的項目會被拉伸以填滿容器的上下空間。如果沒有指定 <code>align-items</code> 的值，這是預設值。</li><li><code>baseline</code>：將項目對齊到它們的基線。基線是文字的概念，可以想像成字母所坐落的那條線。</li></ul>

# --instructions--

範例有助於展示此屬性的實際效果。為 `#box-container` 元素添加 CSS 屬性 `align-items`，並將其值設為 `center`。

**獎勵**嘗試在程式碼編輯器中使用 `align-items` 屬性的其他選項，以查看它們的差異。但請注意，只有 `center` 的值能通過此挑戰。

# --hints--

`#box-container` 元素的 `align-items` 屬性應設定為 `center`。

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
