---
id: 587d78ab367417b2b2512af0
title: '使用 display: flex 來定位兩個方塊'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

本節使用交替的挑戰風格來展示如何使用 CSS 以靈活的方式定位元素。首先，一個挑戰會解釋理論，接著使用一個簡單的 tweet 組件進行實作挑戰，以應用 flexbox 概念。

在元素上放置 CSS 屬性 `display: flex;` 可讓你使用其他 flex 屬性來建置響應式頁面。

# --instructions--

為你的 `#box-container` 添加 CSS 屬性 `display` 並將其值設為 `flex`。

# --hints--

`#box-container` 應該將 `display` 屬性設定為 `flex`。

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
