---
id: 587d78ae367417b2b2512afe
title: 使用 flex 簡寫屬性
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

有一個捷徑可以一次設定多個 flex 屬性。可以使用 `flex` 屬性同時設定 `flex-grow`、`flex-shrink` 和 `flex-basis` 屬性。

例如，`flex: 1 0 10px;` 將會將該項目設定為 `flex-grow: 1;`、`flex-shrink: 0;` 和 `flex-basis: 10px;`。

預設屬性設定為 `flex: 0 1 auto;`。

# --instructions--

為你的 `#box-1` 和 `#box-2` 添加 CSS 屬性 `flex`。給 `#box-1` 設定值，使其 `flex-grow` 為 `2`、`flex-shrink` 為 `2`，且 `flex-basis` 為 `150px`。給 `#box-2` 設定值，使其 `flex-grow` 為 `1`、`flex-shrink` 為 `1`，且 `flex-basis` 為 `150px`。

當容器大於 300px 時，這些值會使 `#box-1` 以 `#box-2` 兩倍的速率增長以填滿額外空間，當容器小於 300px 時，則以 `#box-2` 兩倍的速率縮小。300px 是兩個盒子的 `flex-basis` 值的總和大小。

# --hints--

`#box-1` 元素的 `flex` 屬性應設定為 `2 2 150px`。

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

`#box-2` 元素的 `flex` 屬性應設定為值 `1 1 150px`。

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

你的程式碼應該為 `#box-1` 和 `#box-2` 使用 `flex` 屬性。

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
