---
id: 587d78ad367417b2b2512afb
title: 使用 flex-shrink 屬性縮小項目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

到目前為止，挑戰中的所有屬性都適用於彈性容器（彈性項目的父容器）。然而，彈性項目有幾個有用的屬性。

第一個是 `flex-shrink` 屬性。當使用它時，若彈性容器太小，會允許項目縮小。當父容器的寬度小於其中所有彈性項目寬度總和時，項目會縮小。

`flex-shrink` 屬性接受數字作為值。數字越大，與容器中其他項目相比，它會縮小得越多。<dfn>例如</dfn>，如果一個項目的 `flex-shrink` 值為 `1`，另一個項目的 `flex-shrink` 值為 `3`，則值為 `3` 的項目會縮小為另一個的三倍。

# --instructions--

為你的 `#box-1` 和 `#box-2` 添加 CSS 屬性 `flex-shrink`。將 `#box-1` 的值設為 `1`，將 `#box-2` 的值設為 `2`。

# --hints--

`#box-1` 元素的 `flex-shrink` 屬性應設定為值 `1`。

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

`#box-2` 元素的 `flex-shrink` 屬性應設定為值 `2`。

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
