---
id: 587d78ac367417b2b2512af6
title: 使用 `justify-content` 屬性對齊元素
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c43gnHm'
forumTopicId: 301102
dashedName: align-elements-using-the-justify-content-property
---

# --description--

有時候，彈性容器內的彈性項目不會填滿容器中的所有空間。通常會想告訴 CSS 如何以特定方式對齊和分配彈性項目的間距。幸運的是，`justify-content` 屬性有幾個選項，可用來達成這個目的。但在檢視這些選項之前，有一些重要的術語需要先了解。

<a href="https://www.freecodecamp.org/news/flexbox-the-ultimate-css-flex-cheatsheet/" target="_blank" rel="noopener noreferrer nofollow">欲了解更多關於 flex-box 屬性的資訊，請點此閱讀</a>

請記住，將彈性容器設定為列會使彈性項目從左到右並排排列。將彈性容器設定為行會使彈性項目從上到下垂直堆疊。對於每種情況，彈性項目排列的方向稱為**主軸**。對於列，這是一條水平線，穿過每個項目。對於行，主軸是一條垂直線，穿過這些項目。

有幾種選項可用於設定沿主軸線排列的彈性項目的間距。其中最常用的是 `justify-content: center;`，它會將所有彈性項目對齊到彈性容器的中央。其他選項包括：

<ul><li><code>flex-start</code>: aligns items to the start of the flex container. For a row, this pushes the items to the left of the container. For a column, this pushes the items to the top of the container. This is the default alignment if no <code>justify-content</code> is specified.</li><li><code>flex-end</code>: aligns items to the end of the flex container. For a row, this pushes the items to the right of the container. For a column, this pushes the items to the bottom of the container.</li><li><code>space-between</code>: aligns items to the center of the main axis, with extra space placed between the items. The first and last items are pushed to the very edge of the flex container. For example, in a row the first item is against the left side of the container, the last item is against the right side of the container, then the remaining space is distributed evenly among the other items.</li><li><code>space-around</code>: similar to <code>space-between</code> but the first and last items are not locked to the edges of the container, the space is distributed around all the items with a half space on either end of the flex container.</li><li><code>space-evenly</code>: Distributes space evenly between the flex items with a full space at either end of the flex container.</li></ul>

# --instructions--

範例有助於展示此屬性的實際效果。為你的 `#box-container` 元素添加 CSS 屬性 `justify-content`，並將其值設為 `center`。

**獎勵**嘗試在程式碼編輯器中使用 `justify-content` 屬性的其他選項，可選方案以查看它們的差異。但請注意，只有 `center` 的值能通過此挑戰。

# --hints--

`#box-container` 元素的 `justify-content` 屬性應設定為 `center`。

```js
const boxContainer = document.querySelector('#box-container');
const justifyDirection =
  window.getComputedStyle(boxContainer)['justify-content'];
assert.strictEqual(justifyDirection, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    justify-content: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 100%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 100%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
