---
id: 587d78ad367417b2b2512af8
title: 使用 `align-items` 屬性對齊元素
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c8aggtk'
forumTopicId: 301101
dashedName: align-elements-using-the-align-items-property
---

# --description--

`align-items` 屬性類似於 `justify-content`。回想一下，`justify-content` 屬性沿主軸對齊彈性項目。對於列來說，主軸是水平線，對於行來說則是垂直線。

Flex 容器也有一個 **交叉軸**，它與主軸相反。對於列來說，交叉軸是垂直的；對於行來說，交叉軸是水平的。

CSS 提供 `align-items` 屬性來沿著交叉軸對齊彈性項目。對於列，它告訴 CSS 如何將整列的項目在容器內向上或向下推動。對於行，它則告訴 CSS 如何將所有項目在容器內向左或向右推動。

`align-items` 可用的不同值包括：

<ul><li><code>flex-start</code>: aligns items to the start of the flex container. For rows, this aligns items to the top of the container. For columns, this aligns items to the left of the container.</li><li><code>flex-end</code>: aligns items to the end of the flex container. For rows, this aligns items to the bottom of the container. For columns, this aligns items to the right of the container.</li><li><code>center</code>: align items to the center. For rows, this vertically aligns items (equal space above and below the items). For columns, this horizontally aligns them (equal space to the left and right of the items).</li><li><code>stretch</code>: stretch the items to fill the flex container. For example, rows items are stretched to fill the flex container top-to-bottom. This is the default value if no <code>align-items</code> value is specified.</li><li><code>baseline</code>: align items to their baselines. Baseline is a text concept, think of it as the line that the letters sit on.</li></ul>

# --instructions--

範例有助於展示此屬性的實際效果。為 `#box-container` 元素添加 CSS 屬性 `align-items`，並將其值設為 `center`。

**獎勵**嘗試在程式碼編輯器中使用 `align-items` 屬性的其他選項，以查看它們的差異。但請注意，只有 `center` 的值能通過此挑戰。

# --hints--

`#box-container` 元素的 `align-items` 屬性應設定為 `center`。

```js
const boxContainer = document.querySelector('#box-container');
const alignment = window.getComputedStyle(boxContainer)['align-items'];
assert.strictEqual(alignment, 'center');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 500px;
    align-items: center;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 200px;
    font-size: 24px;
  }

  #box-2 {
    background-color: orangered;
    width: 200px;
    font-size: 18px;
  }
</style>

<div id="box-container">
  <div id="box-1"><p>Hello</p></div>
  <div id="box-2"><p>Goodbye</p></div>
</div>
```
