---
id: 587d78ae367417b2b2512afc
title: 使用 flex-grow 屬性來擴展項目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

`flex-shrink` 的相反是 `flex-grow` 屬性。請記得 `flex-shrink` 控制容器縮小時項目的大小。`flex-grow` 屬性控制父容器擴大時項目的大小。

使用上一個挑戰中類似的範例，如果一個項目有 `flex-grow` 值為 `1`，另一個項目有 `flex-grow` 值為 `3`，則值為 `3` 的項目會比另一個項目成長三倍。

# --instructions--

為你的 `#box-1` 和 `#box-2` 添加 CSS 屬性 `flex-grow`。將 `#box-1` 的值設為 `1`，將 `#box-2` 的值設為 `2`。

# --hints--

`#box-1` 元素的 `flex-grow` 屬性應設定為值 `1`。

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

`#box-2` 元素的 `flex-grow` 屬性應設定為 `2`。

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
