---
id: 66e7ee20b79186306fc12da5
title: 构建一套足球队卡牌
challengeType: 25 
dashedName: lab-football-team-cards
demoType: onClick
---

# --description--

在本实验中，你将构建一组足球队卡片。用户应能够使用下拉菜单并根据球员的位置在不同球员之间进行过滤。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个具有以下属性的 `footballTeam` 对象：`team`、`year`、`headCoach`、`players`。
1. `team` 属性应包含团队名称作为字串。
1. `year` 属性应包含作为数字的年份。
1. `headCoach` 属性应包含主教练的名字，类型为字串。
1. `players` 属性应为包含至少四个元素的数组。
1. `players` 数组中的每个元素都应该是一个具有属性 `name`、`position`、`isCaptain` 的对象。
1. `name` 属性应包含玩家的名称，类型为字串。
1. `position` 属性应具有以下值之一：`"forward"`、`"midfielder"`、`"defender"` 或 `"goalkeeper"`。
1. `isCaptain` 属性应具有布尔值。一个玩家的 `isCaptain` 属性应设置为 `true`。
1. 你应该在页面上显示 `headCoach`、`team` 和 `year` 的值。这些值应显示在具有 `id` 值为 `head-coach`、`team` 和 `year` 的超文本标记语言元素中。
1. 你应该在页面的 `#player-cards` 元素内显示玩家的数据，每个玩家应显示在一个类名为 `player-card` 的 `div` 中，并且在其中嵌套一个包含玩家名字的 `h2`，如果玩家是队长，则在名字后显示 `(Captain)`，以及一个包含 `Position:` 和玩家位置的 `p`。以下是 HTML 应该呈现的示例：

    ```html
    <div class="player-card">
      <h2>Sergio Batista</h2>
      <p>Position: midfielder</p>
    </div>
    <div class="player-card">
      <h2>(Captain) Diego Maradona</h2>
      <p>Position: midfielder</p>
    </div>
    ```

1. 当使用下拉菜单选择其中一个位置时，页面上应仅显示该位置的 `.player-card` 元素。如果选择了 `"All Players"` 选项，则页面上应显示所有球员。

# --hints--

你应该有一个 `footballTeam` 变量。

```js
assert.exists(footballTeam);
```

`footballTeam` 变量应是一个具有四个属性的对象：`team`、`year`、`headCoach` 和 `players`。

```js
assert.isObject(footballTeam);
assert.containsAllKeys(footballTeam, ['team', 'year', 'headCoach', 'players']);
```

`team` 属性应为字串。

```js
assert.isString(footballTeam.team);
```

`year` 属性应为数字。

```js
assert.isNumber(footballTeam.year);
```

`headCoach` 属性应为字串。

```js
assert.isString(footballTeam.headCoach);
```

`players` 属性应为至少包含四个对象的数组，每个对象应具有密钥 `name`、`position`、`isCaptain`。

```js
assert.isArray(footballTeam.players);
assert.isAtLeast(footballTeam.players.length, 4);
footballTeam.players.forEach(player => {
  assert.isObject(player);
  assert.containsAllKeys(player, ['name', 'position', 'isCaptain']);
})
```

`name` 属性应具有字串值。

```js
footballTeam.players.forEach(({name}) => assert.isString(name));
```

`position` 属性应具有 `"forward"`、`"midfielder"`、`"defender"` 或 `"goalkeeper"` 之一的值。

```js
footballTeam.players.forEach(({position}) => {
  assert.isString(position);
  assert.oneOf(position, ["forward", "midfielder", "defender", "goalkeeper"]);
});
```

`isCaptain` 属性的值应为布尔类型，并且只能有一个船长。

```js
footballTeam.players.forEach(({isCaptain}) => assert.isBoolean(isCaptain));
const listOfCaptains = footballTeam.players.filter(({isCaptain}) => isCaptain);
assert.lengthOf(listOfCaptains, 1);
```

你应该在具有 `id` 值为 `head-coach`、`team` 和 `year` 的超文本标记语言元素中显示 `footballTeam` 对象的 `headCoach`、`team` 和 `year` 值。 

```js
const teamElement = document.querySelector('.team-stats #team');
const yearElement = document.querySelector('.team-stats #year');
const headCoachElement = document.querySelector('.team-stats #head-coach');
assert.equal(teamElement?.innerText.trim(), footballTeam.team);
assert.equal(yearElement?.innerText.trim(), footballTeam.year);
assert.equal(headCoachElement?.innerText.trim(), footballTeam.headCoach);
```

当选择 `All Players` 选项时，所有玩家都应出现在 `#player-cards` 内。

```js
const select = document.querySelector('#players')
select.value = 'all';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, footballTeam.players);
```

当选择 `Position Forward` 选项时，`#player-cards` 中应仅包含转发球员。

```js
const forwards = footballTeam.players.filter(({position}) => position === 'forward')
const select = document.querySelector('#players')
select.value = 'forward';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, forwards);
```

当选择 `Position Midfielder` 选项时，`#player-cards` 中应仅包含中场球员。

```js
const midfielders = footballTeam.players.filter(({position}) => position === 'midfielder')
const select = document.querySelector('#players')
select.value = 'midfielder';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, midfielders);
```

当选择 `Position Defender` 选项时，`#player-cards` 中应仅存在后卫球员。

```js
const defenders = footballTeam.players.filter(({position}) => position === 'defender')
const select = document.querySelector('#players')
select.value = 'defender';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, defenders);
```

当选择 `Position Goalkeeper` 选项时，只应存在守门员球员。

```js
const goalkeepers = footballTeam.players.filter(({position}) => position === 'goalkeeper')
const select = document.querySelector('#players')
select.value = 'goalkeeper';
select.dispatchEvent(new Event('change'))
const playerCards = document.querySelectorAll('.player-card');
const arrayFromPage = Array.from(playerCards).map(el => ({
    name: el.querySelector('h2').innerText.replace('(Captain)', '').trim(),
    position: el.querySelector('p').innerText.replace('Position:', '').trim(),
    isCaptain: /Captain/.test(el.querySelector('h2').innerText)
  }))

assert.sameDeepMembers(arrayFromPage, goalkeepers);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      Build a Set of Football Team Cards
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">Team stats</h1>
    <main>
      <div class="team-stats">
        <p>Team: <span id="team"></span></p>
        <p>Year: <span id="year"></span></p>
        <p>Head coach: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">Filter Teammates:</label>
      <select name="players" id="players">
        <option value="all">All Players</option>
        <option value="forward">Position Forward</option>
        <option value="midfielder">Position Midfielder</option>
        <option value="defender">Position Defender</option>
        <option value="goalkeeper">Position Goalkeeper</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>
      Build a Set of Football Team Cards
    </title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1 class="title">Team stats</h1>
    <main>
      <div class="team-stats">
        <p>Team: <span id="team"></span></p>
        <p>Year: <span id="year"></span></p>
        <p>Head coach: <span id="head-coach"></span></p>
      </div>
      <label class="options-label" for="players">Filter Teammates:</label>
      <select name="players" id="players">
        <option value="all">All Players</option>
        <option value="forward">Position Forward</option>
        <option value="midfielder">Position Midfielder</option>
        <option value="defender">Position Defender</option>
        <option value="goalkeeper">Position Goalkeeper</option>
      </select>
      <div class="cards" id="player-cards"></div>
    </main>
    <footer>&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #0a0a23;
  --light-grey: #f5f6f7;
  --white: #ffffff;
  --black: #000;
}

body {
  background-color: var(--dark-grey);
  text-align: center;
  padding: 10px;
}

.title,
.options-label,
.team-stats,
footer {
  color: var(--white);
}

.title {
  margin: 1.3rem 0;
}

.team-stats {
  display: flex;
  justify-content: space-around;
  flex-wrap: wrap;
  font-size: 1.3rem;
  margin: 1.2rem 0;
}

.options-label {
  font-size: 1.2rem;
}

.cards {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
}

.player-card {
  background-color: var(--light-grey);
  padding: 1.3rem;
  margin: 1.2rem;
  width: 300px;
  border-radius: 15px;
}

@media (max-width: 768px) {
  .team-stats {
    flex-direction: column;
  }
}

```

```js
const teamName = document.getElementById("team");
const worldCupYear = document.getElementById("year");
const headCoach = document.getElementById("head-coach");
const playerCards = document.getElementById("player-cards");
const playersDropdownList = document.getElementById("players");
const footballTeam = {
  team: "Argentina",
  year: 1986,
  headCoach: "Carlos Bilardo",
  players: [
    {
      name: "Sergio Almirón",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Sergio Batista",
      position: "midfielder",
      isCaptain: false,
    },
     {
      name: "Ricardo Bochini",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Claudio Borghi",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "José Luis Brown",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Daniel Passarella",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Jorge Burruchaga",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Néstor Clausen",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "José Luis Cuciuffo",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Diego Maradona",
      position: "midfielder",
      isCaptain: true,
    },
    {
      name: "Jorge Valdano",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Héctor Enrique",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Oscar Garré",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Ricardo Giusti",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Luis Islas",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "Julio Olarticoechea",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Pedro Pasculli",
      position: "forward",
      isCaptain: false,
    },
    {
      name: "Nery Pumpido",
      position: "goalkeeper",
      isCaptain: false,
    },
    {
      name: "Oscar Ruggeri",
      position: "defender",
      isCaptain: false,
    },
    {
      name: "Carlos Tapia",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Marcelo Trobbiani",
      position: "midfielder",
      isCaptain: false,
    },
    {
      name: "Héctor Zelada",
      position: "goalkeeper",
      isCaptain: false,
    },
  ],
};

const { sport, team, year, players } = footballTeam;
const coachName = footballTeam.headCoach;

teamName.textContent = team;
worldCupYear.textContent = year;
headCoach.textContent = coachName;

const setPlayerCards = (arr = players) => {
  playerCards.innerHTML += arr
    .map(
      ({ name, position, isCaptain }) =>
        `
        <div class="player-card">
          <h2>${isCaptain ? "(Captain)" : ""} ${name}</h2>
          <p>Position: ${position}</p>
        </div>
      `
    )
    .join("");
};

playersDropdownList.addEventListener("change", (e) => {
  playerCards.innerHTML = "";

  switch (e.target.value) {
    case "forward":
      setPlayerCards(players.filter((player) => player.position === "forward"));
      break;
    case "midfielder":
      setPlayerCards(
        players.filter((player) => player.position === "midfielder")
      );
      break;
    case "defender":
      setPlayerCards(
        players.filter((player) => player.position === "defender")
      );
      break;
    case "goalkeeper":
      setPlayerCards(
        players.filter((player) => player.position === "goalkeeper")
      );
      break;
    default:
      setPlayerCards();
  }
});

setPlayerCards();

```
