---
id: 641da5462576784453146ec2
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

必要なデータが揃ったので、それを使ってUIを埋めていきましょう。ただし、取得したデータには26人の著者の配列が含まれており、すべてを一度にページに追加するとパフォーマンスが低下する可能性があります。

代わりに、一度に8人の著者を追加し、さらに8人追加するボタンを設置して、表示するデータがなくなるまで続けるようにしてください。

`let`を使って`startingIndex`と`endingIndex`という名前の2つの変数を作成し、それぞれに数値の`0`と`8`を代入してください。また、`authorDataArr`という変数を`let`で宣言し、空の配列に設定してください。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`let`を使って`startingIndex`という名前の変数を宣言してください。

```js
assert.match(code, /let\s+startingIndex/)
```

`startingIndex`変数に`0`を代入してください。

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

`let`を使って`endingIndex`という名前の変数を宣言してください。

```js
assert.match(code, /let\s+endingIndex/)
```

`endingIndex`変数に`8`を代入してください。

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

`let`を使って`authorDataArr`という名前の変数を宣言してください。

```js
assert.match(code, /let\s+authorDataArr/)
```

`authorDataArr`変数に空の配列（`[]`）を代入してください。

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
