---
id: 641da7bfbc7f0f477438ad8a
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

いよいよ`displayAuthors`関数を呼び出す時です。しかし、またしてもすべての著者を一度にページに表示したくはありません。代わりに、`startingIndex`と`endingIndex`変数を使って著者の一部を抽出できます。これを行う最適な方法は`.slice()`配列メソッドです。

まず、`authorDataArr`を表示しているコンソールログの文を削除してください。それから、`displayAuthors`配列と`authorDataArr`を引数にして`.slice()`関数を呼び出してください。開始点には`startingIndex`変数を、終了点には`endingIndex`変数を使います。

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

`data`を`authorDataArr`変数に代入してください

```js
assert.match(code, /authorDataArr\s*=\s*data\s*;?/)
```

`authorDataArr`を表示しているコンソールログの文を削除してください

```js
assert.notMatch(code, /console\.log\(\s*("|'|`)Author\s+Data\s+Array:\s*\1\s*,\s*authorDataArr\s*\)\s*;?/)
```

`displayAuthors`関数を呼び出してください。

```js
assert.match(code, /displayAuthors\(/)
```

`displayAuthors`関数を呼び出し、`authorDataArr.slice()`を引数として渡してください。

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(/)
```

`slice()`メソッドの最初のパラメータは`startingIndex`にしてください。

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex/)
```

`slice()`メソッドの2番目のパラメータは`endingIndex`にしてください。

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex\s*,\s*endingIndex\s*\)\s*\)\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    --fcc-editable-region--
    authorDataArr = data;
    console.log("Author Data Array:", authorDataArr);

    --fcc-editable-region-- 
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
    </div>
  `;
  });
};
```
