---
id: 68a47dcdbce12193caf77012
title: 實作一個排序索引尋找器
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

在本實驗中，你將創建一個函式（程式），該函式（程式）傳回一個值在陣列經過遞增排序後應該插入的最低索引。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `getIndexToIns` 函式（程式），它接受兩個引數：一個陣列和一個數字。
2. 你應該使用 `sort` 方法將陣列排序為遞增順序。
3. 你的 `getIndexToIns` 函式應該使用 `findIndex` 方法傳回應該插入該數字的最低索引。
4. 你的 `getIndexToIns` 函式應該始終傳回一個數字。

**提示：**`findIndex` 方法是 JavaScript 中內建的陣列方法。它接受一個回呼函式，並傳回第一個符合條件的元素索引。`findIndex` 和 `sort` 都是高階函式。

**範例：**

- `getIndexToIns([1, 2, 3, 4], 1.5)` 應該傳回 `1`，因為 `1.5` 大於 `1`（索引 `0`）且小於 `2`（索引 `1`）。
- `getIndexToIns([20, 3, 5], 19)` 應該傳回 `2`，因為排序後為 `[3, 5, 20]`，`19` 小於 `20`（索引為 `2`）且大於 `5`（索引為 `1`）。


# --hints--

你應該有一個 `getIndexToIns` 函式（程式）。

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` 應該永遠傳回一個數字。

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

你的函式（程式）應該使用 `sort` 方法。

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


你的函式（程式）應該使用 `findIndex` 方法。

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` 應該傳回 `3`。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` 應該傳回 `2`。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` 應該傳回 `1`。

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` 應該傳回 `0`。

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` 應該傳回 `2`。

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` 應該傳回 `2`。

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` 應該傳回 `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` 應該傳回 `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

