---
id: 660eebd83100d37862268781
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

La tua variabile `character` attualmente non ha un valore. Puoi assegnare un valore usando l’operatore di <dfn>assegnazione</dfn> `=`. Per esempio:

```js
let hello = "Hello";
```

Assegnare un valore a una variabile al momento della sua dichiarazione è noto come <dfn>inizializzazione</dfn>.

Inizializza la tua variabile `character` assegnandole il valore `"Hello"` durante la sua dichiarazione.

# --hints--

Dovresti usare l’operatore di assegnazione `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Dovresti usare la stringa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Dovresti usare le virgolette doppie per la tua stringa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Dovresti assegnare `"Hello"` alla tua variabile `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
