---
id: 660ef5105b8ba095307a0e50
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora registra di nuovo la tua variabile `character` nella console. Dovresti vedere la stringa `"Hello"` e poi la stringa `"World"` nella console.

# --hints--

Dovresti usare `console.log` una seconda volta.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Dovresti registrare `character` nella console una seconda volta.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


La tua nuova `console.log()` dovrebbe venire dopo la tua riassegnazione.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
