---
id: 660ef31a5be625914a0102cd
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Quando dichiari una variabile con la parola chiave `let`, puoi <dfn>riassegnare</dfn> (o cambiare il valore di) quella variabile in seguito. In questo esempio, il valore di `programmer` viene cambiato da `"Naomi"` a `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Nota che quando riassegni una variabile, non devi usare di nuovo la parola chiave `let`.

Dopo la tua `console.log`, assegna il valore `"World"` alla tua variabile `character`.

# --hints--

Dovresti usare `character` due volte nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Dovresti usare l’operatore di assegnazione per riassegnare `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Dovresti usare la stringa `"World"` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

La tua variabile `character` dovrebbe avere il valore `"World"`.

```js
assert.equal(character, "World");
```

La tua riassegnazione non dovrebbe usare `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
