---
id: 660ef6355e8f5a9e67fe5f46
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Con il tipo di dato `number` puoi eseguire operazioni matematiche, come l'addizione. Prova a stampare `count + 1` nella console.

# --hints--

Dovresti consultare il `console` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Dovresti consultare il `.log()` del `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Dovresti stampare `count + 1` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
