---
id: 660ef0f7c4b8e68ccd1f0786
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript ha sette tipi di dati primitivi, e String è uno di questi. In JavaScript, una <dfn>stringa</dfn> rappresenta una sequenza di caratteri e può essere racchiusa sia tra virgolette singole (`'`) sia tra virgolette doppie (`"`).

Nota che le stringhe sono <dfn>immutabili</dfn>, il che significa che una volta create non possono essere modificate. La variabile può comunque essere riassegnata a un altro valore.

Modifica la tua stringa `"Hello"` per usare le virgolette singole.

# --hints--

Non dovresti avere virgolette doppie nel tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Dovresti usare virgolette singole per la tua stringa `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Dovresti continuare a usare `let` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Dovresti ancora dichiarare una variabile `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

La tua variabile `character` dovrebbe ancora avere la stringa `"Hello"` come valore.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
