---
id: 660ee6e3a242da6bd579de69
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Uno dei concetti più importanti nella programmazione sono le variabili. Una <dfn>variabile</dfn> punta a un indirizzo di memoria specifico che archivia un valore. Le variabili ricevono un nome che può essere usato in tutto il codice per consultare quel valore.

Dichiarare una variabile significa darle un nome. In JavaScript, questo si fa spesso con la parola chiave `let`. Per esempio, ecco come dichiareresti una variabile `hello`:

```js
let hello;
```

La denominazione delle variabili segue regole specifiche: i nomi possono includere lettere, numeri, simboli del dollaro e underscore, ma non possono contenere spazi e non devono iniziare con un numero.

Usa la parola chiave `let` per dichiarare una variabile chiamata `character`.

_Nota_: È pratica comune terminare le istruzioni in JavaScript con un punto e virgola. `;`

# --hints--

Dovresti usare `let` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Dovresti usare `character` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Dovresti usare `let` per dichiarare una variabile `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

La tua dichiarazione dovrebbe terminare con un punto e virgola.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
