---
id: 660ef5c1904955978a986a5c
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Usa `let` per dichiarare una variabile `count`. Assegna a essa il <dfn>numero</dfn> `8`. Quando usi un valore numerico, non devi usare le virgolette. Per esempio:

```js
let money = 100;
```

# --hints--

Dovresti avere una seconda parola chiave `let` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Dovresti usare `let` per dichiarare una variabile `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Dovresti assegnare il numero `8` alla tua variabile `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
