---
id: 660ef857f2806aa626d29d17
title: Passaggio 17
challengeType: 1
dashedName: step-17
---

# --description--

Puoi anche eseguire sottrazione (`-`), moltiplicazione (`*`) e divisione (`/`). Sentiti libero di sperimentare con gli operatori e i numeri nel tuo `console.log`. Quando sei pronto per andare avanti, rimuovi il `console.log`.

# --hints--

Non dovresti avere un `console.log()` nel tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

La tua variabile `character` dovrebbe ancora avere il valore `"Hello"`.

```js
assert.equal(character, "Hello");
```

La tua variabile `count` dovrebbe ancora avere il valore `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
