---
id: 660ef19b95d3308e7dd31bb6
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

La <dfn>console</dfn> ti permette di stampare e visualizzare l’output di JavaScript. Puoi inviare informazioni alla console usando `console.log()`. Per esempio, questo codice stamperà `"Naomi"` nella console:

```js
let developer = "Naomi";
console.log(developer);
```

Il codice sopra consulta la variabile `developer` con il suo nome in `console.log()`. Nota che il valore tra le parentesi tonde è il valore che verrà stampato.

Stampa il valore della variabile `character` nella console. Poi, clicca sul pulsante "Console" per visualizzare la console di JavaScript.

# --hints--

L’inizializzazione della variabile `character` non deve essere modificata.

```js
assert.strictEqual(character, 'Hello');
```

Dovresti consultare `console` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Dovresti consultare la proprietà `log` di `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Dovresti usare le parentesi tonde per chiamare il metodo `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Dovresti stampare la variabile `character` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
