---
id: 660ef55dd468079679ee0092
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Ora sei pronto per dichiarare la tua prossima variabile. Rimuovi entrambe le istruzioni `console.log` e la riassegnazione `character`.

Rimuovi anche la variabile `secondCharacter`, ma lascia invariata l'inizializzazione `character`.

# --hints--

Non dovresti avere istruzioni `console.log()` nel tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Non dovresti riassegnare il valore di `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

La tua variabile `character` dovrebbe avere il valore `"Hello"`.

```js
assert.equal(character, "Hello");
```

Non dovresti avere una variabile `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
