---
id: 660f38c34a4de6141c0c369f
title: Passaggio 71
challengeType: 1
dashedName: step-71
---

# --description--

Il tuo generatore di piramidi ora funziona come previsto. Ma questa è un’ottima occasione per esplorare meglio il codice che hai scritto.

L’operatore di addizione non è l’unico modo per aggiungere valori a una variabile. L’operatore <dfn>assegnazione con addizione</dfn> può essere usato come scorciatoia per dire “prendi il valore originale della variabile, aggiungi questo valore e assegna il risultato di nuovo alla variabile.” Per esempio, queste due istruzioni darebbero lo stesso risultato:

```js
test = test + 1;
test += 1;
```

Aggiorna l’istruzione di iterazione nel loop `for` per usare l’assegnazione con addizione.

# --hints--

Il tuo loop `for` non dovrebbe usare `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Il tuo loop `for` dovrebbe usare l’assegnazione con addizione con `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
