---
id: 660f280dda5040f707c76b4a
title: Passaggio 45
challengeType: 1
dashedName: step-45
---

# --description--

Ora hai una serie di caratteri `#`, ma manca ancora la forma a piramide. Fortunatamente, la variabile `i` rappresenta il numero della "riga" corrente nel tuo loop, permettendoti di usarla per creare una struttura simile a una piramide.

Per farlo, userai il metodo `.repeat()` disponibile per le stringhe. Questo metodo accetta un numero come argomento, che specifica quante volte ripetere la stringa obiettivo. Per esempio, usando `.repeat()` per generare la stringa `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Usa il metodo `.repeat()` sulla tua `character` e passagli `i` come numero.

# --hints--

Dovresti usare il metodo `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Dovresti usare il metodo `.repeat()` sulla tua variabile `character`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Dovresti passare `i` al tuo metodo `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Dovresti usare il metodo `.repeat()` nel metodo `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
