---
id: 660f45ccf4ca5c31f253005a
title: Passaggio 86
challengeType: 1
dashedName: step-86
---

# --description--

Al momento, se cambi `continueLoop` in true, il tuo loop `while` girerà all'infinito. Questo si chiama <dfn>loop infinito</dfn>, e dovresti fare attenzione a evitarli. Un loop infinito può bloccare il sistema, costringendo a un riavvio completo per uscirne.

Per evitarlo, inizia usando l'operatore di incremento per aumentare il valore della variabile `done` all'interno del tuo loop.

# --hints--

Il tuo loop `while` dovrebbe incrementare la variabile `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
