---
id: 660f44f10ea40f300b896a5e
title: Passaggio 84
challengeType: 1
dashedName: step-84
---

# --description--

Ora che hai fatto pratica con le istruzioni `if...else if...else`, puoi rimuoverle dal tuo codice.

Una volta completato, usa `let` per dichiarare una variabile `continueLoop` e assegnarle il valore booleano `false`. Poi usa `let` per dichiarare una variabile `done` e assegnarle il valore `0`.

# --hints--

Non dovresti avere un’istruzione `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

Non dovresti avere un’istruzione `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Non dovresti avere un’istruzione `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Dovresti usare `let` per dichiarare una variabile `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

La tua variabile `continueLoop` dovrebbe avere il valore `false`.

```js
assert.isFalse(continueLoop);
```

Dovresti usare `let` per dichiarare una variabile `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

La tua variabile `done` dovrebbe avere il valore `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
