---
id: 660f4a1472f8e63d76162ce5
title: Passaggio 99
challengeType: 1
dashedName: step-99
---

# --description--

Che succede se realizzi la tua piramide capovolta, o <dfn>invertita</dfn>? È ora di provarci!

Inizia creando un nuovo loop `for`. Dichiara il tuo iteratore `i` e assegnagli il valore `count`, poi usa il booleano `false` per le istruzioni di condizione e iterazione.

# --hints--

Il tuo codice dovrebbe avere un loop `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Il tuo loop `for` dovrebbe inizializzare `i` con il valore `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Il tuo loop `for` dovrebbe usare `false` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


Il tuo loop `for` dovrebbe usare `false` come iterazione.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
