---
id: 660f48e1d3682f39e81843c4
title: Passaggio 94
challengeType: 1
dashedName: step-94
---

# --description--

Il tuo generatore di piramidi funziona ancora. Tuttavia, potrebbe succedere di finire di nuovo in un loop infinito.

Poiché stai controllando solo se `done` non è uguale a `count`, se `done` fosse **maggiore** di `count` il tuo loop continuerebbe all’infinito.

Aggiorna la condizione del tuo loop per verificare se `done` è minore o uguale a `count`.

# --hints--

Il tuo loop `while` dovrebbe controllare se `done` è minore o uguale a `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
