---
id: 660f0a55847d6cc485f29ba5
title: Passaggio 26
challengeType: 1
dashedName: step-26
---

# --description--

Un <dfn>metodo</dfn> in JavaScript è una funzione associata a certi valori o oggetti. Un esempio che hai già incontrato è il metodo `.log()`, che fa parte dell'oggetto `console`.

Gli array hanno i loro metodi, e il primo che esplorerai è il metodo `.push()`. Questo ti permette di "aggiungere" un valore alla fine di un array. Ecco un esempio per aggiungere il numero `12` alla fine di un array:

```js
array.push(12);
```

Usa `.push()` per aggiungere la stringa `"freeCodeCamp"` alla fine del tuo array `rows`. Aggiungi questo codice prima del tuo `console.log` così puoi vedere la modifica che hai fatto al tuo array.

# --hints--

Dovresti usare `.push()` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Dovresti usare il metodo `.push()` del tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Dovresti passare la stringa `"freeCodeCamp"` al tuo metodo `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
