---
id: 660f447efc0e722f016c1be0
title: Passaggio 81
challengeType: 1
dashedName: step-81
---

# --description--

Il testo è ricomparso! Questo perché `"false"` è una stringa che, quando valutata come booleano, diventa `true`. Questo significa che `"false"` è un valore truthy.

Un <dfn>valore truthy</dfn> è un valore considerato vero quando valutato come booleano. La maggior parte dei valori che incontri in JavaScript saranno truthy.

Un <dfn>valore falsy</dfn> è l'opposto: un valore considerato falso quando valutato come booleano. JavaScript ha una lista definita di valori falsy. Alcuni di questi includono `false`, `0`, `""`, `null`, `undefined` e `NaN`.

Prova a cambiare la condizione del tuo `if` in una stringa vuota `""`, che è un valore falsy.

# --hints--

La tua istruzione `if` dovrebbe avere `""` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
