---
id: 660f540c2176ea5dec01306d
title: Passaggio 118
challengeType: 1
dashedName: step-118
---

# --description--

Ottimo lavoro! Sperimenta con valori diversi per le tue variabili `character`, `count` e `inverted`.

Quando sei pronto per passare al tuo prossimo progetto, imposta `character` su `"!"`, `count` su `10` e `inverted` su `false` per continuare.

Congratulazioni per aver completato il tuo primo progetto JavaScript!

# --hints--

Dovresti impostare `character` su `"!"`.

```js
assert.equal(character, "!");
```

Dovresti impostare `count` su `10`.

```js
assert.equal(count, 10);
```

Dovresti impostare `inverted` su `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
