---
id: 660f50a21fe7645252804f2b
title: Passaggio 112
challengeType: 1
dashedName: step-112
---

# --description--

E se avessi un modo per alternare tra una piramide invertita e una piramide standard?

Inizia dichiarando una variabile `inverted` e assegnandole il valore `true`. Non modificherai questa variabile nel tuo codice, ma dovrai usare `let` così i nostri test potranno modificarla in seguito.

# --hints--

Dovresti dichiarare una variabile `inverted` con `let`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Dovresti inizializzare `inverted` con il valore `true`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
