---
id: 660f1e3f047bf4e403268713
title: Passaggio 39
challengeType: 1
dashedName: step-39
---

# --description--

Purtroppo ora non puoi vedere cosa sta facendo il tuo loop.

Usa `let` per dichiarare una variabile `result` e assegnarle una stringa vuota. Una stringa vuota è rappresentata da virgolette senza nulla all’interno, come `""`.

# --hints--

Dovresti dichiarare una variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Dovresti usare `let` per dichiarare la tua variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

La tua variabile `result` dovrebbe essere una stringa vuota.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
