---
id: 660f18f059fe0fda192ce394
title: Passaggio 34
challengeType: 1
dashedName: step-34
---

# --description--

Il tuo loop ora ha bisogno di un iteratore adeguato. L'<dfn>iteratore</dfn> è una variabile che puoi dichiarare specificamente nel tuo `for` loop per controllare come il loop itera o attraversa la tua logica.

È una convenzione comune usare `i` come variabile iteratore in un loop. Un loop `for` ti permette di dichiararlo nelle parentesi tonde `()`. Per esempio, ecco un loop `for` che dichiara una variabile `index` e le assegna il valore `100`.

```js
for (let index = 100; "second"; "third") {

}
```

Sostituisci la stringa `"iterator"` con una dichiarazione `let` per la variabile `i`. Assegna il valore `0` all'inizio. Questo darà alla variabile `i` il valore `0` la **prima volta** che il tuo loop viene eseguito.

# --hints--

Dovresti usare `let` per dichiarare una variabile `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Dovresti assegnare `0` alla tua variabile `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Il tuo loop `for` dovrebbe iniziare un iteratore `i` a `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
