---
id: 660f359af3e32e0f1a6880b7
title: Passaggio 68
challengeType: 1
dashedName: step-68
---

# --description--

Ora dovresti vedere lo stesso gruppo di caratteri nella tua console. La tua funzione `padRow` sta facendo esattamente la stessa cosa che facevi prima, ma ora è in una sezione riutilizzabile a parte.

Usa l'operatore di aggiunta per concatenare uno spazio singolo `" "` all'inizio e alla fine della tua stringa `character` ripetuta.

Ricorda che puoi usare l'operatore `+` per concatenare stringhe in questo modo:

```js
" " + "string"
```

# --hints--

Dovresti concatenare uno spazio singolo all'inizio del valore restituito.

```js
assert.match(padRow(1, 1), /^\s/);
```

Dovresti concatenare uno spazio singolo alla fine del valore restituito.

```js
assert.match(padRow(1, 1), /\s$/);
```

La tua funzione `padRow()` dovrebbe restituire la serie `character` ripetuta con uno spazio prima e dopo la serie.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
