---
id: 660f53ad3d39175c5d4335ac
title: Passaggio 117
challengeType: 1
dashedName: step-117
---

# --description--

Il tuo generatore di piramidi è ora in uno stato completo, con più funzionalità di quelle che avevi inizialmente previsto! Il passaggio successivo è pulire il tuo codice.

Rimuovi tutti i commenti, sia a riga singola che multi-linea, dal tuo codice.

# --hints--

Non dovresti avere commenti a riga singola nel tuo codice.

```js
assert.notMatch(code, /\/\//);
```

Non dovresti avere commenti multi-linea nel tuo codice.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Non dovresti avere commenti nel tuo codice.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Dovresti rimuovere il codice che era commentato con commenti multi-linea.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
