---
id: 660f0ee51d7460ce88cd248d
title: Passaggio 29
challengeType: 1
dashedName: step-29
---

# --description--

Ti aspettavi di vedere `4` nella console? `.push()` restituisce la nuova lunghezza dell'array, dopo aver aggiunto il valore che gli fornisci.

È importante sapere quali valori restituisce un metodo. Prenditi un po' di tempo per sperimentare con `.push()` e `.pop()`. Quando sei pronto, rimuovi tutte le tue chiamate `.push()` e `.pop()` e le tue istruzioni `console.log`.

# --hints--

Non dovresti avere una chiamata `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Non dovresti avere una chiamata `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Non dovresti avere istruzioni di log.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Non dovresti avere una variabile `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Non dovresti avere una variabile `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
