---
id: 660f4e74f7fd3f4a99ac2e50
title: Passaggio 108
challengeType: 1
dashedName: step-108
---

# --description--

Gli array hanno anche un metodo `.shift()`. Questo rimuove il **primo** elemento dell'array, a differenza di `.pop()` che rimuove l'ultimo elemento. Ecco un esempio del metodo `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

L'array `numbers` diventerebbe `[2, 3]`.

Subito sotto il tuo array `numbers`, dichiara una variabile `shifted` e assegnale il risultato della chiamata a `.shift()` sull'array `numbers`. Nella riga successiva, stampa la variabile `shifted` nella console.

# --hints--

Dovresti usare `const` per dichiarare una variabile `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Dovresti chiamare `.shift()` sul tuo array `numbers`.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Dovresti assegnare il risultato della chiamata a `.shift()` alla tua variabile `shifted`.

```js
assert.equal(shifted, 1);
```

Dovresti stampare la tua variabile `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
