---
id: 660f455b044d3230ed971e98
title: Passaggio 85
challengeType: 1
dashedName: step-85
---

# --description--

Un loop <dfn>`while`</dfn> continuerà a eseguirsi finché la condizione `condition` specificata rimane vera. Ha la seguente sintassi:

```js
while (condition) {
  logic;
}
```

Usa questa sintassi per dichiarare un loop `while` con `continueLoop` come condizione. Il corpo deve essere vuoto.

# --hints--

Dovresti usare un loop `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Il tuo loop `while` dovrebbe usare `continueLoop` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
