---
id: 660f1bf673487ae0bb25b900
title: Passaggio 37
challengeType: 1
dashedName: step-37
---

# --description--

Il tuo loop dovrebbe ora eseguire otto volte. All'interno del corpo del loop, stampa il valore dell'iteratore `i` e osserva cosa succede.

# --hints--

Dovresti usare `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Dovresti stampare il valore di `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Dovresti stampare il valore di `i` nel tuo loop `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
