---
id: 660f374d532dc41189cc9cc2
title: Passaggio 69
challengeType: 1
dashedName: step-69
---

# --description--

Ora è il momento di un po’ di matematica. Considera una piramide a tre righe. Se vogliamo che sia centrata, dovrebbe assomigliare a qualcosa del genere:

```md
··#··
·###·
#####
```

Gli spazi vuoti sono stati sostituiti con interpunti, o punti centrali, per rendere più leggibile. Se estrapoli il pattern, puoi vedere che gli spazi all’inizio e alla fine di una riga seguono un pattern.

Aggiorna le tue stringhe di spazi vuoti in modo che vengano ripetute `rowCount - rowNumber` volte.

Apri la console per vedere il risultato.

# --hints--

Dovresti chiamare `.repeat()` sulle tue stringhe `" "` per ripeterle `rowCount - rowNumber` volte.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
