---
id: 660f229d2dbe09ef2954a4a1
title: Passaggio 42
challengeType: 1
dashedName: step-42
---

# --description--

Ricorda che nel tuo precedente loop hai usato l'operatore di aggiunta `+` per aumentare il valore di `i` di `1`.

Puoi fare una cosa simile con un valore stringa, aggiungendo una nuova stringa a una stringa esistente. Per esempio, `hello = hello + " World";` aggiungerebbe la stringa `" World"` alla stringa esistente archiviata nella variabile `hello`. Questo si chiama <dfn>concatenazione</dfn>.

Nel tuo loop `for...of`, usa l'operatore di aggiunta per concatenare il valore `row` al valore `result`.

# --hints--

Dovresti usare l'operatore di concatenazione sulla tua variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Dovresti concatenare `row` alla tua variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Dovresti assegnare il risultato della tua concatenazione di nuovo alla variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
