---
id: 660f47afe4c98536715d5fa4
title: Passaggio 91
challengeType: 1
dashedName: step-91
---

# --description--

L’operatore di <dfn>disuguaglianza stretta</dfn> `!==` ti permette di verificare se due valori non sono uguali o non hanno lo stesso tipo. La sintassi è simile a quella dell’operatore di uguaglianza: `value !== 4`.

Attualmente il loop `while` viene eseguito solo se `continueLoop` è true. Aggiorna la condizione del loop `while` per verificare se `done` non è uguale a `count`.

# --hints--

Il tuo loop `while` dovrebbe verificare se `done` e `count` non sono uguali.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
