---
id: 660f2b6fd54ac1fc142804dd
title: Passaggio 47
challengeType: 1
dashedName: step-47
---

# --description--

La logica per formattare questa piramide probabilmente diventerà complicata, il che significa che è un ottimo momento per estrarre quel codice in una funzione.

Una <dfn>funzione</dfn> è un blocco di codice che può essere riutilizzato in tutta la tua applicazione. Le funzioni si dichiarano con la seguente sintassi:

```js
function name(parameter) {

}
```

La parola chiave `function` dice a JavaScript che la variabile `name` sarà una funzione. `parameter` è una variabile che rappresenta un valore passato alla funzione quando viene usata. Una funzione può avere tanti o pochi <dfn>parametri</dfn> quanto vuoi. Come in un loop `for`, lo spazio tra le parentesi graffe è il <dfn>corpo della funzione</dfn>.

Dichiara una funzione `padRow`. Non creare ancora nessuna variabile parametro. Il corpo della funzione deve essere vuoto. Ricorda che devi usare la convenzione camel case per i nomi.

# --hints--

Dovresti usare la parola chiave `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Dovresti dichiarare una funzione `padRow`.

```js
assert.isFunction(padRow);
```

La tua funzione `padRow()` non deve avere parametri.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

La tua funzione `padRow()` deve avere un corpo vuoto.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
