---
id: 660f4cffb1459d45e34902d1
title: Passaggio 106
challengeType: 1
dashedName: step-106
---

# --description--

Puoi effettivamente costruire la piramide invertita senza dover iterare "all'indietro" come hai fatto.

Per farlo, dovrai imparare un paio di nuovi metodi per gli array. Inizia usando `const` per dichiarare una variabile `numbers`. Assegna a questa un array con gli elementi `1`, `2` e `3`. Poi stampa l'array `numbers`.

# --hints--

Dovresti usare `const` per dichiarare una variabile `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

La tua variabile `numbers` dovrebbe essere un array.

```js
assert.isArray(numbers);
```

Il tuo array `numbers` dovrebbe avere gli elementi `1`, `2` e `3` in questo ordine.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Dovresti stampare il tuo array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
