---
id: 660f0f980e98e8cf77f1ce31
title: Passaggio 30
challengeType: 1
dashedName: step-30
---

# --description--

Cambia la tua dichiarazione `rows` in modo che venga assegnato di nuovo un array vuoto.

Inoltre, modifica la tua stringa `'Hello'` per usare di nuovo le virgolette doppie. In generale, non importa quale delle due preferisci, ma dovresti mantenere coerenza in questa scelta in tutto il progetto.

# --hints--

Il tuo array `rows` dovrebbe essere vuoto.

```js
assert.empty(rows);
```

La tua stringa `"Hello"` dovrebbe usare le virgolette doppie.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
