---
id: 660f061d259bbebc37461080
title: Passaggio 20
challengeType: 1
dashedName: step-20
---

# --description--

Puoi consultare i valori all’interno di un array usando l'<dfn>indice</dfn> del valore. Un indice è un numero che rappresenta la posizione del valore nell’array, a partire da `0` per il primo valore.

Puoi consultare il valore usando la <dfn>notazione a parentesi</dfn>, come `array[0]`.

Usa `console.log` e la notazione a parentesi per stampare il primo valore nel tuo array `rows`.

# --hints--

Dovresti avere un’istruzione `console.log()` nel tuo codice.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Dovresti consultare il tuo array `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Dovresti usare la notazione a parentesi con il tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Dovresti usare la notazione a parentesi per consultare il primo elemento del tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Dovresti stampare il primo elemento del tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
