---
id: 660f3ba3cceef11b6ba08b59
title: Passaggio 75
challengeType: 1
dashedName: step-75
---

# --description--

Purtroppo ora la base della piramide è scomparsa. Questo perché hai creato un altro <dfn>errore off-by-one</dfn>.

Il tuo loop originale iterava per `i` valori da `0` a `7`, perché `count` è `8` e la tua condizione richiede che `i` sia minore di `count`. Ora il tuo loop sta iterando per `i` valori da `1` a `7`.

Il tuo loop deve essere aggiornato per iterare anche quando `i` è `8`. Guardando la tua logica, questo significa che il loop dovrebbe iterare quando `i` è <dfn>minore o uguale</dfn> a `count`. Puoi usare l'operatore minore o uguale `<=` per questo.

Aggiorna la condizione del tuo loop per farlo iterare mentre `i` è minore o uguale a `count`.

# --hints--

Il tuo loop `for` non dovrebbe controllare se `i` è minore di `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Il tuo loop `for` dovrebbe controllare se `i` è minore o uguale a `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
