---
id: 660f3915b41a441537ec9f5e
title: Passaggio 72
challengeType: 1
dashedName: step-72
---

# --description--

Poiché stai solo aumentando `i` di `1`, puoi usare l'<dfn>operatore di incremento</dfn> `++`. Questo operatore aumenta il valore di una variabile di 1, aggiornando l'assegnazione per quella variabile. Per esempio, `test` diventerebbe `8` qui:

```js
let test = 7;
test++;
```

Sostituisci la tua assegnazione con addizione con l'operatore di incremento per l'iterazione del tuo loop.

# --hints--

Il tuo loop `for` non dovrebbe usare l'assegnazione con addizione con `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Il tuo loop `for` dovrebbe usare l'operatore di incremento su `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
