---
id: 660f1b6e60bd9edf902c81fd
title: Passaggio 36
challengeType: 1
dashedName: step-36
---

# --description--

La tua istruzione di <dfn>iterazione</dfn> dirà al loop cosa fare con l'iteratore dopo ogni esecuzione.

Quando riassegni una variabile, puoi usare la variabile per fare riferimento al valore precedente prima della riassegnazione. Questo ti permette di fare cose come aggiungere tre a un numero esistente. Per esempio, `bees = bees + 3` aumenterebbe il valore di `bees` di tre.

Usa quella sintassi per sostituire la tua stringa `"iteration"` con un'istruzione di riassegnazione che aumenti `i` di uno.

# --hints--

Dovresti aggiungere uno alla tua variabile `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Dovresti assegnare `i + 1` di nuovo alla tua variabile `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Il tuo loop `for` dovrebbe aumentare `i` di `1` a ogni iterazione.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
