---
id: 660f1cedf3676fe26122ebf6
title: Passaggio 38
challengeType: 1
dashedName: step-38
---

# --description--

Dovresti vedere i numeri da zero a sette stampati nella console, uno per riga. Questo servirà come base per generare la tua piramide.

Sostituisci la tua istruzione di log con un’istruzione per aggiungere `i` al tuo array `rows`.

# --hints--

Non dovresti avere una chiamata `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Dovresti chiamare `.push()` sul tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Dovresti aggiungere `i` al tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

La tua `.push()` dovrebbe avvenire nel tuo loop `for`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
