---
id: 660f3ce51f70571e1c5227c8
title: Passaggio 76
challengeType: 1
dashedName: step-76
---

# --description--

I commenti possono essere utili per spiegare perché il tuo codice segue un certo approccio o per lasciare note da fare per il tuo futuro io.

In JavaScript, puoi usare `//` per lasciare un commento su una singola riga nel tuo codice.

Aggiungi un commento su una singola riga sopra il tuo loop `for` per ricordarti di cambiare il codice in un tipo diverso di loop.

# --hints--

Dovresti iniziare un commento su una singola riga con `//`.

```js
assert.match(code, /\/\//);
```

Il tuo commento su una singola riga dovrebbe essere lungo almeno cinque caratteri.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
