---
id: 660f46460f9c36330ebc07d8
title: Passaggio 88
challengeType: 1
dashedName: step-88
---

# --description--

L'operatore di uguaglianza può causare comportamenti strani in JavaScript. Per esempio, `"0" == 0` è true, anche se uno è una stringa e l'altro un numero.

L'<dfn>operatore di uguaglianza stretta</dfn> `===` serve a verificare se due valori sono uguali e hanno lo stesso tipo. Come regola generale, questo è l'operatore di uguaglianza che dovresti sempre usare. Con l'operatore di uguaglianza stretta, `"0" === 0` diventa false, perché anche se possono avere lo stesso valore zero, non sono dello stesso tipo.

Aggiorna la tua condizione `done == count` per usare l'operatore di uguaglianza stretta.

# --hints--

La tua condizione `if` dovrebbe usare l'uguaglianza stretta.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
