---
id: 660f07d231941bc11719f664
title: Passaggio 21
challengeType: 1
dashedName: step-21
---

# --description--

Gli array sono speciali perché sono considerati <dfn>mutabili</dfn>. Questo significa che puoi cambiare direttamente il valore a un indice.

Per esempio, questo codice assegnerebbe il numero `25` al secondo elemento dell'array:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Aggiorna il **terzo** elemento del tuo array `rows` con il numero `10`. Poi stampa l'array `rows` nella console.

# --hints--

Dovresti usare di nuovo la notazione a parentesi sull'array `rows`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Dovresti consultare il terzo elemento dell'array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Dovresti usare l'operatore di assegnazione sul terzo elemento dell'array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Dovresti assegnare il valore `10` al terzo elemento del tuo array `rows`.

```js
assert.equal(rows[2], 10);
```

Dovresti avere una seconda istruzione `console.log` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Dovresti stampare l'array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
