---
id: 660f20473aef47e9b8c9afc6
title: Passaggio 40
challengeType: 1
dashedName: step-40
---

# --description--

Aggiungi un’istruzione di log per stampare il valore di `result`. A seconda della console che usi, potresti non vedere nulla stampato.

# --hints--

Dovresti aggiungere una chiamata a `console.log`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Dovresti stampare la tua variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
