---
id: 660f4934fb48f63abd5ae371
title: Passaggio 95
challengeType: 1
dashedName: step-95
---

# --description--

Usare `done` per tenere traccia del numero di righe generate funziona, ma puoi semplificare ulteriormente la logica.

Gli array hanno una proprietà speciale `length` che ti permette di vedere quanti valori, o <dfn>elementi</dfn>, ci sono nell'array. Puoi consultare questa proprietà usando una sintassi come `myArray.length`.

Nota che `rows.length` nella chiamata `padRow` ti darebbe un errore off-by-one, perché `done` viene incrementato *prima* della chiamata.

Aggiorna la tua condizione per verificare se `rows.length` è minore di `count`.

# --hints--

Il tuo loop `while` dovrebbe verificare se `rows.length` è minore di `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
