---
id: 660f2a70ad6225fa503e71c3
title: Passaggio 46
challengeType: 1
dashedName: step-46
---

# --description--

Stai andando alla grande! A questo punto, stai incontrando quello che si chiama un <dfn>errore off-by-one</dfn>, un problema frequente nei linguaggi con indicizzazione a base zero come JavaScript.

Il primo indice del tuo array `rows` è `0`, ecco perché inizi il tuo loop `for` con `i = 0`. Ma ripetere una stringa zero volte non produce nulla da stampare.

Per risolvere questo, aggiungi `1` al valore di `i` nella tua chiamata `.repeat()`. Non assegnarlo di nuovo a `i` come hai fatto nelle condizioni del loop.

# --hints--

Dovresti aggiungere `1` a `i` nel tuo metodo `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
