---
id: 660f3dd626be3a1ffe27e5d1
title: Passaggio 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript supporta anche i commenti su più righe. Un commento su più righe inizia con `/*` e termina con `*/`.

A differenza di un commento su una singola riga, un commento su più righe racchiude più righe.

Usa `/*` e `*/` per trasformare il tuo attuale loop `for`, inclusa la sua parte interna, in un commento su più righe.

# --hints--

Dovresti iniziare un commento su più righe con `/*`.

```js
assert.match(code, /\/\*/);
```

Dovresti terminare un commento su più righe con `*/`.

```js
assert.match(code, /\*\//);
```

L'intero loop `for` dovrebbe essere commentato.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
