---
id: 660f2eccfe3f820304af1b39
title: Passaggio 64
challengeType: 1
dashedName: step-64
---

# --description--

Per sapere come formattare una riga, la tua funzione `padRow` dovrà conoscere il numero della riga su cui ti trovi e quante righe in totale vengono generate.

Il modo migliore per farlo è creare dei parametri per la funzione. Dai alla tua funzione `padRow` un parametro `rowNumber` e un parametro `rowCount`. I parametri multipli sono separati da una virgola:

```js
function name(first, second) {

}
```

# --hints--

La tua funzione `padRow` dovrebbe avere un parametro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Dovresti aggiungere una virgola dopo il parametro `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

La tua funzione `padRow` dovrebbe avere un parametro `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
