---
id: 660f415b76859a2736771607
title: Passaggio 78
challengeType: 1
dashedName: step-78
---

# --description--

La tua piramide è scomparsa di nuovo. Va bene — è normale che succeda.

Prima di creare il nuovo loop, devi imparare a usare le istruzioni `if`. Un’<dfn>istruzione `if`</dfn> ti permette di eseguire un blocco di codice solo quando una condizione è soddisfatta. Usano la seguente sintassi:

```js
if (condition) {
  logic
}
```

Crea un’istruzione `if` con il booleano `true` come condizione. Nel corpo, stampa la stringa `"Condition is true"`.

# --hints--

Dovresti creare un’istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

La tua istruzione `if` dovrebbe avere `true` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Il corpo della tua istruzione `if` dovrebbe stampare `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
