---
id: 660f4455f457ef2e3ec6920f
title: Passaggio 80
challengeType: 1
dashedName: step-80
---

# --description--

Ora la stringa non viene più stampata, perché `false` non è `true`. Ma per quanto riguarda altri valori?

Prova a cambiare la condizione nella stringa `"false"`.

# --hints--

La tua istruzione `if` dovrebbe avere come condizione la stringa `"false"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
