---
id: 660f039ff313dbb696b007ca
title: Passaggio 19
challengeType: 1
dashedName: step-19
---

# --description--

Quando un array contiene valori, o <dfn>elementi</dfn>, questi valori sono separati da virgole. Ecco un array che contiene due stringhe:

```js
let array = ["first", "second"];
```

Modifica la tua dichiarazione `rows` in modo che sia un array con le stringhe `"Naomi"`, `"Quincy"` e `"CamperChan"`. L’ordine dei valori in un array è importante, quindi segui quell’ordine. Ricorda che le stringhe fanno distinzione tra maiuscole e minuscole.

# --hints--

Il primo elemento del tuo array dovrebbe essere la stringa `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

Il secondo elemento del tuo array dovrebbe essere la stringa `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

Il terzo elemento del tuo array dovrebbe essere la stringa `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
