---
id: 660f4ae5b3924c3fc3373973
title: Passaggio 101
challengeType: 1
dashedName: step-101
---

# --description--

La tua istruzione di iterazione sarà anche diversa. Invece di aggiungere `1` a `i` ad ogni loop, devi sottrarre `1`.

Come hai fatto prima con `i = i + 1`, aggiorna la tua istruzione di iterazione per assegnare a `i` il valore di sottrarre `1` da se stesso.

# --hints--

Il tuo loop `for` dovrebbe usare `i = i - 1` come iterazione.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
