---
id: 660f34626216270c682e2f7b
title: Passaggio 66
challengeType: 1
dashedName: step-66
---

# --description--

Una <dfn>chiamata di funzione</dfn> ti permette di usare effettivamente una funzione. Forse non te ne sei accorto, ma i metodi come `.push()` che hai usato sono stati chiamate di funzione.

Una funzione si chiama facendo riferimento al nome della funzione e aggiungendo `()`. Ecco come chiamare una funzione `test`:

```js
test();
```

Sostituisci il `character.repeat(i + 1)` nella tua chiamata `.push()` con una chiamata di funzione per la tua funzione `padRow`. Non aggiungere ancora argomenti.

# --hints--

Non dovresti usare `i + 1` nella tua chiamata `push`.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Non dovresti usare `character.repeat` nella tua chiamata `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Dovresti chiamare `padRow` nella tua chiamata `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

Non dovresti avere argomenti nella tua chiamata `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
