---
id: 660f2fbd45b520046cac68e8
title: Passaggio 65
challengeType: 1
dashedName: step-65
---

# --description--

Ricorda che in un passaggio precedente hai imparato i valori di ritorno. Una funzione può <dfn>restituire</dfn> un valore che la tua applicazione può utilizzare separatamente.

In una funzione, la parola chiave `return` serve a specificare un valore di ritorno. Per esempio, questa funzione restituirebbe il valore passato al primo parametro:

```js
function name(parameter) {
  return parameter;
}
```

Usa la parola chiave `return` per restituire il valore della variabile `character`, ripetuto `rowNumber` volte.

# --hints--

Dovresti usare il metodo `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Dovresti usare il metodo `.repeat()` sulla tua variabile `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Dovresti passare `rowNumber` alla tua chiamata `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Dovresti usare la parola chiave `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Dovresti restituire il risultato della tua chiamata `.repeat()`.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
