---
id: 660f17d4e9f227d86e834abd
title: Passaggio 33
challengeType: 1
dashedName: step-33
---

# --description--

Per generare una piramide, dovrai creare più righe. Quando devi eseguire un'attività ripetutamente finché non si verifica una condizione, userai un <dfn>loop</dfn>. Ci sono molti modi per scrivere un loop.

Inizierai con un loop `for` di base. I loop `for` usano la seguente sintassi:

```js
for (iterator; condition; iteration) {
  logic;
}
```

Nei passaggi successivi esplorerai in dettaglio ogni componente di un loop. Per ora, costruisci un loop `for` che includa i termini `"iterator"`, `"condition"` e `"iteration"` per i tre componenti. Mantieni il <dfn>corpo</dfn> del loop, la sezione all'interno delle parentesi graffe `{}`, vuoto.


# --hints--

Dovresti avere un loop `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Il primo componente del tuo loop `for` dovrebbe essere la stringa `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Il secondo componente del tuo loop `for` dovrebbe essere la stringa `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Il terzo componente del tuo loop `for` dovrebbe essere la stringa `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Il corpo del tuo loop `for` dovrebbe essere vuoto.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
