---
id: 660f23b53db70af0f2620e78
title: Passaggio 43
challengeType: 1
dashedName: step-43
---

# --description--

Ora tutti i tuoi numeri appaiono sulla stessa riga. Questo non funziona per creare una piramide.

Dovrai aggiungere una nuova riga a ogni riga. Tuttavia, premere il tasto invio per inserire un'interruzione di riga tra le virgolette in JavaScript genera un errore di parsing. Invece, devi usare la speciale <dfn>sequenza di escape</dfn> `\n`, che viene interpretata come una nuova riga quando la stringa viene mostrata nel log. Per esempio:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Usa un secondo operatore di aggiunta per aggiungere una nuova riga dopo i valori `result` e `row`.

# --hints--

Dovresti usare la sequenza di escape `\n`. Ricorda che deve essere una stringa, quindi va racchiusa tra virgolette.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Dovresti concatenare la tua variabile `row` alla variabile `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Dovresti concatenare la sequenza di escape `\n` alla variabile `row`.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Dovresti assegnare l'intera concatenazione di nuovo alla variabile `result`.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
