---
id: 660f09a2694b59c3a10ee304
title: Passaggio 23
challengeType: 1
dashedName: step-23
---

# --description--

Per ora, rimuovi il primo log della console e la tua assegnazione `rows[rows.length - 1]`. Lascia la seconda istruzione di log `rows` per dopo.

# --hints--

Dovresti rimuovere la tua istruzione `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Dovresti rimuovere la tua riassegnazione `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Non dovresti rimuovere l’istruzione `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
