---
id: 660f0da9bf1035c9097af20a
title: Passaggio 28
challengeType: 1
dashedName: step-28
---

# --description--

Dovresti aver visto `"freeCodeCamp"` stampato nella console. Questo perché `.pop()` restituisce il valore che è stato rimosso dall'array — e in precedenza hai aggiunto `"freeCodeCamp"` alla fine dell'array.

Ma cosa restituisce `.push()`? Assegna il tuo `rows.push()` esistente a una nuova variabile `pushed` e stampala.

# --hints--

Dovresti dichiarare una variabile `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Dovresti usare `let` per dichiarare la tua variabile `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Dovresti assegnare `rows.push("freeCodeCamp")` alla tua variabile `pushed`.

```js
assert.equal(pushed, 4);
```

Dovresti stampare la tua variabile `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
