---
id: 660f1a00ac619ddc1e259a66
title: Passaggio 35
challengeType: 1
dashedName: step-35
---

# --description--

La <dfn>condizione</dfn> di un loop `for` indica al loop quante volte deve iterare. Quando `condition` diventa falso, il loop si fermerà.

In JavaScript, un valore booleano può essere `true` o `false`. Questi non sono stringhe — imparerai di più sulla differenza più avanti.

Per ora, userai l’<dfn>operatore minore</dfn> (`<`). Questo ti permette di verificare se il valore a sinistra è minore del valore a destra. Per esempio, `count < 3` sarà valutato come `true` se `count` è `2`, e `false` se `count` è `4`.

Sostituisci la tua stringa `"condition"` con una condizione per verificare se `i` è minore di `count`.

# --hints--

Dovresti usare l’operatore minore.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Dovresti usare l’operatore minore per verificare se `i` è minore di `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Il tuo loop `for` dovrebbe usare `i < count` come condizione.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
