---
id: 660f39b444fd6f16d1e49c1f
title: Passaggio 73
challengeType: 1
dashedName: step-73
---

# --description--

Invece di dover passare `i + 1` alla tua chiamata `padRow`, potresti iniziare il tuo loop da `1`. Questo ti permetterebbe di creare un loop indicizzato a partire da uno.

Aggiorna il tuo iteratore per iniziare da `1` invece che da `0`.

# --hints--

Il tuo loop `for` dovrebbe inizializzare `i` a `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
