---
id: 660f033cf051ebb50ea3bf48
title: Passaggio 18
challengeType: 1
dashedName: step-18
---

# --description--

Nella programmazione, spesso dovrai lavorare con molti dati. Esistono molte strutture dati che possono aiutarti a organizzare e gestire i tuoi dati. Una delle strutture dati più semplici è un <dfn>array</dfn>.

Un <dfn>array</dfn> è un tipo di dato non primitivo che può contenere una serie di valori. I tipi di dato non primitivi si differenziano dai tipi primitivi perché possono contenere dati più complessi. I tipi primitivi come le stringhe e i numeri possono contenere solo un valore alla volta.

Gli array sono indicati usando le parentesi quadre (`[]`). Ecco un esempio di una variabile con il valore di un array vuoto:

```js
let array = [];
```

Dichiara una variabile `rows` e assegnale un array vuoto.

# --hints--

Dovresti avere una variabile `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Dovresti usare `let` per dichiarare la tua variabile `rows`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Dovresti assegnare un array vuoto alla tua variabile `rows`.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
