---
id: 660f17294346b7d69e79db3d
title: Passaggio 32
challengeType: 1
dashedName: step-32
---

# --description--

Ora sei pronto per iniziare a costruire il tuo generatore di piramidi. La tua variabile `character` servirà come blocco di costruzione per la piramide.

`"Hello"` potrebbe non funzionare molto bene per questo. Cambia il valore di `character` in modo che sia il carattere cancelletto `"#"`.

# --hints--

La tua variabile `character` dovrebbe essere un simbolo cancelletto.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
