---
id: 660f49e32001983c90b75850
title: Passaggio 98
challengeType: 1
dashedName: step-98
---

# --description--

È un loop molto pulito e funzionale. Ottimo lavoro! Ma c’è ancora altro da esplorare.

Usa un commento multilinea per commentare il tuo loop `while`.

# --hints--

Il tuo loop `while` dovrebbe essere commentato.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Il corpo del tuo loop `while` dovrebbe essere commentato.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
