---
id: 660f4b641290da41b2cf0dd9
title: Passaggio 103
challengeType: 1
dashedName: step-103
---

# --description--

Proprio come per l'addizione, ci sono diversi operatori che puoi usare per la sottrazione. L'<dfn>operatore di assegnazione per la sottrazione</dfn> `-=` sottrae il valore dato dal valore attuale della variabile, quindi assegna il risultato di nuovo alla variabile.

Sostituisci la tua istruzione di iterazione con l'istruzione corretta usando l'operatore di assegnazione per la sottrazione.

# --hints--

Il tuo `for` loop non dovrebbe usare `i = i - 1`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

Il tuo `for` loop dovrebbe usare l'operatore di assegnazione per la sottrazione per ridurre `i` di `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
