---
id: 660f5179b3b0ca558f6b4d4f
title: Passaggio 113
challengeType: 1
dashedName: step-113
---

# --description--

Usa un'istruzione `if` per verificare se `inverted` è true. Ricorda che qui non è necessario usare un operatore di uguaglianza.

# --hints--

Dovresti usare un'istruzione `if`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

La tua istruzione `if` dovrebbe verificare se `inverted` è `true`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

La tua condizione `if` non dovrebbe usare operatori di confronto.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
