---
id: 660f255022991ef34ed0ee88
title: Passaggio 44
challengeType: 1
dashedName: step-44
---

# --description--

Stampare numeri non produrrà una piramide visivamente gradevole. Ora che stai mostrando il contenuto formattato del tuo array `rows`, è il momento di aggiornare il tuo loop originale.

Invece di aggiungere `i` all'array, aggiungi il valore della tua variabile `character`.

# --hints--

Non dovresti più aggiungere la tua variabile `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Dovresti aggiungere la tua variabile `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
