---
id: 660f207334fabaeac3269c38
title: Passaggio 41
challengeType: 1
dashedName: step-41
---

# --description--

Per manipolare la stringa `result` userai un tipo diverso di loop. In particolare, un loop `for...of`, che itera su ogni elemento di un oggetto iterabile e lo assegna temporaneamente a una variabile.

La sintassi per un loop `for...of` è la seguente:

```js
for (const value of iterable) {

}
```

Nota che puoi usare `const` perché la variabile esiste solo per una singola iterazione, non per tutto il loop.

Crea un loop `for...of` per iterare sul tuo array `rows`, assegnando ogni valore a una variabile `row`.

# --hints--

Dovresti usare un'altra parola chiave `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Dovresti dichiarare una variabile `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Dovresti usare `const` per dichiarare la tua variabile `row`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Il tuo loop `for...of` dovrebbe dichiarare la tua variabile `row`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

La tua variabile `row` dovrebbe essere estratta da `rows` usando la parola chiave `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Il corpo del tuo loop `for...of` dovrebbe essere vuoto.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
