---
id: 660f34e99571070d56d2f231
title: Passaggio 67
challengeType: 1
dashedName: step-67
---

# --description--

La tua funzione `padRow` ha due parametri che hai definito. I valori vengono forniti a questi parametri quando una funzione viene chiamata.

I valori che fornisci a una chiamata di funzione sono chiamati <dfn>argomenti</dfn>, e tu <dfn>passi</dfn> gli argomenti a una chiamata di funzione. Ecco una chiamata di funzione con `"Hello"` passato come argomento:

```js
test("Hello");
```

Passa `i + 1` e `count` come argomenti alla tua chiamata `padRow`. Come per i parametri, gli argomenti sono separati da una virgola.

# --hints--

Dovresti passare `i + 1` alla tua chiamata `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Dovresti avere una virgola dopo il tuo argomento `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Dovresti passare `count` come secondo argomento.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
