---
id: 660f505d02b2bd513a1c3468
title: Passaggio 111
challengeType: 1
dashedName: step-111
---

# --description--

La tua piramide non è più capovolta. Questo succede perché stai aggiungendo nuove righe alla **fine** dell'array.

Aggiorna il corpo del tuo loop per aggiungere nuove righe all'inizio dell'array.

# --hints--

Dovresti usare il metodo `unshift` di `rows`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Dovresti passare una chiamata `padRow()` come argomento per il tuo metodo `.unshift()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Dovresti passare `i` come primo argomento alla tua chiamata `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Dovresti passare `count` come secondo argomento alla tua chiamata `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
