---
id: 660f4de78f775e480ba2e451
title: Passaggio 107
challengeType: 1
dashedName: step-107
---

# --description--

Il metodo `.unshift()` di un array ti permette di aggiungere un valore all'**inizio** dell'array, a differenza di `.push()` che aggiunge il valore alla fine dell'array. `.unshift()` restituisce la nuova lunghezza dell'array su cui è stato chiamato.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Usa `const` per dichiarare una variabile `unshifted` e assegnarle il risultato della chiamata a `.unshift()` sul tuo array `numbers`. Passa `5` come argomento. Poi stampa la tua variabile `unshifted`.

# --hints--

Dovresti dichiarare la tua variabile `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Dovresti chiamare il metodo `.unshift()` sul tuo array `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Dovresti passare `5` come argomento alla tua chiamata `.unshift()`.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Dovresti assegnare il risultato della tua chiamata `.unshift()` alla tua variabile `unshifted`.

```js
assert.equal(unshifted, 4);
```

Dovresti stampare la tua variabile `unshifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
