---
id: 660f0c34aad72dc712b97624
title: Passaggio 27
challengeType: 1
dashedName: step-27
---

# --description--

Un altro metodo essenziale per questo progetto è il metodo `.pop()`. Rimuove l’ultimo elemento da un array e <dfn>restituisce</dfn> quell’elemento.

Quando un metodo restituisce un valore, puoi pensarlo come se ti restituisse il valore, rendendolo disponibile per l’uso in altre parti del tuo codice.

Crea una nuova variabile chiamata `popped` e assegnale il risultato di `rows.pop()`. Poi, stampa `popped` nella console.

# --hints--

Dovresti dichiarare una variabile chiamata `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Dovresti usare `let` per dichiarare la tua variabile chiamata `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Dovresti chiamare il metodo `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Dovresti chiamare il metodo `.pop()` sul tuo array `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Dovresti stampare la tua variabile `popped`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
