---
id: 660f383d4c772c12ff59904b
title: Passaggio 70
challengeType: 1
dashedName: step-70
---

# --description--

Puoi passare espressioni complete come argomento. La funzione riceverà il risultato della valutazione di quell'espressione. Per esempio, queste due chiamate di funzione qui sotto produrrebbero lo stesso risultato. (Nota che vale la regola dell'ordine delle operazioni PEMDAS—Parentheses, Exponents, Multiplication, Division, Addition, Subtraction):

```js
test(2 * 3 + 1);
test(7);
```

Riguardando il pattern:

```md
··#··
·###·
#####
```

Aggiorna il valore `character` in modo che venga ripetuto `2 * rowNumber - 1` volte.

Apri di nuovo la console per vedere il risultato aggiornato.

# --hints--

La tua funzione dovrebbe passare `2 * rowNumber - 1` alla chiamata `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
