---
id: 660f4f79e2a82a4e92290f44
title: Passaggio 110
challengeType: 1
dashedName: step-110
---

# --description--

A volte potresti voler ripristinare del codice precedente che hai commentato. Puoi farlo rimuovendo `/*` e `*/` attorno a quel codice. Questo si chiama <dfn>decommentare</dfn>.

Decommenta solo il tuo primo `for` loop. Lascia al loro posto il commento su singola riga e gli altri due commenti su più righe.

# --hints--

Dovresti decommentare il tuo primo `for` loop.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Non dovresti rimuovere il commento su singola riga.

```js
assert.match(code, /\/\//);
```

Non dovresti decommentare il tuo `while` loop.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Non dovresti decommentare il tuo secondo `for` loop.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
