---
id: 660f165270622fd4ec0da3f7
title: Passaggio 31
challengeType: 1
dashedName: step-31
---

# --description--

Dichiarare una variabile con la parola chiave `let` permette di riassegnarla. Questo significa che potresti cambiare `character` in seguito con un valore completamente diverso.

Per questo progetto, non vorrai cambiare questi valori delle variabili. Quindi invece, dovresti usare `const` per dichiararle. Le variabili `const` sono speciali.

Innanzitutto, una variabile `const` non può essere riassegnata come una variabile `let`. Questo codice genererebbe un errore:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Una variabile `const` inoltre non può essere non inizializzata. Questo codice genererebbe un errore:

```js
const firstName;
```

Sostituisci le tue parole chiave `let` con `const`.

# --hints--

Dovresti usare `const` per dichiarare la tua variabile `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Dovresti usare `const` per dichiarare la tua variabile `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Dovresti usare `const` per dichiarare la tua variabile `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Non dovresti usare `let` nel tuo codice.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
