---
id: 6610c21b3ef82015573ffbbe
title: Passaggio 51
challengeType: 1
dashedName: step-51
---

# --description--

La tua variabile `call` ha un valore `undefined`, anche se l'hai definita! Questo succede perché la tua funzione `padRow` attualmente non restituisce alcun valore. Per impostazione predefinita, le funzioni restituiscono `undefined` come valore.

Per restituire qualcos'altro, devi usare la parola chiave `return`. Ecco un esempio di funzione che restituisce la stringa `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Usa la parola chiave `return` per far restituire alla tua funzione la stringa `"Hello!"`.

# --hints--

La tua funzione `padRow` dovrebbe usare la parola chiave `return`.

```js
assert.match(padRow.toString(), /return/);
```

La tua funzione `padRow` dovrebbe restituire la stringa `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
