---
id: 6610c87eac0f0b256d7b037e
title: Passaggio 62
challengeType: 1
dashedName: step-62
---

# --description--

Poiché la tua funzione non usava più il parametro, cambiare l'argomento non l'ha influenzata.

Procedi a rimuovere la dichiarazione `test` e l'istruzione `return` dalla tua funzione `padRow`, così la funzione sarà di nuovo vuota.

# --hints--

La tua funzione `padRow` non dovrebbe avere una variabile `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

La tua funzione `padRow` non dovrebbe restituire un valore.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
