---
id: 6610c71600966a2191d3a64a
title: Passaggio 59
challengeType: 1
dashedName: step-59
---

# --description--

I valori restituiti da una funzione vengono usati chiamando la funzione. Puoi usare la chiamata della funzione direttamente come valore che restituisce, oppure catturare il valore restituito in una variabile. In questo modo, puoi usare il valore assegnato a una variabile con ambito locale, all’esterno della funzione in cui è stata creata.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Per usare il tuo valore `"Testing"`, restituiscilo dalla funzione `padRow` aggiornando la tua istruzione `return` per restituire solo la variabile `test`.


# --hints--

La tua funzione `padRow` dovrebbe restituire la variabile `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
