---
id: 6610c538372aa61cc0f5b122
title: Passaggio 57
challengeType: 1
dashedName: step-57
---

# --description--

Le variabili in JavaScript sono disponibili in un <dfn>ambito</dfn> specifico. In altre parole, il luogo in cui una variabile viene dichiarata determina dove nel tuo codice può essere usata.

Il primo ambito è l'ambito globale. Le variabili dichiarate al di fuori di qualsiasi "blocco" come una funzione o un loop `for` si trovano nell'<dfn>ambito globale</dfn>. Le tue variabili `character`, `count` e `rows` sono tutte nell'ambito globale.

Quando una variabile si trova nell'ambito globale, una funzione può consultarla nella sua definizione. Ecco un esempio di una funzione che usa una variabile globale `title`:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Questo esempio restituirebbe `"Professor Naomi"`. Aggiorna la tua funzione `padRow` per restituire il valore della concatenazione della tua variabile `character` all'inizio del parametro `name`.

# --hints--

La tua funzione `padRow` dovrebbe concatenare `character` all'inizio di `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

La tua funzione `padRow` dovrebbe restituire il risultato di `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
