---
id: 6610c6541c82551f95e765ab
title: Passaggio 58
challengeType: 1
dashedName: step-58
---

# --description--

Le variabili possono essere dichiarate anche all'interno di una funzione. Queste variabili sono considerate appartenenti all'<dfn>ambito locale</dfn> o <dfn>ambito di blocco</dfn>. Una variabile dichiarata all'interno di una funzione può essere usata solo all'interno di quella funzione. Se provi a consultarla all'esterno della funzione, otterrai un errore di riferimento.

Per vedere questo in azione, usa `const` per dichiarare una variabile `test` nella tua funzione `padRow`. Inizializzala con il valore `"Testing"`.

Poi, sotto la tua funzione, prova a stampare `test` nella console. Vedrai un errore perché non è definita fuori dall'ambito locale della funzione. Rimuovi quel `console.log` per superare i test e continuare.

# --hints--

La tua funzione dovrebbe dichiarare una variabile `test`.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Dovresti inizializzare `test` con il valore `"Testing"`. Non dimenticare il punto e virgola.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

La tua variabile `test` dovrebbe venire prima della parola chiave `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
