---
id: 6610c16c4fa0df12c0e30675
title: Passaggio 49
challengeType: 1
dashedName: step-49
---

# --description--

Stai chiamando la tua funzione `padRow`, ma non stai facendo nulla con quella chiamata di funzione. Tutte le funzioni in JavaScript <dfn>ritornano</dfn> un valore, cioè forniscono il risultato definito della loro chiamata che puoi usare altrove.

Per vedere il risultato della chiamata della tua funzione `padRow`, dichiara una variabile `call` e assegna a quella variabile la tua chiamata `padRow` esistente.

# --hints--

Dovresti dichiarare una variabile `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Dovresti usare `const` per dichiarare la tua variabile `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Dovresti assegnare a `call` il risultato della tua chiamata `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
