---
id: 6610b9f7619764fad5fd516d
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Puoi anche assegnare il valore di una variabile a un'altra variabile. Per esempio:

```js
let first = "One";
let second = "Two";
second = first;
```

La variabile `second` ora avrebbe il valore `"One"`.

Per vedere questo in azione, cambia la tua assegnazione `secondCharacter` da `"Test"` alla tua variabile `character`.

Poi apri la console per vedere cosa viene registrato.

# --hints--

Non dovresti assegnare il valore `"Test"` alla tua variabile `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

Dovresti assegnare il valore della variabile `character` alla tua variabile `secondCharacter`. Non dimenticare il punto e virgola.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

La tua variabile `secondCharacter` ora dovrebbe avere il valore `"World"`.

```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
