---
id: 6610b8017d1671f2814e8c77
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Quando dichiari una variabile senza inizializzarla, viene considerata <dfn>non inizializzata</dfn>. Attualmente, la tua variabile `secondCharacter` è non inizializzata.

Aggiungi un `console.log()` per vedere qual è il valore della tua variabile `secondCharacter`.

# --hints--

Dovresti aggiungere un secondo `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Dovresti registrare la tua variabile `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
