---
id: 6610c2d8d67563174fcf96dc
title: Passaggio 52
challengeType: 1
dashedName: step-52
---

# --description--

Quando hai un valore scritto esplicitamente nel tuo codice, come la stringa `"Hello!"` nella tua funzione, si considera che sia <dfn>hard-coded</dfn>. Hardcodificare un valore all'interno di una funzione potrebbe non renderla riutilizzabile come vorresti.

Invece, puoi definire dei <dfn>parametri</dfn> per la funzione. I parametri sono variabili speciali a cui viene assegnato un valore quando chiami la funzione, e possono essere usati nella funzione per modificare dinamicamente il risultato del codice della funzione.

Per aggiungere un parametro alla tua funzione, devi aggiungere un nome di variabile all'interno delle parentesi tonde. Per esempio, questa funzione `demo` ha un parametro `name`:

```js
function demo(name) {

}
```

`name` sembra un parametro utile, quindi aggiungilo pure alla tua funzione `padRow`.

# --hints--

La tua funzione `padRow` dovrebbe avere un parametro `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
