---
id: 6610c105bbdacc114d6cdc44
title: Passaggio 48
challengeType: 1
dashedName: step-48
---

# --description--

Per usare una funzione, devi chiamarla. Una <dfn>chiamata di funzione</dfn> dice alla tua applicazione di eseguire il codice della funzione ovunque tu scelga di chiamarla. La sintassi per una chiamata di funzione è il nome della funzione seguito dalle parentesi tonde. Per esempio, questo codice definisce e chiama una funzione `test`.

```js
function test() {

}
test();
```

Chiama la tua funzione `padRow`.

# --hints--

Dovresti chiamare la funzione `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Il corpo della tua funzione `padRow` dovrebbe essere vuoto.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
