---
id: 6610c77d50636722e5b6be17
title: Passaggio 60
challengeType: 1
dashedName: step-60
---

# --description--

Sotto l'istruzione `return`, registra la stringa `"This works!"` nella console.

Dopo averlo fatto, vedrai che la stringa `"This works!"` non viene mostrata nella console e la riga `console.log("This works!")` è grigia.

Copia il log della console e incollalo sopra l'istruzione `return`. Ora la stringa `"This works!"` dovrebbe apparire nella console.

Una cosa importante da sapere sulla parola chiave `return` è che non definisce solo un valore da restituire dalla tua funzione, ma ferma anche l'esecuzione del codice all'interno di una funzione o di un blocco di istruzioni. Questo significa che qualsiasi codice dopo un'istruzione `return` non verrà eseguito.

# --hints--

La tua funzione `padRow` dovrebbe restituire la variabile `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

La tua prima `console.log` dovrebbe venire dopo la parola chiave `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

La tua seconda `console.log` dovrebbe venire prima della parola chiave `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
