---
id: 6610bf6fa14d700beed1b109
title: Passaggio 87
challengeType: 1
dashedName: step-87
---

# --description--

L'operatore di <dfn>uguaglianza</dfn> `==` serve per verificare se due valori sono uguali. Per confrontare due valori, useresti un'istruzione come `value == 8`.

Sotto `done++` all'interno del tuo loop, aggiungi un'istruzione `if`. L'istruzione dovrebbe verificare se `done` è uguale a `count` usando l'operatore di uguaglianza.

# --hints--

Dovresti usare un'istruzione `if` nel tuo loop. Deve essere aggiunta dopo `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

La tua istruzione `if` dovrebbe usare l'operatore di uguaglianza per confrontare `done` e `count` nella condizione.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
