---
id: 6610b741b54b90f0c0fb3d58
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Quando i nomi delle variabili sono composti da più parole, esistono convenzioni specifiche su come scrivere le maiuscole. In JavaScript, la convenzione da usare è il <dfn>camel case</dfn>.

Il camel case significa che la prima parola del nome è tutta in minuscolo, mentre le parole successive iniziano con la lettera maiuscola. Ecco alcuni esempi di camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Usa il camel case per dichiarare una nuova variabile `secondCharacter`.

# --hints--

Dovresti dichiarare una variabile `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Dovresti usare `let` per dichiarare la tua variabile `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Non dovresti assegnare un valore alla tua variabile `secondCharacter`. Non dimenticare il punto e virgola.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
