---
id: 6610b8f6a98d25f4d485a94d
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Il valore predefinito di una variabile non inizializzata è `undefined`. Si tratta di un tipo di dato speciale che rappresenta un valore che non ha ancora una definizione.

Puoi comunque assegnare un valore a una variabile non inizializzata. Ecco un esempio:

```js
let uninitialized;
uninitialized = "assigned";
```

Assegna la stringa `"Test"` alla tua variabile `secondCharacter` sotto la sua dichiarazione. Apri la console per vedere come è cambiato il tuo log.

# --hints--

Non dovresti inizializzare `secondCharacter`. Ricorda che inizializzare significa assegnare un valore quando dichiari la variabile.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Dovresti usare l'operatore di assegnazione su `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Dovresti assegnare la stringa `"Test"` alla tua variabile `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
