---
id: 6610c8cfe4cf4d278e35c156
title: Passaggio 63
challengeType: 1
dashedName: step-63
---

# --description--

Come previsto, la tua funzione ora restituisce di nuovo `undefined`. La tua variabile `call` non è più necessaria, quindi rimuovi la dichiarazione `call` e il `console.log` per la variabile `call`.

# --hints--

Non dovresti avere una dichiarazione `call`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Non dovresti fare il log della tua variabile `call`.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
