---
id: 6610c424b7119919b62932f4
title: Passaggio 53
challengeType: 1
dashedName: step-53
---

# --description--

Una funzione non deve restituire un valore fisso. Può restituire il valore archiviato in una variabile. I parametri sono variabili speciali per una funzione, quindi possono essere restituiti anch’essi.

Modifica la tua funzione `padRow` per `return` direttamente il parametro `name`.

# --hints--

La tua funzione `padRow` dovrebbe restituire il valore del parametro `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
