---
id: 6610c83b52583e245a079217
title: Passaggio 61
challengeType: 1
dashedName: step-61
---

# --description--

Ora la tua variabile `call` ha il valore `"Testing"`. Ma la tua funzione non usa più il parametro `name`.

Rimuovi il parametro `name` dalla dichiarazione della tua funzione, poi rimuovi la stringa `"CamperChan"` dalla chiamata `padRow`.

Inoltre, rimuovi entrambi `console.log` dalla funzione `padRow`.
# --hints--

La tua funzione `padRow` non dovrebbe avere un parametro `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

Non dovresti passare `"CamperChan"` alla chiamata `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Dovresti comunque chiamare la funzione `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Non dovresti avere un `console.log` prima della parola chiave `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

Non dovresti avere un `console.log` dopo la parola chiave `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
