---
id: 6610c48c4ea0891afa7c4696
title: Passaggio 54
challengeType: 1
dashedName: step-54
---

# --description--

Se apri di nuovo la console, vedrai che la tua funzione `padRow` restituisce `undefined`, anche se hai definito un valore di ritorno! Questo succede perché i parametri devono ricevere un valore quando **chiami** la funzione.

Quando passi un valore a una chiamata di funzione, quel valore si chiama <dfn>argomento</dfn>. Ecco un esempio di chiamata a una funzione `demo` passando `"Naomi"` come argomento per il parametro `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Passa il tuo nome come argomento per il parametro `name` nella tua chiamata `padRow`. Ricorda che il tuo nome è una stringa, quindi devi usare le virgolette.

# --hints--

Dovresti passare una stringa alla tua chiamata `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
