---
id: 66458f0a05df478aa627629e
title: Passaggio 82
challengeType: 1
dashedName: step-82
---

# --description--

Il testo è sparito di nuovo! Le stringhe vuote vengono valutate come `false`, che è un <dfn>valore falsy</dfn>. Imparerai di più sui valori truthy e falsy in progetti futuri.

Oltre alle istruzioni `if`, JavaScript ha anche le istruzioni <dfn>else if</dfn>. Le istruzioni `else if` ti permettono di controllare più condizioni in un unico blocco di codice.

Ecco la sintassi per un'istruzione `else if`:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Se la prima condizione è `false`, JavaScript controllerà la condizione successiva nella catena. Se la seconda condizione è `false`, JavaScript controllerà la terza condizione, e così via.

Sotto la tua istruzione `if`, aggiungi un'istruzione `else if` che controlli se `5` è minore di `10`. Poi, all'interno del corpo dell'istruzione `else if`, stampa nella console la stringa `"5 is less than 10"`.

Controlla la console per vedere i risultati.

# --hints--

Dovresti avere un'istruzione `else if`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

La tua istruzione `else if` dovrebbe controllare se `5` è minore di `10`.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Dovresti stampare nella console la stringa `"5 is less than 10"`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
