---
id: 664599653fcd6e97104f9261
title: Passaggio 83
challengeType: 1
dashedName: step-83
---

# --description--

A volte vorrai eseguire codice diverso quando tutte le condizioni `if...else if` sono `false`. Puoi farlo aggiungendo un blocco `else`.

Un blocco `else` verrà valutato solo se le condizioni nei blocchi `if` e `else if` non sono soddisfatte.

Qui il blocco `else` è aggiunto al blocco `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Aggiungi un blocco `else` al blocco `else if`. All’interno del blocco `else`, registra la stringa `"This is the else block"` nella console.

Per vedere i risultati nella console, puoi modificare manualmente il `<` nell’istruzione `else if` in `>`. Questo farà sì che la condizione `false` sia vera e il blocco `else` venga eseguito.

# --hints--

Dovresti avere un blocco `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Il tuo blocco `else` dovrebbe registrare la stringa `"This is the else block"` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
