---
id: 664ed97c55a99e5ffab759dc
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Prima di andare avanti, prenditi un momento per rivedere i concetti che hai imparato.

Usa la parola chiave `let` per dichiarare una variabile `profession` e una variabile `age`. Inizializza `profession` con la stringa `"teacher"`, ma non inizializzare `age` con alcun valore.

Registra entrambe le tue variabili nella console per vedere i risultati.

# --hints--

Dovresti dichiarare una variabile `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Dovresti usare `let` per dichiarare la variabile `profession`.

```js
assert.match(code, /let\s+profession/);
```

Dovresti assegnare la stringa `"teacher"` alla variabile `profession`.

```js
assert.strictEqual(profession, "teacher");
```

Dovresti dichiarare una variabile `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Dovresti usare `let` per dichiarare la variabile `age`.

```js
assert.match(code, /let\s+age/);
```

Non dovresti assegnare un valore alla variabile `age`.

```js
assert.isUndefined(age);
```

Dovresti registrare `age` nella console.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Dovresti registrare `profession` nella console.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
