---
id: 664ed9344ba3495e1f67f8ea
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Ora che hai fatto la revisione di dichiarazione e inizializzazione, rimuovi il codice che hai scritto per la revisione.

Non rimuovere la tua variabile `character`.

# --hints--

Non dovresti avere una variabile `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Non dovresti avere una variabile `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Non dovresti avere un’istruzione `console.log` per `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Non dovresti avere un’istruzione `console.log` per `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Dovresti ancora avere la tua variabile `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
