---
id: 665732e62eb8b30c249b5c2c
title: Passaggio 25
challengeType: 1
dashedName: step-25
---

# --description--

Ora sei pronto per passare al prossimo insieme di lezioni sugli array.

Rimuovi tutto il tuo codice dal passaggio precedente.

# --hints--

Dovresti rimuovere tutti i riferimenti alla tua variabile `cities`.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
