---
id: 665732da4815b70bb083915e
title: Passaggio 24
challengeType: 1
dashedName: step-24
---


# --description--

Negli ultimi passaggi, hai imparato tutto su come lavorare con gli array. Prenditi un momento per rivedere ciò che hai appreso.

Inizia dichiarando una variabile `cities` e inizializzandola come un array delle stringhe `"London"`, `"New York"` e `"Mumbai"`. Poi registra quella variabile nella console.

Dopo aver registrato, modifica l’ultimo elemento di `cities` con la stringa `"Mexico City"`, quindi registra di nuovo la variabile `cities`.

Se fatto correttamente, dovresti vedere questo output nella console.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Dovresti usare `let` per dichiarare una variabile `cities`.

```js
assert.match(code, /let\s+cities/);
```

Dovresti assegnare un array delle stringhe `"London"`, `"New York"` e `"Mumbai"` alla variabile `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Dovresti usare `console.log()` per registrare l’intero array `cities` nella console.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Dovresti aggiornare l’ultimo elemento dell’array `cities` con la stringa `"Mexico City"`. Ricorda che puoi consultare l’ultimo elemento di un array usando `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Dovresti avere due istruzioni `console.log(cities)` nel tuo codice.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
