---
id: 68e6a99a0fd7386fb66e156f
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 게임 설정 패널을 만들어 체크박스 스타일링을 연습합니다.

모든 HTML 보일러플레이트(`DOCTYPE`, `html`, `head` 및 `body`)가 제공되어 있습니다.

`body` 안에 `div` 요소를 만들고 그 안에 `h1` 요소를 중첩하세요. `h1` 요소에 텍스트 `Game Settings`를 넣으세요.

# --hints--

`div` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

`h1` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

`h1` 요소는 `div` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

`h1` 요소에는 텍스트 `Game Settings`가 있어야 합니다.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
