---
id: 68e9a4a47d8a9802cba886b7
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

body에 스타일이 적용되었으니, 이제 Game Settings 패널용 컨테이너 카드를 만들겠습니다.

먼저, `settings-card` 요소에 `div` 클래스를 추가하세요.

# --hints--

`div`에는 `settings-card` 클래스가 있어야 합니다.

```js
assert.strictEqual(document.querySelector("div")?.className, "settings-card");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <div>
--fcc-editable-region--
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}
```
