---
id: 68e9a834dfef752dc35f6f90
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

사용자가 체크박스를 클릭할 때 색상이 바뀌도록 했으니, 체크박스가 선택되었음을 좀 더 명확하게 보여주도록 스타일을 적용해 보세요.

CSS 가상 클래스는 연속해서 사용할 수 있습니다. 여러 조건을 만족하는 요소를 선택하려면 가상 클래스나 가상 요소를 차례로 추가하세요.

`input[type="checkbox"]`에 대한 새로운 선택자를 만드세요. `:checked` 가상 클래스를 추가한 다음 `::after` 가상 요소를 붙이세요. 이렇게 하면 체크박스가 선택되었을 때 시각적 표시를 추가할 수 있습니다.

다음으로 `content`를 `"✓"` 값으로 추가하세요.

# --hints--

`input[type="checkbox"]:checked::after` 가상 클래스 선택자와 가상 요소 선택자가 결합된 타입 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

`content` 가상 클래스 선택자와 가상 요소 선택자가 결합된 타입 선택자에 `input[type="checkbox"]:checked::after` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

`content` 속성에 `"✓"` 값을 지정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```
