---
id: 68e9a3194160f8f5620ddf21
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 요소 바로 다음에 `label` 요소 네 개를 만드세요.

각 `label` 요소 안에 `input` 요소를 만들고 `type` 속성을 `"checkbox"`로 설정하세요.

각 `input` 요소 뒤에 `label` 요소에 다음 텍스트를 순서대로 넣으세요: `Sound Effects`, `Background Music`, `Hard Mode`, `Haptic Feedback`.

# --hints--

`label` 요소 네 개가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

첫 번째 `label` 요소 안에 `input` 요소가 있어야 합니다.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

첫 번째 `input` 요소에 `type` 속성이 `checkbox`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

첫 번째 `label` 요소에 텍스트 `Sound Effects`가 있어야 합니다.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

`Sound Effects` 텍스트는 `input` 요소 뒤에 와야 합니다.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

두 번째 `label` 요소 안에 `input` 요소가 있어야 합니다.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

두 번째 `input` 요소에 `type` 속성이 `checkbox`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

두 번째 `label` 요소에 텍스트 `Background Music`가 있어야 합니다.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

`Background Music` 텍스트는 `input` 요소 뒤에 와야 합니다.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

세 번째 `label` 요소 안에 `input` 요소가 있어야 합니다.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

세 번째 `input` 요소에 `type` 속성이 `checkbox`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

세 번째 `label` 요소에 텍스트 `Hard Mode`가 있어야 합니다.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

`Hard Mode` 텍스트는 `input` 요소 뒤에 와야 합니다.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

네 번째 `label` 요소 안에 `input` 요소가 있어야 합니다.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

네 번째 `input` 요소에 `type` 속성이 `checkbox`로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

네 번째 `label` 요소에 텍스트 `Haptic Feedback`가 있어야 합니다.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

`Haptic Feedback` 텍스트는 `input` 요소 뒤에 와야 합니다.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
