---
id: 68e9a73075dc26207320811a
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

체크박스를 다시 볼 수 있게 되었으니, 스타일 옵션을 마무리하세요. `border-radius` 선택자에 `4px`의 `input[type="checkbox"]`를 추가해 둥근 모서리를 만드세요.

그다음 체크박스의 중앙이 컨테이너 배경과 구분되도록 `background-color`의 `white`를 주세요.

사용자가 체크박스를 클릭할 때 사용자 정의 전환 효과를 설정할 것이므로, 전환이 즉시 일어나지 않고 0.3초 동안 부드럽게 진행되도록 `transition`와 `all` 값을 가진 `0.3s`를 설정하세요.

# --hints--

`border-radius` 선택자에 `input[type="checkbox"]` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius);
```

`border-radius` 속성에 `4px` 값을 넣으세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius, "4px");
```

`background-color` 선택자에 `input[type="checkbox"]` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor);
```

`background-color` 속성에 `white` 값을 넣으세요.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor, ["white", "rgb(255,255,255)", "#ffffff"]);
```

`transition` 선택자에 `input[type="checkbox"]` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.transition);
```

`transition`에 `all 0.3s` 값을 넣으세요.

```js
/**
 * Because CSS is simply the most lovely thing to ever exist, the behaviour of the `all` shorthand does wonky things and will not show up
 * in the computed stylesheet. Instead, we need to query the actual style element and extract the selector, checking the properties as written
 * by the camper.
 * I really hate CSS.
 * <3 Naomi
 */
const { props } = document.querySelector('.fcc-injected-styles')?.innerText?.match(/input\[type=('|")checkbox\1\]\s*\{(?<props>.*)\}/s)?.groups;
assert.match(props, /transition\s*\:\s*all\s+0.3\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
