---
id: 68e9a6201e2efa142d62bfe1
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

다음으로, 가시성을 높이기 위해 체크박스 크기를 키우겠습니다.

`input`에 대한 선택자를 설정하되, 특히 `[type="checkbox"]`를 대상으로 하세요.

선택자 안에서 `width`와 `height`를 `20px`로 설정하세요. 이렇게 하면 이전보다 더 커집니다.

그리고 레이블에 설정한 `cursor` 설정과 일치하도록 `cursor`를 추가하고 값으로 `pointer`를 할당하세요. 그 후 체크박스 위에 마우스를 올리면 포인터가 표시됩니다.

# --hints--

CSS에 `input[type="checkbox"]` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

`width` 선택자에 `input[type="checkbox"]` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

값이 `width`인 `20px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

`height` 선택자에 `input[type="checkbox"]` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

값이 `height`인 `20px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

`cursor` 선택자에 `input[type="checkbox"]` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

값이 `cursor`인 `pointer` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
