---
id: 68e9a7bd60c971278d2d4ffc
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

이전 수업에서 CSS의 가상 클래스와 가상 요소에 대해 배웠습니다. 이제 그 지식을 적용하여 가상 클래스 선택자가 결합된 타입 선택자를 만들어 보세요.

먼저, 타입 선택자 형식을 `input[type="checkbox"]`로 설정하고 그 끝에 `:checked`을 붙이세요.

다음으로, `background-color`에 `#f1be32` 값을 주는 프로프를 추가하세요.

마지막으로, `border-color`에 `#e2a60d` 값을 설정하세요.

# --hints--

`input[type="checkbox"]:checked`의 결합된 타입과 가상 클래스 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

결합된 타입과 가상 클래스 선택자에 `background-color` 프로프를 `input[type="checkbox"]:checked` 값으로 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

`background-color` 프로프에 `#f1be32` 값을 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

결합된 타입과 가상 클래스 선택자에 `border-color` 프로프를 `input[type="checkbox"]:checked` 값으로 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

`border-color` 프로프에 `#e2a60d` 값을 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
