---
id: 68e9a54dd97fd9091e18acce
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 `h1` 선택자를 만들고 `text-align` 속성에 `center` 값을 설정하세요. 이렇게 하면 `h1` 요소가 가운데 정렬됩니다.

# --hints--

CSS에 `h1` 요소에 대한 타입 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1"));
```

`text-align` 선택자에 `h1` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign);
```

`text-align` 속성에 `center` 값을 지정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

--fcc-editable-region--

--fcc-editable-region--

```
