---
id: 68e9a6bf1a74721a4e43e217
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

이제 브라우저가 기본으로 적용하는 체크박스를 지울 거예요. `input[type='checkbox']` 선택자 안에 `appearance` 속성을 `none` 값으로 추가하세요.

`appearance` 속성을 `none`로 설정하면 브라우저가 체크박스에 적용하는 기본 모양이 지워져서 원하는 스타일을 보여줄 수 있습니다.

그렇게 한 후에는 체크박스가 더 이상 보이지 않으니, `border`를 `2px` 두께, `solid` 스타일, 그리고 `#f1be32` 헥스 코드로 설정하세요.

# --hints--

`appearance` 속성을 `input[type="checkbox"]` 선택자에 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

`appearance` 속성을 `none` 값으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

`border` 속성을 `input[type="checkbox"]` 선택자에 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

`border` 속성을 `2px solid #f1be32` 값으로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
