---
id: 68ee86e3ee07b01626a392cc
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`.settings-card` 선택자 내에서 `margin` 속성을 `auto`로 설정하세요.

`margin` 속성을 `auto`로 설정하면 컨테이너 내 남은 공간을 균등하게 분배하도록 요소의 여백이 자동으로 조정되어, 일반적으로 블록 레벨 요소를 수평으로 가운데 정렬할 때 사용됩니다.

마지막으로, `text-align` 속성을 값 `left`으로 설정하세요. 이렇게 하면 텍스트 같은 인라인 콘텐츠가 포함된 요소의 왼쪽에 정렬됩니다.

# --hints--

`margin` 클래스 선택자에 `settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

`margin` 속성을 값 `auto`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

`text-align` 클래스 선택자에 `settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

`text-align` 속성을 값 `left`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
