---
id: 68f15cd397503429f0345262
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`input[type="checkbox"]:checked::after` 선택자의 선언에서 `display`을 `block`로 설정하세요. `display` 속성을 `block`로 설정하면 요소가 블록 레벨 요소가 되어 컨테이너의 전체 너비를 차지하고 새 줄에서 시작하므로 너비, 높이, 간격을 더 쉽게 제어할 수 있습니다.

다음으로 `text-align`를 `center`로 설정하세요. 이렇게 하면 블록 내에서 인라인 콘텐츠(텍스트나 인라인 요소 등)가 수평으로 가운데 정렬됩니다. 블록은 여전히 컨테이너의 전체 너비를 차지하지만 내부의 모든 내용이 가운데 정렬됩니다.

# --hints--

`display`의 가상 클래스 선택자와 가상 요소 선택자가 결합된 타입 선택자에 `input[type="checkbox"]:checked::after` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

`display` 값을 가진 `block` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

`text-align`의 가상 클래스 선택자와 가상 요소 선택자가 결합된 타입 선택자에 `input[type="checkbox"]:checked::after` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

`text-align` 값을 가진 `center` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
