---
id: 68f15ef0446464325ce1aa45
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

체크박스와 레이블 사이에 간격이 필요합니다. CSS에서 `label` 선택자를 사용하여 `display` 속성을 `block`로 설정하세요.

다음으로, 요소들 사이에 수직 간격을 추가하기 위해 `margin` 속성을 `8px auto`로 설정하세요. 이렇게 하면 각 요소 위아래에 일정한 간격이 생겨 콘텐츠를 구분하고 가독성을 높일 수 있습니다.

마지막으로, `cursor` 요소에 `pointer` 속성을 `label`로 설정하세요. 사용자가 `label` 위에 마우스를 올리면 커서가 손 모양으로 바뀌어 해당 요소가 클릭 가능함을 알리고 전반적인 사용자 경험을 향상합니다.

# --hints--

CSS에 `label` 타입 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

`display` 선택자에 `label` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

값이 `display`인 `block` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

`margin` 선택자에 `label` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

값이 `margin`인 `8px auto` 속성이 있어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

`cursor` 선택자에 `label` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

값이 `cursor`인 `pointer` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
