---
id: 68f1560aff1214103ab6b59b
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이제 CSS 스타일링으로 페이지를 꾸미기 시작하세요. `body` 선택자를 먼저 만드세요.

`body`에 `height` 속성을 `100vh` 값으로, 그리고 `background-color` 속성을 `#f0f0f0` 값으로 설정하세요.

`height`의 `100vh`는 `body`가 브라우저 뷰포트의 전체 높이를 차지하게 하며, 연한 회색 배경색은 페이지에 은은한 바탕을 제공합니다.

마지막으로 `text-align` 속성을 `center` 값으로 설정하세요. 이 설정은 자식 요소가 자신의 정렬로 재정의하지 않는 한 페이지 내 모든 인라인 콘텐츠를 가운데 정렬합니다.

# --hints--

`body` 타입 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

`height` 타입 선택자 내에 `body` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

`height` 속성을 `100vh` 값으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

`background-color` 타입 선택자 내에 `body` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

`background-color` 속성을 `#f0f0f0` 값으로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

`text-align` 타입 선택자 내에 `body` 속성을 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

`text-align` 속성을 `center` 값으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
