---
id: 68f29dbf9125476b4808d6fb
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

이제 `settings-card` 컨테이너에 스타일을 적용할 차례입니다. `settings-card`에 대한 클래스 선택자를 만드세요.

여기에 컨테이너에 대한 모든 서식을 작성할 것입니다. 컨테이너의 전체 크기를 정의하려면 `max-width`를 `250px`로 설정하세요.

다음으로, 컨테이너 테두리와 내용 사이에 공간이 있도록 `padding`를 `20px`로 설정하세요.

그다음, `border-radius`를 `10px`로 설정하여 둥근 모서리를 만드세요.

마지막으로 `box-shadow`를 `0 2px 6px rgba(0,0,0,0.2)` 값으로 설정하세요. 이렇게 하면 컨테이너에 미묘한 "떠 있는" 효과가 생겨 깊이감을 줄 수 있습니다.

# --hints--

CSS에 `settings-card` 클래스 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

`max-width` 클래스 선택자에 `.settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

값이 `max-width`인 `250px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

`padding` 클래스 선택자에 `.settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

값이 `padding`인 `20px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

`border-radius` 클래스 선택자에 `.settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

값이 `border-radius`인 `10px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

`box-shadow` 클래스 선택자에 `.settings-card` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

값이 `box-shadow`인 `0 2px 6px rgba(0,0,0,0.2)`가 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
