---
id: 68f29753fd5524530c6e65c8
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

다음으로, 값이 `vertical-align`인 `middle` 속성을 설정하세요.

`vertical-align` 속성은 `inline` 또는 `inline-block` 요소가 주변 텍스트나 다른 인라인 요소와 수직으로 어떻게 정렬되는지 제어합니다. 이미지, 아이콘, 또는 한 줄 내 텍스트 같은 요소의 수직 위치를 조정할 때 자주 사용됩니다.

그다음 `margin`인 `15px` 속성으로 체크박스를 마무리하세요.

# --hints--

`vertical-align` 선택자에 `input[type="checkbox"]` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

값이 `vertical-align`인 `middle` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

`margin-right` 선택자에 `input[type="checkbox"]` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

값이 `margin-right`인 `15px` 속성이 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
