---
id: 6894c1b715b8d18a18ef3203
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 음악 악기 재고를 만드는 과정을 통해 Python의 클래스와 객체를 연습합니다.

Python에서 클래스를 정의하려면 `class` 키워드 다음에 클래스 이름과 콜론을 사용하세요.

```python
class ClassName:
    # class contents go here
```

Python에서 클래스 이름은 일반적으로 PascalCase를 사용한다는 점을 기억하세요. 즉, 각 단어가 대문자로 시작하며 단어 사이에 공백이 없습니다.

`MusicalInstrument`이라는 클래스를 만드는 것부터 시작하세요. 지금은 `pass`를 클래스 본문으로 추가하세요.

# --hints--

`MusicalInstrument`이라는 클래스를 만들어야 합니다.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
