---
id: 6895e356fd877efd6e3657ee
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--


`MusicalInstrument` 클래스 내부에 `play` 매개변수만 있는 `self`이라는 메서드를 정의하세요. `self` 매개변수는 메서드가 특정 인스턴스의 속성에 접근할 수 있게 합니다.

# --hints--

`play` 매개변수를 가진 `self`라는 메서드를 정의해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
