---
id: 6895dffc64e97df0e62e9ef9
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

메서드는 클래스 내부에 정의된 함수로, 클래스의 데이터를 사용해 작업을 수행할 수 있습니다. 메서드를 만드는 방법은 다음과 같습니다:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

Python에서 `__init__` 메서드는 초기화 메서드로 알려진 특별한 메서드이며, 클래스의 새 인스턴스를 생성할 때 호출됩니다.

`MusicalInstrument` 클래스 내부에 `__init__` 메서드를 세 개의 매개변수 `self`, `name`, `instrument_type`와 함께 추가하세요.

`self` 매개변수는 클래스 인스턴스를 가리키는 표준 방식이며, 클래스의 모든 인스턴스 메서드에서 첫 번째 매개변수로 반드시 필요합니다.

# --hints--

`__init__` 클래스 내부에 `MusicalInstrument` 메서드를 정의해야 합니다.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

`__init__` 메서드는 세 개의 매개변수 `self`, `name`, `instrument_type`를 가져야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
