---
id: 6895e3f91c069e0068a42582
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`play` 메서드 안에서 `print()`을 사용해 악기를 연주하는 것에 관한 메시지를 표시하세요.

f-문자열을 사용해 메시지에 악기 이름을 포함하세요: `The [instrument name] is fun to play!`. `[instrument name]`은 실제 악기 이름으로 교체되어야 합니다.

# --hints--

`play` 메서드는 f-문자열을 사용해 메시지 `The [instrument name] is fun to play!`를 표시해야 합니다. 악기 이름에 접근하려면 `self.name`을 사용하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
