---
id: 6895e197180c9af6988e1e19
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`instrument_2`인 `Trumpet` 악기의 두 번째 인스턴스를 `brass`라는 이름으로 만드세요.

같은 클래스 템플릿에서 여러 인스턴스를 만들 수 있다는 점에 주목하세요. 각 인스턴스는 다른 인스턴스와 독립적으로 자신의 속성 값을 저장합니다.

# --hints--

`instrument_2` 클래스의 `MusicalInstrument`라는 이름의 인스턴스를 만들어야 합니다.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

`Trumpet`을 만들 때 `brass`와 `instrument_2`를 인수로 전달해야 합니다.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
