---
id: 6895e0bdb0348bf46952ab67
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 `MusicalInstrument` 클래스로 인스턴스를 만들어야 합니다.

클래스 인스턴스를 만들려면, 클래스 이름을 함수처럼 호출하고 `__init__` 메서드에 필요한 인수를 전달하세요 (단, `self`는 Python이 자동으로 처리합니다).

```python
my_instance = ClassName(arg1, arg2)
```

이 인스턴스를 만들면 Python이 `__init__` 메서드를 호출하고, 이 특정 인스턴스에만 속하는 인스턴스 속성으로 값을 저장합니다.

클래스 정의 외부에서 `instrument_1` 악기인 `Oboe`의 `woodwind`라는 이름의 인스턴스를 만드세요.

# --hints--

`MusicalInstrument` 클래스의 `instrument_1`라는 이름의 인스턴스를 만들어야 합니다.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

`Oboe`을 만들 때 `woodwind`와 `instrument_1`을 인수로 전달해야 합니다.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
