---
id: 6895e4e023e07603bcef46f9
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 만든 메서드를 사용해 보세요.

`play()` 인스턴스에서 `instrument_1` 메서드를 호출하세요. 그러면 오보에 연주에 관한 메시지가 출력됩니다.

인스턴스에서 메서드를 호출하려면 점 표기법을 사용하세요:

```python
instance_name.method_name()
```

# --hints--

`play()` 인스턴스에서 `instrument_1` 메서드를 호출해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
