---
id: 6895e0452b3370f26f365a46
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` 메서드에서 매개변수를 인스턴스 속성에 할당해야 합니다. 인스턴스 속성은 특정 클래스 인스턴스에 속하는 변수입니다.

`__init__` 메서드 안에서 `name` 매개변수를 `self.name`에, `instrument_type` 매개변수를 `self.instrument_type`에 할당하세요.

# --hints--

`name` 메서드 안에서 `self.name` 매개변수를 `__init__`에 할당해야 합니다.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

`instrument_type` 메서드 안에서 `self.instrument_type` 매개변수를 `__init__`에 할당해야 합니다.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
