---
id: 6895e572df8596062234f571
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

마지막으로 `instrument_2`에도 동일하게 하세요. `play()`에서 `instrument_2` 메서드를 호출하고, 그다음 `get_fact()`에서 `instrument_2`을 호출한 결과를 출력하세요.

이로써 음악 악기 워크숍이 완성됩니다!

# --hints--

`play()` 인스턴스에서 `instrument_2` 메서드를 호출해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_2.play()")`))
});
```

`get_fact()`에서 `instrument_2`을 호출하고 결과를 출력해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()
print(instrument_1.get_fact())

--fcc-editable-region--
```

# --solutions--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')


instrument_1.play()              # The Oboe is fun to play!
print(instrument_1.get_fact())   # The Oboe is part of the woodwind family of instruments.

instrument_2.play()              # The Trumpet is fun to play!
print(instrument_2.get_fact())   # The Trumpet is part of the brass family of instruments.
```
