---
id: 6895e5337b1914050b0af7bd
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 `get_fact()` 인스턴스에서 `instrument_1` 메서드를 호출하고 결과를 출력하세요.

`get_fact()`가 직접 출력하지 않고 문자열을 반환하므로, 반환된 문자열을 표시하려면 `print()`를 사용해야 합니다.

# --hints--

`get_fact()`에서 `instrument_1`를 호출하고 결과를 출력해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
