---
id: 6895e4940e39950288e406c2
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

다음으로, 악기에 관한 사실을 반환하는 `get_fact`라는 이름의 메서드를 클래스에 추가하세요.

이 메서드는 `self` 매개변수만 가져야 하며 `The [instrument name] is part of the [instrument type] family of instruments.`라고 말하는 f-문자열을 반환해야 합니다.

메시지를 직접 출력하는 `play` 메서드와 달리, 이 메서드는 코드의 다른 곳에서 사용할 수 있는 문자열을 반환해야 합니다.

# --hints--

`get_fact`라는 이름의 메서드를 `self` 매개변수와 함께 정의하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

`get_fact` 메서드는 악기의 이름과 유형을 포함하는 f-문자열을 반환해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
