---
id: 6895e235935106f88281b648
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

점 표기법을 사용해 인스턴스의 속성에 접근할 수 있습니다:

```py
instance_name.attribute_name
```

이렇게 하면 초기화 중에 저장된 값을 가져올 수 있습니다.

두 인스턴스의 `name`와 `instrument_type` 속성을 출력해 보세요. 생성할 때 할당한 값을 확인할 수 있습니다.

# --hints--

`name`의 `instrument_1` 속성을 출력하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

`instrument_type`의 `instrument_1` 속성을 출력하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

`name`의 `instrument_2` 속성을 출력하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

`instrument_type`의 `instrument_2` 속성을 출력하세요.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
