---
id: 6895e2becb7890fad818ad4d
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

객체 속성에 접근하는 방법을 이해했으니, 기존의 `print` 호출을 제거하세요.

# --hints--

`print(instrument_1.name)` 문장을 제거해야 합니다.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

`print(instrument_1.instrument_type)` 문장을 제거해야 합니다.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

`print(instrument_2.name)` 문장을 제거해야 합니다.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

`print(instrument_2.instrument_type)` 문장을 제거해야 합니다.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)

--fcc-editable-region--
```
