---
id: 5a24c314108439a4d4036159
title: 在陣列上使用展開運算子
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

ES6 中的一個解決方案，用來幫助在 Redux 中強制執行狀態不變性，是展開運算子：`...`。展開運算子有多種應用，其中一種非常適合先前從現有陣列產生新陣列的挑戰。這是相對較新的，但常用的語法。例如，如果你有一個陣列 `myArray` 並寫道：

```js
let newArray = [...myArray];
```

`newArray` 現在是 `myArray` 的複製。兩個陣列仍然分別存在於記憶體中。如果你執行像 `newArray.push(5)` 這樣的可變操作，`myArray` 不會改變。`...` 有效地將 `myArray` 中的值 *展開* 到一個新的陣列中。要複製一個陣列但在新陣列中添加額外的值，你可以寫成 `[...myArray, 'new value']`。這會傳回一個由 `myArray` 中的值和字串 `new value` 作為最後一個值組成的新陣列。展開語法可以在陣列組合中多次使用，但重要的是要注意它只會做淺層複製。也就是說，它只為一維陣列提供不可變的陣列操作。

# --instructions--

當新增待辦事項時，使用展開運算子來傳回狀態的新複本。

# --hints--

Redux store 應該存在並以狀態相等於 `["Do not mutate state!"]` 初始化。

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo` 和 `immutableReducer` 都應該是函式（程式）。

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

在 Redux 儲存庫上分派型別為 `ADD_TO_DO` 的動作應該會添加一個 `todo` 項目，且不應該變更狀態。

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

應該使用展開運算子來傳回新的狀態。

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
