---
id: 5a24c314108439a4d403614e
title: 定義一個 Action Creator
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

在創建一個 action 之後，下一步是將該 action 發送到 Redux store，以便它可以更新其狀態。在 Redux 中，你定義 action creators 來完成這個工作。action creator 簡單來說就是一個傳回 action 的 JavaScript 函式（程式）。換句話說，action creators 會創建表述 action 事件的物件。

# --instructions--

定義一個名為 `actionCreator()` 的函式（程式），當被呼叫時傳回 `action` 物件。

# --hints--

函式（程式）`actionCreator` 應該存在。

```js
assert(typeof actionCreator === 'function');
```

執行 `actionCreator` 函式應該會傳回 `action` 物件。

```js
assert(typeof action === 'object');
```

傳回的 `action` 應該有一個鍵屬性 `type`，其值為 `LOGIN`。

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
