---
id: 5a24c314108439a4d403615b
title: 使用 Object.assign 複製物件
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

過去幾個挑戰是使用陣列，但當狀態是 `object` 時，也有方法可以幫助強制執行狀態不變性。一個處理物件的有用工具是 `Object.assign()` 函式。`Object.assign()` 接受一個標的物件和多個原始碼物件，並將原始碼物件的屬性映射到標的物件。任何相符的屬性都會被原始碼物件中的屬性覆寫。這種行為通常用來透過將空物件作為第一個引數，接著傳入你想複製的物件，來製作物件的淺層複製。以下是一個範例：

```js
const newObject = Object.assign({}, obj1, obj2);
```

這會創建 `newObject` 作為一個新的 `object`，其中包含目前存在於 `obj1` 和 `obj2` 的屬性。

# --instructions--

Redux 狀態和動作已修改為控制代碼中 `state` 的 `object`。編輯程式碼以傳回一個新的 `state` 物件，針對型別為 `ONLINE` 的動作，將 `status` 屬性設為字串 `online`。嘗試使用 `Object.assign()` 來完成此挑戰。

# --hints--

Redux store 應該存在並以等同於第 1 行宣告的 `defaultState` 物件的狀態進行初始化。

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` 和 `immutableReducer` 都應該是函式（程式）。

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

分派型別為 `ONLINE` 的動作應該更新狀態中的屬性 `status` 為 `online`，且不應該變更狀態。

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

應該使用 `Object.assign` 來傳回新的狀態。

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
