---
id: 5a24c314108439a4d4036151
title: 使用 Switch 述語來控制多個動作
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

你可以告訴 Redux store 如何控制代碼多個 action 類型。假設你正在你的 Redux store 中管理使用者身份驗證。你想要有一個狀態表現，表示使用者已登入和已登出。你用一個具有屬性 `authenticated` 的單一狀態物件來表述這個狀態。你還需要產生對應使用者登入和使用者登出的 action 的 action 創建者，以及 action 物件本身。

# --instructions--

程式碼編輯器已為你設定好 store、actions 和 action creators。請填寫 `reducer` 函式以控制多個身份驗證 actions。請在 `reducer` 中使用 JavaScript 的 `switch` 述語來回應不同的 action 事件。這是在撰寫 Redux reducers 時的標準樣式。switch 述語應該針對 `action.type` 進行切換，並傳回適當的身份驗證狀態。

**注意：** 此時，不用擔心狀態的不變性，因為在此範例中它很小且簡單。對於每個動作，你可以傳回一個新的物件 — 例如，`{authenticated: true}`。此外，別忘了在你的 switch 述語中撰寫一個 `default` 狀況，該狀況會傳回目前的 `state`。這很重要，因為一旦你的應用程式有多個 reducer，當任何動作被分派時，它們都會被執行，即使該動作與該 reducer 無關。在這種情況下，你要確保傳回目前的 `state`。

# --hints--

呼叫函式（程式）`loginUser` 應該傳回一個具有型別屬性且設定為字串 `LOGIN` 的物件。

```js
assert(loginUser().type === 'LOGIN');
```

呼叫函式（程式）`logoutUser` 應該傳回一個物件，其屬性 `type` 設定為字串 `LOGOUT`。

```js
assert(logoutUser().type === 'LOGOUT');
```

該 store 應該使用一個 `authenticated` 屬性設為 `false` 的物件來初始化。

```js
assert(store.getState().authenticated === false);
```

分派 `loginUser` 應該會將商店狀態中的 `authenticated` 屬性更新為 `true`。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

分派 `logoutUser` 應該將儲存狀態中的 `authenticated` 屬性更新為 `false`。

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

`authReducer` 函式（程式）應該使用 `switch` 述語來控制多個動作類型。

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
