---
id: 5a24c314108439a4d4036155
title: 將 Action 資料發送到 Store
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

到目前為止，你已經學會如何將動作分派到 Redux store，但到目前為止這些動作除了 `type` 外並未包含任何資訊。你也可以隨著動作一起傳送特定的資料。事實上，這非常常見，因為動作通常源自某些使用者互動，並且往往會攜帶一些資料。Redux store 通常需要知道這些資料。

# --instructions--

程式碼編輯器中定義了一個基本的 `notesReducer()` 和一個 `addNoteText()` 動作建立器。完成 `addNoteText()` 函式的主體，使其傳回一個 `action` 物件。該物件應包含一個 `type` 屬性，值為 `ADD_NOTE`，以及一個 `text` 屬性，設定為傳入動作建立器的 `note` 資料。當你呼叫動作建立器時，你會傳入可供物件存取的特定筆記資訊。

接著，完成在 `notesReducer()` 中撰寫 `switch` 述語。你需要為處理 `addNoteText()` 動作添加一個 case。每當有型別為 `ADD_NOTE` 的動作時，這個 case 就會被觸發，並且它應該傳回傳入 `action` 上的 `text` 屬性作為新的 `state`。

該動作在程式碼底部被分派。完成後，執行程式碼並觀察主控台。這就是將動作特定的資料傳送到 store 並在更新 store `state` 時使用它所需的全部步驟。

# --hints--

動作建立者 `addNoteText` 應該傳回一個具有 `type` 和 `text` 鍵的物件。

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

分派一個型別為 `ADD_NOTE` 的動作，並使用 `addNoteText` 動作建立器，應該會將 `state` 更新為傳遞給該動作建立器的字串。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
