---
id: 5a24c314108439a4d4036153
title: 註冊一個 Store 監聽器
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

你可以存取 Redux `store` 物件上的另一個方法是 `store.subscribe()`。這允許你為 store 訂閱監聽器函式，當有動作被分派到 store 時，這些函式會被呼叫。這個方法的一個簡單用途是為你的 store 訂閱一個函式，每次收到動作並且 store 被更新時，該函式會記錄一則訊息。

# --instructions--

撰寫一個回呼函式，每當 store 接收到一個動作時，便遞增全域變數 `count`，並將此函式傳入 `store.subscribe()` 方法。你會看到 `store.dispatch()` 連續呼叫三次，每次直接傳入一個動作物件。觀察動作分派之間的主控台輸出，以查看更新的發生。

# --hints--

在 store 上分派 `ADD` 動作應該會將狀態增加 `1`。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

應該有一個監聽器函式使用 `store.subscribe` 訂閱到 store。

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` 應該接收一個函式（程式）。

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

傳遞給 `store.subscribe` 的函式（程式）不應該被呼叫。

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

當 store 更新時，`store.subscribe` 的回呼也應該遞增全域的 `count` 變數。

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --before-user-code--

```js
count = 0;
```

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
