---
id: 5a24c314108439a4d4036152
title: 對於 Action Types 使用 const
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

在使用 Redux 時，一個常見的做法是將動作型別指定為唯讀常數，然後在使用的任何地方參考這些常數。你可以重構你正在處理的程式碼，將動作型別寫成 `const` 宣告式。

# --instructions--

將 `LOGIN` 和 `LOGOUT` 宣告為 `const` 值，並分別指定為字串 `'LOGIN'` 和 `'LOGOUT'`。然後，編輯 `authReducer()` 和動作建立者以取用這些常數，而非字串值。

**注意：** 通常慣例是將常數全部以大寫字母書寫，這也是 Redux 的標準做法。

# --hints--

呼叫函式（程式）`loginUser` 應該傳回一個 `type` 屬性設定為字串 `LOGIN` 的物件。

```js
assert(loginUser().type === 'LOGIN');
```

呼叫函式（程式）`logoutUser` 應該傳回一個 `type` 屬性設定為字串 `LOGOUT` 的物件。

```js
assert(logoutUser().type === 'LOGOUT');
```

該 store 應該以一個 `login` 屬性設為 `false` 的物件初始化。

```js
assert(store.getState().authenticated === false);
```

分派 `loginUser` 應該會將商店狀態中的 `login` 屬性更新為 `true`。

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

分派 `logoutUser` 應該將儲存狀態中的 `login` 屬性更新為 `false`。

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

`authReducer` 函式（程式）應該使用 switch 述語來控制多個動作型別。

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` 和 `LOGOUT` 應該宣告為 `const` 值，並且應該被指定為 `LOGIN` 和 `LOGOUT` 的字串。

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

動作建立者和 reducer 應該取用 `LOGIN` 和 `LOGOUT` 常數。

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
