---
id: 5a24c314108439a4d403614c
title: 從 Redux Store 取得狀態
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Redux store 物件提供數個方法，讓你能與它互動。例如，你可以使用 `getState()` 方法來取得 Redux store 物件中目前持有的 `state`。

# --instructions--

前一個挑戰的程式碼在程式碼編輯器中被更簡潔地重寫。使用 `store.getState()` 從 `store` 取得 `state`，並將此指定給新的變數 `currentState`。

# --hints--

Redux store 的初始狀態應該有一個值為 5。

```js
assert(store.getState() === 5);
```

應該存在一個變數 `currentState`，並且應該將 Redux store 的當前狀態指定給它。

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
