---
id: 5a24c314108439a4d4036157
title: 使用 Redux 撰寫計數器
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

現在你已經學會了 Redux 的所有核心原則！你已經看到如何創建 actions 和 action creators、創建一個 Redux store、對 store 分派你的 actions，以及使用純 reducers 設計狀態更新。你甚至已經看到如何透過 reducer 複合來管理複雜的狀態，以及控制代碼非同步 actions。這些範例雖然簡單，但這些概念是 Redux 的核心原則。如果你能夠充分理解它們，你就準備好開始建置你自己的 Redux 應用程式。接下來的挑戰將涵蓋一些關於 `state` 不變性的細節，但首先，這裡是你迄今為止所學內容的複習。

# --instructions--

在本課程中，你將從頭開始使用 Redux 實作一個簡單的計數器。基礎內容已在程式碼編輯器中提供，但你必須填寫細節！使用提供的名稱並定義 `incAction` 和 `decAction` 動作建立者、`counterReducer()`、`INCREMENT` 和 `DECREMENT` 動作型別，最後是 Redux 的 `store`。完成後，你應該能夠分派 `INCREMENT` 或 `DECREMENT` 動作來增加或減少 `store` 中的狀態。祝你建置第一個 Redux 應用程式順利！

# --hints--

動作建立者 `incAction` 應該傳回一個動作物件，其 `type` 與 `INCREMENT` 的值相等。

```js
assert(incAction().type === INCREMENT);
```

動作建立函式 `decAction` 應該傳回一個動作物件，其 `type` 與 `DECREMENT` 的值相等。

```js
assert(decAction().type === DECREMENT);
```

執行 `store.getState()` 應該會傳回一個數字

```js
assert(typeof store.getState() === 'number');
```

Redux 儲存庫應該以 `state` 為 0 初始化。

```js
assert(_store.getState() === 0);
```

在 Redux store 上分派 `incAction` 應該會將 `state` 增加 1。

```js
assert(
  (function () {
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

在 Redux store 上分派 `decAction` 應該會將 `state` 減少 1。

```js
assert(
  (function () {
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` 應該是一個函式（程式）。

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

## --after-user-code--

```js
const _store = Redux.createStore(counterReducer)
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
