---
id: 5a24c314108439a4d403615a
title: 從陣列中移除一個項目
challengeType: 6
forumTopicId: 301447
dashedName: remove-an-item-from-an-array
---

# --description--

是時候練習從陣列中移除條款了。這裡也可以使用展開運算子。其他有用的 JavaScript 方法包括 `slice()` 和 `concat()`。

# --instructions--

reducer 和 action creator 已修改為根據項目的索引從陣列中移除一個項目。完成撰寫 reducer，使其傳回一個新的狀態陣列，該陣列已移除指定索引的項目。

# --hints--

Redux store 應該存在並以狀態相等於 `[0,1,2,3,4,5]` 初始化。

```js
const initialState = store.getState();
assert.isArray(initialState) 
assert.deepEqual(initialState, [0, 1, 2, 3, 4, 5])
```

`removeItem` 和 `immutableReducer` 都應該是函式（程式）。

```js
assert(
  typeof removeItem === 'function' && typeof immutableReducer === 'function'
);
```

分派 `removeItem` action creator 應該從狀態中移除項目，且不應該變更狀態。

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(removeItem(3));
const state_1 = store.getState();
store.dispatch(removeItem(2));
const state_2 = store.getState();
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
const state_3 = store.getState();

assert(isFrozen)
assert.deepEqual(state_1, [0, 1, 2, 4, 5]) 
assert.deepEqual(state_2, [0, 1, 4, 5]) 
assert.deepEqual(state_3, [5])

```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      return [
        ...state.slice(0, action.index),
        ...state.slice(action.index + 1)
      ];
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```
