---
id: 5a24c314108439a4d403614b
title: 創建 Redux Store
challengeType: 6
forumTopicId: 301439
dashedName: create-a-redux-store
---

# --description--

Redux 是一個狀態管理框架，可用於多種不同的網頁技術，包括 React。

在 Redux 中，有一個單一的狀態物件負責你應用程式的整個狀態。這表示如果你有一個包含十個組件的 React 應用程式，且每個組件都有自己的區域狀態，那麼你的應用程式的整個狀態將由一個位於 Redux `store` 中的單一狀態物件所定義。這是學習 Redux 時需要理解的第一個重要原則：Redux `store` 是應用程式狀態的唯一真實來源。

這也表示，當你的應用程式中的任何部分想要更新狀態時，**必須**透過 Redux store 來進行。單向資料流使得追蹤你應用程式中的狀態管理變得更容易。

# --instructions--

Redux 的 `store` 是一個持有並管理應用程式 `state` 的物件。Redux 物件上有一個名為 `createStore()` 的方法，你用它來創建 Redux `store`。此方法需要一個 `reducer` 函式作為必填引數。`reducer` 函式會在後續挑戰中涵蓋，並且已在程式碼編輯器中為你定義。它簡單地以 `state` 作為引數並傳回 `state`。

宣告一個 `store` 變數並將其指定給 `createStore()` 方法，傳入 `reducer` 作為引數。

**注意：** 編輯器中的程式碼使用 ES6 預設引數語法來初始化此狀態，使其持有 `5` 的值。如果你不熟悉預設引數，可以參考 <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/es6/set-default-parameters-for-your-functions" target="_blank" rel="noopener noreferrer nofollow">課程中的 ES6 部分</a>，該部分涵蓋此主題。

# --hints--

Redux store 應該存在。

```js
assert(typeof store.getState === 'function');
```

Redux store 的狀態值應該是 5。

```js
assert(store.getState() === 5);
```

# --seed--

## --seed-contents--

```js
const reducer = (state = 5) => {
  return state;
}

// Redux methods are available from a Redux object
// For example: Redux.createStore()
// Define the store here:
```

# --solutions--

```js
const reducer = (state = 5) => {
  return state;
}

const store = Redux.createStore(reducer);
```
