---
id: 5a24c314108439a4d403614d
title: 定義 Redux Action
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

由於 Redux 是一個狀態管理框架，更新狀態是其核心任務之一。在 Redux 中，所有狀態更新都是透過分派動作來觸發的。動作只是包含已發生動作事件資訊的 JavaScript 物件。Redux 儲存庫接收這些動作物件，然後相應地更新其狀態。有時 Redux 動作也會攜帶一些資料。例如，使用者登入後，動作會攜帶使用者名稱。雖然資料是選用的，但動作必須攜帶指定發生動作 `type` 屬性的 `type`。

將 Redux actions 想像成傳遞有關你應用程式中發生事件的資訊到 Redux store 的信使。然後，store 根據發生的 action 進行狀態更新的業務。

# --instructions--

撰寫 Redux action 就像宣告一個具有 type 屬性的物件一樣簡單。宣告一個物件 `action` 並給它一個屬性 `type`，其值設定為字串 `'LOGIN'`。

# --hints--

應該存在一個 `action` 物件。

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

`action` 物件應該有一個關鍵屬性 `type`，其值為 `LOGIN`。

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
