---
id: 684cff0bf0a4ca0ad7e42990
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Hatua inayofuata itakuwa kuzungusha pembe za mpaka. Ili kuweka sifa ya `border-radius` kwa kutumia Tailwind CSS, unaweza kutumia darasa la vitu la huduma lenye muundo ufuatao: `rounded-size`, ambapo `size` inawakilisha mduara wa pembe. 

Hapa kuna mfano wa kutumia mduara mkubwa sana wa mpaka kwa kipengele:

```html
<p class="rounded-xl">This is border radius.</p>
```

Ndani ya lebo yako ya ufunguzi ya `div`, tumia mduara wa mpaka wenye ukubwa wa `md`. 

# --hints--

Leboa yako ya ufunguzi ya `div` inapaswa kuwa na darasa la vitu la `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
