---
id: 684cd59ec94b5e063bf776f0
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Itakuwa vizuri kuwa na mpaka kuzunguka ile asili nyekundu yenye uangavu.

Ili kuweka `border-width` katika Tailwind CSS, unaweza kutumia darasa la huduma la `border-number` ambapo `number` itawakilisha uangavu au unene wa mpaka wako.

Hapa kuna mfano wa kutumia `border-width` ya `4px` kwenye pande zote:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Ndani ya lebo yako ya kufungua ya `div`, ongeza darasa sahihi la huduma kwa `border-width` iliyowekwa kuwa `2px` kwenye pande zote.

# --hints--

Leboa yako ya kufungua ya `div` inapaswa kuwa na darasa la `border-2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
