---
id: 684ccbdc1ef89e042b744e09
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sehemu inayofuata ya warsha ni kuongeza kitufe cha kufunga.

Chini ya kipengele chako cha `p`, ongeza kipengele cha `button` chenye sifa ya `aria-label` iliyowekwa kuwa `"Dismiss"`. Ndani ya kipengele hicho cha `button`, ongeza kipengele cha `span` chenye sifa ya `aria-hidden`. Maandishi ya kipengele cha `span` yanapaswa kuwa `&times;`.

*TAFADHALI KUMBUKA*: Kwa kuwa warsha hii itazingatia tu kazi na mitindo ya Tailwind CSS, `button` haitakuwa na aina yoyote ya mwingiliano uliowekwa.

# --hints--

Unapaswa kuwa na kipengele cha `button`.

```js
assert.exists(document.querySelector("button"));
```
 
Kipengele chako cha `button` kinapaswa kuwa na sifa ya `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya `aria-label` iliyowekwa kuwa `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Kipengele chako cha `button` kinapaswa kuwa na kipengele cha `span` kilichopangwa ndani.

```js
assert.exists(document.querySelector("button span"));
```

Kipengele chako cha `span` kinapaswa kuwa na sifa ya `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Kipengele chako cha `span` kinapaswa kuwa na maandishi ya `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>
  </body>
</html>
```
