---
id: 684cd8bf5445de089af9c837
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Sasa ni wakati wa kutumia rangi ya mpaka. Ili kuweka sifa ya `border-color` kwa kutumia Tailwind CSS, unaweza kutumia darasa la vitu la huduma lenye muundo ufuatao: `border-color-number`, ambapo `number` inawakilisha kivuli cha rangi ya mpaka.

Hapa kuna mfano wa kutumia rangi ya mpaka ya zambarau kwa kipengee:

```html
<p class="border-purple-400">This is border-color.</p>
```

Ndani ya lebo yako ya kufungua ya `div`, tumia rangi nyekundu ya mpaka yenye kivuli cha `300`.

# --hints--

Leboa yako ya kufungua ya `div` inapaswa kuwa na darasa la vitu la `border-red-300`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
