---
id: 684cc59eed2706fa50e8933c
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kuanza kufanya kazi na misingi ya Tailwind CSS kwa kujenga sehemu ya ujumbe wa kosa.

Kwa hatua chache za mwanzo, utazingatia kujenga alama za uandishi za awali za HTML.

Ndani ya kipengele cha `body`, ongeza kipengele cha `div`. Kisha ndani ya kipengele cha `div`, ongeza kipengele cha `p`.

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert.exists(document.querySelector("div"));
```

Unapaswa kuwa na kipengele cha `p`.

```js
assert.exists(document.querySelector("p"));
```

Kipengele chako cha `p` kinapaswa kuwa ndani ya kipengele cha `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
