---
id: 684cd027744c2a0520ba4d58
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Tailwind CSS hutumia madarasa ya vitu ya huduma ambayo huwekwa moja kwa moja kwenye alama za uandishi. Madarasa haya yanajumuisha mitindo ya fonti, rangi, mipangilio ya Flexbox, Grid na mengineyo.

Ili kutumia rangi ya nyuma katika Tailwind CSS, unaweza kutumia kiambishi awali cha `bg`, kifuatiwe na jina la rangi na nambari. Nambari hiyo inawakilisha kivuli cha rangi.

Hapa kuna mfano wa kutumia rangi ya nyuma ya buluu ya giza kwenye kipengele cha aya.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

Katika lebo yako ya ufunguzi ya `div`, ongeza darasa la vitu lenye darasa la huduma la `bg-red-100` ambalo linawakilisha rangi ya nyuma nyekundu nyepesi.

# --hints--

Leboa yako ya ufunguzi ya `div` inapaswa kuwa na sifa ya kipengele cha darasa.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Sifa yako ya kipengele cha darasa inapaswa kuwa na darasa la vitu la huduma la `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
