---
id: 684cca09aac0fc035db4099f
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele chako cha `p`, ongeza maandishi ya `Error! Something went wrong. Please try again.` ambapo `Error!` inapaswa kufunikwa ndani ya kipengele cha `strong`.

# --hints--

Unapaswa kuwa na kipengele cha `strong` ndani ya kipengele chako cha `p`.

```js
assert.exists(document.querySelector("p strong"));
```

Kipengele chako cha `strong` kinapaswa kuwa na maandishi ya `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi sahihi yanayomalizika na `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
