---
id: 684d03df2a0ddf0c735865f0
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa unaona kuwa upana wa kipengele ni nusu ya ukubwa wa `body`. Hii inafanya kazi vizuri kwa vifaa vikubwa lakini si vizuri sana kwa vifaa vidogo.

Tailwind CSS inafuata mbinu ya kwanza simu katika muundo. Hivyo kama unataka kuunda miundo inayojibadilisha kulingana na kifaa kwa vifaa vingine, unaweza kutumia aina za muundo unaobadilika za Tailwind kama `md` na `lg` hivi:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

Mfano ufuatao utaweka upana kuwa `66%` ya kipengele mzazi kwa ukubwa wa vifaa wa `1024px` na zaidi.

Ndani ya lebo yako ya ufunguzi ya `div`, badilisha `w-1/2` yako ili itumike tu kwa ukubwa wa vifaa wa `768px` na zaidi. Unaweza kuweka kiambishi awali cha darasa hili na aina ya `md`.

# --hints--

Leboa ya ufunguzi ya `div` inapaswa kuwa na darasa la huduma la `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

Hautapaswa tena kuwa na darasa la `w-1/2` katika lebo yako ya ufunguzi ya `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
