---
id: 684d009994a3840b98c9aa3a
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Kwa sasa, upana wa kipengele cha `div` unashughulikia upana mzima wa ukurasa. Ili kutumia upana katika Tailwind CSS, una chaguzi kadhaa, kama kutumia muundo wa `w-number` kwa upana wa kudumu au muundo wa `w-fraction` kwa asilimia.  

Hapa kuna mfano wa kutumia upana uliowekwa kuwa `66%` ya kontena la mzazi wake:

```html
<p class="w-2/3">This is a percentage width.</p>
```

Ndani ya lebo yako ya ufunguzi ya `div`, tumia upana uliowekwa kuwa `50%` ya kontena la mzazi wake.

# --hints--

Lebo yako ya ufunguzi ya `div` inapaswa kuwa na darasa la huduma la `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
