---
id: 684d0d369e2b3b0f288d2c20
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Itakuwa vizuri kuweka kipengele cha `div` katikati kwa usawa wa wima. Ili kutumia `margin` kwa pande mbili kwa wakati mmoja, unaweza kutumia mifumo ifuatayo: `mx-number` au `my-number`.

`mx` huweka ukingo wa kushoto na kulia, wakati `my` huweka ukingo wa juu na chini.

Hapa kuna mfano wa kutumia ukingo wa juu na chini kwa kipengele:

```html
<p class="my-8">This is margin.</p>
```

Ndani ya lebo yako ya kufungua ya `div`, tumia darasa la huduma sahihi kuweka ukingo wa kushoto na kulia kuwa `auto`.

# --hints--

Leboa yako ya kufungua ya `div` inapaswa kuwa na darasa la huduma la `mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
