---
id: 684d105e488dbd1166a717ae
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Ukiongeza ukubwa wa dirisha la onyesho, utaona kuwa vipande vyote bado vimepangwa kushoto ndani ya kontena. Itakuwa nzuri kama vipande vya kubadilika vingekuwa katikati. 

Ili kuweka vipande vya kubadilika kando ya mhimili mkuu wa kontena kwa kutumia Tailwind CSS, unaweza kutumia muundo ufuatao: `justify-value`. `value` inaweza kuwa `start`, `end`, `center`, `between` na zaidi. Thamani yoyote halali inayotumika katika CSS ya kawaida kwa `justify-content` inaweza kutumika hapa kwa Tailwind CSS. 

Ndani ya lebo yako ya ufunguzi ya `div`, tumia darasa la huduma sahihi la kuweka sifa ya `justify-content` kuwa `center`.

# --hints--

Leboa yako ya ufunguzi ya `div` inapaswa kuwa na darasa la huduma la `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
