---
id: 684d160a2370b01310ae0fda
title: Hatua 18
challengeType: 0
dashedName: step-18
---

# --description--

Rangi mpya nyekundu inaonekana nzuri, lakini itakuwa bora zaidi kama maandishi yangekuwa makubwa zaidi. Ili kuweka `font-size` kwa kipengele kwa kutumia Tailwind CSS, unaweza kutumia muundo ufuatao: `text-value`.

Hapa kuna mfano wa kuweka `font-size` kwa kipengele kuwa `0.875rem` au (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Ndani ya kipengele chako cha kufungua `p`, tumia darasa la huduma sahihi kuweka `font-size` kuwa `xl` ambayo inawakilisha `1.25rem` au (`20px`).

# --hints--

Leboa ya kufungua `p` inapaswa kuwa na darasa la huduma la `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
