---
id: 684d0eeb9c81e5109e52fe8b
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Kipengele cha `div` kinaanza kuonekana vizuri zaidi. Lakini sasa kitufe cha kufunga cha `button` kiko karibu sana na kipengele cha aya.

Ili kuunda nafasi kati ya vipande vya kubadilika, unaweza kutumia muundo ufuatao: `gap-number`, ambapo kila kitengo cha `number` kinawakilisha 0.25rem (au 4px) ya nafasi thabiti kati ya safu na nguzo katika mpangilio wa flex.

Hapa kuna mfano wa kutumia `gap` kwa kipengele:

```html
<p class="flex gap-8">This is gap.</p>
```

Ndani ya lebo yako ya kufungua `div`, tumia darasa la vitu sahihi la huduma kwa kuweka `gap` kuwa `16px`.

# --hints--

Lebo yako ya kufungua ya `div` inapaswa kuwa na darasa la huduma la `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
