---
id: 684d1942a1904a14a3709c6c
title: Hatua 20
challengeType: 0
dashedName: step-20
---

# --description--

Ulinganifu wa mtindo wa sehemu yako ya ujumbe wa kosa karibu umekamilika, lakini kuna jambo moja lililobaki kurekebisha. Kwa kuwa umeweka ukubwa tofauti wa fonti kwa aya na maandishi ya kitufe, ulinganifu wima kati ya vipande vya kubadilika hauko sawa.

Ili kuweka vipande vya kubadilika kando ya mhimili wa msalaba wa kontena kwa kutumia Tailwind CSS, unaweza kutumia muundo ufuatao: `items-value`. `value` inaweza kuwa chochote kutoka `start`, `end`, `center`, `baseline` na zaidi.

Ndani ya lebo yako ya kufungua `div`, tumia darasa la huduma sahihi la kuweka sifa ya `align-items` kuwa `center`.

Na kwa mabadiliko hayo ya mwisho, sehemu yako ya ujumbe wa kosa imemalizika!

# --hints--

Leboa yako ya kufungua `div` inapaswa kuwa na darasa la huduma la `items-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
    --fcc-editable-region--
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto flex gap-4 justify-center items-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss" class="text-red-700 text-2xl">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
