---
id: 684d09cad4ec900e41db2e74
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Kwa sasa kipengele cha `div` kimewekwa kona ya juu kushoto ya skrini ya awali. Lakini itakuwa vizuri kuongeza nafasi kidogo juu ya kipengele hicho. Ili kutumia `margin` kwa kipengele kwa kutumia Tailwind CSS, unaweza tumia muundo ufuatao: `m-number`, ambapo kila kitengo cha `number` kinawakilisha 0.25rem (au pikseli 4) kutumika kwa pande zote nne.

Ikiwa unataka kutumia ukingo kwa upande mmoja tu, unaweza kutumia mojawapo ya mifano ifuatayo: `mt-number` (ukingo wa juu), `mb-number` (ukingo wa chini), `mr-number` (ukingo wa kulia) na `ml-number` (ukingo wa kushoto).

Hapa kuna mfano wa kutumia `margin-bottom` ya 32px kwa kipengele:

```html
<p class="mb-8">This is margin.</p>
```

Ndani ya lebo yako ya kufungua `div`, tumia darasa la huduma sahihi ili kutumia `margin-top` ya `16px`.

# --hints--

Leboa yako ya ufunguzi ya `div` inapaswa kuwa na darasa la huduma la `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
