---
id: 684d11ffd15a531225d081e9
title: Hatua 17
challengeType: 0
dashedName: step-17
---

# --description--

Sasa baada ya mitindo kutumika kwa kipengele cha `div`, ni wakati wa kutumia mitindo kwa kipengele cha `p`. 

Kwa sasa, rangi ya maandishi kwa kipengele cha aya ni nyeusi. Lakini itakuwa nzuri kama rangi ingekuwa nyekundu yenye giza badala yake.

Katika hatua iliyopita, ulijifunza jinsi ya kufanya kazi na rangi za nyuma kama hii: `bg-red-100`. Unaweza kufanya kitu kinachofanana na rangi ya maandishi kwa kufuata muundo huu: `text-color-shade`. 

Hapa kuna mfano wa jinsi ya kutumia rangi ya maandishi ya buluu kwa kipengele:

```html
<p class="text-blue-500">This has blue text.</p>
```

Anza kwa kuongeza sifa ya kipengele cha darasa la vitu (class) kwa kipengele chako cha `p`. Kisha, tumia darasa la huduma sahihi la kuweka rangi ya maandishi nyekundu yenye kivuli kilichowekwa kuwa `700`.

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na sifa ya kipengele cha darasa la vitu (class).

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.hasAttribute("class"));
```

Lebo yako ya kufungua ya `p` inapaswa kuwa na darasa la huduma la `text-red-700`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl?.classList.contains("text-red-700"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p>
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
