---
id: 684d17a3b350bb13d3c19c0f
title: Hatua 19
challengeType: 0
dashedName: step-19
---

# --description--

Sasa ni wakati wa kutumia mitindo fulani kwa kipengele cha `button`. 

Anza kwa kuongeza sifa ya kipengele ya darasa la vitu kwenye kipengele chako cha `button`. Kisha, tumia darasa la vitu la huduma sahihi kuweka rangi ya maandishi kuwa nyekundu na kivuli kikiwa `700`.

Kisha, ndani ya lebo yako ya ufunguzi ya `button`, tumia darasa la vitu la huduma sahihi kuweka `font-size` kuwa `2xl` ambayo inawakilisha `1.5rem` au (`24px`).

# --hints--

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele ya darasa la vitu.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Lebo yako ya ufunguzi ya `button` inapaswa kuwa na darasa la vitu la huduma la `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Lebo yako ya ufunguzi ya `button` inapaswa kuwa na darasa la vitu la huduma la `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
