---
id: 684d0e11094c3e0fe1fb98f5
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Kwa sasa, kipengele cha aya kimewekwa juu ya kipengele cha dismiss `button`. Lakini itakuwa vizuri kuwa na kipengele cha `button` kuwa katika safu moja na aya. Hapa ndipo Flexbox inapoingia.

Ili kuweka kipengele cha `div` katika mpangilio wa Flexbox katika Tailwind CSS, unaweza kutumia darasa la huduma la `flex`.

# --hints--

Leboleo lako la kufungua la `div` linapaswa kuwa na darasa la huduma la `flex`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("flex"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
