---
id: 684d07cbfa43a20d6a002ed7
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Mitindo ya kipengele cha `div` inaanza kuungana vizuri. Lakini itakuwa nzuri kama kungekuwa na nafasi kidogo kuzunguka maandishi. 

Ili kutumia `padding` kwa kipengele kwa kutumia Tailwind CSS, unaweza tumia muundo ufuatao: `p-number`, ambapo `number` inawakilisha pikseli unazotaka kutumia pande zote nne. 

Hapa kuna mfano wa kutumia `padding` ya `8px` kwa kipengele:

```html
<p class="p-8">This is padding.</p>
```

Ndani ya lebo yako ya kufungua `div`, tumia darasa la huduma sahihi ili kutumia `padding` ya `4px`.

# --hints--

Lebo yako ya kufungua `div` inapaswa kuwa na darasa la huduma la `p-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
