---
id: 68501d32b69f9d7756334ffc
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Ikiwa utaongeza ukubwa wa dirisha la onyesho, unapaswa kuona kipengele cha `div` kikiwa katikati kwa usawa kwenye ukurasa. Lakini kwa vifaa vidogo, upana wa kipengele cha `div` unachukua upana mzima wa ukurasa.

Sahihisha darasa la `mx-auto` ili litumike tu kwa vifaa vya `md` na vikubwa zaidi.

# --hints--

Leboa ya ufunguzi ya `div` inapaswa kuwa na darasa la huduma la `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

Hautaki tena kuwa na darasa la `mx-auto` kwenye leboa yako ya ufunguzi ya `div`.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
