---
id: 6140827cff96e906bd38fc2b
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

如 <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp 风格指南</a> 中所述，徽标应保持 `35 / 4` 的宽高比，并在文本周围留有内边距。

首先，将 `background-color` 设置为 `#0a0a23`，以便你能看见 logo。 然后，使用 `aspect-ratio` 属性将所需的长宽比设置为 `35 / 4`。 最后，在周围添加 `0.4rem` 的 `padding`。

# --hints--

你应该给 `#logo` 一个 `background-color`，值为 `#0a0a23`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

你应使用 `aspect-ratio` 属性。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

你不应使用 `height` 属性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

你不应该修改`width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

你应该给 `img` 一个 `aspect-ratio`，值为 `35 / 4`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

你应该给 `img` 一个 `padding`，值为 `0.4rem`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(10rem, 18vw);

}
--fcc-editable-region--

```
