---
id: 614ccc21ea91ef1736b9b578
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

本工作坊为你提供了一个基本的模板。

通过为你的 `html` 元素添加一个 `lang` 属性来开始这段无障碍之旅。 这将帮助屏幕阅读器识别页面的语言。

此外，记得将样式表链接到页面。

# --hints--

你应该为 `html` 元素添加一个 `lang` 属性。*提示：你可以使用值 `en` 来表示英语。*

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

你的 `head` 元素中应该有一个 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素应包含 `rel="stylesheet"` 和 `href="styles.css"`。

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
