---
id: 66c644c08ebcd2ef90c750bd
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie mit der `push`-Methode Elemente am Ende eines Arrays hinzufügen, so wie hier:

```js
fruits.push('pear');
```

Verwenden Sie die `push`-Methode, um den String `"Apples"` zum `shoppingList`-Array hinzuzufügen.

# --hints--

Ihr `shoppingList`-Array sollte ein Element enthalten.

```js
assert.lengthOf(shoppingList, 1);
```

Ihr `shoppingList`-Array sollte den String `"Apples"` enthalten. Überprüfen Sie die Rechtschreibung und Groß-/Kleinschreibung sorgfältig.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
