---
id: 66c64210413532ee9d3bd342
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Für diesen nächsten Teil des Workshops üben Sie, Lebensmittel zur Einkaufsliste hinzuzufügen.

Beginnen Sie damit, `console.log` zu verwenden, um die Nachricht `"It will be nice to have some fruit to eat."` zu protokollieren.

# --hints--

Sie sollten die Nachricht `"It will be nice to have some fruit to eat."` in der Konsole protokollieren.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
