---
id: 66c64095f61166ede6e39a84
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Für diese Einkaufsliste verwenden Sie ein Array, um die Artikel darzustellen, die Sie kaufen müssen.

In vorherigen Lektionen haben Sie gelernt, wie man Arrays so erstellt:

```js
const fruits = ['apple', 'banana', 'orange'];
```

Erstellen Sie in diesem Schritt eine Variable namens `shoppingList` und weisen Sie ihr ein leeres Array zu.

# --hints--

Sie sollten eine Variable namens `shoppingList` haben.

```js
assert.isNotNull(shoppingList);
```

Sie sollten der Variable `shoppingList` ein Array zuweisen.

```js
assert.isArray(shoppingList);
```

Ihr Array sollte leer sein.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
