---
id: 66c63bfa4685e4e3843fa0dc
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie weiter über Arrays, indem Sie eine Einkaufsliste erstellen.

Beginnen Sie damit, eine `console.log` hinzuzufügen, die den String `"Grocery shopping list"` in der Konsole ausgibt.

# --hints--

Sie sollten eine `console.log` in Ihrem Code haben.

```js
assert.match(code, /console\.log(.*)/);
```

Ihre `console.log` sollte die Nachricht `"Grocery shopping list"` ausgeben. Überprüfen Sie Rechtschreibung und Groß-/Kleinschreibung sorgfältig.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
