---
id: 66c6491fe4c8e0f16845425f
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, eine weitere Frucht zur Liste hinzuzufügen.

Verwenden Sie dieselbe Array-Methode wie zuvor, um den String `"Grapes"` ans Ende des `shoppingList`-Arrays anzufügen.

Fügen Sie dann eine `console.log` hinzu und rufen Sie die `getShoppingListMsg`-Funktion mit dem `shoppingList`-Array als Argument innerhalb von `console.log` auf, um die aktualisierte Liste in der Konsole auszugeben.

# --hints--

Ihr `shoppingList`-Array sollte jetzt zwei Elemente enthalten.

```js
assert.lengthOf(shoppingList, 2);
```

Ihr `shoppingList`-Array sollte den String `"Grapes"` enthalten.

```js
assert.equal(shoppingList[1], "Grapes");
```

Sie sollten die `getShoppingListMsg(shoppingList)`-Funktion innerhalb von `console.log` aufrufen.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
