---
id: 66c645b345a39bf04864dc50
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Um die aktualisierte Einkaufsliste zu sehen, müssen Sie die aktuelle Einkaufsliste und eine kurze Nachricht in der Konsole ausgeben.

Da diese Nachricht im gesamten Workshop mehrfach verwendet wird, ist es am besten, eine wiederverwendbare Funktion zu erstellen.

Erstellen Sie eine Funktion namens `getShoppingListMsg`, die ein Array als Parameter entgegennimmt und den String `"Current Shopping List: "` gefolgt vom Inhalt des übergebenen Arrays zurückgibt.

Sie können Template-Literale oder String-Verkettung mit dem `+`-Operator verwenden, um den String mit dem an die Funktion übergebenen Array zu kombinieren.

# --hints--

Sie sollten eine Funktion namens `getShoppingListMsg` in Ihrem Code haben.

```js
assert.isFunction(getShoppingListMsg);
```

Ihre `getShoppingListMsg`-Funktion sollte einen String zurückgeben.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Ihre `getShoppingListMsg`-Funktion sollte ein Array akzeptieren und die Nachricht `"Current Shopping List: "` gefolgt vom Inhalt dieses Arrays zurückgeben. Überprüfen Sie sorgfältig Leerzeichen und Rechtschreibfehler.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
