---
id: 66c72f4d0528bd268a82107b
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Jetzt ist es Zeit, das aktualisierte `shoppingList`-Array in der Konsole auszugeben.

Fügen Sie einen weiteren `console.log` hinzu und rufen Sie die `getShoppingListMsg`-Funktion mit dem `shoppingList`-Array als Argument innerhalb von `console.log` auf, um die aktualisierte Liste in der Konsole anzuzeigen.

# --hints--

Sie sollten `getShoppingListMsg(shoppingList)` innerhalb von `console.log` aufrufen.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
