---
id: 66c726c34ecf1e238aa9d7d5
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie Elemente mit der `unshift()`-Methode am Anfang eines Arrays hinzufügen.

Hier eine Erinnerung daran, wie Sie die `unshift()`-Methode verwenden:

```js
array.unshift(item1, item2, ..., itemX);
```

Verwenden Sie die `unshift()`-Methode, um den String `"Vegetable Oil"` am Anfang des `shoppingList`-Arrays hinzuzufügen.

# --hints--

Ihr `shoppingList`-Array sollte drei Elemente enthalten.

```js
assert.lengthOf(shoppingList, 3);
```

Sie sollten die `unshift()`-Methode verwenden, um `"Vegetable Oil"` am Anfang des `shoppingList`-Arrays hinzuzufügen.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
