---
id: 66c730183f4020275cbf0611
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Für diesen nächsten Abschnitt des Workshops überprüfen Sie, wie Sie Elemente vom Ende des Arrays entfernen.

Beginnen Sie damit, einen `console.log` hinzuzufügen, der die Nachricht `"This looks like too much junk food."` protokolliert.

# --hints--

Sie sollten die Nachricht `"This looks like too much junk food."` protokollieren.

```js
assert.match(code, /console\.log\((['"])(This\s+looks\s+like\s+too\s+much\s+junk\s+food.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
