---
id: 66c74079c30b1c2d166cb9a4
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie mit der `shift`-Methode ein Element vom Anfang des Arrays entfernen.

Hier eine Erinnerung, wie Sie die `shift`-Methode verwenden:

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

Verwenden Sie die `shift`-Methode, um das erste Element aus dem `shoppingList`-Array zu entfernen.

# --hints--

Ihr `shoppingList`-Array sollte insgesamt `5` Elemente enthalten.

```js
assert.lengthOf(shoppingList, 5);
```

Sie sollten die `shift`-Methode verwenden, um das erste Element aus dem `shoppingList`-Array zu entfernen.

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
