---
id: 66c73fa7433e082c4be096b1
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Im letzten Teil des Workshops überprüfen Sie, wie Sie ein Element vom Anfang eines Arrays entfernen.

Beginnen Sie damit, einen `console.log` hinzuzufügen, der die Nachricht `"On second thought, maybe we should be more health conscious."` ausgibt.

# --hints--

Sie sollten einen `console.log` haben, der die Nachricht `"On second thought, maybe we should be more health conscious."` ausgibt.

```js
assert.match(code, /console\.log\((["'])(On\s+second\s+thought,\s+maybe\s+we\s+should\s+be\s+more\s+health\s+conscious\.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
