---
id: 66c742d045c9fc2e09fa64b1
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Die letzte Änderung an der Einkaufsliste besteht darin, den ersten Eintrag in der Liste zu aktualisieren.

In vorherigen Lektionen haben Sie gelernt, wie Sie einen Eintrag mit der Klammernotation und dem Index des zu aktualisierenden Eintrags ändern.

Hier eine Erinnerung daran, wie man einen Eintrag in einem Array aktualisiert:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

Aktualisieren Sie den ersten Eintrag im `shoppingList`-Array zu `"Canola Oil"`.

# --hints--

Sie sollten den ersten Eintrag im `shoppingList`-Array zu `"Canola Oil"` aktualisieren.

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
