---
id: 66c72b0ffbc5522525768558
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

In früheren Schritten haben Sie überprüft, wie Sie mit der `push`-Methode ein Element am Ende des Arrays hinzufügen.

Aber die `push`-Methode akzeptiert mehrere Argumente, sodass Sie mehrere Elemente auf diese Weise am Ende des Arrays hinzufügen können:

```js
array.push(item1, item2, item3);
```

Verwenden Sie in diesem Schritt die `push`-Methode, um die Strings `"Popcorn", "Beef Jerky", "Potato Chips"` zum `shoppingList`-Array hinzuzufügen.

Die Reihenfolge ist wichtig, also fügen Sie die Elemente in der angegebenen Reihenfolge hinzu.

# --hints--

Ihr `shoppingList`-Array sollte insgesamt `6` Strings enthalten.

```js
assert.lengthOf(shoppingList, 6);
```

Ihr `shoppingList` sollte den String `"Popcorn"` enthalten.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Ihr `shoppingList` sollte den String `"Beef Jerky"` enthalten.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Ihr `shoppingList` sollte den String `"Potato Chips"` enthalten.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
