---
id: 66c73a7798f6f62b2ae58f22
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Jetzt ist es Zeit, weitere Elemente am Anfang der Einkaufsliste hinzuzufügen.

Beginnen Sie damit, eine `console.log`-Anweisung hinzuzufügen, die die Nachricht `"It might be nice to get a dessert."` ausgibt.

Unterhalb dieser `console`-Anweisung verwenden Sie die richtige Array-Methode, um den String `"Chocolate Cake"` am Anfang des `shoppingList`-Arrays hinzuzufügen.

Fügen Sie schließlich eine `console.log` hinzu und rufen Sie die `getShoppingListMsg`-Funktion mit dem `shoppingList`-Array als Argument innerhalb von `console.log` auf, um die aktualisierte Liste in der Konsole auszugeben.

# --hints--

Sie sollten eine `console.log` haben, die die Nachricht `"It might be nice to get a dessert."` ausgibt.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Ihr `shoppingList`-Array sollte insgesamt `6` Elemente enthalten.

```js
assert.lengthOf(shoppingList, 6);
```

Sie sollten die `unshift`-Methode verwenden, um den String `"Chocolate Cake"` am Anfang des `shoppingList`-Arrays hinzuzufügen.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Sie sollten `getShoppingListMsg(shoppingList)` innerhalb von `console.log` aufrufen.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
