---
id: 66c730ee6ae076281721d0b9
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie Sie Elemente am Ende eines Arrays mit der `pop`-Methode entfernen.

Hier eine Erinnerung daran, wie Sie die `pop`-Methode verwenden:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Verwenden Sie die `pop`-Methode, um das letzte Element aus dem `shoppingList`-Array zu entfernen.

# --hints--

Ihr `shoppingList`-Array sollte jetzt `5` Elemente enthalten.

```js
assert.lengthOf(shoppingList, 5);
```

Ihr `shoppingList`-Array sollte den String `"Potato Chips"` nicht mehr enthalten.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
