---
id: 66cbe2319d3845545a293a0b
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Jetzt ist es Zeit, die Nachricht zu sehen, die in der Konsole ausgegeben wird.

Fügen Sie einen `console.log` hinzu und rufen Sie die `getShoppingListMsg`-Funktion mit dem `shoppingList`-Array als Argument innerhalb von `console.log` auf, um die Nachricht in der Konsole auszugeben.

# --hints--

Sie sollten die `getShoppingListMsg(shoppingList)`-Funktion innerhalb von `console.log` aufrufen.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
