---
id: 65823dde36392f505a39f7c7
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

最后，将你的 `map()` 调用传递给 `sum()` 函数以获取总费用并返回结果。

# --hints--

你应该调用 `sum()` 函数，并将你当前的 `map()` 调用作为参数传递给它，然后从 `total_expenses` 函数返回结果。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
