---
id: 65823040bdb32949d4c1a96f
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

你将需要一个函数来计算总支出金额。

定义一个名为 `total_expenses` 的函数，它接受一个参数 `expenses`。 现在用 `pass` 语句填充函数主体。

# --hints--

你应该定义一个名为 `total_expenses` 的函数。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

你的 `total_expenses` 函数应该采用单个 `expenses` 参数。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
