---
id: 65823634a308cb4d6bcc5fa6
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

Lambda 函数可以与 `map()` 函数进行有价值的结合，该函数对对象集合（例如列表）中的每个元素执行指定的函数：

```py
map(lambda x: x * 2, [1, 2, 3])
```

要执行的函数作为第一个参数传递，可迭代对象作为第二个参数传递。

上述示例的结果将是 `[2, 4, 6]`，其中传递给 `map()` 的列表中的每个项目都通过 lambda 函数的操作加倍。

修改你的 `print()` 调用以打印调用 `map()` 的结果，其中 `test` 作为第一个参数，`[2, 3, 5, 8]` 作为第二个参数。 你还无法看到可读的输出。 

# --hints--

你的代码中应该有 `print(map(test, [2, 3, 5, 8]))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
