---
id: 65824dfdb6815d563b2d3256
title: 步骤 37
challengeType: 20
dashedName: step-37
---

# --description--

`input()` 函数接受用户输入并以字符串的形式返回用户输入。

在 `while` 循环内，调用 `input()` 函数，传递字符串 `'Enter your choice: '` 作为参数，并将结果分配给名为 `choice` 的变量。

# --hints--

你应该在 `while` 循环中声明一个名为 `choice` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

你应该将 `input()` 函数的结果（以字符串 `'Enter your choice: '` 作为其参数）分配给 `choice` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

--fcc-editable-region--
```
