---
id: 658252f6b1526d57b103d48a
title: 步骤 39
challengeType: 20
dashedName: step-39
---

# --description--

在 `if` 语句中，你应该要求用户输入费用金额并将其存储在变量中。

将字符串 `'Enter amount: '` 传递给空的 `input()` 调用，这样你就可以存储费用。

# --hints--

你应该将 `'Enter amount: '` 传递给分配给 `amount` 变量的 `input()` 调用。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input('Enter amount: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = input()
--fcc-editable-region--
```
