---
id: 6582324b5d37fe4b622bc3c0
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda 函数是 Python 中的简短、匿名函数，非常适合简单的一次性任务。 它们由 `lambda` 关键字定义，并使用以下语法：

```py
lambda x: expr
```

在上面的例子中，`x` 表现为将在表达式 `expr` 中使用的参数，它的作用就像传统函数中的任何参数一样。`expr` 是在调用 lambda 函数时被计算并返回的表达式。

创建一个名为 `test` 的变量，并为其分配一个 lambda 函数，该函数接受 `x` 参数并返回 `x * 2`。

# --hints--

你应该声明一个名为 `test` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

你应该将 `lambda x: x * 2` 分配给你的 `test` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
