---
id: 658212ba98182f3e855e85f9
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

在这个项目中，你将使用 Python 构建一个简单但实用的费用跟踪器。

首先定义一个名为 `add_expense` 的函数，该函数接受三个参数：`expenses`、`amount` 和 `category`。 使用 `pass` 关键字填充函数主体。

# --hints--

你应该使用 `def` 关键字来声明一个名为 `add_expense` 的函数。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

你的 `add_expense` 函数应该采用三个参数：`expenses`、`amount` 和 `category`。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
