---
id: 65822bd82d708c4895080c35
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

在 Python 中，需要了解的一件重要事情是，用于定义字符串的相同类型的引号不能在字符串内部使用。 例如，字符串 `'I'm a string!'` 无效。 要在该字符串中使用单引号，你应该：

- 通过在引文前添加反引号来逃避引文：`'I\'m a string!'`
- 或者使用双引号定义字符串：`"I'm a string!"`（首选）。

你可以通过字典的键访问字典中的值。 你需要使用括号表示法并在方括号之间包含键：

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

你目前正在 f 字符串中插入 `expense` 字典。 修改 f 字符串表达式以访问 `expense` 字典中的 `'amount'` 键和 `'category'` 键的值。

# --hints--

你应该将 `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` 传递给你的 `print()` 调用。 请记住在单引号 f 字符串中使用双引号，反之亦然。

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
