---
id: 65822773f4c60a479ec89f66
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

在 `print_expenses` 函数内部，创建一个 `for` 循环，迭代 `expenses` 列表中的每个项目。 使用 `expense` 作为循环变量，并将 `pass` 移动到循环体内。

# --hints--

你应该在 `print_expenses` 函数中创建一个 `for` 循环，该循环遍历 `expenses`。 记住在循环体内使用 `pass`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

你应该使用 `expense` 作为循环变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

你应该只在循环体内有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
