---
id: 65823e87c9741750a22085a7
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

接下来，定义一个名为 `filter_expenses_by_category` 的函数，它接受两个参数：`expenses` 和 `category`。 使用 `pass` 填充函数体。

# --hints--

你应该定义一个名为 `filter_expenses_by_category` 的函数。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

你的 `filter_expenses_by_category` 函数应该以 `expenses` 和 `category` 作为参数。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
