---
id: 65823b0167fd894f4a7ea60a
title: 步骤 24
challengeType: 20
dashedName: step-24
---

# --description--

接下来，你将在 `total_expenses` 函数中实现相同的逻辑。

现在，删除 `test` 函数和 `print()` 调用。

# --hints--

你应该删除你的 `test` 变量和相关的 `print()` 调用。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
