---
id: 6582267ef39f5047411c1743
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

接下来，定义一个名为 `print_expenses` 的函数，它接受一个参数 `expenses`。 此功能稍后将用于显示列表中的每项支出。

使用 `pass` 语句填充新函数的主体。

# --hints--

你应该定义一个名为 `print_expenses` 的新函数。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

你的 `print_expenses` 函数应该采用单个参数 `expenses`。

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
