---
id: 658246c28575d653d1f89b59
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

`while` 循环是另一种循环，只要指定条件为 `True`，它就会运行一部分代码。 当条件变为 `False` 时循环终止：

```py
while condition:
    <code>
```

在 `expenses` 列表下方，创建一个 `while` 循环。 使用 `True` 作为条件，并在循环体内部打印字符串 `'\nExpense Tracker'` 以显示程序的标题。

# --hints--

你应该创建一个 `while` 循环，并使用 `True` 作为 `main()` 函数中的条件。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

你应该在 `while` 循环中打印 `'\nExpense Tracker'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
