---
id: 6582507654b3ed5712341382
title: 步骤 38
challengeType: 20
dashedName: step-38
---

# --description--

你将使用条件语句来检查用户的选择。 如果选择的是 `'1'`，则表示用户想要添加一项费用。

仍然在 `while` 循环中，在 `choice` 变量下，编写一个 `if` 语句来检查 `choice` 是否等于字符串 `'1'`。 如果这是真的，这将成为增加一项新支出的起点。

在 `if` 语句主体内，声明一个变量 `amount` 并为其分配一个空的 `input()` 调用。

# --hints--

你应该创建一个 `if` 语句来检查 `while` 循环中的 `choice` 是否等于 `'1'`。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

你应该在 `if` 语句主体中声明一个名为 `amount` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

你应该为你的 `amount` 变量分配一个空的 `input()` 调用。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
