---
id: 65822934c8a0454803e3ef12
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

接下来，你将显示每项费用的详细信息。

在 `for` 循环内，用 `print()` 调用替换 `pass`，并将以下 f 字符串传递给它：`f'Amount: {expense}, Category: {expense}'`。

# --hints--

你应该在 `for` 循环中打印 `f'Amount: {expense}, Category: {expense}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

循环体内不应有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
