---
id: 658258c0e5fbe85c14c060cf
title: 第 43 步
challengeType: 20
dashedName: step-43
---

# --description--

要列出所有费用，你可以在 `if` 语句后使用 `elif` 子句。 `elif` 检查附加条件并且仅在 `if` 语句之后起作用。

创建一个 `elif` 子句来检查用户的选择是否等于字符串 `'2'`。 在 `elif` 子句中，打印字符串 `'\nAll Expenses:'`。

# --hints--

你应该在 `if` 语句之后创建一个 `elif` 子句来检查 `choice` 是否等于 `'2'`。

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1]
cond.is_equivalent("choice == '2'") or cond.is_equivalent("'2' == choice")
`)) })
```

你应该在 `elif` 子句内打印 `'\nAll Expenses:'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("print('\\\\nAll Expenses:')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
        
--fcc-editable-region--
```
