---
id: 658225d560369446ece5518b
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

将另一个键值对添加到你要附加到 `expense` 列表的字典中。 使用字符串 `'category'` 作为键，使用 `category` 参数作为值。

# --hints--

你应该将第二个键值对添加到作为参数传递给 `.append()` 方法的字典中。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

你应该将 `{'amount': amount, 'category': category}` 作为参数传递给 `.append()` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
