---
id: 65821c1c40ccac44adf38e6b
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

创建一个名为 `expenses` 的空列表。 你将使用它来存储你的每笔开支。

# --hints--

你应该声明一个名为 `expenses` 的新变量。

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

你应该为 `expenses` 变量分配一个空列表。

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
