---
id: 65821fcc010c3245718f2a06
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

`add_expense` 函数的 `expenses` 参数将是费用列表。 你希望能够在列表末尾添加项目。 为此，你将使用 `.append()` 列表方法。

在 `expenses` 列表中添加对 `.append()` 方法的调用。 暂时不要向 `.append()` 传递任何参数。

# --hints--

你的函数中应该有 `expenses.append()`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

你的函数中不应该有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
