---
id: 65824111a09164518320088d
title: 第 30 步
challengeType: 20
dashedName: step-30
---

# --description--

`filter()` 函数允许你根据函数的输出从可迭代对象（例如列表）中选择项目： 

```py
filter(my_function, my_list)
```

`filter()` 以函数作为其第一个参数，以可迭代对象作为其第二个参数。 它返回一个迭代器，这是一个特殊的对象，使你能够迭代集合（如列表）的元素。

上述示例的结果是一个迭代器，包含 `my_list` 的元素，而 `my_function` 返回 `True`。 

在 `filter_expenses_by_category` 函数中，调用 `filter()`，将你在上一步中编写的 `lambda` 函数作为第一个参数传递，将 `expenses` 列表作为第二个参数传递。

# --hints--

你应该在 `filter_expenses_by_category` 函数内调用 `filter()`。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

你应该将 `lambda expense: expense['category'] == category` 作为第一个参数传递给 `filter()` 调用。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

你应该将 `expenses` 作为第二个参数传递给 `filter()` 调用。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
