---
id: 65824872894f59546e3084e2
title: 步骤 34
challengeType: 20
dashedName: step-34
---

# --description--

你在上一步中创建的 `while` 循环是一个无限循环，它允许程序不断显示菜单选项，直到用户决定退出。 

在 `print()` 调用之后，添加另一个 `print()` 调用以打印字符串 `'1. Add an expense'`。

# --hints--

你应该打印 `'1. Add an expense'` 在 `while` 循环中现有的 `print()` 调用之后。

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        
--fcc-editable-region--
```
