---
id: 658238f7604f154ea9a23e1e
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

`sum()` 函数返回作为其参数传递的可迭代对象中的项目的总和。 你将使用 `sum()` 以及 `map()` 和 `lambda` 函数来获取总费用金额。

现在，做一个小测试并修改你当前的 `print()` 调用，将 `list()` 调用替换为对 `sum()` 函数的调用，并将当前 `map()` 调用作为参数传递给它。

# --hints--

你应该打印 `sum(map(test, [2, 3, 5, 8]))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
