---
id: 65823ff0d4b991510fade1a8
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

在 `filter_expenses_by_category` 函数中，将 `pass` 替换为 `lambda` 函数。 使用 `expense` 作为参数，并使用相等运算符评估 `expense` 字典的 `'category'` 键的值与 `category` 之间的比较。

# --hints--

你应该创建一个 `lambda` 函数，该函数使用名为 `expense` 的参数并评估 `filter_expenses_by_category` 函数内的表达式 `expense['category'] == category`。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

`filter_expenses_by_category` 函数内不应有 `pass`。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
