---
id: 65822308ded1ba4632f66aa3
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

字典是 Python 中的另一种内置数据类型。 字典是以*键*-*值*对形式存在的数据集合。 字典用花括号（`{}`）定义，并且包含用逗号分隔的键值对。 每个键后面跟着一个冒号（`:`）和值：

```py
{'amount': 50.0, 'category': 'Food'}
```

在上面的例子中，`'amount'` 和 `'category'` 是键，`50.0` 和 `'Food'` 是它们对应的值。

创建一个带有键 `'amount'` 和值 `amount` 参数的字典，并将新字典传递给 `.append()` 调用。

# --hints--

你应该将字典作为参数传递给 `.append()` 方法。

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

你应该将 `{'amount': amount}` 作为参数传递给 `.append()` 方法。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
