---
id: 65824561f3478e5371a33ae5
title: 步骤 32
challengeType: 20
dashedName: step-32
---

# --description--

下一步是定义主函数，它将成为交互式费用跟踪程序的入口点。

定义一个名为 `main` 的函数，没有参数。 使用你在此项目开始时创建的 `expenses` 列表填充函数主体。 你将使用此列表来存储费用记录。

# --hints--

你应该定义一个名为 `main()` 的函数，没有参数。

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

你应该将 `expenses` 列表移到 `main()` 函数体内部。

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
