---
id: 6582352cbdbcde4d0c4f7b0b
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

要调用 lambda 函数，你可以使用通常的函数语法，并在变量名后加上一对括号。

调用你的 `test` lambda 函数并传递 `3` 作为参数。 然后，打印结果。

# --hints--

你的代码中应该有 `print(test(3))`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
