---
id: 66314f2add0a8b57da5d6d91
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

你应该在控制台看到类似 `<map object at 0xd273a8>` 的内容被打印出来，这就是 `map()` 返回的 map 对象的字串表现形式。

为了获得可读的输出，你需要将地图对象转换为列表。 通过将 `map()` 调用作为参数传递给 `list()` 函数来实现。

# --hints--

你应该将 `map()` 调用作为参数传递给 `list()` 函数。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
