---
id: 66694eaaeb7d253049fa7883
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

要将项目添加到列表中，你可以将它们写在方括号之间，并用逗号分隔值。

在列表中添加数字 `1` 和 `2`。

# --hints--

`my_list` 的第一个元素应该是 `1`。

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

`my_list` 的第二个元素应该是 `2`。

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` 里面应该有两个元素。

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

你应该直接将 `[1, 2]` 分配给 `my_list`。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
