---
id: 66694fc4bba24f33ca01fa5b
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 有少量的列表方法。 例如添加或删除列表项的方法。

你可以使用 `append()` 方法将项目添加到列表末尾。 例如：

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` 现在是 `[4, 5, 6, 7]`。

尝试使用 `append()` 方法将 `3` 添加到 `my_list`。 然后打印该列表。

# --hints--

`my_list` 的最后一个元素应该是 `3`。

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` 的长度应为 3。

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

你不应该更改列表声明。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

你应该打印这份清单。

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

你应该在使用 `append()` 后打印列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
