---
id: 6669539c1379793f9cb8917c
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

Python 列表是可变的，这意味着列表项的值可以改变。 你可以使用括号符号来更改元素的值。

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

这将使 `example_list` 具有 `[4, 'oh', 6, 7]` 的值。

将 `my_list` 的第一个元素更改为 `0`，然后打印列表以检查其值。

# --hints--

不要更改列表声明。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

`my_list` 的第一个元素应该是 0。

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

你应该使用 `my_list[0] = 0` 更改第一个元素。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

你应该打印这份清单。

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

你应该在更改第一个元素后打印列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
