---
id: 66694bfa435cef241f6f9a68
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在开始项目之前，你将了解列表。

与数字和字符串不同，<dfn>列表</dfn>是一种可用作其他值的容器的数据类型：

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

该列表的特点是所有值都用方括号括起来，值之间用逗号分隔，例如：`["A", "happy", "list"]`。 如果列表不包含任何值，那么它是一个空列表：`[]`。

列表可以包含不同的数据类型：`[1, "Up", ["Down", "Twice"]]`。 其中包括所有可能的数据类型。 它还可以包含另一个列表！

创建一个名为 `my_list` 的变量并为其分配一个空列表。

# --hints--

你应该有一个名为 `my_list` 的变量。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

变量 `my_list` 应该具有空列表的值。

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
