---
id: 666951d2e4b27e3a3c72ca5f
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

你还可以访问单个元素来获取其值。

列表与字符串一样都是从零开始索引的。 这意味着第一个元素位于索引 `0`，第二个元素位于索引 `1`，依此类推。

要访问元素，请使用括号表示法。 例如，`example_list[1]` 正在访问索引 `1` 处的元素，即 `example_list` 的第二个元素。

打印 `my_list` 的第一个元素。

# --hints--

你应该使用 `print(my_list[0])` 打印列表的第一个元素。

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
