---
id: 666aae6a5d30a71f1fd7749f
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

`insert`方法可以在列表的任意位置插入一个元素。 第一个参数是要插入元素的位置，第二个参数是实际要插入的元素。 下面的例子说明了如何在`example_list`的第三个位置插入一个新的元素：

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

使用 `insert()`方法，在`my_list` 合适的位置添加元素 `1`，以便于让计数递增，然后打印列表。

# --hints--

你需要使用`insert()`方法在`my_list`变量索引 `1`的位置添加元素`1`。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

这样`my_list`中会有 4个元素。

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` 应该是`[0, 1, 2, 3]`。

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

你需要在增加了新元素之后打印列表。

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
