---
id: 666ab9013cd5563f2b39947a
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在你已经探索了列表，请删除你为审查而编写的所有代码，以便可以开始该项目。

# --hints--

你应该删除所有内容。

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
