---
id: 666ab82c874add3c95145213
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

`pop()` 方法可以用于从列表中删除元素。 默认情况下，它会删除列表中的最后一个元素。 你可以将一个索引作为参数传递给方法，然后它会删除给定索引位置的元素。

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

在这个例子中，`fruits_list.pop(2)` 删除了列表中索引为 `2` 的元素。

使用 `pop()` 删除 `my_list`中的最后一个元素，然后打印 `my_list`。

# --hints--

你应该使用 `pop()`。

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` 应该是 `[0, 1, 2]`。

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

在使用 `pop()` 之后，你要打印 `my_list`。

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
