---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了させてください。

**User Stories:**

1. `fibonacci` という名前の関数を作成してください。
2. `sequence` 関数内に `fibonacci` という配列を定義し、`[0, 1]` の値で初期化してください。
3. `fibonacci` 関数は、正の整数 `n` を1つのパラメータとして受け取る必要があります。
4. `fibonacci(n)` を呼び出すと、動的計画法のアプローチを使ってフィボナッチ数列の `n` 番目の数を計算し返す必要があります。フィボナッチ数列は、各数が直前の2つの数の合計です。
5. フィボナッチ数列の `n` の位置で計算された各数は、インデックス `sequence` の `n - 1` 配列に格納される必要があります。

# --hints--

`fibonacci` という名前の関数を用意してください。

```js
assert.isFunction(fibonacci);
```

`fibonacci` 関数は1つのパラメータを受け取る必要があります。

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

`sequence` 関数内に `fibonacci` という配列を用意し、`[0, 1]` で初期化してください。

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

コード内で再帰を使わないでください。

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
