---
id: 7a97b79c096ecb000d03bd36
title: N番目のフィボナッチ数計算機を作成する
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `fibonacci` という名前の関数を作成してください。
2. `sequence` 関数内に `fibonacci` という配列を定義し、`[0, 1]` の値で初期化してください。
3. `fibonacci` 関数は、非負整数の `n` を1つのパラメータとして受け取る必要があります。
4. `fibonacci(n)` を呼び出すと、動的計画法のアプローチを使ってフィボナッチ数列の `n` 番目の数を計算し返す必要があります。フィボナッチ数列は、各数が直前の2つの数の合計です。
5. 計算された各フィボナッチ数は、`sequence` 配列に追加される必要があります。

# --hints--

`fibonacci` という名前の関数を用意してください。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.fibonacci);
```

`fibonacci` 関数は1つのパラメータを受け取る必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
assert.equal(fibonacci?.parameters.length, 1);
```

`sequence` 関数内に `fibonacci` という配列を用意し、`[0, 1]` で初期化してください。

```js
const explorer = await __helpers.Explorer(code);
const { fibonacci } = explorer.allFunctions;
const { sequence } = fibonacci?.variables ?? {};
assert.isTrue(sequence?.value.matches('[0, 1]'));
```

`fibonacci(0)` は `0` を返す必要があります。

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` は `1` を返す必要があります。

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` は `1` を返す必要があります。

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` は `2` を返す必要があります。

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` は `5` を返す必要があります。

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` は `55` を返す必要があります。

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` は `610` を返す必要があります。

```js
assert.equal(fibonacci(15), 610);
```

コード内で再帰を使わないでください。

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
