---
id: 686daa7ed79ceacd0b264e7d
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

يجب أن تُسكن عنصر `path` داخل عنصر `svg` لإعطاء الصورة شكلها.

أنشئ عنصر `path`.

# --hints--

يجب أن يكون لديك عنصر `path` مسكن داخل عنصر `svg`.

```js
const path = document.querySelector('svg path');
assert.exists(path);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    <svg>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </svg>
  </body>
</html>
```
