---
id: 686d749d6f32f708fdf2b7ad
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في درس سابق، تعلّمت عن عناصر `svg` وكيف تُستخدم غالبًا كأيقونات في المشاريع. في قواعد الشفرة الحقيقية، عادةً ما تعتمد على مكتبات الأيقونات، لذلك لا تحتاج إلى إنشاء عناصر `svg` من الصفر. مع ذلك، في هذه الورشة، ستبني أيقونة قلب لتتعلم عن الخصائص الأساسية المستخدمة داخل عنصر `svg`.

ابدأ بإنشاء عنصر `svg` على الصفحة.

# --hints--

يجب أن يكون لديك عنصر `svg` على الصفحة.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
