---
id: 6853e58d52ec64349a5760ca
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、最短経路アルゴリズムを実装します。グラフ内のノード間の最短経路を計算し、通った経路も返すPythonの関数を書いていきます。

たとえば、都市が異なる距離の道路でつながっているグラフがあるとき、ある都市から別の都市への最短ルートをアルゴリズムが見つけます。都市Aから都市Dへ移動したい場合、A ⇨ B ⇨ C ⇨ D（合計15km）と進むほうが、直接A ⇨ D（20km）より短いと判明するかもしれません。

まずは、`INF`という名前の変数を定義し、正の無限大を表す`float('inf')`を代入してください。後で、2つのノード間の距離が無限大であることを示すのに使います。

# --hints--

ファイルの先頭で`INF`という変数を定義してください。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

変数`float('inf')`に`INF`を代入してください。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
