---
id: 686253f1b051998ad4904e3e
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

これから、隣接行列、開始ノード、そしてオプションのターゲットノードという3つのパラメータを受け取るメインの関数を作成します。

`shortest_path` という名前の関数を作成し、`matrix`、`start_node`、`target_node` の3つのパラメータを受け取るようにしてください。`None` にはデフォルト値として `target_node` を代入してください。

`target_node` パラメータはデフォルトでは `None` に設定されており、これはターゲットノードが指定されなかった場合に、開始ノードからグラフ内のすべての他のノードへの最短経路を計算することを意味します。

関数の本体には、今のところ `pass` 文を追加してください。

# --hints--

`shortest_path` という名前の関数を作成してください。

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

関数は3つのパラメータを受け取る必要があります：`matrix`、`start_node`、そして `target_node` で、`target_node` は `None` をデフォルト値として持ちます。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
