---
id: 686251fb49ebc489ba60ccb5
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

グラフの隣接行列を表すために2次元リストを作成する必要があります。この行列は、グラフ内のノード間の辺の重みを表すために使われます。

`adj_matrix` という名前の変数を作成し、以下の重みを持つグラフを表す2次元リストを代入してください。

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

`adj_matrix` という名前の変数を作成してください。

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` は指定された値を含む2次元リストである必要があります。

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
