---
id: 686b78bc13d518e1267448e4
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

距離を追跡するだけでなく、各ノードに到達するまでに通った実際の経路も記録する必要があります。

各エントリーがそのノードに到達するまでの経路を保存するリストを作成します。最初は、各ノードの経路はそのノード自身だけを含みます。

リスト内包表記はリストを簡潔に作成する方法を提供します。例えば：

```py
[x * 2 for x in range(3)]
```

`paths`の各`[node_no]`に対して`node_no`を含むリストを作成するリスト内包表記を使って、`range(n)`という変数を初期化してください。

# --hints--

`paths`という変数を用意してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

`paths`はリスト内包表記を使って初期化してください。

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

リスト内包表記では`node_no`を繰り返すために`range(n)`を使ってください。

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

リスト内包表記は`[node_no]`の各`node_no`に対して`range(n)`を評価する必要があります。

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
