---
id: 686b82af323098e7a5203bf4
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

アルゴリズムを実行する際、すでに訪れたノードを記録しておく必要があります。そうすることで、同じノードを複数回処理しないようにできます。

これを行うために、`visited` という名前のリストを作成し、すべてのノードに対して `False` で初期化してください。

# --hints--

`visited` 関数の中に `shortest_path` という名前のリストがあるはずです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

`visited` を使って、すべてのノードに対して `False` で `[False] * n` リストを初期化してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
