---
id: 68760c4c2648e7832c5eef16
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

最も近い未訪問のノードを見つけるループの後に、有効なノードが実際に見つかったかどうかを確認する必要があります。

もしそのようなノードが存在しなければ、残りのノードは開始ノードから到達不可能であり、アルゴリズムは早期に停止する必要があります。

ネストされた`for`ループと同じレベルに、`if`をチェックして真の場合にループを抜ける`current == -1`文を追加してください。

# --hints--

内側の`if`ループの後に、`current`がまだ`-1`かどうかをチェックする`for`文があるはずです。

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

ループを抜けるには`break`文を使うべきです。

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
