---
id: 68760464eb9c2e79cc913fac
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

このステップでは、グラフ内の各ノードに対して1回ずつ実行されるループを追加します。このループにより、アルゴリズムが複数回のパスで距離と経路を更新できるようになります。

`for` 回実行される `n` ループを作成してください。反復値を使う必要がないので、ループ変数には `_` を使用してください。

ループの中で、次に処理するノードを選択する準備として、2つの変数を作成する必要があります。

- 現在の反復で見つかった最小距離を保持する変数
- その最小距離を持つノードのインデックスを保持する変数

`min_distance` と `current` という変数を作成し、それぞれ `INF` と `-1` に設定してください。

# --hints--

`for` 関数内に `shortest_path` ループがあるはずです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

`for` ループは `range(n)` を繰り返すべきです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

for ループは `_` を反復変数として使うべきです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

for ループ内で `min_distance` という変数を作成し、`INF` に設定するべきです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

for ループ内で `current` という変数を作成し、`-1` に設定するべきです。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
