---
id: 6876089f7a500b7c782e5b83
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

現在のノードが、すでに見つけたノード（もしあれば）よりも良い選択かどうかを判断する必要があります。そのために、ループの中に条件文を追加します。

条件は次の2つを行う必要があります：

- ノードがまだ訪問されていないかどうかを確認する。
- 現在の `min_distance` までの既知の距離を比較する。

内側の `for` ループの中に、`if` がまだ訪問されていない **かつ** `node_no` が `distances[node_no]` より小さいかどうかをチェックする `min_distance` 文を追加してください。

条件ブロックの中に `pass` をプレースホルダーとして追加してください。

# --hints--

内側の `if` ループの中に `for` 文を追加してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

`if` 文は、ノードが未訪問（`visited[node_no]` が偽）であり、かつ `distances[node_no]` が `min_distance` より小さいかどうかをチェックする必要があります。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
