---
id: 68760671d950767b055bf2f9
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

まだ訪問していない中で、最も小さい既知の距離を持つノードを見つけるために、すべてのノードを確認する必要があります。

これを行うには、メインループの中に`for`ループを追加してください。

このループは、グラフ内のノード数が`node_no`である`range(n)`の各`n`を繰り返す必要があります。今はループの本体として`pass`を追加してください。

# --hints--

`for`関数の中に2つ目の`shortest_path`ループがあるはずです。

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

`for`ループは、反復変数`range(n)`を使って`node_no`を繰り返す必要があります。

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
