---
id: 68777816e95d4936f129819f
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

隣接ノードへの距離を更新しようとする前に、その隣接ノードが到達可能であり、かつ未訪問であることを確認する必要があります。その後、現在のノードを経由してその隣接ノードに到達するための合計距離を計算します。

`for` ループの中に、次のことをチェックする `if` 文を追加してください。

- `distance` が `INF` と等しくない（つまりノード間にエッジがある）こと
- 隣接ノード `node_no` がまだ訪問されていないこと

`if` ブロックの中に、`new_distance` という変数を作成し、`distances[current]`（現在のノードまでの最短距離）と `distance`（現在のノードから隣接ノードまでの距離）の合計を代入してください。

# --hints--

`if` が `distance` と等しくないこと、かつ隣接ノードが未訪問であることをチェックする `INF` 文があるはずです。

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

`new_distance` ブロックの中に `if` という変数を作成してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

`new_distance` と `distances[current]` の合計を `distance` に代入してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
