---
id: 687b3926419dec576850b814
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

隣接ノードへの新しい可能な距離を計算したので、それが現在`distances`リストに保存されている距離より良いかどうかを確認してください。良ければ、その距離を更新します。

既存の`if`ブロック内に、`if`が`new_distance`より小さいかどうかをチェックする`distances[node_no]`文を追加してください。

この新しい条件ブロック内で、`distances[node_no]`を更新して`new_distance`を保存してください。

# --hints--

`if`が`new_distance`より小さいかどうかをチェックするネストされた`distances[node_no]`文があるはずです。

```js
({
    test: () => {
      runPython(`  
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]  
      assert cond.is_equivalent("new_distance < distances[node_no]") or cond.is_equivalent("distances[node_no] > new_distance")  
      `);
    }
});
```

`distances[node_no]`を更新して`new_distance`を保存してください。

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("distances[node_no] = new_distance")
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

        if current == -1:  
            break  
        visited[current] = True  

--fcc-editable-region--  
        for node_no in range(n):
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance

--fcc-editable-region--
```
