---
id: 68e4c0ff9edcc92771d7549a
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

現在、`distances` リストには要素が1つだけです。しかし、グラフ内の各ノードに対して1つの距離の値が必要です。

Pythonでは、リストを整数で乗算すると繰り返すことができます。たとえば、`[0] * 3` は `[0, 0, 0]` を作成します。

`distances` リストを、`n` 個の `INF` を含むように、`[INF]` を `n` で乗算して更新してください。

# --hints--

`distances` を `n` 個の `INF` を含むように更新してください。

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
