---
id: 564944c91be2204b269d51e3
title: jQueryを使って要素内のテキストを変更する
challengeType: 6
forumTopicId: 16773
dashedName: change-text-inside-an-element-using-jquery
---

# --description--

jQueryを使うと、要素の開始タグと終了タグの間のテキストを変更できます。HTMLマークアップを変更することも可能です。

jQueryには`.html()`という関数があり、要素内にHTMLタグやテキストを追加できます。この関数を使うと、要素内に以前あった内容はすべて、指定した内容に置き換えられます。

見出しのテキストをどのように書き換えて強調するか、以下の例をご覧ください。

```js
$("h3").html("<em>jQuery Playground</em>");
```

また、jQueryには`.text()`という似た関数もありますが、こちらはタグを追加せずテキストだけを変更します。つまり、この関数は渡されたHTMLタグを評価せず、単に既存の内容を置き換えるテキストとして扱います。

idが`target4`のボタンのテキストを強調してください。

<a href="https://www.freecodecamp.org/news/html-elements-explained-what-are-html-tags/#em-element" target="_blank" rel="noopener noreferrer nofollow">こちらのニュース記事の&lt;em&gt;について</a>を参照して、`<i>`と`<em>`の違いや使い方を学んでください。

なお、`<i>`タグは伝統的にテキストの強調に使われてきましたが、現在はアイコン用のタグとしても採用されています。`<em>`タグは強調用タグとして広く受け入れられています。どちらを使ってもこのチャレンジはクリアできます。

# --hints--

`target4`ボタンのテキストをHTMLタグを追加して強調してください。

```js
assert.isTrue(
  /<em>|<i>\s*#target4\s*<\/em>|<\/i>/gi.test($('#target4').html())
);
```

テキストの他の部分は変更しないでください。

```js
assert($('#target4') && $('#target4').text().trim() === '#target4');
```

他のテキストは一切変更しないでください。

```js
assert.isFalse(/<em>|<i>/gi.test($('h3').html()));
```

`.html()`を使い、`.text()`は使わないでください。

```js
assert(code.match(/\.html\(/g));
```

jQueryで`button id="target4"`を選択してください。

```js
assert(code.match(/\$\(\s*?(\"|\')#target4(\"|\')\s*?\)\.html\(/));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target4").html('<em>#target4</em>');
  });
</script>

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
