---
id: bad87fee1348bd9acdd08826
title: スクリプトタグとドキュメントレディの仕組みを学ぼう
challengeType: 6
forumTopicId: 18224
dashedName: learn-how-script-tags-and-document-ready-work
---

# --description--

これで、史上最も人気のあるJavaScriptツールであるjQueryを学ぶ準備ができました。

jQueryを使い始める前に、HTMLにいくつか追加する必要があります。

まず、ページの一番上に`script`要素を追加してください。次の行で必ず終了タグを閉じてください。

ブラウザーは、jQueryを含む`script`要素内のJavaScriptを実行します。

`script`要素の中に、このコード`$(document).ready(function() {`を`script`に追加してください。そして、次の行で（まだ`script`要素の中で）`});`で閉じてください。

`functions`については後で詳しく学びます。重要なのは、この`function`の中に入れたコードは、ブラウザーがページを読み込んだ直後に実行されるということです。

これは重要です。なぜなら、`document ready function`がなければ、HTMLがレンダリングされる前にコードが実行されてしまい、バグが発生する可能性があるからです。

# --hints--

`script`要素を作成し、有効で終了タグがあることを確認してください。

```js
assert(
  code.match(/<\/script\s*>/g) &&
    code.match(
      /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
    ) &&
    code.match(/<\/script\s*>/g).length ===
      code.match(
        /<script(\sasync|\sdefer)*(\s(charset|src|type)\s*=\s*["\"]+[^"\"]*["\"]+)*(\sasync|\sdefer)*\s*>/g
      ).length
);
```

`$(document).ready(function() {`を`script`要素の先頭に追加してください。

```js
assert(
  code.match(
    /\$\s*?\(\s*?document\s*?\)\.ready\s*?\(\s*?function\s*?\(\s*?\)\s*?\{/g
  )
);
```

`$(document).ready(function() {`関数を`});`で閉じてください。

```js
assert(code.match(/\n*?\s*?\}\s*?\);/g));
```

# --seed--

## --seed-contents--

```html
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
  });
</script>
<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
