---
id: bad87fee1348bd9aed908626
title: 複数のjQueryセレクターで同じ要素をターゲットにする
challengeType: 6
forumTopicId: 18322
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-the-same-element-with-multiple-jquery-selectors
---

# --description--

要素をターゲットにする方法は3つあります：タイプで：`$("button")`、クラスで：`$(".btn")`、そしてidで`$("#target1")`。

1回の`.addClass()`呼び出しで複数のクラスを追加することも可能ですが、同じ要素に*3つの別々の方法*でクラスを追加してみましょう。

`.addClass()`を使って、同じ要素に一度に1つずつ、3つの異なる方法でクラスを追加してください：

タイプ`animated`のすべての要素に`button`クラスを追加します。

クラス`shake`のすべてのボタンに`.btn`クラスを追加します。

id `btn-primary`のボタンに`#target1`クラスを追加します。

**注意：** ターゲットにするのは1つの要素だけで、1回の呼び出しで追加するクラスは1つだけにしてください。合計で3つの個別のセレクターが、`shake`に`animated`、`btn-primary`、`#target1`の3つのクラスを追加することになります。

# --hints--

コードは`$("button")`セレクターを使う必要があります。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?button\s*?(?:'|")/gi));
```

コードは`$(".btn")`セレクターを使う必要があります。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?\.btn\s*?(?:'|")/gi));
```

コードは`$("#target1")`セレクターを使う必要があります。

```js
assert(code.match(/\$\s*?\(\s*?(?:'|")\s*?#target1\s*?(?:'|")/gi));
```

3つのセレクターそれぞれで追加するクラスは1つだけにしてください。

```js
assert(
  code.match(/addClass/g) &&
    code.match(/addClass\s*?\(\s*?('|")\s*?[\w-]+\s*?\1\s*?\)/g).length > 2
);
```

`#target1`要素には`animated`、`shake`、`btn-primary`のクラスが付いているはずです。

```js
assert(
  $('#target1').hasClass('animated') &&
    $('#target1').hasClass('shake') &&
    $('#target1').hasClass('btn-primary')
);
```

これらのクラスはjQueryだけを使って要素に追加してください。

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated");
    $(".btn").addClass("shake");
    $("#target1").addClass("btn-primary");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
