---
id: bad87fee1348bd9bedc08826
title: jQueryを使ってセレクターでHTML要素を指定する
challengeType: 6
forumTopicId: 18319
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-html-elements-with-selectors-using-jquery
---

# --description--

今、`document ready` 関数があります。

では、最初のjQuery文を書いてみましょう。すべてのjQuery関数は`$`で始まります。これは通常ドル記号オペレーター、またはブリングと呼ばれます。

jQueryはよく<dfn>セレクター</dfn>でHTML要素を選択し、その要素に対して何かを行います。

例えば、すべての`button`要素をバウンスさせてみましょう。ドキュメントのready関数内にこのコードを追加してください。

```js
$("button").addClass("animated bounce");
```

jQueryライブラリとAnimate.cssライブラリはすでに背景で読み込まれているので、エディター内で使えます。つまり、jQueryを使ってAnimate.cssの`bounce`クラスを`button`要素に適用しています。

# --hints--

jQueryの`addClass()`関数を使って、`animated`と`bounce`のクラスを`button`要素に付けてください。

```js
assert($('button').hasClass('animated') && $('button').hasClass('bounce'));
```

これらのクラスはjQueryで要素に追加するだけにしてください。

```js
assert(!code.match(/class.*animated/g));
```

jQueryのコードは`$(document).ready();`関数の中に書いてください。

```js
assert(
  code.replace(/\s/g, '').match(/\$\(document\)\.ready\(function\(\)\{\$/g)
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
