---
id: bad87fee1348bd9aedc08826
title: jQueryを使ってクラスで要素を指定する
challengeType: 6
forumTopicId: 18316
required:
  - link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

`button` 要素すべてをバウンスさせたのがわかりますか？`$("button")` でそれらを選択し、`.addClass("animated bounce");` でCSSクラスを追加しました。

jQueryの`.addClass()` 関数を使いましたね。これは要素にクラスを追加できる関数です。

まず、`div` セレクターを使って、クラス`well` を持つ`$(".well")` 要素を指定しましょう。

CSSの宣言と同じように、クラス名の前に`.` を付けることに注意してください。

次に、jQueryの`.addClass()` 関数を使って、クラス`animated` と`shake` を追加します。

例えば、クラス`text-primary` を持つすべての要素を揺らすには、次のコードを`document ready function` に追加できます。

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

jQueryの`addClass()` 関数を使って、クラス`animated` と`shake` をクラス`well` を持つすべての要素に付けてください。

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

これらのクラスはjQueryで要素に追加するだけにしてください。

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
