---
id: 68c15a2a1f0e70ca7154aba5
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、従業員の給与追跡システムを作成します。

まず、`Employee`という名前のクラスを作成してください。その中に、`__init__`、`self`、`name`というパラメータを持つ`level`メソッドを作成します。`__init__`メソッド内で、`name`と`level`を同じ名前のインスタンス属性に代入してください。

# --hints--

`Employee`という名前のクラスが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

`Employee`クラスには、`__init__`、`self`、`name`という3つのパラメータを持つ`level`メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

`__init__`メソッドは`self.name`を`name`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

`__init__`メソッドは`self.level`を`level`に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
