---
id: 68c18b1d0fce081f85312d97
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

前のレッスンで、単一のアンダースコアが接頭辞として付いた属性は慣例として内部使用を意図していることを学びました。

`name` と `level` の両方の属性を `_name` と `_level` に変更してください。これらはクラスの外部から変更されるべきではないためです。

ただし、これは属性がクラスの外部からアクセスまたは変更されるのを防ぐものではありません。また、Pythonでは、プライベート属性（ダブルアンダースコアが接頭辞として付く）にも常にアクセスする方法があります。

# --hints--

`__init__` メソッドは `self._name` を `name` に設定する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

`__init__` メソッドは `self._level` を `level` に設定する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
