---
id: 68c16924498728d259e4781c
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

`Employee` クラスのインスタンスを作成し、文字列 `Charlie Brown` と `trainee` を渡してください。インスタンスは `charlie_brown` という名前の変数に代入してください。

# --hints--

`charlie_brown` という名前の変数があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

`Employee('Charlie Brown', 'trainee')` を `charlie_brown` に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
