---
id: 68c27fa1697ceb7c578b2c9f
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`@property` デコレーターは、Pythonでメソッドをプロパティに変えるために使われます。通常、属性の値を取得するためのゲッターメソッドを定義する際に使われます。

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

`name` という名前のメソッドを `self` パラメータで作成し、`@property` でデコレートしてください。メソッドの中で `self._name` を返してください。
    
# --hints--

`Employee` クラスには `name` パラメータを持つ `self` メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

`name` メソッドは `@property` でデコレートしてください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

`name` メソッドは `self._name` を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
