---
id: 68c28c6517e7de5e85f8ebe5
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`_name` と `_level` の両方の属性に対するゲッターができたので、`__str__` が返す文字列を `self.name` と `self.level` を使うように更新してください。これにより、属性に直接アクセスするのではなくゲッターを呼び出すことになります。

# --hints--

`__str__` メソッドは `f'{self.name}: {self.level}'` を返す必要があります

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
