---
id: 68c2d2414f5e111c02f4db5e
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

`__repr__` メソッドは、`repr` 関数を呼び出したときに内部で呼び出されます。実際に動作を確認するには、コードの最後で `repr(charlie_brown)` を呼び出した結果を出力してください。

# --hints--

`repr(charlie_brown)` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
