---
id: 68c27eb292bfec6b9402d58f
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

`Employee` クラスに `__str__` メソッドを追加してください。`name: level` と `name` を対応する属性に置き換えた形式の `level` の f-string を返すようにしてください。

その後、`charlie_brown` をコンソールに出力してください。

# --hints--

`Employee` クラスには `self` パラメータを持つ `__str__` メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

`__str__` メソッドは `name: level` と `name` の値をそれぞれ `level` と `self._name` に置き換えた形式の `self._level` の f-string を返す必要があります。

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

`charlie_brown` をコンソールに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

--fcc-editable-region--
    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--

```
