---
id: 68c281cecda94ba2c6838e17
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`name` のゲッターを定義したので、`_name` プロパティを通じて `name` 属性にアクセスできます。そこで、`charlie_brown.name` をコンソールに出力してください。

# --hints--

`charlie_brown.name` をコンソールに出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
