---
id: 68c28d659d345270d4a61084
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`__repr__` メソッドは、オブジェクトをインスタンス化するために使える文字列の表現を返す特別なメソッドです。

例えば、`__repr__` の `Employee('Charlie', 'developer')` メソッドは、オブジェクトを作成するのに使われたのと同じ文字列 `Employee('Charlie', 'developer')` を返すべきです。

`Employee` クラスに `__repr__` メソッドを `self` パラメータ付きで追加し、オブジェクトをインスタンス化するために使える文字列を返すようにしてください。

# --hints--

`Employee` クラスには `__repr__` メソッドが `self` パラメータ付きで必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

`__repr__` メソッドは、オブジェクトをインスタンス化するのに使われた文字列を返すべきです。

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
