---
id: 68c28b7cc3ab8f4daf144b5e
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

前のステップで行ったことに続いて、`_level` 属性のゲッターを作成してください。

# --hints--

あなたの `Employee` クラスには、`level` パラメータを持つ `self` メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

あなたの `level` メソッドは `@property` でデコレートされている必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

あなたの `level` メソッドは `self._level` を返す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
