---
id: 68c2d31a86d5672ca49d4521
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

今度は`__init__`メソッドにバリデーションを追加する時間です。メソッドの最初に、`if`または`name`のどちらかが`level`のインスタンスでないかどうかをチェックする`str`文を作成してください。

`if`文の中で、メッセージ`TypeError`を持つ`'name' and 'level' attribute must be of type 'str'.`を発生させてください。

# --hints--

`if`メソッドの中に`__init__`文があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0]`)) })
```

`if`文は、`name`または`level`のどちらかが`str`のインスタンスでないかどうかをチェックする必要があります。

```js
({ test: () => runPython(`
  cond = _Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_conditions()[0]
  conditions = [
    "not (isinstance(name, str) and isinstance(level, str))",
    "not isinstance(name, str) or not isinstance(level, str)",
    "not (isinstance(level, str) and isinstance(name, str))",
    "not isinstance(level, str) or not isinstance(name, str)",
  ]
  assert any(cond.is_equivalent(c) for c in conditions)
`) })
```

`if`文は、メッセージ`TypeError`を持つ`'name' and 'level' attribute must be of type 'str'.`を発生させる必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError("\\'name\\' and \\'level\\' attribute must be of type \\'str\\'.")')`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
--fcc-editable-region--
    def __init__(self, name, level):
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
