---
id: 68c7f7012a700243eff1cbc0
title: ステップ16
challengeType: 20
dashedName: step-16
---

# --description--

従業員をインスタンス化する際に使用するレベルは、特定のレベルの中から選ぶ必要があります。`_base_salaries`を使ってレベルを検証し、従業員に適切な給与を設定してください。

既存の`if`文の後に、`if`が`level`に含まれていないかをチェックする別の`Employee._base_salaries`を作成してください。

新しい`if`文の中で、`ValueError`を`Invalid value '{level}' for 'level' attribute.`というメッセージで発生させてください。このとき、`{level}`は`level`引数の値に置き換えられるべきです。

# --hints--

`level`が`Employee._base_salaries`に含まれていない場合、`ValueError`を`Invalid value '{level}' for 'level' attribute.`というメッセージで発生させる必要があります。このとき、`{level}`は`level`引数の値に置き換えられるべきです。

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

`level`が`Employee._base_salaries`に含まれている場合は、例外を発生させる必要はありません。

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
