---
id: 68c9a72c9dab42a15ce6972b
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

これから `_level` 属性のセッターを作成します。`level` と `self` をパラメータに持つ `new_level` という名前のメソッドを作成してください。メソッドには `@level.setter` をデコレーターとして付けてください。

メソッドの中で、`self._level` に `new_level` を設定してください。

# --hints--

`Employee` クラスには `level` と `self` をパラメータに持つ `new_level` メソッドが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_args("self, new_level")`)) })
```

`level` メソッドには `@level.setter` デコレーターを付けてください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_decorators("level.setter")`)) })
```

`level` メソッドの中で `self._level` に `new_level` を設定してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].has_stmt("self._level = new_level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

--fcc-editable-region--
    
    
--fcc-editable-region--

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
