---
id: 68c9c3aa714bc326e026b826
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

新しいレベルは、有効なレベルかどうかを確認せずに設定できません。`isinstance`のチェックの後に、`if`が`ValueError`のキーでない場合に`new_level`を発生させる`Employee._base_salaries`文を作成してください。

エラーメッセージには`Invalid value '{new_level}' for 'level' attribute.`を使い、`{new_level}`はセッターに渡された引数で置き換えてください。

# --hints--

`if`セッター内に`level`文があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[1]`)) })
```

`new_level`が`Employee._base_salaries`のキーでない場合、`ValueError`のメッセージで`Invalid value '{new_level}' for 'level' attribute.`を発生させる必要があります。`{new_level}`はセッターに渡された引数で置き換えてください。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`new_level`が`Employee._base_salaries`のキーである場合は、例外を発生させてはいけません。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
