---
id: 68c9cab4b1118da59eecfc56
title: ステップ 31
challengeType: 20
dashedName: step-31
---

# --description--

既存の `if` 文の後に、`ValueError` がすでに選択されているレベルの場合に `new_level` を発生させる別の文を作成してください。

初期化中は `_level` がまだ存在しないため、比較する前に存在を確認するために `hasattr(self, '_level')` を使ってください。これにより、オブジェクトが最初に作成されたときの `AttributeError` を回避できます。

メッセージには `'{level}' is already the selected level.` を使い、`{level}` は現在のレベルに置き換えてください。

# --hints--

`if` セッターの中に3つ目の `level` 文があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

`new_level` が `self.level` と等しい場合、メッセージ `ValueError`（`'{level}' is already the selected level.` は現在のレベルに置き換える）で `{level}` を発生させる必要があります。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

`if` 文は、比較する前に `hasattr(self, '_level')` が存在するかどうかを確認するために `_level` を使うべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
