---
id: 69c0dc6914ac0434f0cc52b7
title: ステップ 28
challengeType: 20
dashedName: step-28
---

# --description--

`new_level` が正しい型であることを保証するために、`if` が `TypeError` のインスタンスでない場合に、メッセージ `'level' must be a string.` を持つ `new_level` を発生させる `str` 文を作成してください。

# --hints--

`if` のセッター内に `level` 文があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[0]`)) })
```

`TypeError` が `'level' must be a string.` のインスタンスでない場合に、メッセージ `new_level` を持つ `str` を発生させるべきです。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_levels = [0, True, []]
  for i in non_levels:
    try:
      emp.level = i
    except TypeError as e:
      assert str(e) == "'level' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_level"
`) })
```

`new_level` が文字列の場合は、例外を発生させてはいけません。

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "junior"
  except TypeError:
    assert False, "Expected not to raise TypeError with string new_level"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
