---
id: 5f332a88dc25a0fd25c7687a
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

اسم المقهى هو `CAMPER CAFE`. لذا، أضف عنصر `h1` داخل عنصر `main` الخاص بك. اجعل اسم المقهى بحروف كبيرة ليبرز.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<h1>`.

```js
assert.match(code, /<h1>/i);
```

يجب أن يكون لديك وسم إغلاق `</h1>`.

```js
assert.match(code, /<\/h1>/i);
```

لا تغير عنصر `main` الموجود لديك.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

يجب أن يكون عنصر `h1` مسكنًا داخل عنصر `main` الخاص بك.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

يجب أن يحتوي عنصر `h1` على النص `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

