---
id: 5f33294a6af5e9188dbdb8f3
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرب على أساسيات CSS (أوراق الأنماط المتتالية) من خلال بناء قائمة مقهى.

ابدأ بإضافة بعض محتوى القائمة. أضف عنصر `main` داخل عنصر `body` الموجود. سيحتوي في النهاية على معلومات الأسعار عن القهوة والحلويات التي يقدمها المقهى.

# --hints--

يجب أن تحتوي شفرتك على وسم افتتاحي `<main>`.

```js
assert.match(code, /<main>/i);
```

يجب أن تحتوي شفرتك على وسم إغلاق `</main>`.

```js
assert.match(code, /<\/main>/i);
```

لا تغير عنصر `body` الخاص بك. تأكد من عدم حذف وسم الإغلاق عن طريق الخطأ.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

يجب أن يكون وسم `main` داخل وسم `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
