---
id: 5f3477cb2e27333b1ab2b955
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

الآن عليك ربط ملف `styles.css`، لكي تُطبَّق الأنماط مرة أخرى. داخل عنصر `head`، أضف عنصر `link`. أعطه خاصية `rel` بقيمة `"stylesheet"` وخاصية `href` بقيمة `"styles.css"`.

**ملاحظة:** عنصر `link` هو عنصر فارغ، مما يعني أنه لا يحتوي على وسم إغلاق. يجب كتابة العناصر الفارغة كـ `<link>` بدلاً من `<link></link>`.

# --hints--

يجب أن يحتوي شفرتك على عنصر `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

عنصر `link` هو عنصر فارغ، فلا يجب أن يحتوي على وسم إغلاق `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

لا تغير عنصر `head` الموجود لديك. تأكد من أنك لم تحذف وسم الإغلاق.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

يجب أن يكون لديك عنصر `link` واحد.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
