---
id: 5f344fc1520b6719f2e35605
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

سيكون هناك قسمان في القائمة، واحد للقهوة وآخر للحلويات. أضف عنصر `section` داخل عنصر `main` لتوفر مكانًا لوضع كل أنواع القهوة المتاحة.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

يجب أن يكون لديك وسم إغلاق `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

لا تغير عنصر `main` الموجود لديك. تأكد من أنك لم تحذف وسم الإغلاق.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

يجب أن يكون عنصر `section` داخل عنصر `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
