---
id: 5f344f9c805cd193c33d829c
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

في الدروس السابقة، تعلمت كيف تضيف خصائص وقيم `CSS` مثل هذه:

```css
element {
 property: value;
}
```

وسّط محتوى عنصر `h1` بتعيين خاصية `text-align` إلى القيمة `center`.

# --hints--

يجب أن يكون لديك محدِّد `h1` في عنصر `style` الخاص بك.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

يجب أن تكون خاصية `text-align` لديك بقيمة `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

يجب أن يعيّن محدِّد `h1` الخاص بك خاصية `text-align` إلى `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

