---
id: 5f344fbc22624a2976425065
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أنشئ عنصر `h2` داخل عنصر `section` وأعطه النص `Coffee`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<h2>`.

```js
assert.match(code, /<h2\s*>/i);
```

يجب أن يكون لديك وسم إغلاق `</h2>`.

```js
assert.match(code, /<\/h2\s*>/i);
```

لا تغير عنصر `section` الموجود لديك. تأكد من أنك لم تحذف وسم الإغلاق.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

يجب أن يكون عنصر `h2` داخل عنصر `section`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

يجب أن يحتوي عنصر `h2` على النص `Coffee`.

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
