---
id: 5f3477cbcb6ba47918c1da92
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

لكي يظهر تنسيق الصفحة مشابهًا على الهاتف المحمول كما هو على سطح المكتب أو الحاسوب المحمول، تحتاج إلى إضافة عنصر `meta` مع خاصية `content` خاصة.

تعلمت عن عنصر `meta` الخاص بإطار العرض في الدروس السابقة.

إليك مثالًا:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

يجب أن يحتوي شفرتك على عنصرين `meta`.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

يجب أن يحتوي عنصر `meta` الخاص بك على خاصية `name` بقيمة `viewport`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

يجب أن يحتوي عنصر `meta` الخاص بك على خاصية `content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
