---
id: 5f344fad8bf01691e71a30eb
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

حتى الآن، كان تحكمك محدودًا في عرض ومظهر المحتوى الخاص بك. لتغيير ذلك، أضف عنصر `style` داخل عنصر `head`.

# --hints--

يجب أن تحتوي شفرتك على وسم افتتاحي `<style>`.

```js
assert.match(code, /<style\s*>/i);
```

يجب أن تحتوي شفرتك على وسم إغلاق `</style>`.

```js
assert.match(code, /<\/style\s*>/);
```

يجب أن يكون عنصر `style` مسكنًا داخل عنصر `head`.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
