---
id: 5f3477aefa51bfc29327200b
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

لقد نسّقت ثلاثة عناصر بكتابة CSS داخل وسم `style`. هذا يعمل، لكن بما أن هناك المزيد من التنسيقات، من الأفضل وضع كل التنسيقات في ملف منفصل وربطه.

تم إنشاء ملف `styles.css` منفصل لك. يمكنك التنقل بين الملفات باستخدام علامات التبويب في أعلى المحرر.

ابدأ بإعادة كتابة التنسيقات التي أنشأتها في ملف `styles.css`. تأكد من استبعاد وسمَي `style` الافتتاحي والختامي.

# --hints--

يجب أن يحتوي ملف `styles.css` على محدِّد النوع `h1, h2, p`.

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```
 
يجب أن يضبط المحدِّد الخاص بك خاصية `text-align` إلى `center`.

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

يجب ألا تحتوي شفرتك على وسمَي `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
