---
id: 5f35e5c4321f818cdc4bed30
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

الآن بعد أن أصبح كل شيء جيدًا، حان الوقت لبدء إضافة بعض عناصر القائمة.

أضف عنصر `article` فارغًا تحت عنوان `Coffee`. سيحتوي على نكهة وسعر كل نوع قهوة تقدمها حاليًا.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<article>`.

```js
assert.match(code,/<article>/i);
```

يجب أن يكون لديك وسم إغلاق `</article>`.

```js
assert.match(code, /<\/article>/i);
```

لا يجب أن تغيّر عنصر `h2` الموجود.

```js
assert.lengthOf(document.querySelectorAll('h2'), 1);
```

يجب أن يأتي عنصر `article` بعد عنصر `h2`.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

