---
id: 5f356ed60785e1f3e9850b6e
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

الآن من السهل رؤية أن النص في وسط عنصر `#menu`. حاليًا، عرض عنصر `#menu` محدد بوحدة البيكسل (`px`).

غيّر قيمة خاصية `width` لتصبح `80%`، لجعلها تمثل `80%` من عرض العنصر الأصل (`body`).

# --hints--

يجب تعيين خاصية `width` إلى `80%`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert.isTrue(hasWidth);
```

يجب ألا تحتوي على خاصية `width` بقيمة `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isFalse(hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
--fcc-editable-region--
  width: 300px;
--fcc-editable-region--
  background-color: burlywood;
}
```

