---
id: 5f356ed63c7807a4f1e6d054
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

الهدف الآن هو جعل `div` لا يشغل العرض الكامل للصفحة. خاصية `width` في CSS مثالية لهذا الغرض.

يمكنك استخدام محدِّد `id` لاستهداف عنصر معين يحمل خاصية `id`.

تعلمت كيفية العمل مع محدِّد `id` في دروس سابقة مثل هذا:

```css 
#cat {
  width: 250px;
}
```

استخدم محدِّد `#menu` لإعطاء العنصر عرضًا قدره `300px`.

# --hints--

يجب أن يكون لديك محدِّد `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert.exists(hasDiv);
```

يجب أن تضبط خاصية `width` على `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isTrue(hasWidth);
```

يجب أن يكون عرض `div` الخاص بك 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert.equal(divWidth, '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```

