---
id: 5f356ed656a336993abd9f7c
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

بعد ذلك، تريد توسيط العنصر `#menu` أفقيًا. يمكنك فعل ذلك بتعيين خاصيتي `margin-left` و`margin-right` إلى القيمة `auto`. فكر في الهامش كمساحة غير مرئية حول العنصر. باستخدام هاتين الخاصيتين، وسّط العنصر `#menu` داخل العنصر `body`.

# --hints--

يجب تعيين الخاصية `margin-left` إلى القيمة `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert.isTrue(hasMargin);
```

يجب تعيين الخاصية `margin-right` إلى القيمة `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert.isTrue(hasMargin);
```

يجب تعيين الخاصيتين `margin-left` و`margin-right` للعنصر `#menu` إلى القيمة `auto`.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert.equal(divMarginRight, 'auto');
assert.equal(divMarginLeft, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 80%;
  background-color: burlywood;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
