---
id: 5f716bee5838c354c728a7c5
title: خطوة 48
challengeType: 0
dashedName: step-48
---

# --description--

أسفل الحلوى التي أضفتها للتو، أضف بقية الحلويات والأسعار باستخدام ثلاثة عناصر `article` إضافية، كل منها يحتوي على عنصرين `p` مسكنين. يجب أن يحتوي كل عنصر على نص الحلوى والسعر الصحيحين، ويجب أن تحتوي جميعها على الأصناف الصحيحة.

```md
Cherry Pie 2.75
Cheesecake 3.00
Cinnamon Roll 2.50
```

# --before-each--

```js
const section = [...document.querySelectorAll('section')].at(-1);
const paragraphs = section?.querySelectorAll('p');
const prices = section?.querySelectorAll('p.price');
const desserts = section?.querySelectorAll('p.dessert');
```

# --hints--

يجب أن يكون لديك أربعة عناصر `article` في `section` الثانية.

```js
const articles = section?.querySelectorAll('article');
assert.lengthOf(articles, 4);
```

يجب أن يكون لديك أربعة عناصر `.dessert`.

```js
assert.lengthOf(desserts, 4);
```

يجب أن يكون لديك أربعة عناصر `.price` جديدة.

```js
assert.lengthOf(prices, 4);
```

يجب أن يحتوي العنصر الثاني من نوع `section` على ثمانية عناصر من نوع `p`.

```js
assert.lengthOf(paragraphs, 8);
```

يجب أن تحتوي عناصر `.dessert` على النصوص `Donut`، `Cherry Pie`، `Cheesecake`، و`Cinnamon Roll`.

```js
assert.match(desserts?.[0]?.innerText.trim(), /donut/i);
assert.match(desserts?.[1]?.innerText.trim(), /cherry pie/i);
assert.match(desserts?.[2]?.innerText.trim(), /cheesecake/i);
assert.match(desserts?.[3]?.innerText.trim(), /cinnamon roll/i);
```

يجب أن تحتوي عناصر `.price` الجديدة على النصوص `1.50`، `2.75`، `3.00`، و`2.50`.

```js
assert.match(prices?.[0]?.innerText.trim(), /1\.50/);
assert.match(prices?.[1]?.innerText.trim(), /2\.75/);
assert.match(prices?.[2]?.innerText.trim(), /3\.00/);
assert.match(prices?.[3]?.innerText.trim(), /2\.50/);
```

يجب ألا توجد أي مسافات بين عناصر `p` الخاصة بك.

```js
assert.lengthOf(document.querySelectorAll('.item'), 9);
document.querySelectorAll('.item').forEach(item => {
  assert.lengthOf(item.children, 2);
  const [first, second] = item.children;
  assert.strictEqual(first.nextSibling, second);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%;
}
```
