---
id: 65c4f300da28d8da361bfa93
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Dentro de tu cláusula `elif`, reemplaza `pass` con una llamada al método `_insert` usando el hijo derecho del nodo actual como primer argumento y `key` como segundo argumento. Asigna el resultado al atributo `right` del nodo actual.

# --hints--

Debes eliminar la palabra clave `pass` del bloque `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Debes llamar al método `self._insert()` pasando `node.right` y `key` como argumentos.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Debes asignar el resultado de tu llamada a `self._insert()` al atributo derecho (`node.right`) del nodo actual.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
