---
id: 65c4f37ba6cbcfdc77e24165
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

El método `insert` será llamado por el usuario. Además del parámetro `self`, también necesitará un parámetro `key`. Este parámetro será el valor clave que se insertará en el árbol binario de búsqueda.

Agrega `key` como el segundo parámetro en la definición de la función.

# --hints--

El método `insert` debe contener el parámetro `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

El método `insert` debe contener el parámetro `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
