---
id: 65c4f195de7e2ad5932be717
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

A continuación, necesitas definir un mecanismo para insertar nodos en el árbol. Para eso, necesitas definir un método `_insert`, que es una función auxiliar y será usado por el método real `insert` más adelante. 

Este método es recursivo, lo que significa que se llama a sí mismo para recorrer el árbol hasta que se encuentra la ubicación adecuada para el nuevo nodo.

Define un método `_insert` con los parámetros `self`, `node` y `key`.

# --hints--

Debes definir un método `_insert` dentro de la clase `BinarySearchTree`. Recuerda usar `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Tu método `_insert` debe tomar tres parámetros: `self`, `node` y `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
