---
id: 65c4f3258d2e4cdacc919dfd
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Al final de tu método `_insert`, después de que el proceso de inserción esté completo, devuelve el nodo actual para actualizar la estructura del árbol en los niveles superiores de la pila de llamadas recursivas.


# --hints--

Debes devolver el nodo actual fuera de los bloques condicionales.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```

