---
id: 65c4f09e074dd8d37830ea00
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Crea otra clase vacía llamada `BinarySearchTree` que representa un árbol binario de búsqueda.

# --hints--

Debes tener una clase llamada `BinarySearchTree`. Recuerda usar la palabra clave `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
