---
id: 65c4f0b82db314d3fcc0b8e1
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Dentro de la clase `BinarySearchTree`, reemplaza `pass` con un método `__init__` y añade un parámetro `self` a este método.

# --hints--

Debes eliminar la palabra clave `pass` de la clase `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Debes definir un método `__init__` y agregar un parámetro `self` a este método. Recuerda usar la palabra clave `pass` dentro del cuerpo del método.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
