---
id: 65c4f22498d22ed775ef8efb
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Ahora necesitas verificar si el parámetro `node` es `None`. Si lo es, esto significa que el método ha llegado a un nodo hoja o a un lugar vacío en el árbol donde se debe insertar el nuevo nodo.

Dentro del cuerpo del método `_insert`, reemplaza `pass` con una estructura `if` que verifique si `node is None`.

Dentro del nuevo bloque `if`, retorna `TreeNode(key)` para crear una nueva instancia de `TreeNode` con la clave proporcionada. Esto se convertirá en el nuevo nodo hoja, insertando efectivamente la clave en el árbol.

# --hints--

No debes tener `pass` en tu método `_insert`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Debes escribir una estructura `if` para verificar si `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Debes retornar `TreeNode(key)` desde el bloque `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
