---
id: 65c4f2851645e8d84d14f5d1
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora necesitas recorrer recursivamente el árbol e insertar los valores usando el principio para árboles binarios:

- Los valores menores que la clave se colocan en el subárbol izquierdo
- Los valores mayores que la clave se colocan en el subárbol derecho

Después de tu estructura condicional existente, escribe otra estructura `if` para verificar si `key` es menor que `node.key`.

# --hints--

Debes escribir otra estructura if para verificar si `key` es menor que `node.key`. Recuerda usar la palabra clave `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

