---
id: 65c4f2b7178afed8e88f782f
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Si `key` < `node.key` evalúa a `True`, entonces el nuevo nodo debe colocarse en el subárbol izquierdo.

Elimina `pass` y llama recursivamente al método `_insert` con el hijo izquierdo como primer argumento y `key` como segundo argumento. Asigna el resultado al atributo `left` del nodo actual.

# --hints--

Debes eliminar la palabra clave `pass` del bloque `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Debes llamar al método `self._insert` pasando `node.left` y `key` como argumentos.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Debes asignar el resultado de tu llamada a `self._insert()` al atributo `left` del nodo actual.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
