---
id: 65c4eb814cc977c95cd7df0e
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, vas a crear un Árbol Binario de Búsqueda (BST). Un BST es una estructura de datos en la que cada nodo tiene como máximo dos hijos, con el hijo izquierdo que contiene valores menores que el nodo padre y el hijo derecho que contiene valores mayores que el nodo padre, lo que permite operaciones eficientes de búsqueda y ordenamiento.

Así es como se ve un Binary Search Tree:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Comienza definiendo una clase vacía `TreeNode`. La clase `TreeNode` representa un nodo en un árbol binario de búsqueda. Usa la palabra clave `pass` para completar el cuerpo de la clase y evitar un error.

# --hints--

Debes usar la palabra clave `class` para declarar una clase vacía llamada `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
