---
id: 65c4f0d32ca17ad4b1636b0e
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Dentro del método `__init__`, elimina `pass` e inicializa `root` con el valor `None`. 

El atributo `root` representa el nodo raíz del árbol binario de búsqueda. Dado que este es el constructor cuando se crea un nuevo objeto `BinarySearchTree`, comienza con un árbol vacío, por lo que el atributo `root` se establece en `None`.

# --hints--

Debes eliminar la sentencia `pass` del método `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Debes inicializar el atributo `root` a `None` usando `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
