---
id: 65c4f02dbd3557d23f12db62
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Dentro del método `__init__`, elimina `pass` y asigna el valor del parámetro `key` al atributo `key` del nodo usando `self.key`.

Esto significa que el atributo `key` de la instancia `TreeNode` se establecerá con el valor pasado durante la creación del objeto.

# --hints--

Debes eliminar la sentencia `pass` del método `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Debes asignar el valor del parámetro `key` al atributo `key` del nodo usando `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
