---
id: 65c4f013851cefd1a4fe4c96
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

El método `__init__` toma dos parámetros: `self` (que representa la instancia de la clase que se está creando) y `key` (el valor que se almacenará en el nodo). Agrega esos dos parámetros al método `__init__()`.

# --hints--

Debes agregar el parámetro `self` al método.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Debes agregar el parámetro `key` al método.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
