---
id: 65c4ef8463f869d0eaf87c67
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro de la clase `TreeNode`, reemplaza `pass` con un método `__init__` para que puedas inicializar los atributos del objeto. No agregues ningún parámetro por ahora.

# --hints--

Debes eliminar la palabra clave `pass` de la clase `TreeNode` y moverla dentro del método `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Debes definir un método `__init__` dentro de tu clase `TreeNode`. Recuerda usar la palabra clave `pass` dentro.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
