---
id: 65c4f07da204f4d2f325fbd3
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro del método `__init__`, inicializa los atributos `left` y `right` del nodo a `None`. Esto es porque cuando un nodo se crea por primera vez, no tiene hijos izquierdos ni derechos. Recuerda usar la palabra clave `self`. 

# --hints--

Debes asignar `None` al atributo `left` del nodo usando `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Debes asignar `None` al atributo `right` del nodo usando `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
