---
id: 65c63a4c4da62e9ae18e321a
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Es hora de trabajar en la funcionalidad de búsqueda. Así como creaste un método auxiliar `_insert` para el método `insert`, necesitas crear un método auxiliar `_search` para el método `search`.

Define el método `_search` con tres parámetros, a saber, `self`, `node` y `key`.

# --hints--

Debes declarar un método `_search`. Nota el guion bajo inicial y recuerda usar `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Tu método `_search` debe tomar tres parámetros: `self`, `node` y `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
