---
id: 65c6409418806da8d0636ffc
title: Paso 26
challengeType: 20
dashedName: step-26
---

# --description--

Las funcionalidades `insert` y `search` están completas, es hora de probarlas.

Ten en cuenta que, en este punto, los nodos no están ordenados y solo se insertan. Más adelante trabajarás en el ordenamiento usando el recorrido inorder.

Puedes crear una instancia de una clase en Python así:

```python
object_name = ClassName()
```

Fuera de las definiciones de clase, crea una instancia de la clase `BinarySearchTree` y asígnala a la variable `bst`.


# --hints--

Debes crear una instancia de la clase `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Debes asignar la nueva instancia de `BinarySearchTree` a la variable `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
