---
id: 65c646d4148ae3b2d1cbcac4
title: Paso 30
challengeType: 20
dashedName: step-30
---

# --description--

Ten en cuenta que tu búsqueda devuelve algo como `80: <__main__.TreeNode object at 0x108b3e0>`. Esta es la representación de cadena predeterminada al imprimir una instancia de una clase.

Para cambiar eso y mostrar un valor útil, define otro método llamado `__str__` en la clase `TreeNode`. Este recibe un solo argumento `self`.

Después de definir `__str__` obtendrás una excepción en la consola porque el método `__str__` aún no devuelve nada. Trabajarás en el cuerpo del método en el siguiente paso.

# --hints--

Debes definir un método `__str__` que tome un solo argumento `self`. Recuerda usar `pass`.

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
