---
id: 65c63e6962159fa372ecdbec
title: Paso 22
challengeType: 20
dashedName: step-22
---

# --description--

Si la segunda sentencia `if` no es `True`, significa que la clave objetivo es mayor o igual que la clave del nodo actual. 

En un árbol binario de búsqueda, si la clave objetivo es mayor que la clave del nodo actual, la búsqueda continúa en el subárbol derecho.

Después del bloque `if`, devuelve el resultado de llamar al método `_search` con el hijo derecho del nodo actual y la clave como argumentos.

# --hints--

Debes llamar recursivamente al método `_search` usando `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Debes pasar dos argumentos a la llamada del método `_search`: `node.right` y `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Debes devolver el resultado de la llamada al método `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
