---
id: 65c63da2ed6769a10e141341
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Ahora vas a definir un caso base para la búsqueda recursiva. Elimina el `pass` actual y escribe una sentencia `if` que verifique dos condiciones:

- Si `node` es `None`: Esto indica que la búsqueda ha llegado al final de una rama sin encontrar la clave.
- Si `node.key == key`: Esto significa que la clave ha sido encontrada en el nodo actual.

Combina las dos condiciones con el operador `or` y devuelve el nodo actual dentro del bloque `if`.

# --hints--

Debes eliminar la palabra clave `pass` del método `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Debes escribir una estructura `if` que verifique si `node` es `None` o si `node.key` es igual a `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Debes retornar `node` desde el bloque `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
