---
id: 65c63fd9b387f0a6c5bf6a72
title: Pasos 24
challengeType: 20
dashedName: step-24
---

# --description--

Dentro del método `search`, elimina `pass` y llama al método auxiliar `_search` con los siguientes argumentos.

- `self.root`: Este es el nodo raíz del árbol binario de búsqueda. La búsqueda comienza desde la raíz.
- `key`: Este es el valor que el usuario quiere encontrar en el árbol binario de búsqueda.

Internamente, `search` delega la lógica de búsqueda al método auxiliar `_search`, que realiza la búsqueda recursiva dentro del árbol binario de búsqueda.

# --hints--

Debes eliminar la palabra clave `pass` del método `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Debes llamar al método `_search` dentro del método `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Debes llamar al método `\_search` pasando `self.root` y `key` como argumentos.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
