---
id: 65c63df529bd15a24c187c62
title: Paso 21
challengeType: 20
dashedName: step-21
---

# --description--

Escribe otra estructura `if` que verifique si la clave objetivo es menor que la clave del nodo actual.

Dentro del bloque `if`, devuelve el resultado de llamar al método `_search` con el hijo izquierdo del nodo actual y `key` como argumentos.

# --hints--

Debes escribir otra estructura `if` para verificar si `key` < `node.key`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

Debes retornar `self._search(node.left, key)` desde tu nuevo bloque `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
