---
id: 66d75dd0aa65a71600dc669b
title: 构建一个库存管理程序
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

在本实验中，你将构建一个库存管理程序，允许你添加、更新、查找和移除库存中的产品。你将使用一个对象数组来表现你的库存，其中每个对象将包含 name 和 quantity 作为密钥。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该声明一个名为 `inventory` 的空数组，用于保存具有键 `name` 且值为小写字串，以及键 `quantity` 且值为整数的产品对象。
1. 你应该创建一个名为 `findProductIndex` 的函数，该函数以产品名称作为其参数，并返回对应产品对象在 `inventory` 数组中的索引。该函数应始终使用产品名称的小写形式来执行查找。如果未找到产品，函数应返回 `-1`。
1. 你应该创建一个名为 `addProduct` 的函数，该函数以产品对象作为其参数。
1. 如果产品已经存在于库存中，`addProduct` 函数应通过将传入函数的数量添加到当前的 `quantity` 来更新其 `quantity` 值，并在控制台记录产品名称，后跟一个空格和 `quantity updated`。
1. 如果产品不在库存中，`addProduct` 函数应将产品推送到 `inventory` 数组，并将产品名称记录到控制台，后跟一个空格和 `added to inventory`。
1. 你应该创建一个名为 `removeProduct` 的函数，该函数以产品名称和数量作为其参数。
1. `removeProduct` 函数应从库存中对应的产品对象中减去传入的数量，并将字串 `Remaining <product-name> pieces: <product-quantity>` 记录到控制台，其中 `<product-name>` 应替换为产品名称，`<product-quantity>` 应替换为产品数量。
1. 如果减法后的数量为零，`removeProduct` 应该从库存中移除该产品对象。如果库存中的数量不足以执行减法，`removeProduct` 函数应将字串 `Not enough <product-name> available, remaining pieces: <product-quantity>` 记录到控制台日志。
1. 如果要移除的产品不在库存中，`removeProduct` 函数应在控制台记录 `<product-name> not found`。

**注意：** 为防止冲突，运行测试时仅保留用户故事中提到的日志。

# --hints--

你应该声明一个名为 `inventory` 的空数组。

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

你应该有一个名为 `findProductIndex` 的函数。

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` 应返回 `inventory` 数组中 `name` 相等于 `"flour"` 的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` 应该返回 `inventory` 数组中 `name` 与 `"flour"` 相等的对象的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

当在 `inventory` 数组中未找到 `name` 与 `"flour"` 相等的对象时，`findProductIndex("Flour")` 应返回 `-1`。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

你应该有一个名为 `addProduct` 的函数。

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` 应该为 `inventory` 数组中 `name` 等于 `"flour"` 的对象的 `quantity` 值添加 `5`。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` 应该在 `inventory` 数组中找到 `name` 与 `"flour"` 相等的对象时记录 `flour quantity updated` 日志。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` 应该将 `{name: "flour", quantity: 5}` 添加到 `inventory` 数组中，当库存中没有 `name` 相等于 `"flour"` 的对象时。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` 应该在库存中未找到 `name` 相等于 `"flour"` 的对象时记录日志 `flour added to inventory`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

你应该创建一个名为 `removeProduct` 的函数。

```js
assert.isFunction(removeProduct);
```

当在 `inventory` 数组中未找到 `name` 与 `"flour"` 相等的对象时，`removeProduct("FLOUR", 5)` 应记录 `flour not found` 日志。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` 应该从 `inventory` 数组中 `name` 等于 `"flour"` 的对象的 `quantity` 值中减去 `5`。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` 应该在 `inventory` 等于 `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]` 时向控制台记录 `Remaining flour pieces: 15`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

如果减法后的 `quantity` 为零，`removeProduct` 应该从库存中移除该产品对象。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` 应该在 `inventory` 等于 `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]` 时记录 `Not enough flour available, remaining pieces: 5` 日志。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
