---
id: 686b9720ee1d032bd77a480a
title: 调试 `ISBN Validator`
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

ISBN（国际标准书号）是分配给商业图书的唯一标识符。它可以是 10 位或 13 位数字，最后一位是根据其他数字计算得出的校验位。

Camperbot 试图构建他们自己的 ISBN 验证器。然而，他们在过程中犯了一些错误。

在本实验中，你将修复现有代码并使其正常运行。

预期行为： 当用户运行程序时，会显示提示符 `Enter ISBN and length: `。用户可以以 `ISBN,length` 形式输入他们想要验证的 ISBN 代码。ISBN 代码不应包含连字符，后跟其长度（`10` 或 `13`），两者之间用逗号分隔。

示例输入：  
`1530051126,10` 用于 ISBN-10 代码。  
`9781530051120,13` 用于 ISBN-13 代码。

如何查找校验位：

- 你不必了解此实验中的详细计算逻辑。函数 `calculate_check_digit_10` 和 `calculate_check_digit_13` 会负责计算并返回预期的校验位字串。
- ISBN-10 代码的校验位可以是从 `0` 到 `9` 的数字或大写字母 `X`。
- ISBN-13 代码的校验位可以是从 `0` 到 `9` 的数字。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该修复当前代码中的 `IndentationError`。
1. 即使用户未输入逗号分隔值，程序也应处理 `IndexError` 而不会崩溃。
1. 当用户未输入逗号分隔值时，他们应在控制台看到消息 `Enter comma-separated values.`，并且程序应终止。
1. 即使用户输入非数字的长度值，程序也应处理 `ValueError` 而不会崩溃。
1. 当用户输入非数字的长度值时，他们应在控制台看到消息 `Length must be a number.`，并且程序应终止。
1. 你应该修复 `validate_isbn` 函数中的一位偏差误差。
1. 你应该修复当前代码中当用户输入有效 ISBN 代码时发生的 `TypeError`。
1. 当用户输入有效的 ISBN 代码时，你应该修复当前代码中的 `IndexError`。
1. 即使用户输入了除数字以外的字符组成的错误 ISBN 代码，程序也应处理发生的 `ValueError`，而不会崩溃。
1. 当用户输入包含数字以外字符的错误 ISBN 代码时，他们应在控制台看到消息 `Invalid character was found.`。
1. 当用户输入 `1530051126,10` 时，他们应该看到消息 `Valid ISBN Code.`
1. 当用户输入 `9781530051120,13` 时，他们应该看到消息 `Valid ISBN Code.`

**重要**：你需要在全局空间中注释掉 `main()` 调用，测试才能正确运行。

当你完成该项目时，用户应根据他们输入的值看到以下消息。

| ISBN Code | Length | Message | Example input |
|-----------|--------|---------| ------------- |
| Valid | Valid | `Valid ISBN Code.` | `1530051126,10` |
| Invalid Number | Valid | `Invalid ISBN Code.` | `1530051125,10` |
| Does not match specified length or left blank | Valid | `ISBN-10 code should be 10 digits long.` or `ISBN-13 code should be 13 digits long.`, depending on the length they entered. | `9781530051120,10` or `1530051126,13` |
| Contains non-numeric characters (except for the check digit) | Valid | `Invalid character was found.` | `15-0051126,10` |
| Any | Invalid Number | `Length should be 10 or 13.` | `1530051126,9` |
| Any | Contains non-numeric characters or left blank | `Length must be a number.` | `1530051125,A` |
| Not comma-separated | Not comma-separated | `Enter comma-separated values.` | `1530051125` |

如果你愿意，可以使用以下值手动测试你的代码。

有效 ISBN-10 代码的示例输入：  
`1530051126,10`  
`9971502100,10`  
`080442957X,10`

有效 ISBN-13 代码的示例输入：  
`9781530051120,13`  
`9781947172104,13`

# --hints--

你应该注释掉对 `main` 函数的调用，以使其余的测试能够正常工作。

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

你应该有一个 `validate_isbn` 函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

你应该有一个 `calculate_check_digit_10` 函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

你应该有一个 `calculate_check_digit_13` 函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

你应该有一个 `main` 函数。

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

当用户输入的值不是逗号分隔值时，你应该在控制台看到消息 `Enter comma-separated values.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

当用户为长度输入非数字值时，你应该在控制台看到消息 `Length must be a number.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

当用户输入包含数字以外字符的错误 ISBN 代码时，你应该在控制台看到消息 `Invalid character was found.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

当用户输入 `1530051126,10` 时，你应该在控制台看到消息 `Valid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入 `9781530051120,13` 时，你应该看到消息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入 `1530051125,10` 时，你应该看到消息 `Invalid ISBN Code.`。

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

当用户输入 `9781530051120,10` 时，你应该看到消息 `ISBN-10 code should be 10 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

当用户输入 `1530051126,13` 时，你应该看到消息 `ISBN-13 code should be 13 digits long.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

当用户输入 `15-0051126,10` 时，你应该看到消息 `Invalid character was found.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

当用户输入 `1530051126,9` 时，你应该看到消息 `Length should be 10 or 13.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

当用户输入 `1530051125,A` 时，你应该看到消息 `Length must be a number.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

当用户输入 `1530051125` 时，你应该看到消息 `Enter comma-separated values.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

当用户输入 `9971502100,10` 时，你应该看到消息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入 `080442957X,10` 时，你应该看到消息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

当用户输入 `9781947172104,13` 时，你应该看到消息 `Valid ISBN Code.`

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
