---
id: 66c06d618d075c7f7f1b890a
title: بناء عرّاف
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

في هذا المختبر، ستُعرّف خمسة نبوءات وتختار واحدة منها عشوائيًا لعرض النبوءة على المستخدم.

**الهدف:** أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تهيئ المتغيرات الخمسة `fortune1`، `fortune2`، `fortune3`، `fortune4`، و`fortune5` بقيمة سلسلة نصية من اختيارك. يمكنك استخدام الخيارات أدناه إذا أردت:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. يجب أن تختار رقمًا عشوائيًا بين 1 و5، شاملًا، وتُعيّنه إلى المتغير `randomNumber`.

3. يجب أن تنشئ متغير `selectedFortune` وتُعيّن النبوءة المناسبة بناءً على القواعد التالية:

- إذا كان `randomNumber` يساوي 1، عيّن قيمة `fortune1` إلى `selectedFortune`.
   - إذا كان `randomNumber` يساوي 2، عيّن قيمة `fortune2` إلى `selectedFortune`.
   - إذا كان `randomNumber` يساوي 3، عيّن قيمة `fortune3` إلى `selectedFortune`.
   - إذا كان `randomNumber` يساوي 4، عيّن قيمة `fortune4` إلى `selectedFortune`.
   - إذا كان `randomNumber` يساوي 5، عيّن قيمة `fortune5` إلى `selectedFortune`.

4. يجب أن تسجل `selectedFortune` في وحدة التحكم.

# --hints--

يجب أن تهيئ `fortune1` بقيمة سلسلة نصية.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

يجب أن تهيئ `fortune2` بقيمة سلسلة نصية.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

يجب أن تهيئ `fortune3` بقيمة سلسلة نصية.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

يجب أن تهيئ `fortune4` بقيمة سلسلة نصية.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

يجب أن تهيئ `fortune5` بقيمة سلسلة نصية.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

يجب أن تستخدم الطريقة `Math.random()` لتوليد رقم عشوائي.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

يجب أن تولد رقمًا عشوائيًا بين 1 و5، شاملًا، وتُعيّنه إلى المتغير `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

يجب أن يكون لديك متغير `selectedFortune` يُعيّن له قيمة بناءً على قيمة `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

يجب أن يتوافق `randomNumber` مع نبوءته. على سبيل المثال، إذا كان `randomNumber` يساوي 1، يجب أن تكون قيمة `selectedFortune` مساوية لـ `fortune1` وهكذا.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

يجب أن تُخرج `selectedFortune` إلى وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
