---
id: 66be24cb4144f955b6bcc550
title: Створіть сторінку гральних карт
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Вам потрібно створити вебсторінку з принаймні трьома гральними картами.
1. Ви повинні мати елемент `main` з `id` зі значенням `playing-cards`.
1. В межах елемента `#playing-cards` має бути принаймні три елементи `div`, кожен з яких має клас `card`.
1. В межах кожного елемента `.card` має бути принаймні три елементи `div`: перший з класом `left`, другий з класом `middle` і третій з класом `right`.
1. Елемент `#playing-cards` має використовувати flexbox, щоб горизонтально відцентрувати дочірні елементи, дозволити їм переноситись на інший рядок і додати `20px` між ними.
1. Елементи `.card` мають використовувати flexbox, щоб рівномірно розподілити дочірні елементи за допомогою `space-between`.
1. Елементи `.left` мають використовувати flexbox-властивості, щоб бути вирівняними за початком батьківського контейнера.
1. Елементи `.middle` мають використовувати flexbox-властивості, щоб бути вирівняними за центром батьківського контейнера.
1. Елементи `.right` мають використовувати flexbox-властивості, щоб бути вирівняними за кінцем батьківського контейнера.
1. Елементи `.middle` мають використовувати flexbox, щоб дочірні елементи утворювали стовпець.

Ось декілька символів, які за бажанням можна скопіювати і вставити на картах: `♠`, `♣`, `♥`, `♦`.

# --hints--

Ви повинні мати елемент `main` з `id` зі значенням `playing-cards`.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

В межах елемента `#playing-cards` має бути принаймні три елементи `div` з класом `card`.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

Всі елементи `.card` повинні мати ширину і висоту.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

Кожен елемент `.card` повинен мати три дочірні елементи `div`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

В межах кожного елемента `.card` має бути `div` з класом `left`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

В межах кожного елемента `.card` має бути `div` з класом `middle`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

В межах кожного елемента `.card` має бути `div` з класом `right`.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

Ви повинні мати селектор `#playing-cards`, який встановлює властивість `display` своїх елементів на `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

Ви повинні мати селектор `#playing-cards`, який встановлює властивість `justify-content` своїх елементів на `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

Ви повинні мати селектор `#playing-cards`, який встановлює властивість `flex-wrap` своїх елементів на `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

Ви повинні мати селектор `#playing-cards`, який встановлює властивість `gap` своїх елементів на `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

Ви повинні мати селектор `.card`, який встановлює властивість `display` своїх елементів на `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

Ви повинні мати селектор `.card`, який встановлює властивість `justify-content` своїх елементів на `space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

Ви повинні мати селектор `.left`, який встановлює властивість `align-self` своїх елементів на `flex-start`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

Ви повинні мати селектор `.middle`, який встановлює властивість `align-self` своїх елементів на `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

Ви повинні мати селектор `.right`, який встановлює властивість `align-self` своїх елементів на `flex-end`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

Ви повинні мати селектор `.middle`, який встановлює властивість `flex-direction` своїх елементів на `column`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
