---
id: 587d78a9367417b2b2512ae9
title: Tumia mkondo wa Bezier kusogeza mchoro
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

Changamoto iliyopita ilijadili neno kuu la `ease-out` linaloelezea mabadiliko ya uhuishaji yanayoongeza kasi mwanzoni kisha kupungua mwishoni mwa uhuishaji. Kushoto, tofauti kati ya neno kuu la `ease-out` (kwa kipengele cha buluu) na neno kuu la `linear` (kwa kipengele cha nyekundu) inaonyeshwa. Mfululizo wa uhuishaji unaofanana na neno kuu la `ease-out` unaweza kufanikishwa kwa kutumia kitendakazi maalum cha mkondo wa Bezier wa daraja la tatu.

Kwa ujumla, kubadilisha pointi za tagi ya nanga za `p1` na `p2` huendesha uundaji wa mikondo tofauti ya Bezier, ambayo hudhibiti jinsi uhuishaji unavyosogea kwa muda. Hapa kuna mfano wa mkondo wa Bezier ukitumia thamani kuiga mtindo wa `ease-out`:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

Kumbuka kuwa vitendakazi vyote vya `cubic-bezier` huanza na `p0` kwa (0, 0) na kumalizika na `p3` kwa (1, 1). Katika mfano huu, mkondo unasogea kwa kasi zaidi kupitia mhimili wa Y (huanza kwa 0, kwenda kwa thamani ya y ya `p1` ya 0, kisha kwenda kwa thamani ya y ya `p2` ya 1) kuliko unavyosogea kupitia mhimili wa X (0 kuanzia, kisha 0 kwa `p1`, hadi 0.58 kwa `p2`). Matokeo yake, mabadiliko katika kipengele kinachohamishwa husogea kwa kasi zaidi kuliko muda wa uhuishaji kwa sehemu hiyo. Mwishoni mwa mkondo, uhusiano kati ya mabadiliko ya thamani za x na y hubadilika - thamani ya y inasogea kutoka 1 hadi 1 (hakuna mabadiliko), na thamani za x husogea kutoka 0.58 hadi 1, na kufanya mabadiliko ya uhuishaji kusogea polepole ikilinganishwa na muda wa uhuishaji.

# --instructions--

Ili kuona athari ya mkondo huu wa Bezier ukiwa hai, badilisha `animation-timing-function` ya kipengele chenye kitambulisho cha `red` kuwa kitendakazi cha `cubic-bezier` chenye x1, y1, x2, y2 zikiwa zimewekwa kwa mtiririko wa `0, 0, 0.58, 1` kama thamani. Hii itafanya vipengele vyote viendelee kupitia uhuishaji kwa njia sawa.

# --hints--

Thamani ya sifa ya `animation-timing-function` ya kipengele chenye kitambulisho cha `red` inapaswa kuwa kitendakazi cha `cubic-bezier` chenye thamani za x1, y1, x2, y2 zikiwa mtiririko 0, 0, 0.58, 1 .

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

Kipengele chenye kitambulisho cha `red` hakipaswi tena kuwa na sifa ya `animation-timing-function` yenye thamani ya `linear`.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

Thamani ya sifa ya `animation-timing-function` kwa kipengele chenye kitambulisho cha `blue` haipaswi kubadilika.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
