---
id: 587d78a5367417b2b2512ada
title: Tumia sifa ya CSS Transform scale kupanua kipengele unapoelekeza kipanya juu yake
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

Sifa ya `transform` ina kazi mbalimbali zinazokuwezesha kupanua, kusogeza, kuzungusha, kupinda, n.k., vipengele vyako. Inapotumika pamoja na madarasa ya bandia kama `:hover` yanayobainisha hali fulani ya kipengele, sifa ya `transform` inaweza kuongeza urahisi mwingiliano kwa vipengele vyako.

Hapa kuna mfano wa kupanua vipengele vya aya mara 2.1 ya ukubwa wao wa awali wakati mtumizi anapoelekeza kipanya juu yao:

```css
p:hover {
  transform: scale(2.1);
}
```

**Kumbuka:** Kutumia transform kwa kipengele cha `div` kutakuwa na athari pia kwa vipengele vyote vya watoto vilivyomo ndani ya div.

# --instructions--

Ongeza sheria ya CSS kwa hali ya `hover` ya `div` na tumia sifa ya `transform` kupanua kipengele cha `div` mara 1.1 ya ukubwa wake wa awali wakati mtumizi anapoelekeza kipanya juu yake.

# --hints--

Ukubwa wa kipengele cha `div` unapaswa kupanuka mara 1.1 wakati mtumizi anapoelekeza kipanya juu yake.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
