---
id: 587d78a4367417b2b2512ad3
title: Badilisha rangi za vipengele mbalimbali kuwa rangi za nyongeza
challengeType: 0
videoUrl: 'https://scrimba.com/c/cWmPpud'
forumTopicId: 301033
dashedName: adjust-the-color-of-various-elements-to-complementary-colors
---

# --description--

Changamoto ya rangi za nyongeza ilionyesha kuwa rangi zinazopingana kwenye gurudumu la rangi zinaweza kufanya kila moja ionekane angavu zaidi zinapowekwa kando kwa kando. Hata hivyo, tofauti kali ya kuona inaweza kuwa ya kushtua ikiwa itatumika kupita kiasi kwenye tovuti, na wakati mwingine inaweza kufanya maandishi kuwa magumu kusoma ikiwa yamewekwa kwenye mandhari yenye rangi ya nyongeza. Katika matumizi halisi, moja ya rangi huwa ya kuongoza na rangi ya nyongeza hutumika kuleta mvuto wa kuona kwa maudhui fulani kwenye ukurasa.

# --instructions--

Ukurasa huu utatumia kivuli cha teal (`#09A7A1`) kama rangi kuu, na rangi yake ya machungwa (`#FF790E`) ya nyongeza kuangazia kwa kuona vifungo vya usajili. Badilisha `background-color` ya `header` na `footer` kutoka nyeusi kuwa rangi ya teal. Kisha badilisha maandishi ya `h2` `color` kuwa teal pia. Mwishowe, badilisha `background-color` ya `button` kuwa rangi ya machungwa.

# --hints--

Kipengele cha `header` kinapaswa kuwa na `background-color` ya #09A7A1.

```js
const headerElement = document.querySelector('header'); 
const headerStyle = window.getComputedStyle(headerElement);
assert.equal(headerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

Kipengele cha `footer` kinapaswa kuwa na `background-color` ya #09A7A1.

```js
const footerElement = document.querySelector('footer'); 
const footerStyle = window.getComputedStyle(footerElement);
assert.equal(footerStyle?.backgroundColor, 'rgb(9, 167, 161)');
```

Kipengele cha `h2` kinapaswa kuwa na `color` ya #09A7A1.

```js
const h2Element = document.querySelector('h2'); 
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.color, 'rgb(9, 167, 161)');
```

Kipengele cha `button` kinapaswa kuwa na `background-color` ya #FF790E.

```js
const buttonElement = document.querySelector('button'); 
const buttonStyle = window.getComputedStyle(buttonElement);
assert.equal(buttonStyle?.backgroundColor, 'rgb(255, 121, 14)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: black;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: black;
  }
  button {
    background-color: white;
  }
  footer {
    background-color: black;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```

# --solutions--

```html
<style>
  body {
    background-color: white;
  }
  header {
    background-color: #09A7A1;
    color: white;
    padding: 0.25em;
  }
  h2 {
    color: #09A7A1;
  }
  button {
    background-color: #FF790E;
  }
  footer {
    background-color: #09A7A1;
    color: white;
    padding: 0.5em;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
</header>
<main>
  <article>
    <h2>Machine Learning in the Kitchen</h2>
    <p>Join this two day workshop that walks through how to implement cutting-edge snack-getting algorithms with a command line interface. Coding usually involves writing exact instructions, but sometimes you need your computer to execute flexible commands, like <code>fetch Pringles</code>.</p>
    <button>Sign Up</button>
  </article>
  <article>
    <h2>Bisection Vegetable Chopping</h2>
    <p>This week-long retreat will level-up your coding ninja skills to actual ninja skills. No longer is the humble bisection search limited to sorted arrays or coding interview questions, applying its concepts in the kitchen will have you chopping carrots in O(log n) time before you know it.</p>
    <button>Sign Up</button>
  </article>
</main>
<br>
<footer>&copy; 2018 FCC Kitchen</footer>
```
