---
id: 587d781e367417b2b2512aca
title: Hamisha kipengele chenye nafasi maalum kwa uondoaji wa nafasi wa CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

Uondoaji wa nafasi wa `top` au `bottom`, na `left` au `right` hueleza kivinjari ni kwa kiasi gani kipengele kinapaswa kuondolewa mbali kulingana na mahali kingekuwa katika mtiririko wa kawaida wa hati. Unauondoa kipengele mbali na sehemu fulani, ambayo husogeza kipengele mbali na upande uliorejelea (kwa maana nyingine, upande wa kinyume). Kama ulivyoona katika changamoto iliyopita, kutumia uondoaji wa `top` ulisogeza `h2` chini. Vivyo hivyo, kutumia uondoaji wa `left` husogeza kipengele kwenda kulia.

# --instructions--

Tumia uondoaji wa nafasi wa CSS kusogeza `h2` pikseli 15 kwenda kulia na pikseli 10 juu.

# --hints--

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa CSS kuweka `h2` kwa nafasi maalum kwa kusogeza 10px juu. Kwa maneno mengine, isogeze 10px mbali na `bottom` ya mahali ambapo kawaida kingekaa.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa CSS kuweka `h2` kwa nafasi maalum kwa kusogeza 15px kuelekea kulia. Kwa maneno mengine, isogeze 15px mbali na `left` ya mahali ambapo kawaida kingekaa.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
