---
id: 587d78a7367417b2b2512ae0
title: Tumia uhuishaji wa CSS kubadilisha hali ya hover ya kitufe
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

Unaweza kutumia uhuishaji wa CSS `@keyframes` kubadilisha rangi ya kitufe katika hali yake ya hover.

Hapa kuna mfano wa kubadilisha upana wa picha wakati ya hover:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

Kumbuka kuwa `ms` inasimamia millisekunde, ambapo 1000ms ni sawa na 1s.

Tumia uhuishaji wa CSS `@keyframes` kubadilisha `background-color` ya kipengele cha `button` ili kiwe `#4791d0` wakati mtumizi anapobeba kidole juu yake. Sheria ya `@keyframes` inapaswa kuwa na kipengele kimoja tu cha `100%`.

# --hints--

Sheria ya @keyframes inapaswa kutumia rangi ya nyuma ya `animation-name`.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

Inapaswa kuwepo na sheria moja chini ya `@keyframes` inayobadilisha `background-color` kuwa `#4791d0` kwa 100%.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
