---
id: 587d781b367417b2b2512abc
title: Rekebisha sifa ya background-color ya maandishi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDqwA6'
forumTopicId: 301032
dashedName: adjust-the-background-color-property-of-text
---

# --description--

Badala ya kurekebisha rangi ya nyuma kwa ujumla au rangi ya maandishi ili kufanya mbele kusomeka kwa urahisi, unaweza kuongeza `background-color` kwa kipengele kinachoshikilia maandishi unayotaka kusisitiza. Changamoto hii inatumia `rgba()` badala ya `hex` au `rgb()` za kawaida.

<blockquote>rgba inasimama kwa:<br>  r = nyekundu<br>  g = kijani<br>  b = bluu<br>  a = alfa/kiwango cha uwazi</blockquote>

Thamani za RGB zinaweza kuwa kutoka 0 hadi 255. Thamani ya alfa inaweza kuwa kutoka 1, ambayo ni rangi kamili au isiyo na uwazi kabisa, hadi 0, ambayo ni uwazi kabisa au safi. `rgba()` ni nzuri kutumia katika kesi hii, kwani inakuwezesha kurekebisha uwazi. Hii inamaanisha hauhitaji kufunika kabisa rangi ya nyuma.

Utaitumia `background-color: rgba(45, 45, 45, 0.1)` kwa changamoto hii. Inazalisha rangi ya kijivu giza ambayo karibu haionekani kutokana na thamani ya uwazi ya chini ya 0.1.

# --instructions--

Ili kufanya maandishi yaonekane zaidi, rekebisha `background-color` ya kipengele cha `h4` kwa thamani ya `rgba()` iliyotolewa.

Pia kwa `h4`, ondoa sifa ya `height` na ongeza `padding` ya 10px.

# --hints--

Msimbo wako unapaswa kuongeza sifa ya `background-color` kwa kipengele cha `h4` na kuiweka kuwa `rgba(45, 45, 45, 0.1)`.

```js
assert(
  /(background-color|background):rgba\(45,45,45,0?\.1\)(;?}|;)/gi.test(
    code.replace(/\s/g, '')
  )
);
```

Msimbo wako unapaswa kuongeza sifa ya `padding` kwa kipengele cha `h4` na kuiweka kuwa pikseli 10.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.equal(h4style?.paddingTop, "10px");
assert.equal(h4style?.paddingRight, "10px");
assert.equal(h4style?.paddingBottom, "10px");
assert.equal(h4style?.paddingLeft, "10px");
```

Sifa ya `height` kwenye kipengele cha `h4` inapaswa kufutwa.

```js
const h4Element = document.querySelector("h4"); 
const h4style = window.getComputedStyle(h4Element);
assert.notEqual(h4style?.height, '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;


  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    padding: 10px;
    background-color: rgba(45, 45, 45, 0.1);
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Alphabet</h4>
      <hr>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
