---
id: 587d78a7367417b2b2512ae2
title: Unda mwelekeo wa kuona kwa kupunguza mwonekano wa kipengele kutoka kushoto kwenda kulia
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

Kwa changamoto hii, utabadilisha `opacity` ya kipengele kilicho na michoro ili kipengele hicho kipungue mwonekano taratibu linapofikia upande wa kulia wa skrini.

Katika michoro inayoonyeshwa, kipengele mviringo chenye mandhari ya mchanganyiko wa rangi kinahama kwenda upande wa kulia kwa nambari 50% ya michoro kulingana na kanuni ya `@keyframes`.

# --instructions--

Lenga kipengele chenye kitambulisho cha `ball` na ongeza sifa ya `opacity` iliyowekwa kuwa 0.1 katika `50%`, ili kipengele kipungue mwonekano linapohama kwenda kulia.

# --hints--

Kanuni ya `keyframes` ya kupunguza mwonekano inapaswa kuweka sifa ya `opacity` kuwa 0.1 katika asilimia 50%.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
