---
id: 587d78a6367417b2b2512adc
title: Tumia sifa ya CSS Transform skewY kupinda kipengele kando ya mhimili wa Y
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

Kwa kuwa kitendakazi cha `skewX()` hupinda kipengele kilichochaguliwa kando ya mhimili wa X kwa digrii fulani, si jambo la kushangaza kwamba sifa ya `skewY()` hupinda kipengele kando ya mhimili wa Y (wima).

# --instructions--

Pinda kipengele chenye kitambulisho cha `top` kwa digrii -10 kando ya mhimili wa Y kwa kutumia sifa ya `transform`.

# --hints--

Kipengele chenye kitambulisho cha `top` kinapaswa kupindishwa kwa digrii -10 kando ya mhimili wake wa Y.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
