---
id: 587d78a5367417b2b2512ad8
title: Unda muundo kwa kuongeza mfano mdogo kama picha ya nyuma
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

Njia moja ya kuongeza muundo na mvuto kwenye picha ya nyuma na kuifanya ionekane zaidi ni kuongeza mfano mdogo. Muhimu ni usawa, kwani hutaki picha ya nyuma ionekane sana na kuondoa mvuto kutoka mbele. Sifa ya `background` inaunga mkono kitendakazi cha `url()` ili kuunganisha na picha ya muundo au mfano uliotolewa. Anwani ya kiungo imefungwa kwa nukuu ndani ya mabano ya kawaida.

# --instructions--

Ukichukua url ya `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`, weka `background` ya ukurasa mzima kwa kutumia kichaguzi cha `body`.

# --hints--

Kipengele chako cha `body` kinapaswa kuwa na sifa ya `background` iliyowekwa kuwa `url()` yenye kiungo kilichotolewa.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
