---
id: 587d78a6367417b2b2512adb
title: Tumia sifa ya CSS Transform skewX kupinda kipengele kando ya mhimili wa X
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

Kitendakazi kinachofuata cha sifa ya `transform` ni `skewX()`, ambayo hupinda kipengele kilichochaguliwa kando ya mhimili wake wa X (wima) kwa kiwango kilichotolewa.

Msimbo ufuatao hupinda kipengele cha aya kwa digrii -32 kando ya mhimili wa X.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

Pinda kipengele chenye kitambulisho cha `bottom` kwa digrii 24 kando ya mhimili wa X kwa kutumia sifa ya `transform`.

# --hints--

Kipengele chenye kitambulisho `bottom` kinapaswa kupindishwa kwa digrii 24 kando ya mhimili wake wa X.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
