---
id: 587d78a8367417b2b2512ae6
title: Animashea vitu vingi kwa viwango tofauti
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpWZc9'
forumTopicId: 301042
dashedName: animate-multiple-elements-at-variable-rates
---

# --description--

Katika changamoto iliyopita, ulibadilisha viwango vya uhuishaji kwa vitu viwili vilivyo na uhuishaji sawa kwa kubadilisha sheria za `@keyframes`. Unaweza kufanikisha lengo lile lile kwa kudhibiti `animation-duration` ya vitu vingi.

Katika uhuishaji unaoendelea kwenye mhariri wa msimbo, kuna nyota tatu angani zinazong'aa kwa kiwango kimoja katika mzunguko unaoendelea. Ili kuzifanya ziongoe kwa viwango tofauti, unaweza kuweka sifa ya `animation-duration` kwa thamani tofauti kwa kila kitu.

# --instructions--

Weka sifa ya `animation-duration` kwa vitu vyenye madarasa ya `star-1`, `star-2`, na `star-3` kuwa 1s, 0.9s, na 1.1s, kwa mtiririko huo.

# --hints--

Sifa ya `animation-duration` kwa nyota yenye darasa la `star-1` inapaswa kubaki 1s.

```js
const starOne = document.querySelector('.star-1');
const starOneStyle = window.getComputedStyle(starOne); 
assert.equal(starOneStyle?.animationDuration, '1s');
```

Sifa ya `animation-duration` kwa nyota yenye darasa la `star-2` inapaswa kuwa 0.9s.

```js
const starTwo = document.querySelector('.star-2');
const starTwoStyle = window.getComputedStyle(starTwo); 
assert.equal(starTwoStyle?.animationDuration , '0.9s');
```

Sifa ya `animation-duration` kwa nyota yenye darasa la `star-3` inapaswa kuwa 1.1s.

```js
const starThree = document.querySelector('.star-3');
const starThreeStyle = window.getComputedStyle(starThree); 
assert.equal(starThreeStyle?.animationDuration, '1.1s');
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-duration: 1s;
    animation-name: twinkle;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-duration: 0.9s;
    animation-name: twinkle;
  }

  .star-3 {
    margin-top: 10%;
    margin-left: 50%;
    animation-duration: 1.1s;
    animation-name: twinkle;
  }

  @keyframes twinkle {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
<div class="star-3 stars"></div>
```
