---
id: 587d781e367417b2b2512acb
title: Funga kipengele kwa mzazi wake kwa kutumia nafasi ya absolute
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

Chaguo linalofuata kwa sifa ya CSS `position` ni `absolute`, ambayo inafunga kipengele mahali pake kulingana na kontena la mzazi wake. Tofauti na nafasi ya `relative`, hii huondoa kipengele kutoka kwenye mtiririko wa kawaida wa hati, hivyo vitu vinavyomzunguka havikitii. Sifa za uondoaji wa nafasi za CSS (juu au chini na kushoto au kulia) hutumika kurekebisha nafasi.

Jambo moja la kipekee kuhusu nafasi ya absolute ni kwamba itafungwa kulingana na babu yake wa karibu zaidi aliye na *nafasi* iliyowekwa. Ikiwa utasahau kuongeza sheria ya nafasi kwa kipengele mzazi, (hii kawaida hufanywa kwa kutumia `position: relative;`), kivinjari kitaendelea kutafuta juu kwenye mnyororo na hatimaye itatumia lebo ya `body` kama chaguo la msingi.

# --instructions--

Funga kipengele cha `#searchbar` juu-kulia mwa mzazi wake wa `section` kwa kutangaza sifa yake ya `position` kuwa `absolute`. Mpe uondoaji wa nafasi wa `top` na `right` wa pikseli 50 kila moja.

# --hints--

Kipengele cha `#searchbar` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa `top` wa pikseli 50 kwenye kipengele cha `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa `right` wa pikseli 50 kwenye kipengele cha `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
