---
id: 587d781a367417b2b2512ab8
title: Tumia lebo ya u kuangazia maandishi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN6aQCL'
forumTopicId: 301082
dashedName: use-the-u-tag-to-underline-text
---

# --description--

Ili kuangazia maandishi kwa mstari chini, unaweza kutumia lebo ya `u`. Hii mara nyingi hutumika kuonyesha kuwa sehemu ya maandishi ni muhimu, au kitu cha kukumbuka. Kwa lebo ya `u`, kivinjari hutatua CSS ya `text-decoration: underline;` kwa kipengele hicho.

# --instructions--

Zungusha lebo ya `u` tu kuzunguka maandishi ya `Ph.D. students`.

**Kumbuka:** Jaribu kuepuka kutumia lebo ya `u` wakati inaweza kuchanganyikiwa na kiungo. Lebo za tagi ya nanga pia zina muundo wa mstari chini kwa chaguo-msingi.

# --hints--

Msimbo wako unapaswa kuongeza lebo ya `u` kwenye alama za uandishi.

```js
assert.lengthOf(document.querySelectorAll('u'), 1);
```

Leboa ya `u` inapaswa kuzunguka maandishi ya `Ph.D. students`.

```js
assert.equal(document.querySelector('u')?.textContent , 'Ph.D. students');
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were Ph.D. students at <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
