---
id: 587d78a4367417b2b2512ad5
title: Badilisha sauti ya rangi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDJvT7'
forumTopicId: 301038
dashedName: adjust-the-tone-of-a-color
---

# --description--

Chaguo la `hsl()` katika CSS pia hufanya iwe rahisi kubadilisha sauti ya rangi. Kuchanganya nyeupe na rangi safi huunda rangi ya tinti ya rangi hiyo, na kuongeza nyeusi kutaunda kivuli. Vinginevyo, sauti hutengenezwa kwa kuongeza kijivu au kwa kuchanganya tinti na kivuli. Kumbuka kuwa 's' na 'l' za `hsl()` zinamaanisha saturation na uangavu, mtawalia. Asilimia ya saturation hubadilisha kiasi cha kijivu na asilimia ya uangavu huamua ni kiasi gani cha nyeupe au nyeusi kiko kwenye rangi. Hii ni muhimu wakati una rangi msingi unayopenda, lakini unahitaji mabadiliko tofauti ya rangi hiyo.

# --instructions--

Vipengele vyote vina `background-color` ya chaguo-msingi ya `transparent`. Kipengele chetu cha `nav` kwa sasa kinaonekana kuwa na rangi ya nyuma ya `cyan`, kwa sababu kipengele kilicho nyuma yake kina `background-color` iliyowekwa kuwa `cyan`. Ongeza `background-color` kwa kipengele cha `nav` ili kitumie rangi ya `cyan` sawa, lakini kiwe na thamani za `80%` za saturation na `25%` za uangavu kubadilisha sauti na kivuli chake.

# --hints--

Kipengele cha `nav` kinapaswa kuwa na `background-color` ya sauti ya cyan iliyobadilishwa kwa kutumia sifa ya `hsl()`.

```js
// Computed style of hsl(180, 80%, 25%) results in rgb(13,115,115)
assert.equal(
  new __helpers.CSSHelp(document).getStyle('nav').getPropVal('background-color', true), 
  'rgb(13,115,115)'
)
```

# --seed--

## --seed-contents--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {

  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>

<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```

# --solutions--

```html
<style>
  header {
    background-color: hsl(180, 90%, 35%);
    color: #FFFFFF;
  }

  nav {
    background-color: hsl(180, 80%, 25%);
  }

  h1 {
    text-indent: 10px;
    padding-top: 10px;
  }

  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }

  nav li {
    display: inline;
    margin-right: 20px;
  }

  a {
    text-decoration: none;
    color: inherit;
  }
</style>
<header>
  <h1>Cooking with FCC!</h1>
  <nav>
    <ul>
      <li><a href="#">Home</a></li>
      <li><a href="#">Classes</a></li>
      <li><a href="#">Contact</a></li>
    </ul>
  </nav>
</header>
```
