---
id: 587d78a5367417b2b2512ad9
title: Tumia sifa ya CSS Transform scale kubadilisha ukubwa wa kipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZVSg'
forumTopicId: 301076
dashedName: use-the-css-transform-scale-property-to-change-the-size-of-an-element
---

# --description--

Kubadilisha ukubwa wa kipengele, CSS ina sifa ya `transform`, pamoja na kitendakazi chake cha `scale()`. Mfano wa msimbo ufuatao unaongeza mara mbili ukubwa wa vipengele vyote vya aya kwenye ukurasa:

```css
p {
  transform: scale(2);
}
```

# --instructions--

Ongeza ukubwa wa kipengele chenye kitambulisho cha `ball2` hadi mara 1.5 ya ukubwa wake wa awali.

# --hints--

Sifa ya `transform` kwa `#ball2` inapaswa kuwekwa ili kuibadilisha ukubwa hadi mara 1.5 ya ukubwa wake.

```js
assert.match(code,
    /#ball2\s*?{\s*?left:\s*?65%;\s*?transform:\s*?scale\(1\.5\);\s*?}|#ball2\s*?{\s*?transform:\s*?scale\(1\.5\);\s*?left:\s*?65%;\s*?}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;

  }


</style>

<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>
```

# --solutions--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;
    transform: scale(1.5);
  }
</style>
<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>
```
