---
id: 587d78a4367417b2b2512ad4
title: Badilisha hue ya rangi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

Rangi zina sifa kadhaa ikiwemo hue, saturation, na uangavu. CSS3 ilizindua kitendakazi cha `hsl()` kama njia mbadala ya kuchagua rangi kwa kusema moja kwa moja sifa hizi.

**Hue** ni kile ambacho watu kwa kawaida hufikiria kama 'rangi'. Ukifikiria mfululizo wa rangi unaoanza na nyekundu upande wa kushoto, ukipita kijani katikati, na buluu upande wa kulia, hue ni mahali rangi inapoelea kwenye mstari huu. Katika `hsl()`, hue hutumia dhana ya gurudumu la rangi badala ya mfululizo, ambapo pembe ya rangi kwenye mduara hutolewa kama thamani kati ya 0 na 360.

**Saturation** ni kiasi cha kijivu kilichopo kwenye rangi. Rangi yenye saturation kamili haina kijivu kabisa, na rangi yenye saturation kidogo ni karibu kabisa kijivu. Hii hutolewa kama asilimia ambapo 100% ni saturation kamili.

**Uangavu** ni kiasi cha nyeupe au nyeusi kilichopo kwenye rangi. Asilimia hutolewa kuanzia 0% (nyeusi) hadi 100% (nyeupe), ambapo 50% ni rangi ya kawaida.

Hapa kuna mifano michache ya kutumia `hsl()` na rangi zenye saturation kamili na uangavu wa kawaida:

<table><thead><tr><th>Rangi</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

Badilisha `background-color` ya kila kipengele cha `div` kulingana na majina ya madarasa (`green`, `cyan`, au `blue`) kwa kutumia `hsl()`. Zote tatu zinapaswa kuwa na saturation kamili na uangavu wa kawaida.

# --hints--

Msimbo wako unapaswa kutumia kitendakazi cha `hsl()` kutangaza rangi ya kijani.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Msimbo wako unapaswa kutumia kitendakazi cha `hsl()` kutangaza rangi ya cyan.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Msimbo wako unapaswa kutumia kitendakazi cha `hsl()` kutangaza rangi ya buluu.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Kipengele cha `div` chenye darasa la `green` kinapaswa kuwa na `background-color` ya kijani.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

Kipengele cha `div` chenye darasa la `cyan` kinapaswa kuwa na `background-color` ya cyan.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

Kipengele cha `div` chenye darasa la `blue` kinapaswa kuwa na `background-color` ya buluu.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
