---
id: 587d78a3367417b2b2512acf
title: Badilisha Nafasi ya Vitu Vinavyoviringika kwa Sifa ya z-index
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

Wakati vitu vimewekwa kwa njia ya kuviringika (yaani kwa kutumia `position: absolute | relative | fixed | sticky`), kitu kinachofuata baadaye katika alama za uandishi za HTML kitaonekana kwa chaguo-msingi juu ya vitu vingine. Hata hivyo, sifa ya `z-index` inaweza kubainisha mpangilio wa jinsi vitu vinavyopangwa juu ya vingine. Lazima iwe nambari kamili (yaani nambari nzima isiyo na sehemu ya desimali), na thamani kubwa zaidi za sifa ya `z-index` ya kitu huchukua nafasi ya juu zaidi katika mpangilio kuliko zile zenye thamani ndogo.

# --instructions--

Ongeza sifa ya `z-index` kwa kitu chenye jina la darasa la `first` (mstatili mwekundu) na uiweke thamani ya 2 ili kifunikie kitu kingine (mstatili wa buluu).

# --hints--

Kitu chenye darasa la `first` kinapaswa kuwa na thamani ya `z-index` ya 2.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
