---
id: 587d78a5367417b2b2512ad6
title: "Tengeneza mwelekeo wa hatua kwa hatua wa `linear-gradient()` wa CSS"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

Kuweka rangi kwenye vipengele vya HTML hakuzuiliwi kwa rangi moja tambarare tu. CSS inatoa uwezo wa kutumia mabadiliko ya rangi, yanayojulikana pia kama mwelekeo wa rangi (gradients), kwenye vipengele. Hii hupatikana kupitia sifa ya `background` na kitendakazi cha `linear-gradient()`. Hapa kuna sintaksia ya jumla:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

Hoja ya kwanza inaeleza mwelekeo kutoka ambapo mabadiliko ya rangi huanza - inaweza kuelezwa kwa digrii, ambapo `90deg` huunda mwelekeo wa wima (kutoka kushoto kwenda kulia) na `45deg` huunda mwelekeo wa mwinuko (kutoka chini kushoto hadi juu kulia). Hoja zinazofuata zinaeleza mpangilio wa rangi zinazotumika kwenye mwelekeo huo.

Mfano:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

Tumia `linear-gradient()` kwa sifa ya `background` ya kipengele cha `div`, na uweke mwelekeo wa digrii 35 kubadilisha rangi kutoka `#CCFFFF` hadi `#FFCCCC`.

# --hints--

Kipengele cha `div` kinapaswa kuwa na `linear-gradient` `background` yenye mwelekeo na rangi zilizotajwa.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
