---
id: 587d781b367417b2b2512aba
title: Tumia lebo ya s kufuta mstari juu ya maandishi
challengeType: 0
forumTopicId: 301079
dashedName: use-the-s-tag-to-strikethrough-text
---

# --description--

Ili kufuta mstari juu ya maandishi, ambayo ni wakati mstari wa usawa unavuka herufi, unaweza kutumia lebo ya `s`. Inaonyesha kuwa sehemu ya maandishi haipo tena halali. Kwa lebo ya `s`, kivinjari hutumia CSS ya `text-decoration: line-through;` kwa kipengele hicho.

# --instructions--

Zungusha lebo ya `s` kuzunguka `Google` ndani ya lebo ya `h4` kisha ongeza neno `Alphabet` kando yake bila muundo wa kufuta mstari juu ya maandishi.

# --hints--

Msimbo wako unapaswa kuongeza lebo moja ya `s` kwenye alama za uandishi.

```js
assert.lengthOf(document.querySelectorAll('s'),1);
```

Leboa `s` inapaswa kuzunguka maandishi ya `Google` katika lebo ya `h4`. Haipaswi kuwa na neno `Alphabet`.

```js
assert.match(document.querySelector('h4 > s')?.textContent, /Google/gi);
assert.notMatch(document.querySelector('h4 > s')?.textContent, /Alphabet/gi);
```

Unapaswa kujumuisha neno `Alphabet` katika lebo ya `h4`, bila muundo wa kufuta mstari juu ya maandishi.

```js
assert.match(document.querySelector('h4')?.innerHTML, /Alphabet/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4><s>Google</s> Alphabet</h4>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
