---
id: 587d78a7367417b2b2512ae1
title: Unda Mwendo Kwa Kutumia Uhuishaji wa CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

Wakati vipengele vina `position` iliyobainishwa, kama `fixed` au `relative`, sifa za uondoaji wa nafasi za CSS `right`, `left`, `top`, na `bottom` zinaweza kutumika katika sheria za uhuishaji kuunda mwendo.

Kama inavyoonyeshwa katika mfano hapa chini, unaweza kusukuma kipengee chini kisha juu kwa kuweka sifa ya `top` ya fremu ya msingi ya `50%` kuwa 50px, lakini ikiwa imewekwa 0px kwa fremu ya kwanza (`0%`) na ya mwisho (`100%`).

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

Ongeza mwendo wa usawa kwa uhuishaji wa `div`. Kwa kutumia sifa ya uondoaji wa nafasi ya `left`, ongeza kwenye sheria ya `@keyframes` ili rangi ya upinde wa mvua ianze kwa pikseli 0 katika `0%`, ihamie pikseli 25 katika `50%`, na kumalizika kwa pikseli -25 katika `100%`. Usibadilishe sifa ya `top` katika mhariri - uhuishaji unapaswa kuwa na mwendo wa wima na usawa.

# --hints--

Sheria ya `@keyframes` kwa `0%` inapaswa kutumia uondoaji wa `left` wa 0px.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

Sheria ya `@keyframes` kwa `50%` inapaswa kutumia uondoaji wa `left` wa 25px.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

Sheria ya `@keyframes` kwa `100%` inapaswa kutumia uondoaji wa `left` wa -25px.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
