---
id: 587d781c367417b2b2512ac2
title: Weka ukubwa wa fonti kwa vichwa vingi vya sehemu
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
dashedName: set-the-font-size-for-multiple-heading-elements
---

# --description--

Sifa ya `font-size` hutumika kubainisha ukubwa wa maandishi katika kipengee fulani. Kanuni hii inaweza kutumika kwa vipengee vingi ili kuunda muonekano thabiti wa maandishi kwenye ukurasa. Katika changamoto hii, utaweka thamani kwa lebo zote za `h1` hadi `h6` ili kusawazisha ukubwa wa vichwa vya sehemu.

# --instructions--

Katika lebo za `style`, weka `font-size` ya:

- lebo ya `h1` kuwa 68px.
- lebo ya `h2` kuwa 52px.
- lebo ya `h3` kuwa 40px.
- lebo ya `h4` kuwa 32px.
- lebo ya `h5` kuwa 21px.
- lebo ya `h6` kuwa 14px.

# --hints--

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h1` kuwa pikseli 68.

```js
 const fontSizeOfh1 = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh1 ,'68px');
```

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h2` kuwa pikseli 52.

```js
 const fontSizeOfh2 = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh2 ,'52px');
```

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h3` kuwa pikseli 40.

```js
 const fontSizeOfh3 = new __helpers.CSSHelp(document).getStyle('h3')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh3 ,'40px');
```

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h4` kuwa pikseli 32.

```js
 const fontSizeOfh4 = new __helpers.CSSHelp(document).getStyle('h4')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh4 , '32px');
```

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h5` kuwa pikseli 21.

```js
 const fontSizeOfh5 = new __helpers.CSSHelp(document).getStyle('h5')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh5 ,'21px');
```

Msimbo wako unapaswa kuweka sifa ya `font-size` kwa lebo ya `h6` kuwa pikseli 14.

```js
 const fontSizeOfh6 = new __helpers.CSSHelp(document).getStyle('h6')?.getPropertyValue('font-size');
 assert.equal(fontSizeOfh6 , '14px');
```

# --seed--

## --seed-contents--

```html
<style>






</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```

# --solutions--

```html
<style>
  h1 {
    font-size: 68px;
  }
  h2 {
    font-size: 52px;
  }
  h3 {
    font-size: 40px;
  }
  h4 {
    font-size: 32px;
  }
  h5 {
    font-size: 21px;
  }
  h6 {
    font-size: 14px;
  }
</style>
<h1>This is h1 text</h1>
<h2>This is h2 text</h2>
<h3>This is h3 text</h3>
<h4>This is h4 text</h4>
<h5>This is h5 text</h5>
<h6>This is h6 text</h6>
```
