---
id: 587d78a3367417b2b2512ace
title: Sogeza vipengele kushoto au kulia kwa kutumia sifa ya float
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

Chombo kinachofuata cha kuweka nafasi hakitumi `position`, bali kinaweka sifa ya `float` ya kipengele. Vipengele vinavyotiririka huondolewa kutoka kwenye mtiririko wa kawaida wa hati na kusogezwa ama upande wa `left` au `right` wa kipengele mzazi kinachokijumuisha. Hutatumiwa sana pamoja na sifa ya `width` kuonyesha ni nafasi gani ya usawa kipengele kinachotiririka kinahitaji.

# --instructions--

Alama za uandishi zilizotolewa zingefanya kazi vizuri kama mpangilio wa safu mbili, ambapo vipengele vya `section` na `aside` viko pembeni kwa pembeni. Toa kipengele cha `#left` sifa ya `float` yenye thamani ya `left` na kipengele cha `#right` sifa ya `float` yenye thamani ya `right`.

# --hints--

Kipengele chenye kitambulisho cha `left` kinapaswa kuwa na thamani ya `float` ya `left`.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

Kipengele chenye kitambulisho cha `right` kinapaswa kuwa na thamani ya `float` ya `right`.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
