---
id: 587d78a4367417b2b2512ad2
title: Jifunze kuhusu rangi za ngazi ya tatu
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

Monyeti za kompyuta na skrini za vifaa huunda rangi tofauti kwa kuchanganya kiasi cha mwanga mwekundu, kijani, na buluu. Hii inajulikana kama mfano wa rangi wa kuongeza wa RGB katika nadharia ya kisasa ya rangi. Mwekundu (R), kijani (G), na buluu (B) huitwa rangi za msingi. Kuchanganya rangi mbili za msingi huunda rangi za sekondari cyan (G + B), magenta (R + B) na njano (R + G). Umeona rangi hizi katika changamoto ya rangi za nyongeza. Rangi hizi za sekondari ni nyongeza kwa rangi ya msingi ambayo haikutumika katika uundaji wake, na ziko kinyume na rangi hiyo ya msingi kwenye gurudumu la rangi. Kwa mfano, magenta hutengenezwa kwa mwekundu na buluu, na ni nyongeza kwa kijani.

Rangi za ngazi ya tatu ni matokeo ya kuchanganya rangi ya msingi na moja ya majirani zake wa rangi za sekondari. Kwa mfano, ndani ya mfano wa rangi wa RGB, mwekundu (misingi) na njano (sekondari) hutoa rangi ya chungwa (ngazi ya tatu). Hii inaongeza rangi sita zaidi kwenye gurudumu la rangi rahisi kwa jumla ya rangi kumi na mbili.

Kuna njia mbalimbali za kuchagua rangi tofauti zinazosababisha mchanganyiko mzuri katika muundo. Mfano mmoja unaoweza kutumia rangi za ngazi ya tatu huitwa mpangilio wa rangi ya nyongeza iliyogawanywa. Mpangilio huu huanza na rangi ya msingi, kisha huoanisha na rangi mbili zilizo jirani na nyongeza yake. Rangi hizi tatu hutoa tofauti kali ya kuona katika muundo, lakini ni laini zaidi kuliko kutumia rangi mbili za nyongeza.

Hapa kuna rangi tatu zilizotengenezwa kwa kutumia mpangilio wa nyongeza iliyogawanywa:

<table><thead><tr><th>Rangi</th><th>Kodi ya Hex</th></tr></thead><thead></thead><tbody><tr><td>chungwa</td><td>#FF7F00</td></tr><tr><td>cyan</td><td>#00FFFF</td></tr><tr><td>raspberry</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

Badilisha sifa ya `background-color` ya madarasa ya `orange`, `cyan`, na `raspberry` kuwa rangi zao husika. Hakikisha kutumia kodi za hex badala ya majina ya rangi.

# --hints--

Kipengele cha `div` chenye darasa la `orange` kinapaswa kuwa na sifa ya `background-color` ya rangi ya chungwa.

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

Kipengele cha `div` chenye darasa la `cyan` kinapaswa kuwa na sifa ya `background-color` ya rangi ya cyan.

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

Kipengele cha `div` chenye darasa la `raspberry` kinapaswa kuwa na sifa ya `background-color` ya rangi ya raspberry.

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

Thamani zote za `background-color` kwa madarasa ya rangi zinapaswa kuwa kodi za hex na si majina ya rangi.

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
