---
id: 587d781e367417b2b2512acc
title: Funga kipengele kwenye dirisha la kivinjari kwa kutumia nafasi ya kudumu
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

Mpangilio unaofuata wa muundo unaotolewa na CSS ni nafasi ya `fixed`, ambayo ni aina ya nafasi ya kabisa inayofunga kipengele kulingana na dirisha la kivinjari. Kama ilivyo kwa nafasi ya kabisa, hutumika pamoja na vigezo vya uondoaji wa nafasi vya CSS na pia huondoa kipengele kutoka kwenye mtiririko wa kawaida wa hati. Vitu vingine havitambui tena mahali kilipo, jambo ambalo linaweza kuhitaji marekebisho ya muundo sehemu nyingine.

Tofauti kuu kati ya nafasi za `fixed` na `absolute` ni kwamba kipengele chenye nafasi ya kudumu hakitahamia wakati mtumizi anaposogeza ukurasa.

# --instructions--

Bar ya urambazaji katika msimbo ina lebo ya kitambulisho cha `navbar`. Badilisha `position` yake kuwa `fixed`, na ipe uondoaji wa nafasi wa 0 pikseli kutoka `top` na 0 pikseli kutoka `left`. Baada ya kuongeza msimbo, sogeza dirisha la onyesho kuona jinsi urambazaji unavyobaki mahali pake.

# --hints--

Kipengele cha `#navbar` kinapaswa kuwa na `position` iliyowekwa kuwa `fixed`.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa `top` wa 0 pikseli kwenye kipengele cha `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

Msimbo wako unapaswa kutumia uondoaji wa nafasi wa `left` wa 0 pikseli kwenye kipengele cha `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
