---
id: 587d78a8367417b2b2512ae3
title: Anzisha vipengele kuendelea kwa kutumia idadi isiyo na kikomo ya uhuishaji
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

Changamoto zilizopita zilijumuisha jinsi ya kutumia baadhi ya sifa za uhuishaji na kanuni ya `@keyframes`. Sifa nyingine ya uhuishaji ni `animation-iteration-count`, ambayo inakuwezesha kudhibiti ni mara ngapi ungependa kurudia uhuishaji huo. Hapa kuna mfano:

```css
animation-iteration-count: 3;
```

Katika kesi hii uhuishaji utaacha baada ya kuendeshwa mara 3, lakini inawezekana kufanya uhuishaji uendelee kuendeshwa kwa kuweka thamani hiyo kuwa `infinite`.

# --instructions--

Ili kuendelea kurusha mpira upande wa kulia kwa mzunguko usioisha, badilisha sifa ya `animation-iteration-count` kuwa `infinite`.

# --hints--

Sifa ya `animation-iteration-count` inapaswa kuwa na thamani ya `infinite`.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
