---
id: 587d781a367417b2b2512ab9
title: Tumia lebo ya em kuonyesha maandishi kwa herufi za mviringo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJRBtp'
forumTopicId: 301078
dashedName: use-the-em-tag-to-italicize-text
---

# --description--

Ili kuonyesha msisitizo kwenye maandishi, unaweza kutumia lebo ya `em`. Hii inaonyesha maandishi kwa herufi za mviringo, kwani kivinjari hutumia CSS ya `font-style: italic;` kwa kipengele hicho.

# --instructions--

Zungusha lebo ya `em` kuzunguka yaliyomo ndani ya lebo ya aya ili kuipa msisitizo.

# --hints--

Msimbo wako unapaswa kuongeza lebo ya `em` kwenye alama za uandishi.

```js
assert.lengthOf(document.querySelectorAll('em'),1);
```

Leboa ya `em` inapaswa kuzunguka yaliyomo ndani ya lebo ya `p` lakini si lebo ya `p` yenyewe.

```js
const paragraphElement = document.querySelector('p');
const emElement = document.querySelector('em');
assert.lengthOf(paragraphElement?.children, 1);
assert.lengthOf(emElement?.children, 2);
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p><em>Google was founded by Larry Page and Sergey Brin while they were <u>Ph.D. students</u> at <strong>Stanford University</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
