---
id: 587d781d367417b2b2512ac8
title: Rekebisha hali ya hover ya tagi ya nanga
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

Changamoto hii itagusa matumizi ya madarasa ya bandia. Darasa la bandia ni neno kuu linaloweza kuongezwa kwa vichaguzi, ili kuchagua hali maalum ya kipengee.

Kwa mfano, mtindo wa tagi ya nanga unaweza kubadilishwa kwa hali yake ya hover kwa kutumia kichaguzi cha darasa la bandia `:hover`. Huu ndio msimbo wa CSS kubadilisha `color` ya tagi ya nanga kuwa nyekundu wakati wa hali yake ya hover:

```css
a:hover {
  color: red;
}
```

# --instructions--

Mhariri wa msimbo una sheria ya CSS ya kupangilia tagi zote za `a` kuwa nyeusi. Ongeza sheria ili wakati mtumizi anapobeba kidole juu ya tagi ya `a`, `color` iwe ya buluu.

# --hints--

Taji ya nanga `color` inapaswa kubaki nyeusi, ongeza tu sheria za CSS kwa hali ya `:hover`.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

Taji ya nanga inapaswa kuwa na `color` ya buluu wakati wa hover.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
