---
id: 587d78a8367417b2b2512ae5
title: Anzisha vitu kwa viwango tofauti
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ89WA4'
forumTopicId: 301040
dashedName: animate-elements-at-variable-rates
---

# --description--

Kuna njia mbalimbali za kubadilisha viwango vya uhuishaji wa vitu vinavyohusiana. Hadi sasa, hii imefikiwa kwa kutumia sifa ya `animation-iteration-count` na kuweka sheria za `@keyframes`.

Ili kuonyesha, uhuishaji upande wa kulia unajumuisha nyota mbili ambazo kila moja hupungua ukubwa na uwazi katika alama ya 20% katika sheria ya `@keyframes`, ambayo huunda uhuishaji wa kung'aa. Unaweza kubadilisha sheria ya `@keyframes` kwa moja ya vitu hivyo ili nyota zing'ae kwa viwango tofauti.

# --instructions--

Badilisha kiwango cha uhuishaji kwa kitu chenye jina la darasa la `star-1` kwa kubadilisha sheria yake ya `@keyframes` kuwa 50%.

# --hints--

Sheria ya `@keyframes` kwa darasa la `star-1` inapaswa kuwa 50%.

```js
assert.match(code,/twinkle-1\s*?{\s*?50%/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    50% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```
