---
id: 587d78a9367417b2b2512aea
title: Fanya Miondoko Iwe Asili Zaidi Kutumia Mkondo wa Bezier
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

Changamoto hii huhuisha kipengele kuiga mwendo wa mpira unaochezwa kwa kurusha na kupokea. Changamoto zilizopita zilijumuisha mikondo ya Bezier ya mstatili `linear` na `ease-out`, lakini hakuna hata mmoja aliyeonyesha mwendo wa kurusha mpira kwa usahihi. Unahitaji kubinafsisha mkondo wa Bezier kwa hili.

`animation-timing-function` hujirudia moja kwa moja kila fremu muhimu wakati `animation-iteration-count` imewekwa kuwa isiyo na kikomo. Kwa kuwa kuna kanuni ya fremu muhimu katikati ya muda wa uhuishaji (katika `50%`), inasababisha mfululizo wa uhuishaji unaofanana mara mbili katika mwendo wa juu na chini wa mpira.

Mkondo wa Bezier wa mstatili ufuatao unaiga mwendo wa kurusha mpira:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

Angalia kuwa thamani ya y2 ni kubwa kuliko 1. Ingawa mkondo wa Bezier wa mstatili umewekwa kwenye mfumo wa kuratibu wa 1 kwa 1, na unaweza tu kupokea thamani za x kutoka 0 hadi 1, thamani ya y inaweza kuwekwa kuwa nambari kubwa kuliko moja. Hii husababisha mwendo wa kuruka ambao ni mzuri kwa kuiga mpira unaorushwa.

# --instructions--

Badilisha thamani ya `animation-timing-function` ya kipengele chenye kitambulisho `green` kuwa kitendakazi cha `cubic-bezier` chenye thamani za x1, y1, x2, y2 zilizo sawa na 0.311, 0.441, 0.444, 1.649.

# --hints--

Thamani ya sifa ya `animation-timing-function` kwa kipengele chenye kitambulisho `green` inapaswa kuwa kitendakazi cha `cubic-bezier` chenye thamani za x1, y1, x2, y2 kama zilivyoainishwa.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
