---
id: 587d78a3367417b2b2512ad0
title: "Weka kipengele katikati kwa usawa kwa kutumia sifa ya `margin`"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

Mbinu nyingine ya kuweka nafasi ni kuweka kipengele cha kizuizi katikati kwa usawa. Njia moja ya kufanya hivi ni kuweka `margin` yake kuwa thamani ya auto.

Njia hii pia hufanya kazi kwa picha. Picha ni vipengele vya mstari mmoja kwa chaguo-msingi, lakini vinaweza kubadilishwa kuwa vipengele vya kizuizi unapoweka sifa ya `display` kuwa `block`.

# --instructions--

Weka `div` katikati ya ukurasa kwa kuongeza sifa ya `margin` yenye thamani ya `auto`.

# --hints--

`div` inapaswa kuwa na `margin` iliyowekwa kuwa `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
