---
id: 58a7a6ebf9a6318348e2d5aa
title: Badilisha hali ya kujaza ya uhuishaji
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

Hiyo ni nzuri, lakini bado haifanyi kazi vizuri. Angalia jinsi uhuishaji unavyorejesha hali yake baada ya `500ms` kupita, na kusababisha kitufe kurudi kwenye rangi ya awali. Unataka kitufe kashike rangi iliyobainishwa.

Hii inaweza kufanyika kwa kuweka sifa ya `animation-fill-mode` kuwa `forwards`. `animation-fill-mode` inaeleza mtindo unaotumika kwa kipengele wakati uhuishaji umekamilika. Unaweza kuuiweka hivi:

```css
animation-fill-mode: forwards;
```

# --instructions--

Weka sifa ya `animation-fill-mode` ya `button:hover` kuwa `forwards` ili kitufe kishike rangi wakati mtumizi anapobeba kipanya juu yake.

# --hints--

`button:hover` inapaswa kuwa na sifa ya `animation-fill-mode` yenye thamani ya `forwards`.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
