---
id: 668f08ea07b99b1f4a91acab
title: Crea una página de recetas
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una declaración `!DOCTYPE html`.
1. Debes tener un elemento `html` con el atributo `lang` configurado a `en`.
1. Debes tener un elemento `head` que contenga un elemento `title` con el nombre de tu receta y un elemento `meta` con un atributo `charset` configurado a `UTF-8`.
1. Debes tener un elemento `body`.
1. Debes tener un elemento `h1` con el nombre de tu receta.
1. Debes tener un elemento `p` que introduzca la receta debajo del `h1`.
1. Debes tener un elemento `h2` con el texto `Ingredients` para la sección de ingredientes.
1. Debes tener una lista no ordenada (elemento `ul`) con al menos cuatro elementos de lista (elementos `li`) que enumeren tus ingredientes bajo el primer elemento `h2`.
1. Debes tener un segundo elemento `h2` con el texto `Instructions` para la sección de instrucciones.
1. Debes tener una lista ordenada (elemento `ol`) con al menos cuatro elementos de lista que enumeren los pasos de la receta en orden, debajo del segundo `h2`.
1. Debes tener un elemento `img` con un atributo `src` configurado a una imagen válida, puedes usar `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` si deseas, y un atributo `alt` que describa la imagen.

# --hints--

Tu página de receta debe tener una declaración `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Debes tener un elemento `html` con el atributo `lang` configurado a `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Debes tener un elemento `head` dentro del elemento `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Debes tener un elemento `title` dentro de tu elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Tu elemento `title` debe tener el título de tu receta.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Debes tener un elemento `meta` dentro de tu elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Tu elemento `meta` debe tener su atributo `charset` configurado a `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Debes tener un elemento `body` dentro de tu elemento `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Debes tener un elemento `h1` con el nombre de tu receta.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Debes tener solo un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debes tener un elemento `p` debajo de tu elemento `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Tu primer elemento `p` debe describir tu receta.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Tu primer elemento `h2` debe tener el texto `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Debes tener un elemento de lista no ordenada debajo de tu primer elemento `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Debes tener al menos cuatro elementos de lista en tu lista no ordenada con los ingredientes.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Tu segundo elemento `h2` debe tener el texto `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Debes tener un elemento de lista ordenada debajo de tu segundo elemento `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Debes tener al menos cuatro elementos de lista en tu lista ordenada con las instrucciones.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Debes tener al menos un elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

Todos tus elementos `img` deben tener un atributo `src` y un valor válidos.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Todos tus elementos `img` deben tener un atributo `alt` para describir la imagen.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
