---
id: 68e448ec4ac4205968c0a805
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

文中の `"love"` で `"I love learning!"` という単語が3回繰り返されているといいですね。

前のレッスンで覚えているように、`repeat()` メソッドを使うと文字列を指定した回数だけ繰り返すことができます。例を示します。

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

文字列は不変なので、このメソッドは元の文字列を変更しません。繰り返した内容を含む新しい文字列を返します。

`repeatedLove` という変数を作り、`"love ".repeat(3)` を代入してください。それから結果を確認できるように `repeatedLove` 変数をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`repeatedLove` という変数を用意してください。

```js
assert.exists(repeatedLove);
```

`repeatedLove` 変数は文字列であるべきです。

```js
assert.isString(repeatedLove);
```

`"love ".repeat(3)` を `repeatedLove` 変数に代入してください。

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

`repeatedLove` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
