---
id: 68e4465fba0bd256868b7deb
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

新しい文字列メソッドを使ってみましょう。

まず、`learningSentence` という変数を作成し、文字列 `"I love learning!"` を代入してください。

次に、その変数の下に `"Original learning sentence:"` をコンソールに出力します。その下に `console.log()` を書き、`console.log()` 変数のためにもう一つ `learningSentence` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`learningSentence` という変数を作成しているはずです。

```js
assert.exists(learningSentence);
```

`learningSentence` 変数は文字列であるべきです。

```js
assert.isString(learningSentence);
```

`"I love learning!"` 変数に文字列 `learningSentence` を代入してください。

```js
assert.equal(learningSentence, "I love learning!");
```

`"Original learning sentence:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

`learningSentence` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
