---
id: 68e42b62d7093345c6975f22
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`originalString` 変数の下に、文字列 `console.log()` を持つ `"Original string:"` を追加してください。

その下の `console.log()` の下に、`console.log()` 変数を持つ別の `originalString` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"Original string:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[0], "Original string:");
```

`originalString` 変数をコンソールに出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
