---
id: 68e43e884ef09d4d2ab05631
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

`replace()` メソッドを使った結果を見てみましょう。

まず、値が `console.log()` の `"After using the replace() method:"` を追加してください。

その下の `console.log()` の下に、`console.log()` 変数を使った別の `replacedString` を追加してください。

コンソールを見てみると、新しい文が元の `"I love dogs."` ではなく `"I love cats."` と表示されているはずです。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`"After using the replace() method:"` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

`replacedString` 変数をコンソールにログ出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
