---
id: 68e43f94365a964f26c88f78
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

新しい文字列メソッドを使ってみましょう。

まず、`exampleSentence` という新しい変数を作成し、文字列 `"I love cats and cats are so much fun!"` を代入してください。

次に、その変数の下に `"Original sentence:"` をコンソールに出力します。その下に `console.log()` を書き、`console.log()` 変数を使ってもう一つの `exampleSentence` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`exampleSentence` という変数を用意してください。

```js
assert.exists(exampleSentence);
```

`exampleSentence` 変数は文字列である必要があります。

```js
assert.isString(exampleSentence);
```

`"I love cats and cats are so much fun!"` 変数に文字列 `exampleSentence` を代入してください。

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

`"Original sentence:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[4], "Original sentence:");
```

`exampleSentence` 変数をコンソールに出力してください。

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
