---
id: 68e44519c769bf545bb3f22a
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

`"I love dogs and dogs are so much fun!"` の期待される出力を得るには、`replaceAll()` メソッドの代わりに `replace()` メソッドを使う必要があります。このメソッドは、部分文字列に一致するすべての箇所に対して新しい文字列を返します。

`exampleSentence.replace("cats", "dogs")` を更新して `replaceAll()` メソッドを使うようにしてください。これでコンソールに正しいテキストが表示されるはずです。

# --hints--

`exampleSentence.replace("cats", "dogs")` のコード行を更新して `replaceAll()` メソッドを使うようにしてください。

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
