---
id: 68e42de6396f5c49d90d1951
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

前のレッスンで、文字列の中の部分文字列を見つけて別の値に置き換えるために使う`replace()`メソッドについて学びました。

例を示します：

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

文字列は不変（immutable）であるため、この例では元の`text`は変更されません。代わりに新しい文字列が作成されます。

`replacedString`という新しい変数を作成し、`originalString.replace("cats", "dogs")`を代入してください。

# --hints--

`replacedString`という変数があるはずです。

```js
assert.exists(replacedString);
```

`replacedString`変数は文字列であるべきです。

```js
assert.isString(replacedString);
```

`originalString.replace("cats", "dogs")`を`replacedString`変数に代入してください。

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
