---
id: 68e44abea678a65b877864dc
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

その繰り返された単語が文の中に表示されるといいですね。`newSentence` という新しい変数を作成し、単語 `"I "` と `repeatedLove` 変数を連結し、その後に文字列 `" learning."` を連結した結果を代入してください。それから `newSentence` 変数をコンソールにログ出力してください。文字列の連結には `+` 演算子かテンプレートリテラルのどちらかを使うことができます。

これでテキスト `"I love love love  learning."` が表示されるはずです。

*注意*: 最後の `"love"` の後に余分なスペースがありますが、これは正しいです。最後のステップでその余分なスペースをトリムします。

# --hints--

`newSentence` という変数を用意してください。

```js
assert.exists(newSentence);
```

`newSentence` 変数は文字列である必要があります。

```js
assert.isString(newSentence);
```

単語 `"I "` と `repeatedLove` 変数を連結し、その後に文字列 `" learning."` を連結してください。この結果を `newSentence` 変数に代入します。スペースに注意してください。

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

`"I love love love  learning."` というリテラル文字列を `newSentence` 変数に代入してはいけません。正しい方法については指示をもう一度読んでください。

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
