---
id: 68e41084bb275539f72de9fe
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、`replace()`、`replaceAll()`、`repeat()` メソッドを使って文字列変換アプリを作りながら練習します。

まず、`originalString` という変数を作成し、`"I love cats."` の値を代入してください。

# --hints--

`originalString` という変数を用意してください。

```js
assert.exists(originalString);
```

`originalString` 変数は文字列である必要があります。

```js
assert.isString(originalString);
```

`"I love cats."` 変数に文字列 `originalString` を代入してください。

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
