---
id: 646c5ffef5598d449b52ec12
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Ahora debes trabajar en las orejas del gato. Habrá una oreja derecha y una izquierda, y dentro de cada una, habrá un oído interno.

Dentro de tu elemento `.cat-head`, crea un elemento `div` con la clase `cat-ears`.

# --hints--

No debes cambiar el elemento `div` existente con la clase `cat-head`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-head'), 1);
```

Debes crear un elemento `div` dentro de tu elemento `.cat-head`.

```js
assert.lengthOf(document.querySelectorAll('.cat-head div'), 1);
```

Tu elemento `div` debe tener la clase `cat-ears`.

```js
assert.equal(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class'),'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
