---
id: 646ce9d790d2a44de5f99e04
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Crea dos elementos `div`, el primero dentro del elemento `.cat-left-ear` con una clase de `cat-left-inner-ear`, y el segundo dentro del elemento `.cat-right-ear` con una clase de `cat-right-inner-ear`.

# --hints--

No debes cambiar el elemento `div` existente con la clase `cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-left-ear'), 1);
```

No debes cambiar el elemento `div` existente con la clase `cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('div.cat-right-ear'), 1);
```

Debes tener un elemento `div` dentro de tu elemento `.cat-left-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-left-ear div'), 1);
```

Debes tener un elemento `div` dentro de tu elemento `.cat-right-ear`.

```js
assert.lengthOf(document.querySelectorAll('.cat-right-ear div'), 1);
```

El nuevo elemento `div` dentro de `.cat-left-ear` debe tener la clase `cat-left-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

El nuevo elemento `div` dentro de `.cat-right-ear` debe tener la clase `cat-right-inner-ear`.

```js
assert.isTrue(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
