---
id: 646cf1206cac5f51804f49cf
title: Paso 25
challengeType: 0
dashedName: step-25
---

# --description--

Para posicionar correctamente la oreja izquierda, puedes usar la propiedad CSS transform para rotarla en cierto grado.

La propiedad `transform` te permite modificar la forma, posición y tamaño de un elemento sin cambiar el diseño ni afectar a los elementos circundantes. Tiene funciones como `translate()`, `rotate()`, `scale()`, `skew()` y `matrix()`.

Establece la propiedad `transform` a `rotate(-45deg)` y observa lo que ocurre.

# --hints--

Debes establecer la propiedad `transform` de tu elemento `.cat-left-ear` a `rotate(-45deg)`. No olvides agregar un punto y coma.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform, 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--
  
  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
