---
id: 6476fc5cf14b276e6d04e82a
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Usa un selector de clase para dar a tu elemento `.box` un ancho de `200px`, una altura de `600px` y un color de fondo de `#000`. Además, dale una `posición` de `absolute`, un `top` de `800px` y un `left` de `650px`.

# --hints--

Deberías tener un selector `.box`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.box'));
```

Tu selector `.box` debe tener una propiedad `width` establecida en `200px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.width, '200px');
```

Tu selector `.box` debe tener una propiedad `height` establecida en `600px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.height, '600px');
```

Tu selector `.box` debe tener una propiedad `background-color` establecida en `#000`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor, 'rgb(0, 0, 0)');
```

Tu selector `.box` debe tener una propiedad `position` establecida en `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.position, 'absolute');
```

Tu selector `.box` debe tener una propiedad `left` establecida en `650px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.left, '650px');
```

Tu selector `.box` debe tener una propiedad `top` establecida en `800px`.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.box')?.top, '800px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

