---
id: 6676a8a8d9861319ee901a90
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Ahora puedes comenzar a trabajar en las orejas del gato.

Limpia tu código de revisión eliminando el selector `.cat-right-ear` y todas sus propiedades.

Usando un selector de clase, da al elemento `.cat-left-ear` un borde izquierdo y derecho de `35px solid transparent` a cada uno. Además, establece el borde inferior en `70px solid #5e5e5e`.

# --hints--


Tu selector `.cat-right-ear` no debe tener una propiedad `height`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Tu selector `.cat-right-ear` no debe tener una propiedad `width`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


Tu selector `.cat-right-ear` no debe tener una propiedad `border-left`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Tu selector `.cat-right-ear` no debe tener una propiedad `border-right`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Tu selector `.cat-right-ear` no debe tener una propiedad `border-top`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Tu selector `.cat-right-ear` no debe tener una propiedad `border-bottom`.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

No debes tener un selector `.cat-right-ear`.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Debes tener un selector `.cat-left-ear`.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Tu selector `.cat-left-ear` debe tener una propiedad `border-left` establecida en `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Tu selector `.cat-left-ear` debe tener una propiedad `border-right` establecida en `35px solid transparent`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Tu selector `.cat-left-ear` debe tener una propiedad `border-bottom` establecida en `70px solid #5e5e5e`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
