---
id: 68f890b73bbfc9efa7379e7b
title: معالجة أخطاء إجبار النوع في تطبيق به أخطاء برمجية
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

انضممت للتو إلى شركة تطوير ويب محلية، ومهمتك الأولى هي تنظيف بعض الشفرات التي تحتوي على أخطاء برمجية تركها المطوّرون السابقون.

كانوا يحاولون إجراء عمليات حسابية بسيطة، لكن حدث خطأ ما، والنتائج غير منطقية. مهمتك مراجعة الشفرة، وتصحيح الأخطاء، وتعديلها لتعمل كما هو متوقع وتكون أسهل في القراءة.

**الهدف:** نفّذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك متغير اسمه `firstResult` يجمع الأعداد `5` و `10` بشكل صحيح ليعطي القيمة `15`.
2. يجب أن يكون لديك متغير اسمه `secondResult` يطرح العدد `5` من `8` بشكل صحيح ليعطي القيمة `3`.
3. يجب أن يكون لديك متغير اسمه `thirdResult` ينتج القيمة `6` بجمع عددين. استبدل القيمة المنطقية المستخدمة حاليًا في التعبير بعدد.
4. يجب أن يكون لديك متغير اسمه `fourthResult` ينتج القيمة `8` بجمع عددين. استبدل القيمة المنطقية المستخدمة حاليًا في التعبير بعدد.
5. يجب أن يكون لديك متغير اسمه `fifthResult` يضرب عددين بشكل صحيح ليحصل على حاصل الضرب `20`. استبدل السلسلة النصية المستخدمة حاليًا في التعبير بعدد.
6. يجب أن يكون لديك متغير اسمه `sixthResult` يجمع عددين بشكل صحيح ليعطي القيمة `22`. استبدل القيمة `null` المستخدمة حاليًا في التعبير بعدد.

# --hints--

يجب أن يكون لديك متغير اسمه `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

يجب أن يحتوي `firstResult` على القيمة الرقمية `15`.

```js
assert.strictEqual(firstResult, 15);
```

يجب أن يكون لديك متغير اسمه `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

يجب أن يحتوي `secondResult` على القيمة الرقمية `3`.

```js
assert.strictEqual(secondResult, 3);
```

يجب أن يكون لديك متغير اسمه `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

يجب أن يحتوي `thirdResult` على القيمة الرقمية `6`.

```js
assert.strictEqual(thirdResult, 6);
```

يجب ألا يتضمن `thirdResult` قيمة منطقية (`true` أو `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

يجب أن يكون لديك متغير اسمه `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

يجب أن يحتوي `fourthResult` على القيمة الرقمية `8`.

```js
assert.strictEqual(fourthResult, 8);
```

يجب ألا يتضمن `fourthResult` قيمة منطقية (`true` أو `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

يجب أن يكون لديك متغير اسمه `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

يجب أن يحتوي `fifthResult` على القيمة الرقمية `20`.

```js
assert.strictEqual(fifthResult, 20);
```

يجب ألا يتضمن `fifthResult` ضربًا بسلسلة نصية (مثل "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

يجب أن يكون لديك متغير اسمه `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

يجب أن يحتوي `sixthResult` على القيمة الرقمية `22`.

```js
assert.strictEqual(sixthResult, 22);
```

يجب ألا يتضمن `sixthResult` القيمة `null` في تعبيره.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
