---
id: 687ad7be1ff45032ccf1d92f
title: Створіть перемикач теми
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

У цій лабораторній роботі ви створите застосунок, який перемикається між різними темами. Коли ви перемикаєтеся на іншу тему, колір фону зміниться, а також відобразиться інше повідомлення.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути елемент `button` з текстом `Switch Theme`.
2. Ваш елемент `button` повинен мати такі атрибути:
   - Атрибут `id`, встановлений у `"theme-switcher-button"`.
   - Атрибут `aria-haspopup`, встановлений у `"true"`.
   - Атрибут `aria-expanded`, встановлений у `"false"`.
   - Атрибут `aria-controls`, встановлений у `"theme-dropdown"`.
3. У вас має бути елемент `ul` з такими атрибутами:
   - Атрибут `id`, встановлений у `"theme-dropdown"`.
   - Атрибут `role`, встановлений у `"menu"`.
   - Атрибут `aria-labelledby`, встановлений у `"theme-switcher-button"`.
   - Атрибут `hidden`.
4. Ваш елемент `ul` повинен містити щонайменше два елементи `li` з атрибутом `role`, встановленим у `"menuitem"`, і текстом на ваш вибір, що представляє різні теми.
5. Кожен із ваших елементів `li` повинен мати атрибут `id`, який починається з `theme-` і закінчується темою, яку ви встановили для тексту елемента `li`. Наприклад, якщо одна з ваших тем — `Light`, тоді ваш `id` повинен бути `theme-light`.
6. У вас має бути елемент з атрибутом `aria-live`, встановленим у `"polite"`.
7. У вас має бути масив `themes` з щонайменше двома об’єктами, кожен з яких містить властивості `name` і `message`. `name` представлятиме різну тему, а `message` відображатиметься при перемиканні теми. Ви можете придумати будь-які значення `name` і `message`, але значення `name` мають відповідати текстовому вмісту ваших елементів `li`.
8. Коли користувач натискає на `#theme-switcher-button`, елемент `#theme-dropdown` має відобразити доступні теми.
9. Коли елемент `#theme-dropdown` відображається, атрибут `aria-expanded` має бути встановлений у `"true"` для елемента `button`.
10. Коли користувач натискає на `#theme-switcher-button`, поки відображається елемент `#theme-dropdown`, меню має закритися, додати атрибут `hidden` і встановити атрибут `aria-expanded` у `"false"`.
11. Коли користувач натискає на `#theme-switcher-button` і вибирає тему, відповідний клас `theme-<name>` має бути доданий до елемента `body`, а пов’язане повідомлення теми, яке ви встановили в масиві `themes`, має відобразитися в елементі `aria-live="polite"`.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

У вас має бути елемент `button`.

```js
assert.exists(document.querySelector("button"));
```

Ваш елемент `button` повинен містити текст `Switch Theme`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Ваш елемент `button` повинен мати атрибут `id`, встановлений у `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Ваш елемент `button` повинен мати атрибут `aria-haspopup`, встановлений у `"true"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Ваш елемент `button` повинен мати атрибут `aria-expanded`, встановлений у `"false"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Ваш елемент `button` повинен мати атрибут `aria-controls`, встановлений у `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

У вас має бути елемент `ul`.

```js
assert.exists(document.querySelector("ul"));
```

Ваш елемент `ul` повинен мати атрибут `id`, встановлений у `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Ваш елемент `ul` повинен мати атрибут `role`, встановлений у `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Ваш елемент `ul` повинен мати атрибут `aria-labelledby`, встановлений у `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Ваш елемент `ul` повинен мати атрибут `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Ваш елемент `ul` повинен містити щонайменше два елементи `li`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Кожен із ваших елементів `li` повинен мати атрибут `role`, встановлений у `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Кожен із ваших елементів `li` повинен містити текст, що представляє тему.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Кожен із ваших елементів `li` повинен мати атрибут `id`, який починається з `theme-` і закінчується відповідною темою для цього елемента. Переконайтеся, що значення `id` повністю в нижньому регістрі.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

У вас має бути елемент з атрибутом `aria-live`, встановленим у `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

У вас має бути масив `themes`.

```js
assert.isArray(themes);
```

Ваш масив `themes` повинен містити щонайменше два об’єкти.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Ваш масив `themes` повинен містити щонайменше два об’єкти, кожен з яких має властивості `name` і `message`.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

Значення властивостей `name` і `message` мають бути рядками, які не є порожніми.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Кожне значення властивості `name` має відповідати текстовому вмісту одного з ваших елементів `li`. Переконайтеся, що ці значення повністю в нижньому регістрі.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Коли користувач натискає на `#theme-switcher-button`, щоб відобразити варіанти тем, атрибут `hidden` має бути видалений з елемента `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Коли елемент `#theme-dropdown` відображається, атрибут `aria-expanded` має бути встановлений у `"true"` для елемента `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Коли користувач натискає на `#theme-switcher-button`, поки відображається елемент `#theme-dropdown`, випадаюче меню має повернутися до `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Коли користувач натискає на `#theme-switcher-button`, поки відображається елемент `#theme-dropdown`, атрибут `aria-expanded` має бути встановлений у `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Коли користувач натискає на `#theme-switcher-button` і вибирає тему, відповідний клас `theme-<name>` має бути доданий до елемента `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Коли користувач натискає на `#theme-switcher-button` і вибирає тему, у елементі `message` має відобразитися `theme`, пов’язаний із вибраним `themes` із масиву `aria-live="polite"`.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
