---
id: a3bfc1673c0526e06d3ac698
title: Erstellen Sie einen Rechner für die Summe von Primzahlen
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

In diesem Labor erstellen Sie einen Rechner, der eine Zahl entgegennimmt und die Summe aller Primzahlen zurückgibt, die kleiner oder gleich dieser Zahl sind.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `sumPrimes`-Funktion haben, die eine Zahl als Argument akzeptiert.
1. Die `sumPrimes`-Funktion sollte die Summe aller Primzahlen zurückgeben, die kleiner oder gleich der angegebenen Zahl sind.
1. Wenn die Eingabezahl kleiner als `2` ist, sollte die Funktion `0` zurückgeben.

# --hints--

Sie sollten eine `sumPrimes`-Funktion haben.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` sollte `17` zurückgeben.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` sollte `10` zurückgeben.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` sollte `2` zurückgeben.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` sollte `0` zurückgeben.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` sollte `73156` zurückgeben.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
