---
id: 68f1f8c248306da204655117
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O atributo `muted` silenciará o áudio na reprodução inicial. Se você tiver `controls` ativado, o usuário poderá ativar o áudio. Omitir o atributo `muted` fará o áudio ser reproduzido na reprodução inicial.

O atributo `muted` é um atributo booleano e não precisa de um valor.

Adicione o atributo `muted` ao elemento `video`.

# --hints--

Seu elemento `video` deve ter o atributo `muted`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
