---
id: 68f1f371fefd448b4f70b898
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Em uma lição anterior, você aprendeu sobre diferentes atributos disponíveis para o elemento `video`. O atributo `width` determina a largura do vídeo em pixels.

Adicione o atributo `width` ao elemento `video` com um valor de `640`.

# --hints--

Seu elemento `video` deve ter o atributo `width` com o valor `640`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
