---
id: 68f1f6f97af9749a42260a43
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O atributo `controls` fornece controles de reprodução incluindo reprodução, retrocesso e controle de volume para o elemento `video`. 

O atributo `controls` é um atributo booleano e não precisa de um valor.

Adicione o atributo `controls` ao elemento `video`.

Agora você deve ver o elemento de vídeo exibido na página.

# --hints--

Seu elemento `video` deve ter o atributo `controls`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
  </body>
</html>
```
