---
id: 68f1eb9ad0a7575e4dbc31d5
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Em seguida, crie um elemento `video` abaixo do `h1`. Nos próximos passos, você adicionará os atributos necessários para tornar o player de vídeo funcional. 

# --hints--

Você deve ter um elemento `video`.

```js
assert.exists(document.querySelector('video'));
```

O novo elemento `video` deve estar abaixo do `h1`.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
