---
id: 68f18724d71db1dda8d6fec7
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você irá construir um player de vídeo em HTML. O boilerplate HTML foi fornecido para você.

Crie um elemento `h1` e dê a ele o texto `Working with the HTML Video Element`.

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Seu elemento `h1` deve conter o texto `Working with the HTML Video Element`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
