---
id: 68f1f57723debf9383c75947
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

O atributo `loop` reiniciará o vídeo assim que a reprodução for concluída. Pense em um meme da internet que repete a reprodução. Omitir o atributo `loop` fará com que o vídeo seja reproduzido uma vez.

O atributo `loop` é um atributo booleano e não precisa de um valor.

Adicione o atributo `loop` ao elemento `video`.

# --hints--

Seu elemento `video` deve ter o atributo `loop`.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
  </body>
</html>
```
