---
id: 68f1fac259a4afabb478b960
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Você pode ter notado que não vinculou ao vídeo real. Você fará isso na próxima fase. Quando se trata de tipos de arquivos de vídeo, existem diferenças no suporte dos navegadores. Para acomodar isso, você pode usar elementos `source` dentro do elemento `video` e o navegador selecionará o primeiro `source` compatível.

Aqui está um exemplo de um elemento `source`:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

O elemento `source` é um elemento vazio, portanto não possui uma tag de fechamento. 

Adicione um elemento `source` dentro do seu elemento `video`. 


# --hints--

Você deve ter um elemento `source` dentro do seu elemento `video`.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```
