---
id: 68f1f993460f06a64c033061
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

O atributo `poster` é uma imagem em miniatura do vídeo. Pense nos vídeos que você assiste no YouTube. Ele é exibido enquanto o vídeo está sendo baixado. Se o atributo for omitido, o primeiro quadro do vídeo é mostrado durante a fase de download.

Agora, adicione o atributo `poster` com o valor `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` ao seu elemento `video`.

# --hints--

Seu elemento `video` deve ter o atributo `poster` com o valor `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video
    width="640"
    loop
    controls
    muted
  >
--fcc-editable-region--
  </video>
  </body>
</html>
```
