---
id: 68f26a5f79e97c5032c10642
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Você usou um arquivo de vídeo com a extensão de arquivo `mp4` e precisa informar o navegador para que ele saiba como ler o arquivo.

Você usará o atributo `type` para especificar o tipo MIME `video/mp4`.

MIME (Multipurpose Internet Mail Extensions) é um padrão para descrever documentos em outras formas além de texto ASCII, por exemplo, áudio, vídeo e imagens.

MP4, formalmente conhecido como MPEG-4 Part 14, é um formato de contêiner multimídia digital. Ele é amplamente usado para armazenar vídeo e áudio e também pode incluir outros tipos de dados como legendas e imagens estáticas. Os arquivos MP4 são projetados para streaming pela Internet e são compatíveis com muitos dispositivos e plataformas.

Agora, adicione o atributo `type` e o valor `video/mp4`.

# --hints--

Seu elemento `source` deve ter um atributo `type`.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

Seu elemento `source` deve ter um atributo `type` com o valor `video/mp4`.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"

  >
--fcc-editable-region--
  </video>
</body>
</html>
```
