---
id: 68f26d52b13e0f5b6e2d2e09
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Outro tipo MIME comum é o tipo MIME `video/webm`.

WebM é um formato de arquivo de mídia audiovisual de código aberto desenvolvido pelo Google, projetado principalmente para conteúdo de mídia baseado na web. Ele suporta codecs de vídeo como VP8, VP9 e AV1 e codecs de áudio como Vorbis e Opus, tornando-se uma escolha popular para os elementos de vídeo e áudio do HTML5.

Abaixo do seu primeiro elemento `source`, adicione outro elemento `source` e dê a ele um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` e um atributo `type` com o valor `video/webm`.

# --hints--

Você deve ter um segundo elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

Seu segundo elemento `source` deve ter um atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

Seu segundo elemento `source` deve ter um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

Seu segundo elemento `source` deve ter um atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

Seu segundo elemento `source` deve ter um atributo `type` com o valor `video/webm`.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
