---
id: 68f27139bddf9f6bf2821acf
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

O último elemento `source` que você adicionará será para o tipo MIME `video/quicktime`.

QuickTime é uma arquitetura multimídia extensível criada pela Apple, que suporta reprodução, streaming, codificação e transcodificação de vários formatos de mídia digital. Não tão popular quanto o formato MP4, você pode precisar dele para suporte a aplicações legadas.

Abaixo do seu terceiro elemento `source`, adicione um quarto elemento `source` e dê a ele um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` e um atributo `type` com o valor `video/quicktime`.

Parabéns! Você completou o Workshop do HTML Video Player.

# --hints--

Você deve ter um quarto elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

Seu quarto elemento `source` deve ter um atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

Seu quarto elemento `source` deve ter um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

Seu quarto elemento `source` deve ter um atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

Seu quarto elemento `source` deve ter um atributo `type` com o valor `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
