---
id: 68f2701b2160c765f23b60f8
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Outro tipo MIME comum é o tipo MIME `video/ogg`.

Ogg é um formato de contêiner multimídia digital projetado para fornecer streaming eficiente e manipulação de multimídia digital. É mantido pela Xiph.Org Foundation e é gratuito e aberto, sem restrições de patentes de software. Seu nome é derivado de "ogging", jargão do jogo de computador Netrek.

Abaixo do seu segundo elemento `source`, adicione um terceiro elemento `source` e dê a ele um atributo `src` com o valor `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` e um atributo `type` com o valor `video/ogg`.

# --hints--

Você deve ter um terceiro elemento `source`.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

Seu terceiro elemento `source` deve ter um atributo `src`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

Seu terceiro elemento `source` deve ter um atributo `src` com um valor de `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

Seu terceiro elemento `source` deve ter um atributo `type`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

Seu terceiro elemento `source` deve ter um atributo `type` com o valor `video/ogg`.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
