---
id: 66c07238b01053abaf812065
title: Crie uma calculadora de fatorial
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Um fatorial Ă© o produto de um nĂșmero inteiro e todos os nĂșmeros inteiros abaixo dele. Por exemplo, o fatorial de `5` Ă© `5 * 4 * 3 * 2 * 1 = 120`. Neste laboratĂłrio, vocĂȘ criarĂĄ uma calculadora de fatorial que recebe um nĂșmero do usuĂĄrio e calcula o fatorial desse nĂșmero.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratĂłrio.

**HistĂłrias de UsuĂĄrio:**

1. VocĂȘ deve declarar uma variĂĄvel `num` e atribuir a ela um nĂșmero de sua escolha. O nĂșmero atribuĂ­do deve estar entre `1` e `20`, inclusive.
1. Crie uma funĂ§ĂŁo chamada `factorialCalculator` que recebe um nĂșmero como argumento e retorna o fatorial desse nĂșmero.
1. Dentro da funĂ§ĂŁo, declare uma variĂĄvel `result` e atribua a ela o valor `1`. Usando um loop, percorra todos os nĂșmeros de `1` atĂ© o nĂșmero de entrada (inclusive) e, para cada nĂșmero, multiplique a variĂĄvel `result` pelo nĂșmero atual e atribua o resultado Ă  variĂĄvel `result`. VocĂȘ pode optar por usar um loop `for`, `while` ou `do...while` aqui.
1. VocĂȘ deve chamar a funĂ§ĂŁo `factorialCalculator` com `num` como argumento e atribuir o resultado Ă  variĂĄvel `factorial`.
1. VocĂȘ deve armazenar o resultado final no formato `Factorial of [num] is [factorial]` e atribuĂ­-lo Ă  variĂĄvel `resultMsg`.
1. VocĂȘ deve enviar o valor de `resultMsg` para o console.

# --hints--

VocĂȘ deve ter uma variĂĄvel `num`.

```js
assert.isDefined(num);
```

VocĂȘ deve atribuir um valor Ă  variĂĄvel `num`

```js
assert.isNotNull(num);
```

O valor de `num` deve ser entre `1` e `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

O valor de `num` deve ser um nĂșmero.

```js
assert.isNumber(num);
```

VocĂȘ deve ter uma funĂ§ĂŁo chamada `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

A funĂ§ĂŁo `factorialCalculator` deve receber um nĂșmero como argumento.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

O fatorial de `5` deve retornar `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

O fatorial de `7` deve retornar `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

VocĂȘ deve chamar a funĂ§ĂŁo `factorialCalculator` com a variĂĄvel `num` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

VocĂȘ deve definir uma variĂĄvel `factorial` e atribuir a ela o resultado da funĂ§ĂŁo `factorialCalculator`.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

`factorialCalculator` deve produzir o resultado correto.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Seu `resultMsg` deve conter uma string no formato `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

VocĂȘ deve enviar o valor de `resultMsg` para o console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
