---
id: 66c07238b01053abaf812065
title: Crie uma calculadora de fatorial
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Um fatorial é o produto de um número inteiro e todos os números inteiros abaixo dele. Por exemplo, o fatorial de `5` é `5 * 4 * 3 * 2 * 1 = 120`. Neste laboratório, você criará uma calculadora de fatorial que recebe um número do usuário e calcula o fatorial desse número.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve declarar uma variável `num` e atribuir a ela um número de sua escolha. O número atribuído deve estar entre `1` e `20`, inclusive.
1. Crie uma função chamada `factorialCalculator` que recebe um número como argumento e retorna o fatorial desse número.
1. Dentro da função, declare uma variável `result` e atribua a ela o valor `1`. Usando um loop, percorra todos os números de `1` até o número de entrada (inclusive) e, para cada número, multiplique a variável `result` pelo número atual e atribua o resultado à variável `result`. Você pode optar por usar um loop `for`, `while` ou `do...while` aqui.
1. Você deve chamar a função `factorialCalculator` com `num` como argumento e atribuir o resultado à variável `factorial`.
1. Você deve armazenar o resultado final no formato `Factorial of [num] is [factorial]` e atribuí-lo à variável `resultMsg`.
1. Você deve enviar o valor de `resultMsg` para o console.

# --hints--

Você deve ter uma variável `num`.

```js
assert.isDefined(num);
```

Você deve atribuir um valor à variável `num`

```js
assert.isNotNull(num);
```

O valor de `num` deve ser entre `1` e `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

O valor de `num` deve ser um número.

```js
assert.isNumber(num);
```

Você deve ter uma função chamada `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

A função `factorialCalculator` deve receber um número como argumento.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

O fatorial de `5` deve retornar `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

O fatorial de `7` deve retornar `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

Você deve chamar a função `factorialCalculator` com a variável `num` como argumento.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

Você deve definir uma variável `factorial` e atribuir a ela o resultado da função `factorialCalculator`.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

`factorialCalculator` deve produzir o resultado correto.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Seu `resultMsg` deve conter uma string no formato `Fatorial de [num] é [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

Você deve enviar o valor de `resultMsg` para o console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
