---
id: 6a57c331ae94fa4631b60866
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Manchmal fragen Mitarbeitende nach einem Artefakt, das das Museum nicht besitzt. Wenn unter dieser `id` kein Artefakt gespeichert ist, ist `collection[id]` `undefined`, und das Auslesen von `.title` daraus würde einen Fehler werfen.

Aktualisieren Sie die return-Anweisung, um eine Bedingung zu verwenden. Geben Sie `artifact.title` zurück, wenn ein Artefakt gefunden wird, und den String `"Artifact not found"`, wenn nicht.

# --hints--

`getArtifactTitle(101)` sollte weiterhin `"Golden Mask"` zurückgeben.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` sollte `"Artifact not found"` zurückgeben, wenn kein Artefakt zur `id` passt.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` sollte den Titel jedes gefundenen Artefakts zurückgeben, nicht nur von `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
