---
id: 6a57c331ae94fa4631b60870
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Erstellen Sie schließlich eine lesbare Zusammenfassung für jedes Artefakt. Beginnen Sie damit, den Fall zu behandeln, in dem ein Artefakt fehlt.

Erstellen Sie eine `buildSummary`-Funktion, die einen `id`-Parameter entgegennimmt. Verwenden Sie `collection[id]`, um das Artefakt nachzuschlagen, und speichern Sie es in einer Variablen namens `artifact`. Wenn das Artefakt nicht gefunden wird, geben Sie den String `"Artifact not found"` zurück.

# --hints--

Sie sollten eine Funktion namens `buildSummary` mit einem Parameter namens `id` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Sie sollten `collection[id]` einer Variablen namens `artifact` zuweisen.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.artifact);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.artifact?.value.matches("collection[id]"));
```

`buildSummary` sollte `"Artifact not found"` zurückgeben, wenn kein Artefakt mit der `id` übereinstimmt.

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

`buildSummary` sollte nicht `"Artifact not found"` zurückgeben, wenn das Artefakt existiert.

```js
assert.notStrictEqual(buildSummary(101), "Artifact not found");
assert.notStrictEqual(buildSummary(102), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
