---
id: 6a57c331ae94fa4631b60872
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Geben Sie als Nächstes ein Template Literal zurück, das das Artefakt zusammenfasst. Setzen Sie den Titel in die erste Zeile, gefolgt von jeweils einer Zeile für die Kategorie, den Namen der Kuratorin oder des Kurators, die aktuelle Galerie und den Ausstellungsstatus. Verwenden Sie für die Galerie `currentLocation.gallery`.

Wenn Sie zum Beispiel `buildSummary(101)` aufrufen, sollte es Folgendes zurückgeben:

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)` sollte den Titel des Artefakts, `Golden Mask`, enthalten.

```js
assert.include(buildSummary(101), "Golden Mask");
```

`buildSummary(101)` sollte eine Zeile `Category: Ceremonial` enthalten.

```js
assert.include(buildSummary(101), "Category: Ceremonial");
```

`buildSummary(101)` sollte eine Zeile `Curator: Fran Sinclair` enthalten.

```js
assert.include(buildSummary(101), "Curator: Fran Sinclair");
```

`buildSummary(101)` sollte eine Zeile `Current Gallery: Hall C` enthalten, wobei `currentLocation.gallery` verwendet wird.

```js
assert.include(buildSummary(101), "Current Gallery: Hall C");
```

`buildSummary(101)` sollte eine Zeile `On Display: true` enthalten.

```js
assert.include(buildSummary(101), "On Display: true");
```

`buildSummary(101)` sollte jedes Feld in eine eigene Zeile setzen.

```js
assert.isAtLeast(buildSummary(101).split("\n").length, 5);
```

`buildSummary(102)` sollte den aktuellsten Standort verwenden und `Current Gallery: Hall B` anzeigen.

```js
assert.include(buildSummary(102), "Current Gallery: Hall B");
```

`buildSummary` sollte die Zusammenfassung aus den eigenen Werten eines beliebigen Artefakts erstellen, jeweils ein Feld pro Zeile und in der richtigen Reihenfolge.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [
    { gallery: "Hall Y", year: 2001 },
    { gallery: "Hall Z", year: 2002 },
  ],
  tags: [],
  onDisplay: false,
};
const summary = buildSummary(457);
assert.deepEqual(
  summary.split("\n").map((line) => line.trim()),
  [
    "Test Artifact",
    "Category: Test Category",
    "Curator: Test Curator",
    "Current Gallery: Hall Z",
    "On Display: false",
  ]
);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
