---
id: 6a57c331ae94fa4631b60869
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Neue Forschungen bestätigen, dass die Goldene Maske zur königlichen Familie gehörte. Fügen Sie dem Artefakt mit der `id` `101` den Tag `"royal"` hinzu. Rufen Sie dann `console.log()` mit `collection[101].tags` auf, um die Tags dieses Artefakts in der Konsole anzuzeigen.

# --hints--

Sie sollten `addTag` mit `101` und `"royal"` aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

Sie sollten die `tags` des ersten Artefakts in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Sie sollten die `tags` nach dem Hinzufügen des neuen Tags protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Die `tags` des ersten Artefakts sollten `"royal"` enthalten.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
