---
id: 6a57c331ae94fa4631b6086c
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Im Laufe der Zeit werden Artefakte häufig ein- und ausgeblendet.

Erstellen Sie eine `toggleDisplayStatus`-Funktion, die einen `id`-Parameter entgegennimmt. Suchen Sie das Artefakt mit `collection[id]`. Wenn es existiert, setzen Sie seine `onDisplay`-Eigenschaft auf das Gegenteil seines aktuellen Werts.

# --hints--

Sie sollten eine Funktion namens `toggleDisplayStatus` mit einem Parameter namens `id` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Das Aufrufen von `toggleDisplayStatus` sollte den `onDisplay`-Wert des Artefakts in beide Richtungen umschalten.

```js
const before101 = collection[101].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, !before101);
toggleDisplayStatus(101);
assert.strictEqual(collection[101].onDisplay, before101);
const before102 = collection[102].onDisplay;
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, !before102);
toggleDisplayStatus(102);
assert.strictEqual(collection[102].onDisplay, before102);
```

`toggleDisplayStatus` sollte keinen Fehler werfen oder die Sammlung ändern, wenn kein Artefakt mit der `id` übereinstimmt.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => toggleDisplayStatus(999));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

--fcc-editable-region--

--fcc-editable-region--
```
