---
id: 6a57c331ae94fa4631b60865
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Mitarbeitende suchen dutzende Male am Tag nach Artefakten, daher braucht Ihre Prüferin eine wiederverwendbare Suchfunktion. Die eckigen Klammern können auch eine Variable aufnehmen, sodass `collection[id]` genau das Artefakt zurückgibt, auf das sich `id` bezieht.

Erstellen Sie eine Funktion namens `getArtifactTitle`, die einen `id`-Parameter entgegennimmt. Verwenden Sie `collection[id]`, um das Artefakt zu erhalten, speichern Sie es in einer Variablen namens `artifact` und geben Sie `artifact.title` zurück.

# --hints--

Sie sollten eine Funktion namens `getArtifactTitle` mit einem Parameter namens `id` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Sie sollten `collection[id]` einer Variablen namens `artifact` zuweisen.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` sollte `artifact.title` zurückgeben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` sollte `"Golden Mask"` zurückgeben.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` sollte `"Bronze Tablet"` zurückgeben.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
