---
id: 6a57c331ae94fa4631b60861
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jedes Artefakt hat eine verantwortliche Person.

Fügen Sie dem Artefakt, das unter dem Schlüssel `101` gespeichert ist, eine `curator`-Eigenschaft hinzu, und setzen Sie deren Wert auf ein Objekt mit einer `id` von `201` und einem `name` von `"Earl Sinclair"`.

# --hints--

Das unter `101` gespeicherte Artefakt sollte eine `curator`-Eigenschaft haben.

```js
assert.property(collection[101], "curator");
```

Der `curator` sollte ein Objekt sein.

```js
assert.isObject(collection[101].curator);
```

Der `curator` sollte eine `id` von `201` haben.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

Der `curator` sollte einen `name` mit dem Wert `"Earl Sinclair"` haben.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
