---
id: 6a57c331ae94fa4631b60864
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Die Überprüfung beginnt damit, das bereits gespeicherte zu lesen. Da jedes Artefakt unter seiner ID gespeichert ist, können Sie auf eines zugreifen, indem Sie diese ID in eckige Klammern setzen, zum Beispiel `collection[101]`.

Rufen Sie `console.log()` mit `collection[101].title` auf, um den Titel des ersten Artefakts zu protokollieren. Rufen Sie dann `console.log()` mit `collection[101].curator.name` auf, um den Namen seines Kurators zu protokollieren.

# --hints--

Sie sollten den `title` des ersten Artefakts in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

Sie sollten den Kurator-`name` des ersten Artefakts in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

Sie sollten zuerst den `title` und dann den Kurator-`name` protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
