---
id: 6a57c331ae94fa4631b60862
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Jedes Artefakt verfolgt auch, wo es sich befunden hat, wie es beschrieben wird und ob es ausgestellt ist.

Fügen Sie dem unter `101` gespeicherten Artefakt drei weitere Eigenschaften hinzu:

- `locations`: ein Array mit zwei Objekten, `{ gallery: "Hall A", year: 2020 }` und `{ gallery: "Hall C", year: 2024 }`
- `tags`: ein Array mit den Strings `"gold"` und `"egypt"`
- `onDisplay`: der Boolean `true`

# --hints--

Die `locations` des Artefakts sollten ein Array der beiden angegebenen Standortobjekte sein.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

Die `tags` des Artefakts sollten `["gold", "egypt"]` sein.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

Das `onDisplay` des Artefakts sollte `true` sein.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
