---
id: 6a57c265abbc95dfebc05098
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie einen Auditor für die Herkunft von Artefakten. Sie speichern Museumsartefakte in einem Objekt und schreiben Funktionen, um sie nachzuschlagen, zu taggen, zwischen Galerien zu verschieben und ihre Geschichte zusammenzufassen.

Jedes Artefakt wird unter seiner ID gespeichert. Erstellen Sie ein Objekt namens `collection` und fügen Sie einen Eintrag hinzu. Setzen Sie den Schlüssel auf `101` und den Wert auf ein Objekt mit einem `title` von `"Golden Mask"` und einer `category` von `"Ceremonial"`.

# --hints--

Sie sollten eine Variable namens `collection` deklarieren.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` sollte ein Objekt sein.

```js
assert.isObject(collection);
```

`collection` sollte einen Eintrag haben.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

Das unter `101` gespeicherte Artefakt sollte ein Objekt sein.

```js
assert.isObject(collection[101]);
```

Das unter `101` gespeicherte Artefakt sollte einen `title` von `"Golden Mask"` haben.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

Das unter `101` gespeicherte Artefakt sollte eine `category` von `"Ceremonial"` haben.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
