---
id: 6a57c331ae94fa4631b60868
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Kuratoren versehen Artefakte mit Tags, um sie leichter durchsuchbar zu machen, und ein Tag sollte niemals doppelt hinzugefügt werden. Bevor Sie einen neuen Tag hinzufügen, können Sie die Methode `includes()` verwenden, um zu prüfen, ob er bereits existiert, da sie `true` zurückgibt, wenn ein Array bereits einen Wert enthält.

Erstellen Sie eine Funktion namens `addTag`, die die Parameter `id` und `tag` entgegennimmt. Suchen Sie das Artefakt mit `collection[id]`. Wenn es existiert und sein `tags`-Array `tag` noch nicht enthält, fügen Sie `tag` zu `artifact.tags` hinzu.

# --hints--

Sie sollten eine Funktion namens `addTag` mit den Parametern `id` und `tag` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Ein Aufruf von `addTag` mit einem neuen Tag sollte diesen Tag zum `tags`-Array des Artefakts hinzufügen.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Ein Aufruf von `addTag` mit einem Tag, das das Artefakt bereits hat, sollte keinen doppelten Eintrag hinzufügen.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` sollte den Tag zu dem Artefakt hinzufügen, auf das sich die `id` bezieht.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`addTag` sollte keinen Fehler werfen oder die Sammlung ändern, wenn kein Artefakt zur `id` passt.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

Sie sollten `tag` an das bestehende `tags`-Array des Artefakts anhängen, anstatt das Array zu ersetzen.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
