---
id: 66c057041df6394ca796bf33
title: crea un generatore di frasi
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

In questo laboratorio, creerai due storie diverse usando un modello di frase. Userai variabili per archiviare parti diverse della storia e poi stamperai le storie nella console.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti dichiarare le seguenti variabili usando `let`:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. Dovresti assegnare alle variabili sopra alcuni valori stringa a tua scelta.

1. Dovresti dichiarare una variabile `firstStory`.
   
1. Dovresti usare il seguente modello di storia per creare la prima storia e assegnarla alla variabile `firstStory`: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. Dovresti stampare la tua prima storia nella console usando il messaggio `"First story: [firstStory]"`.

1. Dovresti assegnare nuovi valori alle variabili `adjective`, `noun`, `verb`, `place`, `adjective2` e `noun2`.

1. Dovresti dichiarare una variabile `secondStory`.

1. Crea un’altra storia usando lo stesso modello e assegnala alla variabile `secondStory`.

1. Dovresti stampare la tua seconda storia nella console usando il messaggio `"Second story: [secondStory]"`.


# --hints--

Dovresti dichiarare una variabile `adjective`.

```js
assert.isNotNull(adjective);
```

Dovresti dichiarare una variabile `noun`.

```js
assert.isNotNull(noun);
```

Dovresti dichiarare una variabile `verb`.

```js
assert.isNotNull(verb);
```

Dovresti dichiarare una variabile `place`.

```js
assert.isNotNull(place);
```

Dovresti dichiarare una variabile `adjective2`.

```js
assert.isNotNull(adjective2);
```

Dovresti dichiarare una variabile `noun2`.

```js
assert.isNotNull(noun2);
```

Dovresti assegnare un valore stringa alla variabile `adjective`.

```js
assert.isString(adjective);
```

Dovresti assegnare un valore stringa alla variabile `noun`.

```js
assert.isString(noun);
```

Dovresti assegnare un valore stringa alla variabile `verb`.

```js
assert.isString(verb);
```

Dovresti assegnare un valore stringa alla variabile `place`.

```js
assert.isString(place);
```

Dovresti assegnare un valore stringa alla variabile `adjective2`.

```js
assert.isString(adjective2);
```

Dovresti assegnare un valore stringa alla variabile `noun2`.

```js
assert.isString(noun2);
```

Dovresti dichiarare una variabile `firstStory`.

```js
assert.isNotNull(firstStory);
```

Dovresti usare il formato corretto della storia per la prima storia: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Fai attenzione agli spazi.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Dovresti assemblare la tua prima storia usando le variabili dichiarate nell’ordine corretto.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Dovresti stampare la tua prima storia usando il messaggio `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Dovresti dichiarare una variabile `secondStory`.

```js
assert.isNotNull(secondStory);
```

Dovresti riassegnare la variabile `adjective` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Dovresti riassegnare la variabile `noun` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Dovresti riassegnare la variabile `verb` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Dovresti riassegnare la variabile `place` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Dovresti riassegnare la variabile `adjective2` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Dovresti riassegnare la variabile `noun2` per la seconda storia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Dovresti usare il formato corretto della storia per la seconda storia: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Fai attenzione agli spazi.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Dovresti assemblare la tua seconda storia usando le variabili dichiarate nell’ordine corretto.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Dovresti stampare la tua seconda storia usando il formato `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

La `firstStory` non dovrebbe essere uguale alla `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
