---
id: 6852db755c9c1274d75aae9a
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, crearás un extractor de PIN. Los dígitos del PIN están ocultos en cada línea de un poema.

Comienza creando una función llamada `pin_extractor` con un parámetro `poem`. Una función sin cuerpo genera un error, así que agrega `pass` dentro de la función como un marcador de posición temporal para que sea válida.

# --hints--

Debes tener una función llamada `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

La función `pin_extractor` debe tener un parámetro llamado `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
