---
id: 6852dd8486c7ee0171322950
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Necesitarás una variable para almacenar el código secreto mientras lo descifras, así que elimina `pass` y crea una variable llamada `secret_code` dentro de la función y asígnale una cadena vacía.

# --hints--

Debes tener una variable `secret_code` dentro de la función `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

La variable `secret_code` debe tener un valor de una cadena vacía.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

Ya no deberías tener `pass` en tu código.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
