---
id: 68530231c05b93324e071104
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Mientras trabajas en la función, sería útil tener una llamada a la función para propósitos de depuración, así que crea una variable `poem`, fuera de la función, que contenga este poema:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

Puedes usar una cadena de múltiples líneas para eso.

Y luego, llama a la función con la variable `poem` como argumento.

# --hints--

Debes tener una variable global llamada `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

La variable `poem` debe tener un valor igual a `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

Debes llamar a la función `pin_extractor` con `poem` como argumento.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
