---
id: 68531fbb934d735415053274
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora que has agregado una llamada a función, puedes ver cualquier error creado dentro de la función en la terminal, y siempre puedes usar `print` para ver el valor de las cosas.

Hay un dígito del pin oculto en alguna línea, así que dentro de la función usa el método `split` para dividir la cadena en una lista de líneas. Divide las líneas usando el carácter de nueva línea (`\n`), y asigna la lista resultante a una variable llamada `lines`.

# --hints--

Debes tener una variable llamada `lines` dentro de tu función `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

La variable `lines` debe tener un valor de `poem.split('\n')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    
--fcc-editable-region--
poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
