---
id: 68559f3561d7df287586142e
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Ahora el extractor de pines funciona, ¡pero podría ser más eficiente! ¡Podría extraer el pin de muchos poemas al mismo tiempo!

Agrega dos poemas más en el ámbito global, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` y `poem3 = 'There\nonce\nwas\na\ndragon'`, y por ahora, comenta el `print` con la llamada a `pin_extractor`.


# --hints--

Las variables `poem2` y `poem3` deben existir y tener el valor correcto.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

Debes comentar la llamada `print(pin_extractor(poem))`.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
