---
id: 685535bf527b2c77236b3a9a
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

El dígito n-ésimo del pin está oculto como la longitud de la palabra n-ésima en la línea n-ésima.

Para encontrar la longitud de la palabra n, el siguiente paso es separar la línea del poema en una lista de palabras: dentro del ciclo, crea una variable `words` y asígnale el valor de `line` dividido en palabras usando el método `split`.

Luego, agrega una llamada a `print` con `words` como argumento.

# --hints--

Debes tener una variable `words` dentro del ciclo.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

La variable `words` debe tener un valor de `line.split()`.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

Deberías tener `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
