---
id: 6855918e87c33c6979d21d3f
title: Paso 14
challengeType: 20
dashedName: step-14
---

# --description--

Ahora la función está haciendo un buen trabajo extrayendo el código del poema, pero hay un caso límite que considerar: la línea del poema podría ser más corta de lo esperado, y eso haría que la función falle (puedes probar esto eliminando la palabra `bright` del poema).

Pon la línea `secret_code += str(len(words[line_index]))` dentro de una estructura `if` que verifique que hay suficientes palabras en la lista `words`.

# --hints--

Dentro del ciclo debes tener una sentencia `if`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

La condición de la sentencia if debe ser `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Dentro de la sentencia `if` debes tener `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
