---
id: 685578aae65e4106ad2b4300
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Como aprendiste en una lección anterior, la función `enumerate` permite llevar un registro del índice mientras se itera sobre un iterable:

```py
fruits = ['apple', 'plum', 'bananas'] 

for index, item in enumerate(fruits):  
    print(index, item)  

# 0 apple  
# 1 plum
# 2 bananas
```

El dígito n-ésimo del código secreto proviene de la palabra n-ésima de la línea n-ésima, así que necesitas saber cuál es el número de cada línea.

Cambia el ciclo para que itere sobre `enumerate(lines)`, y agrega otra variable de ciclo antes de `line` llamada `line_index`.

Además, actualiza la llamada `print` a `print(line_index, line)`.

# --hints--

El ciclo debe iterar sobre `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

El ciclo debe tener `line_index, line` como variable del ciclo.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

Debes tener `print(line_index, line)` dentro del ciclo.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
