---
id: 68553423d775cc60134ac8bb
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Necesitas trabajar en cada línea de forma independiente: crea un ciclo sobre `lines` que use `line` como variable del ciclo. Dentro del ciclo, imprime `line`.

# --hints--

Debes tener un ciclo `for` en tu función `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

El ciclo debe iterar sobre `lines`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

El ciclo debe tener `line` como variable del ciclo (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

Debes tener `print(line)` dentro del ciclo.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
