---
id: 6859290fc8e07e6b2a8d8523
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Elimina las dos llamadas a `print` y agrega una sentencia `return` que devuelva `secret_code` como última línea en la función.

# --hints--

No debes tener una llamada a `print`.

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

La función `pin_extractor` debería devolver `secret_code` al final.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)` debería devolver `5262`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
