---
id: 660ee6e3a242da6bd579de69
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

واحد من أهم المفاهيم في البرمجة هو المتغيرات. يشير <dfn>المتغير</dfn> إلى عنوان ذاكرة محدد يخزن قيمة. تُعطى المتغيرات اسمًا يمكن استخدامه في كودك للوصول إلى تلك القيمة.

الإعلان عن متغير يعني إعطاؤه اسمًا. في JavaScript، يتم ذلك غالبًا باستخدام الكلمة المفتاحية `let`. على سبيل المثال، هكذا تعلن عن متغير `hello`:

```js
let hello;
```

تتبع تسمية المتغيرات قواعد محددة: يمكن أن تشمل الأسماء حروفًا وأرقامًا وعلامات الدولار والشرطات السفلية، لكنها لا يمكن أن تحتوي على فراغات ولا يجب أن تبدأ برقم.

استخدم الكلمة المفتاحية `let` لإعلان متغير اسمه `character`.

_ملاحظة_: من الممارسات الشائعة إنهاء العبارات في JavaScript بفاصلة منقوطة `;`

# --hints--

يجب استخدام `let` في كودك.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

يجب استخدام `character` في كودك.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

يجب أن تستخدم `let` لإعلان متغير `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

يجب أن ينتهي إعلانك بفاصلة منقوطة.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
