---
id: 660ef19b95d3308e7dd31bb6
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

تسمح لك <dfn>وحدة التحكم</dfn> بطباعة وعرض ناتج JavaScript. يمكنك إرسال معلومات إلى وحدة التحكم باستخدام `console.log()`. على سبيل المثال، ستطبع هذه الشفرة `"Naomi"` في وحدة التحكم:

```js
let developer = "Naomi";
console.log(developer);
```

تصل الشفرة أعلاه إلى المتغير `developer` باسمه داخل `console.log()`. لاحظ أن القيمة بين الأقواس العادية هي القيمة التي ستُطبع.

اطبع قيمة المتغير `character` في وحدة التحكم. ثم، انقر على زر "وحدة التحكم" لعرض وحدة تحكم JavaScript.

# --hints--

لا يجب تغيير تهيئة المتغير `character`.

```js
assert.strictEqual(character, 'Hello');
```

يجب أن تصل إلى `console` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

يجب أن تصل إلى الخاصية `log` في `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

يجب أن تستخدم الأقواس العادية لاستدعاء الطريقة `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

يجب أن تطبع المتغير `character` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
