---
id: 660eebd83100d37862268781
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

المتغير `character` الخاص بك لا يحتوي حاليًا على قيمة. يمكنك تعيين قيمة باستخدام معامل <dfn>التعيين</dfn> `=`. على سبيل المثال:

```js
let hello = "Hello";
```

تعيين قيمة لمتغير عند لحظة إعلانه يُعرف بـ <dfn>تهيئة</dfn>.

هيئ متغير `character` الخاص بك بتعيين القيمة `"Hello"` أثناء إعلانه.

# --hints--

يجب أن تستخدم معامل التعيين `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

يجب أن تستخدم السلسلة النصية `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

يجب أن تستخدم علامات اقتباس مزدوجة لسلسلة `"Hello"` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

يجب أن تعين القيمة `"Hello"` إلى متغير `character` الخاص بك.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
