---
id: 660ef5c1904955978a986a5c
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

استخدم `let` لإعلان متغير `count`. عيّن له <dfn>العدد</dfn> `8`. عند استخدام قيمة عددية، لا تستخدم علامات الاقتباس. على سبيل المثال:

```js
let money = 100;
```

# --hints--

يجب أن يكون لديك كلمة `let` ثانية في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

يجب أن تستخدم `let` لإعلان متغير `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

يجب أن تعيّن العدد `8` إلى متغير `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
