---
id: 660ef857f2806aa626d29d17
title: الخطوة 17
challengeType: 1
dashedName: step-17
---

# --description--

يمكنك أيضًا إجراء الطرح (`-`) والضرب (`*`) والقسمة (`/`). جرب استخدام العوامل والأرقام في `console.log` الخاص بك. عندما تكون مستعدًا للانتقال، احذف `console.log`.

# --hints--

يجب ألا يحتوي شفرتك على `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

يجب أن تظل قيمة المتغير `character` هي `"Hello"`.

```js
assert.equal(character, "Hello");
```

يجب أن تظل قيمة المتغير `count` هي `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
