---
id: 660ef6355e8f5a9e67fe5f46
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

باستخدام نوع البيانات `number`، يمكنك إجراء عمليات رياضية مثل الجمع. جرّب طباعة `count + 1` في وحدة التحكم.

# --hints--

يجب الوصول إلى `console` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

يجب الوصول إلى `.log()` الخاصة بـ `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

يجب تسجيل `count + 1` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
