---
id: 660ef0f7c4b8e68ccd1f0786
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

يحتوي JavaScript على سبعة أنواع بيانات أولية، وتُعد سلسلة نصية واحدة منها. في JavaScript، تمثل <dfn>السلسلة النصية</dfn> متتالية من الأحرف ويمكن تضمينها بين علامات اقتباس مفردة (`'`) أو مزدوجة (`"`).

لاحظ أن السلاسل النصية <dfn>غير قابلة للتغيير</dfn>، مما يعني أنه بمجرد إنشائها لا يمكن تعديلها. مع ذلك، يمكن إعادة تعيين المتغير لقيمة أخرى.

غيّر سلسلة `"Hello"` لتستخدم علامات اقتباس مفردة.

# --hints--

يجب ألا تحتوي شفرتك على علامات اقتباس مزدوجة.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

يجب أن تستخدم علامات اقتباس مفردة لسلسلة `"Hello"` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

يجب أن تستمر في استخدام `let` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

يجب أن تستمر في إعلان متغير `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

يجب أن يحتوي متغير `character` على السلسلة النصية `"Hello"` كقيمته.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
