---
id: 660ef55dd468079679ee0092
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

أنت الآن جاهز لإعلان متغيرك التالي. احذف كلا تعبيرَي `console.log`، وحذف إعادة تعيين `character`.

احذف أيضًا متغير `secondCharacter`، لكن اترك تهيئة `character` كما هي.

# --hints--

يجب ألا يحتوي شفرتك على أي تعبيرات `console.log()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

يجب ألا تعيد تعيين قيمة `character`.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

يجب أن يحتوي متغير `character` على القيمة `"Hello"`.

```js
assert.equal(character, "Hello");
```

يجب ألا يكون لديك متغير `secondCharacter`.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
