---
id: 660ef5105b8ba095307a0e50
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

سجل متغير `character` في وحدة التحكم مرة أخرى. يجب أن ترى السلسلة النصية `"Hello"`، ثم السلسلة النصية `"World"` في وحدة التحكم.

# --hints--

يجب أن تستخدم `console.log` للمرة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

يجب أن تسجل `character` في وحدة التحكم للمرة الثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


يجب أن يأتي `console.log()` الجديد بعد إعادة التعيين.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
