---
id: 660ef31a5be625914a0102cd
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

عندما تُعلن عن متغير باستخدام الكلمة المفتاحية `let`، يمكنك <dfn>إعادة تعيين</dfn> (أو تغيير قيمة) ذلك المتغير لاحقًا. في هذا المثال، تغيرت قيمة `programmer` من `"Naomi"` إلى `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

لاحظ أنه عند إعادة تعيين متغير، لا تستخدم الكلمة المفتاحية `let` مرة أخرى.

بعد استخدام `console.log`، عيّن القيمة `"World"` إلى المتغير `character`.

# --hints--

يجب أن تستخدم `character` مرتين في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

يجب أن تستخدم معامل التعيين لإعادة تعيين `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

يجب أن تستخدم السلسلة النصية `"World"` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

يجب أن يحتوي متغير `character` على القيمة `"World"`.

```js
assert.equal(character, "World");
```

يجب ألا تستخدم `let` في إعادة التعيين.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
