---
id: 660f5179b3b0ca558f6b4d4f
title: الخطوة 113
challengeType: 1
dashedName: step-113
---

# --description--

استخدم عبارة `if` للتحقق مما إذا كانت `inverted` صحيحة. تذكر أنك لا تحتاج إلى استخدام معامل المساواة هنا.

# --hints--

يجب أن تستخدم عبارة `if`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

يجب أن تتحقق عبارة `if` مما إذا كانت `inverted` تساوي `true`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

يجب ألا تستخدم شرط `if` أي معاملات مقارنة.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
