---
id: 660f3b664421471aa595170f
title: الخطوة 74
challengeType: 1
dashedName: step-74
---

# --description--

يبدو الهرم غريبًا قليلًا الآن. بما أنك تبدأ الحلقة عند `1` بدلًا من `0`، فلا تحتاج إلى إضافة واحد إلى `i` عند تمريره إلى `padRow`.

حدّث المعلمة الأولى في استدعاء `padRow` لتكون `i`.

# --hints--

لا يجب أن تمرر `i + 1` إلى `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

يجب أن تمرر `i` إلى `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
