---
id: 660f38c34a4de6141c0c369f
title: الخطوة 71
challengeType: 1
dashedName: step-71
---

# --description--

مولد الهرم الخاص بك يعمل الآن كما هو متوقع. لكنها فرصة ممتازة لاستكشاف الشفرة التي كتبتها بشكل أعمق.

عامل الجمع ليس الطريقة الوحيدة لإضافة قيم إلى متغير. يمكن استخدام عامل <dfn>التعيين الجمعي</dfn> كاختصار يعني "خذ القيمة الأصلية للمتغير، أضف هذه القيمة، ثم عيّن النتيجة مرة أخرى إلى المتغير." على سبيل المثال، هذان البيانان سيعطيان نفس النتيجة:

```js
test = test + 1;
test += 1;
```

حدّث عبارة التكرار في حلقة `for` لاستخدام التعيين الجمعي.

# --hints--

يجب ألا تستخدم حلقة `for` التعبير `i = i + 1`؛

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

يجب أن تستخدم حلقة `for` التعيين الجمعي مع `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
