---
id: 660f48e1d3682f39e81843c4
title: خطوة 94
challengeType: 1
dashedName: step-94
---

# --description--

مولد الهرم الخاص بك لا يزال يعمل. مع ذلك، قد يحدث أن تدخل في حلقة لا نهائية مرة أخرى.

بما أنك تتحقق فقط مما إذا كان `done` لا يساوي `count`، فإذا كان `done` أكبر من `count` فإن الحلقة ستستمر إلى الأبد.

حدّث شرط الحلقة لتتحقق مما إذا كان `done` أقل من أو يساوي `count`.

# --hints--

يجب أن تتحقق حلقة `while` مما إذا كان `done` أقل من أو يساوي `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
