---
id: 660f1b6e60bd9edf902c81fd
title: الخطوة 36
challengeType: 1
dashedName: step-36
---

# --description--

جملة <dfn>التكرار</dfn> ستخبر الحلقة بما يجب فعله مع المكرر بعد كل دورة.

عندما تعيد تعيين متغير، يمكنك استخدام المتغير كمرجع للقيمة السابقة قبل التعيين الجديد. هذا يسمح لك بفعل أشياء مثل إضافة ثلاثة إلى رقم موجود. على سبيل المثال، `bees = bees + 3` سيزيد قيمة `bees` بمقدار ثلاثة.

استخدم تلك الصيغة لتحل محل سلسلة `"iteration"` بجملة تعيين تزيد `i` بمقدار واحد.

# --hints--

يجب أن تضيف واحدًا إلى متغير `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

يجب أن تعيّن `i + 1` مرة أخرى إلى متغير `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

يجب أن تزيد حلقة `for` قيمة `i` بمقدار `1` في كل تكرار.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
