---
id: 660f46460f9c36330ebc07d8
title: الخطوة 88
challengeType: 1
dashedName: step-88
---

# --description--

يمكن أن يؤدي معامل المساواة إلى سلوك غريب في JavaScript. على سبيل المثال، `"0" == 0` صحيح، رغم أن أحدهما سلسلة نصية والآخر رقم.

يُستخدم <dfn>معامل المساواة الصارمة</dfn> `===` للتحقق مما إذا كانت قيمتان متساويتين وتشتركان في نفس النوع. كقاعدة عامة، هذا هو معامل المساواة الذي يجب أن تستخدمه دائمًا. مع معامل المساواة الصارمة، تصبح `"0" === 0` خاطئة، لأنه رغم أن قيمتهما صفر، إلا أنهما ليسا من نفس النوع.

حدّث شرط `done == count` ليستخدم معامل المساواة الصارمة.

# --hints--

يجب أن يستخدم شرط `if` معامل المساواة الصارمة.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
