---
id: 660f2b6fd54ac1fc142804dd
title: الخطوة 47
challengeType: 1
dashedName: step-47
---

# --description--

من المحتمل أن تصبح منطق تنسيق هذا الهرم معقدًا، مما يعني أن الوقت مناسب لاستخراج تلك الشفرة إلى دالة.

<dfn>الدالة</dfn> هي كتلة شفرة يمكن إعادة استخدامها في جميع أنحاء التطبيق. تُعلن الدوال باستخدام الصيغة النحوية التالية:

```js
function name(parameter) {

}
```

الكلمة المفتاحية `function` تخبر JavaScript أن المتغير `name` سيكون دالة. المتغير `parameter` يمثل قيمة تُمرر إلى الدالة عند استخدامها. يمكن أن تحتوي الدالة على عدد من <dfn>المعلمات</dfn> كما تريد، سواء كثيرة أو قليلة. مثل حلقة `for`، المساحة بين الأقواس المتعرجة هي <dfn>جسم الدالة</dfn>.

أعلن عن دالة `padRow`. لا تنشئ أي متغيرات معلمة بعد. يجب أن يكون جسم الدالة فارغًا. تذكر أن تستخدم نمط camel case في اصطلاح التسمية.

# --hints--

يجب أن تستخدم الكلمة المفتاحية `function`.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

يجب أن تعلن عن دالة `padRow`.

```js
assert.isFunction(padRow);
```

يجب ألا تحتوي دالة `padRow()` على أي معلمات.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

يجب أن يكون جسم دالة `padRow()` فارغًا.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
