---
id: 660f0c34aad72dc712b97624
title: الخطوة 27
challengeType: 1
dashedName: step-27
---

# --description--

طريقة أخرى أساسية لهذا المشروع هي الطريقة `.pop()`. تزيل العنصر الأخير من المصفوفة وتُعيد ذلك العنصر.

عندما تُعيد الطريقة قيمة، يمكنك اعتبارها تعيد القيمة إليك، مما يجعلها متاحة للاستخدام في أجزاء أخرى من شفرتك.

أنشئ متغيرًا جديدًا اسمه `popped` وقم بتعيين نتيجة `rows.pop()` له. ثم، سجّل `popped` في وحدة التحكم.

# --hints--

يجب إعلان متغير اسمه `popped`.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

يجب استخدام `let` لإعلان المتغير المسمى `popped`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

يجب استدعاء الطريقة `.pop()`.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

يجب استدعاء الطريقة `.pop()` على مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

يجب تسجيل المتغير `popped`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
