---
id: 660f3915b41a441537ec9f5e
title: الخطوة 72
challengeType: 1
dashedName: step-72
---

# --description--

بما أنك تزيد `i` فقط بمقدار `1`، يمكنك استخدام <dfn>معامل التزايد</dfn> `++`. هذا المعامل يزيد قيمة المتغير بمقدار 1، محدثًا التعيين لذلك المتغير. على سبيل المثال، `test` ستصبح `8` هنا:

```js
let test = 7;
test++;
```

استبدل تعيين الجمع الخاص بك بمعامل التزايد لتكرار الحلقة.

# --hints--

يجب ألا يستخدم حلقة `for` تعيين الجمع مع `i`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

يجب أن تستخدم حلقة `for` معامل التزايد على `i`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
