---
id: 660f1cedf3676fe26122ebf6
title: الخطوة 38
challengeType: 1
dashedName: step-38
---

# --description--

يجب أن ترى الأرقام من صفر إلى سبعة مطبوعة في وحدة التحكم، رقم واحد في كل سطر. هذا سيكون الأساس لإنشاء الهرم الخاص بك.

استبدل عبارة الطباعة الخاصة بك بعبارة تدفع `i` إلى مصفوفة `rows` الخاصة بك.

# --hints--

يجب ألا يكون لديك استدعاء `console.log`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

يجب أن تستدعي `.push()` على مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

يجب أن تدفع `i` إلى مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

يجب أن يحدث استدعاء `.push()` داخل حلقة `for` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
