---
id: 660f061d259bbebc37461080
title: الخطوة 20
challengeType: 1
dashedName: step-20
---

# --description--

يمكنك الوصول إلى القيم داخل المصفوفة باستخدام <dfn>فهرس</dfn> القيمة. الفهرس هو رقم يمثل موقع القيمة في المصفوفة، ويبدأ من `0` للقيمة الأولى.

يمكنك الوصول إلى القيمة باستخدام <dfn>ترميز الأقواس</dfn>، مثل `array[0]`.

استخدم `console.log` وترميز الأقواس لطباعة القيمة الأولى في مصفوفة `rows` الخاصة بك.

# --hints--

يجب أن تحتوي شفرتك على عبارة `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

يجب أن تصل إلى مصفوفة `rows` الخاصة بك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

يجب أن تستخدم ترميز الأقواس مع مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

يجب أن تستخدم ترميز الأقواس للوصول إلى العنصر الأول في مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

يجب أن تسجل العنصر الأول في مصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
