---
id: 660f0da9bf1035c9097af20a
title: الخطوة 28
challengeType: 1
dashedName: step-28
---

# --description--

يجب أن تكون قد رأيت `"freeCodeCamp"` مطبوعة في وحدة التحكم. هذا لأن `.pop()` تُعيد القيمة التي أُزيلت من المصفوفة - وقد دفعت `"freeCodeCamp"` إلى نهاية المصفوفة سابقًا.

لكن ماذا تُعيد `.push()`؟ عيّن `rows.push()` الموجود لديك إلى متغير جديد اسمه `pushed`، ثم اطبعه.

# --hints--

يجب إعلان متغير `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

يجب استخدام `let` لإعلان متغير `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

يجب تعيين `rows.push("freeCodeCamp")` إلى متغير `pushed`.

```js
assert.equal(pushed, 4);
```

يجب طباعة متغير `pushed`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
