---
id: 660f34626216270c682e2f7b
title: الخطوة 66
challengeType: 1
dashedName: step-66
---

# --description--

يسمح لك <dfn>استدعاء الدالة</dfn> باستخدام الدالة فعليًا. قد لا تكون مدركًا لذلك، لكن الطرق مثل `.push()` التي كنت تستخدمها هي استدعاءات دوال.

تُستدعى الدالة بالإشارة إلى اسم الدالة، وإضافة `()`. إليك كيفية استدعاء دالة `test`:

```js
test();
```

استبدل `character.repeat(i + 1)` في استدعاء `.push()` الخاص بك باستدعاء دالة لـ `padRow`. لا تضف أي معلمات إليها بعد.

# --hints--

يجب ألا تستخدم `i + 1` في استدعاء `push` الخاص بك.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

يجب ألا تستخدم `character.repeat` في استدعاء `.push()` الخاص بك.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

يجب أن تستدعي `padRow` في استدعاء `.push()` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

يجب ألا تحتوي استدعاءات `padRow` على معلمات.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
