---
id: 660f3ce51f70571e1c5227c8
title: الخطوة 76
challengeType: 1
dashedName: step-76
---

# --description--

يمكن أن تساعد التعليقات في شرح سبب اختيار شفرتك نهجًا معينًا، أو ترك ملاحظات للمهام التي يجب إنجازها لنفسك في المستقبل.

في JavaScript، يمكنك استخدام `//` لترك تعليق في سطر واحد داخل شفرتك.

أضف تعليقًا في سطر واحد فوق حلقة `for` لتذكير نفسك بتغيير الشفرة إلى نوع مختلف من الحلقات.

# --hints--

يجب أن تبدأ تعليق السطر الواحد بـ `//`.

```js
assert.match(code, /\/\//);
```

يجب أن يكون تعليق السطر الواحد الخاص بك طوله خمسة أحرف على الأقل.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
