---
id: 660f359af3e32e0f1a6880b7
title: الخطوة 68
challengeType: 1
dashedName: step-68
---

# --description--

يجب أن ترى الآن نفس مجموعة الأحرف في وحدة التحكم. دالة `padRow` تقوم بنفس الشيء الذي كنت تفعله سابقًا، لكنها الآن في قسم قابل لإعادة الاستخدام خاص بها.

استخدم عامل الجمع لربط فراغ واحد `" "` في بداية ونهاية سلسلة `character` المكررة.

تذكر أنه يمكنك استخدام العامل `+` لربط السلاسل النصية بهذا الشكل:

```js
" " + "string"
```

# --hints--

يجب ربط فراغ واحد في بداية القيمة التي تُرجعها.

```js
assert.match(padRow(1, 1), /^\s/);
```

يجب ربط فراغ واحد في نهاية القيمة التي تُرجعها.

```js
assert.match(padRow(1, 1), /\s$/);
```

يجب أن تُرجع دالة `padRow()` سلسلة `character` المكررة مع فراغ قبل وبعد السلسلة.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
