---
id: 660f3dd626be3a1ffe27e5d1
title: الخطوة 77
challengeType: 1
dashedName: step-77
---

# --description--

تدعم JavaScript أيضًا التعليقات متعددة الأسطر. يبدأ التعليق متعدد الأسطر بـ `/*` وينتهي بـ `*/`.

على عكس التعليق أحادي السطر، يغطي التعليق متعدد الأسطر عدة أسطر.

استخدم `/*` و `*/` لتحويل حلقة `for` الحالية، بما في ذلك جسمها، إلى تعليق متعدد الأسطر.

# --hints--

يجب أن تبدأ التعليق متعدد الأسطر بـ `/*`.

```js
assert.match(code, /\/\*/);
```

يجب أن تنهي التعليق متعدد الأسطر بـ `*/`.

```js
assert.match(code, /\*\//);
```

يجب تعليق كامل حلقة `for` الخاصة بك.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
