---
id: 660f255022991ef34ed0ee88
title: الخطوة 44
challengeType: 1
dashedName: step-44
---

# --description--

طباعة الأرقام لن تُنتج هرمًا جذابًا بصريًا. الآن بعد أن تُخرج المحتوى المنسق لمصفوفة `rows`، حان الوقت لتحديث الحلقة الأصلية.

بدلًا من دفع `i` إلى المصفوفة، ادفع قيمة متغير `character`.

# --hints--

يجب ألا تدفع متغير `i` بعد الآن.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

يجب أن تدفع متغير `character`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
