---
id: 660f530d6e33d159e1bf4947
title: خطوة 115
challengeType: 1
dashedName: step-115
---

# --description--

إذا لم تكن الهرم مقلوبًا، فستحتاج إلى وجود كتلة `else` تبني الهرم بالترتيب الطبيعي.

في الخطوات السابقة، تعلمت كيف تتعامل مع جملة `else` بهذا الشكل:

```js
if (condition) {
  // if condition is true, run this code
} else {
  // if condition is false, run this code
}
```

أضف كتلة `else` إلى كتلة `if` الخاصة بك.

# --hints--

يجب أن تضيف كتلة `else`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{/);
```

يجب أن تكون كتلة `else` فارغة.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)\s*\{\s*rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)\s*\);\s*\}\s*else\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  }

}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
