---
id: 660f505d02b2bd513a1c3468
title: الخطوة 111
challengeType: 1
dashedName: step-111
---

# --description--

لم تعد هرمك مقلوبًا. هذا لأنك تضيف صفوفًا جديدة إلى **نهاية** المصفوفة.

حدّث جسم الحلقة لإضافة صفوف جديدة إلى بداية المصفوفة.

# --hints--

يجب أن تستخدم طريقة `unshift` الخاصة بـ `rows`.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

يجب أن تمرر استدعاء `padRow()` كمعلمة لطريقة `.unshift()` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

يجب أن تمرر `i` كالمعلمة الأولى لاستدعاء `padRow()` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

يجب أن تمرر `count` كالمعلمة الثانية لاستدعاء `padRow()` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
