---
id: 660f415b76859a2736771607
title: الخطوة 78
challengeType: 1
dashedName: step-78
---

# --description--

اختفى الهرم الخاص بك مرة أخرى. لا بأس - هذا متوقع.

قبل أن تنشئ الحلقة الجديدة، تحتاج إلى تعلم عن جمل `if`. تسمح لك <dfn>جملة `if`</dfn> بتنفيذ كتلة شفرة فقط عندما يتحقق شرط معين. تستخدم الصيغة النحوية التالية:

```js
if (condition) {
  logic
}
```

أنشئ جملة `if` مع القيمة المنطقية `true` كشرط. في جسم الجملة، اطبع السلسلة النصية `"Condition is true"`.

# --hints--

يجب أن تنشئ جملة `if`.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

يجب أن تحتوي جملة `if` على `true` كشرط.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

يجب أن يسجل جسم جملة `if` السلسلة النصية `"Condition is true"`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
