---
id: 660f4f79e2a82a4e92290f44
title: الخطوة 110
challengeType: 1
dashedName: step-110
---

# --description--

أحيانًا قد ترغب في استرجاع شفرة سابقة علّقت عليها. يمكنك فعل ذلك بإزالة `/*` و`*/` حول تلك الشفرة. يُسمى هذا <dfn>إلغاء التعليق</dfn>.

ألغِ تعليق أول حلقة `for` فقط. اترك تعليق السطر الواحد والتعليقين المتعددين الآخرين كما هما.

# --hints--

يجب أن تلغي تعليق أول حلقة `for` لديك.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

يجب ألا تزيل تعليق السطر الواحد.

```js
assert.match(code, /\/\//);
```

يجب ألا تلغي تعليق حلقة `while`.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

يجب ألا تلغي تعليق الحلقة الثانية `for`.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
