---
id: 660f447efc0e722f016c1be0
title: الخطوة 81
challengeType: 1
dashedName: step-81
---

# --description--

ظهر النص مرة أخرى! هذا لأن `"false"` هي سلسلة نصية، وعند تقييمها كقيمة منطقية تصبح `true`. هذا يعني أن `"false"` هي قيمة صحيحة.

<dfn>القيمة الصحيحة</dfn> هي قيمة تُعتبر صحيحة عند تقييمها كقيمة منطقية. معظم القيم التي تواجهها في JavaScript ستكون صحيحة.

<dfn>القيمة الخاطئة</dfn> هي العكس - قيمة تُعتبر خاطئة عند تقييمها كقيمة منطقية. لدى JavaScript قائمة محددة من القيم الخاطئة. بعض منها تشمل `false`، `0`، `""`، `null`، `undefined`، و `NaN`.

جرّب تغيير شرط `if` إلى سلسلة نصية فارغة `""`، وهي قيمة خاطئة.

# --hints--

يجب أن يحتوي شرط `if` على `""`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
