---
id: 660f4cffb1459d45e34902d1
title: الخطوة 106
challengeType: 1
dashedName: step-106
---

# --description--

يمكنك فعليًا بناء الهرم المقلوب دون الحاجة إلى التكرار "عكسيًا" كما فعلت.

للقيام بذلك، ستحتاج إلى تعلم بعض الطرق الجديدة للمصفوفة. ابدأ باستخدام `const` لإعلان متغير `numbers`. عيّن له مصفوفة تحتوي على العناصر `1` و `2` و `3`. ثم سجّل مصفوفة `numbers`.

# --hints--

يجب أن تستخدم `const` لإعلان متغير `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

يجب أن يكون متغير `numbers` مصفوفة.

```js
assert.isArray(numbers);
```

يجب أن تحتوي مصفوفة `numbers` على العناصر `1` و `2` و `3` بهذا الترتيب.

```js
assert.deepEqual(numbers, [1,2,3]);
```

يجب أن تسجل مصفوفة `numbers`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
