---
id: 660f4de78f775e480ba2e451
title: الخطوة 107
challengeType: 1
dashedName: step-107
---

# --description--

تسمح طريقة `.unshift()` في المصفوفة بإضافة قيمة إلى **بداية** المصفوفة، على عكس `.push()` التي تضيف القيمة في نهاية المصفوفة. تُرجع `.unshift()` الطول الجديد للمصفوفة التي نُفذت عليها.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

استخدم `const` لإعلان متغير `unshifted`، وعيّن له نتيجة استدعاء `.unshift()` على مصفوفة `numbers` الخاصة بك. مرر `5` كمعلمة. ثم اطبع متغير `unshifted`.

# --hints--

يجب إعلان متغير `unshifted` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

يجب استدعاء طريقة `.unshift()` على مصفوفة `numbers` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

يجب تمرير `5` كمعلمة لاستدعاء `.unshift()` الخاص بك.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

يجب تعيين نتيجة استدعاء `.unshift()` إلى متغير `unshifted` الخاص بك.

```js
assert.equal(unshifted, 4);
```

يجب طباعة متغير `unshifted` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
