---
id: 660f17294346b7d69e79db3d
title: خطوة 32
challengeType: 1
dashedName: step-32
---

# --description--

أنت الآن جاهز لبدء بناء مولد الهرم الخاص بك. المتغير `character` سيكون وحدة البناء للهرم.

قد لا يعمل `"Hello"` جيدًا لهذا الغرض. غيّر قيمة `character` لتكون رمز الهاش `"#"`.

# --hints--

يجب أن يكون المتغير `character` رمز هاش.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
