---
id: 660f17d4e9f227d86e834abd
title: الخطوة 33
challengeType: 1
dashedName: step-33
---

# --description--

لإنشاء هرم، تحتاج إلى تكوين عدة صفوف. عندما تضطر إلى تنفيذ مهمة بشكل متكرر حتى يتحقق شرط معين، ستستخدم <dfn>حلقة تكرار</dfn>. هناك عدة طرق لكتابة حلقة تكرار.

ستبدأ بحلقة `for` أساسية. تستخدم حلقات `for` الصيغة النحوية التالية:

```js
for (iterator; condition; iteration) {
  logic;
}
```

في الخطوات القادمة، ستستكشف كل مكوِّن من مكوِّنات الحلقة بالتفصيل. الآن، أنشئ حلقة `for` تتضمن المصطلحات `"iterator"`، و`"condition"`، و`"iteration"` للمكوِّنات الثلاثة. اجعل <dfn>جسم</dfn> الحلقة، القسم داخل الأقواس المتعرجة `{}`، فارغًا.


# --hints--

يجب أن يكون لديك حلقة `for`.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

المكوِّن الأول في حلقة `for` يجب أن يكون السلسلة النصية `"iterator"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

المكوِّن الثاني في حلقة `for` يجب أن يكون السلسلة النصية `"condition"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

المكوِّن الثالث في حلقة `for` يجب أن يكون السلسلة النصية `"iteration"`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

يجب أن يكون جسم حلقة `for` فارغًا.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
