---
id: 660f34e99571070d56d2f231
title: الخطوة 67
challengeType: 1
dashedName: step-67
---

# --description--

دالتك `padRow` تحتوي على معلمتين قمت بتعريفهما. تُزوَّد القيم لتلك المعلمات عند استدعاء الدالة.

القيم التي تزودها لاستدعاء دالة تُسمى <dfn>المعلمات</dfn>، وأنت <dfn>تمرر</dfn> المعلمات إلى استدعاء الدالة. هذا مثال لاستدعاء دالة مع تمرير `"Hello"` كمعلمة:

```js
test("Hello");
```

مرر `i + 1` و `count` كمعلمتين لاستدعاء دالتك `padRow`. مثل المعلمات، تُفصل المعلمات بفاصلة.

# --hints--

يجب أن تمرر `i + 1` إلى استدعاء `padRow()` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

يجب أن تضع فاصلة بعد المعلمة `i + 1`.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

يجب أن تمرر `count` كمعلمتك الثانية.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
