---
id: 660f07d231941bc11719f664
title: الخطوة 21
challengeType: 1
dashedName: step-21
---

# --description--

المصفوفات مميزة لأنها تُعتبر <dfn>قابلة للتغيير</dfn>. هذا يعني أنه يمكنك تغيير القيمة عند فهرس معين مباشرة.

على سبيل المثال، هذا الشفرة ستُعيّن الرقم `25` للعنصر الثاني في المصفوفة:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

حدّث العنصر **الثالث** في مصفوفة `rows` ليصبح الرقم `10`. ثم اطبع مصفوفة `rows` في وحدة التحكم.

# --hints--

يجب أن تستخدم ترميز الأقواس على مصفوفة `rows` مرة أخرى.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

يجب أن تصل إلى العنصر الثالث في مصفوفة `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

يجب أن تستخدم معامل التعيين على العنصر الثالث في مصفوفة `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

يجب أن تُعيّن القيمة `10` للعنصر الثالث في مصفوفة `rows`.

```js
assert.equal(rows[2], 10);
```

يجب أن يكون لديك بيان `console.log` ثاني في شفرتك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

يجب أن تسجل مصفوفة `rows`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
