---
id: 660f165270622fd4ec0da3f7
title: الخطوة 31
challengeType: 1
dashedName: step-31
---

# --description--

إعلان متغير باستخدام الكلمة المفتاحية `let` يسمح بإعادة تعيينه. هذا يعني أنه يمكنك تغيير `character` لاحقًا ليأخذ قيمة مختلفة تمامًا.

لهذا المشروع، لن ترغب في تغيير قيم هذه المتغيرات. لذلك، يجب عليك استخدام `const` لإعلانها. المتغيرات المعلنة بـ `const` خاصة.

أولًا، لا يمكن إعادة تعيين متغير `const` مثل متغير `let`. هذا الشفرة ستُرمي خطأ:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

أيضًا، لا يمكن أن يكون متغير `const` غير مهيأ. هذا الشفرة ستُرمي خطأ:

```js
const firstName;
```

استبدل كلمات `let` المفتاحية بـ `const`.

# --hints--

يجب استخدام `const` لإعلان متغير `character`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

يجب استخدام `const` لإعلان متغير `count`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

يجب استخدام `const` لإعلان متغير `rows`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

يجب عدم استخدام `let` في شفرتك.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
