---
id: 660f3ba3cceef11b6ba08b59
title: الخطوة 75
challengeType: 1
dashedName: step-75
---

# --description--

للأسف، اختفى الآن أسفل الهرم. هذا لأنك أنشأت <dfn>خطأ بمقدار واحد</dfn> آخر.

كان تكرارك الأصلي يعمل لقيم `i` من `0` إلى `7`، لأن `count` تساوي `8` وشرطك يتطلب أن تكون قيمة `i` أصغر من `count`. الآن يعمل التكرار لقيم `i` من `1` إلى `7`.

يجب تحديث التكرار ليعمل أيضًا عندما تكون قيمة `i` هي `8`. بالنظر إلى منطقك، هذا يعني أن التكرار يجب أن يعمل عندما تكون قيمة `i` <dfn>أصغر من أو يساوي</dfn> `count`. يمكنك استخدام معامل أصغر من أو يساوي `<=` لهذا الغرض.

حدّث شرط التكرار ليعمل طالما أن `i` أصغر من أو يساوي `count`.

# --hints--

يجب ألا يتحقق تكرارك `for` مما إذا كانت قيمة `i` أصغر من `count`.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

يجب أن يتحقق تكرارك `for` مما إذا كانت قيمة `i` أصغر من أو تساوي `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
