---
id: 660f50a21fe7645252804f2b
title: الخطوة 112
challengeType: 1
dashedName: step-112
---

# --description--

ماذا لو كان لديك طريقة للتبديل بين هرم مقلوب وهرم عادي؟

ابدأ بإعلان متغير `inverted` وتعيين القيمة `true` له. لن تغير هذا المتغير في شفرتك، لكن عليك استخدام `let` حتى تتمكن اختباراتنا من تعديله لاحقًا.

# --hints--

يجب إعلان متغير `inverted` باستخدام `let`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

يجب تهيئة `inverted` بالقيمة `true`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
