---
id: 660f540c2176ea5dec01306d
title: خطوة 118
challengeType: 1
dashedName: step-118
---

# --description--

عمل جيد! جرّب قيمًا مختلفة لمتغيرات `character` و `count` و `inverted`.

عندما تكون مستعدًا للانتقال إلى مشروعك التالي، عرّف `character` على `"!"`، و `count` على `10`، و `inverted` على `false` للاستمرار.

تهانينا على إكمال مشروع JavaScript الأول لك!

# --hints--

يجب أن تعرّف `character` على `"!"`.

```js
assert.equal(character, "!");
```

يجب أن تعرّف `count` على `10`.

```js
assert.equal(count, 10);
```

يجب أن تعرّف `inverted` على `false`.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
