---
id: 660f0a55847d6cc485f29ba5
title: الخطوة 26
challengeType: 1
dashedName: step-26
---

# --description--

الـ <dfn>طريقة</dfn> في JavaScript هي دالة مرتبطة بقيم أو كائنات معينة. مثال واجهته سابقًا هو طريقة `.log()`، التي هي جزء من كائن `console`.

للمصفوفات طرق خاصة بها، وأول طريقة ستستكشفها هي `.push()`. تتيح لك هذه الطريقة "دفع" قيمة إلى نهاية المصفوفة. إليك مثالًا لإضافة الرقم `12` إلى نهاية مصفوفة:

```js
array.push(12);
```

استخدم `.push()` لإضافة السلسلة النصية `"freeCodeCamp"` إلى نهاية مصفوفة `rows`. أضف هذا الشفرة قبل `console.log` لتتمكن من رؤية التغيير الذي أجريته على المصفوفة.

# --hints--

يجب أن تستخدم `.push()` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

يجب أن تستخدم طريقة `.push()` الخاصة بمصفوفة `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

يجب أن تمرر السلسلة النصية `"freeCodeCamp"` إلى طريقة `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
