---
id: 660f0f980e98e8cf77f1ce31
title: خطوة 30
challengeType: 1
dashedName: step-30
---

# --description--

غيّر إعلان `rows` ليُعيَّن مصفوفة فارغة مرة أخرى.

أيضًا، غيّر سلسلة النص `'Hello'` لاستخدام علامات اقتباس مزدوجة مرة أخرى. عمومًا، لا يهم أيهما تفضل، لكن يجب أن تحافظ على الاتساق في هذا الاختيار طوال مشروعك.

# --hints--

يجب أن تكون مصفوفة `rows` فارغة.

```js
assert.empty(rows);
```

يجب أن تستخدم سلسلة النص `"Hello"` علامات اقتباس مزدوجة.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
