---
id: 660f2a70ad6225fa503e71c3
title: الخطوة 46
challengeType: 1
dashedName: step-46
---

# --description--

أنت تقترب! في هذه المرحلة، تواجه ما يُعرف بـ <dfn>خطأ بمقدار واحد</dfn>، وهو مشكلة شائعة في لغات البرمجة التي تستخدم الفهرسة التي تبدأ من الصفر مثل JavaScript.

أول فهرس في مصفوفة `rows` هو `0`، ولهذا تبدأ حلقة `for` بـ `i = 0`. لكن تكرار سلسلة نصية صفر مرات لا ينتج عنه أي شيء للطباعة.

لإصلاح ذلك، أضف `1` إلى قيمة `i` في استدعاء `.repeat()`. لا تُعيّن القيمة مرة أخرى إلى `i` كما فعلت في شروط الحلقة.

# --hints--

يجب أن تضيف `1` إلى `i` في طريقة `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
