---
id: 660f4e74f7fd3f4a99ac2e50
title: الخطوة 108
challengeType: 1
dashedName: step-108
---

# --description--

تحتوي المصفوفات أيضًا على طريقة `.shift()`. هذه الطريقة تزيل العنصر **الأول** من المصفوفة، على عكس `.pop()` التي تزيل العنصر الأخير. إليك مثالًا على طريقة `.shift()`:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

ستكون مصفوفة `numbers` هي `[2, 3]`.

أسفل مصفوفة `numbers` مباشرة، أعلن عن متغير `shifted` وقم بتعيين نتيجة استدعاء `.shift()` على مصفوفة `numbers`. في السطر التالي، سجّل المتغير `shifted` في وحدة التحكم.

# --hints--

يجب أن تستخدم `const` لإعلان متغير `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

يجب أن تستدعي `.shift()` على مصفوفة `numbers` الخاصة بك.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

يجب أن تعيّن نتيجة استدعاء `.shift()` إلى متغير `shifted`.

```js
assert.equal(shifted, 1);
```

يجب أن تسجل متغير `shifted`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
