---
id: 660f2fbd45b520046cac68e8
title: الخطوة 65
challengeType: 1
dashedName: step-65
---

# --description--

تذكر في خطوة سابقة، تعلمت عن قيم الإرجاع. يمكن للدالة <dfn>إرجاع</dfn> قيمة ليستهلكها تطبيقك بشكل منفصل.

في الدالة، تُستخدم الكلمة المفتاحية `return` لتحديد قيمة الإرجاع. على سبيل المثال، هذه الدالة سترجع القيمة المعطاة للمعلمة الأولى:

```js
function name(parameter) {
  return parameter;
}
```

استخدم الكلمة المفتاحية `return` لإرجاع قيمة المتغير `character`، مكررة `rowNumber` مرة.

# --hints--

يجب أن تستخدم الطريقة `.repeat()`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

يجب أن تستخدم الطريقة `.repeat()` على متغير `character` الخاص بك.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

يجب أن تمرر `rowNumber` إلى استدعاء `.repeat()` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

يجب أن تستخدم الكلمة المفتاحية `return`.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

يجب أن تُرجع نتيجة استدعاء `.repeat()` الخاص بك.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
