---
id: 660f1a00ac619ddc1e259a66
title: الخطوة 35
challengeType: 1
dashedName: step-35
---

# --description--

<dfn>الشرط</dfn> في حلقة `for` يُخبر الحلقة بعدد مرات التكرار التي يجب أن تنفذها. عندما يصبح الـ`condition` خاطئًا، تتوقف الحلقة.

في JavaScript، يمكن أن تكون القيمة المنطقية إما `true` أو `false`. هذه ليست سلاسل نصية - ستتعلم المزيد عن الفرق لاحقًا.

في الوقت الحالي، ستستخدم <dfn>معامل أصغر من</dfn> (`<`). هذا يسمح لك بالتحقق مما إذا كانت القيمة على اليسار أصغر من القيمة على اليمين. على سبيل المثال، `count < 3` ستُقيّم إلى `true` إذا كانت `count` تساوي `2`، و`false` إذا كانت `count` تساوي `4`.

استبدل سلسلة `"condition"` الخاصة بك بشرط يتحقق مما إذا كانت `i` أصغر من `count`.

# --hints--

يجب أن تستخدم معامل أصغر من.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

يجب أن تستخدم معامل أصغر من للتحقق مما إذا كانت `i` أصغر من `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

يجب أن تستخدم حلقة `for` الخاصة بك الشرط `i < count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
