---
id: 660f53ad3d39175c5d4335ac
title: الخطوة 117
challengeType: 1
dashedName: step-117
---

# --description--

مولد الهرم الخاص بك الآن في حالة مكتملة، مع وظائف أكثر مما خططت له في البداية! الخطوة التالية هي تنظيف شفرتك.

احذف كل التعليقات، سواء كانت مفردة السطر أو متعددة الأسطر، من شفرتك.

# --hints--

يجب ألا تحتوي شفرتك على أي تعليقات مفردة السطر.

```js
assert.notMatch(code, /\/\//);
```

يجب ألا تحتوي شفرتك على أي تعليقات متعددة الأسطر.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

يجب ألا تحتوي شفرتك على أي تعليقات.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

يجب حذف الشفرة التي تم تعليقها باستخدام تعليقات متعددة الأسطر.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
