---
id: 660f280dda5040f707c76b4a
title: الخطوة 45
challengeType: 1
dashedName: step-45
---

# --description--

الآن لديك سلسلة من أحرف `#`، لكن شكل الهرم لا يزال مفقودًا. لحسن الحظ، المتغير `i` يمثل رقم "الصف" الحالي في الحلقة، مما يمكنك من استخدامه لصنع هيكل يشبه الهرم.

لتحقيق ذلك، ستستخدم الطريقة `.repeat()` المتاحة للسلاسل النصية. تقبل هذه الطريقة رقمًا كمعلمة، يحدد عدد مرات تكرار السلسلة المستهدفة. على سبيل المثال، استخدام `.repeat()` لإنشاء السلسلة النصية `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

استخدم الطريقة `.repeat()` على `character`، ومرر لها `i` كعدد التكرار.

# --hints--

يجب أن تستخدم الطريقة `.repeat()`.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

يجب أن تستخدم الطريقة `.repeat()` على متغير `character` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

يجب أن تمرر `i` إلى الطريقة `.repeat()` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

يجب أن تستخدم الطريقة `.repeat()` داخل الطريقة `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
