---
id: 660f374d532dc41189cc9cc2
title: الخطوة 69
challengeType: 1
dashedName: step-69
---

# --description--

حان الوقت الآن لبعض الرياضيات. تخيل هرمًا مكوّنًا من ثلاثة صفوف. إذا أردنا أن يكون في الوسط، فسيبدو شيء مثل:

```md
··#··
·###·
#####
```

تم استبدال الفراغات بعلامات نقطة وسطى لتحسين القراءة. إذا استنتجت النمط، سترى أن الفراغات في بداية ونهاية الصف تتبع نمطًا معينًا.

حدّث سلاسل الفراغات لتتكرر `rowCount - rowNumber` مرة.

افتح وحدة التحكم لترى النتيجة.

# --hints--

يجب أن تستدعي `.repeat()` على سلاسل `" "` لتكرارها `rowCount - rowNumber` مرة.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
