---
id: 660f47afe4c98536715d5fa4
title: الخطوة 91
challengeType: 1
dashedName: step-91
---

# --description--

يسمح لك معامل <dfn>عدم المساواة الصارمة</dfn> `!==` بالتحقق مما إذا كانت قيمتان غير متساويتين، أو لا تشتركان في نفس النوع. الصيغة النحوية مشابهة لمعامل المساواة: `value !== 4`.

حاليًا، حلقة `while` تعمل فقط إذا كانت `continueLoop` صحيحة. حدّث شرط حلقة `while` للتحقق مما إذا كانت `done` لا تساوي `count`.

# --hints--

يجب أن تتحقق حلقة `while` من أن `done` و `count` غير متساويتين.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
