---
id: 660f2eccfe3f820304af1b39
title: الخطوة 64
challengeType: 1
dashedName: step-64
---

# --description--

لكي تعرف كيف تُنسق صفًا، يجب أن تعرف دالة `padRow` رقم الصف الذي تعمل عليه، وعدد الصفوف الكلي التي تُولد.

أفضل طريقة لذلك هي إنشاء معلمات دالة لهما. أضف إلى دالة `padRow` معلمتي `rowNumber` و `rowCount`. تُفصل المعلمات المتعددة بفاصلة:

```js
function name(first, second) {

}
```

# --hints--

يجب أن تحتوي دالة `padRow` على معلمة `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

يجب إضافة فاصلة بعد معلمة `rowNumber`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

يجب أن تحتوي دالة `padRow` على معلمة `rowCount`.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
