---
id: 660f033cf051ebb50ea3bf48
title: الخطوة 18
challengeType: 1
dashedName: step-18
---

# --description--

في البرمجة، ستحتاج غالبًا إلى التعامل مع الكثير من البيانات. هناك العديد من هياكل البيانات التي تساعدك على تنظيم وإدارة بياناتك. أحد أبسط هياكل البيانات هو <dfn>المصفوفة</dfn>.

<dfn>المصفوفة</dfn> هي نوع بيانات غير بدائي يمكنه احتواء سلسلة من القيم. تختلف أنواع البيانات غير البدائية عن البدائية في قدرتها على احتواء بيانات أكثر تعقيدًا. أنواع البيانات البدائية مثل السلاسل النصية والأرقام يمكنها احتواء قيمة واحدة فقط في كل مرة.

تُرمز المصفوفات باستخدام الأقواس المربعة (`[]`). إليك مثالًا على متغير بقيمة مصفوفة فارغة:

```js
let array = [];
```

أعلن عن متغير `rows` وقم بتعيين مصفوفة فارغة له.

# --hints--

يجب أن يكون لديك متغير `rows`.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

يجب استخدام `let` لإعلان متغير `rows` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

يجب تعيين مصفوفة فارغة إلى متغير `rows` الخاص بك.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
