---
id: 660f229d2dbe09ef2954a4a1
title: الخطوة 42
challengeType: 1
dashedName: step-42
---

# --description--

تذكر في الحلقة السابقة أنك استخدمت معامل الجمع `+` لزيادة قيمة `i` بمقدار `1`.

يمكنك فعل شيء مشابه مع قيمة سلسلة نصية، بإضافة سلسلة جديدة إلى سلسلة موجودة. على سبيل المثال، `hello = hello + " World";` سيضيف السلسلة النصية `" World"` إلى السلسلة المخزنة في المتغير `hello`. هذا يسمى <dfn>معامل التسلسل</dfn>.

في حلقة `for...of` الخاصة بك، استخدم معامل الجمع لتسلسل قيمة `row` إلى قيمة `result`.

# --hints--

يجب أن تستخدم معامل التسلسل على متغير `result` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

يجب أن تسلسل `row` إلى متغير `result` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

يجب أن تعيّن نتيجة التسلسل مرة أخرى إلى متغير `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
