---
id: 660f18f059fe0fda192ce394
title: الخطوة 34
challengeType: 1
dashedName: step-34
---

# --description--

تحتاج الحلقة الآن إلى مكرر مناسب. <dfn>المكرر</dfn> هو متغير يمكنك إعلانه خصيصًا في حلقة `for` للتحكم في كيفية تكرار الحلقة أو مرورها عبر المنطق الخاص بك.

من الاصطلاحات الشائعة استخدام `i` كمتغير المكرر في الحلقة. تسمح لك حلقة `for` بإعلان هذا داخل الأقواس العادية `()`. على سبيل المثال، هذه حلقة `for` تعلن متغير `index` وتعيّن له القيمة `100`.

```js
for (let index = 100; "second"; "third") {

}
```

استبدل السلسلة النصية `"iterator"` بإعلان `let` للمتغير `i`. عيّن له القيمة `0` كبداية. هذا سيعطي المتغير `i` القيمة `0` في **المرة الأولى** التي تعمل فيها الحلقة.

# --hints--

يجب استخدام `let` لإعلان متغير `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

يجب تعيين القيمة `0` للمتغير `i`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

يجب أن تبدأ حلقة `for` الخاصة بك مكرر `i` عند القيمة `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
