---
id: 660f4774e3e0df35a68bb5f2
title: الخطوة 90
challengeType: 1
dashedName: step-90
---

# --description--

لجعل الهرم يُولد مرة أخرى، ادفع نتيجة استدعاء `padRow` مع `done` و `count` كمعلمتين إلى مصفوفة `rows`، كما فعلت في الحلقة الأولى.

# --hints--

يجب أن تستدعي الحلقة طريقة `.push()` على `rows`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

يجب أن تستدعي دالة `padRow` داخل طريقة `.push()`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

يجب أن تمرر `done` كالمعلمة الأولى لاستدعاء `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

يجب أن تمرر `count` كالمعلمة الثانية لاستدعاء `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
