---
id: 660f0ee51d7460ce88cd248d
title: الخطوة 29
challengeType: 1
dashedName: step-29
---

# --description--

هل توقعت رؤية `4` في وحدة التحكم؟ تعيد `.push()` الطول الجديد للمصفوفة بعد إضافة القيمة التي تعطيها لها.

من المهم أن تعرف القيم التي تعيدها الطريقة. خذ بعض الوقت لتجربة `.push()` و `.pop()`. عندما تكون جاهزًا، احذف كل استدعاءات `.push()` و `.pop()`، وكذلك عبارات `console.log` الخاصة بك.

# --hints--

يجب ألا يكون لديك استدعاء `.push()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

يجب ألا يكون لديك استدعاء `.pop()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

يجب ألا يكون لديك أي عبارات تسجيل.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

يجب ألا يكون لديك متغير `popped`.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

يجب ألا يكون لديك متغير `pushed`.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
