---
id: 660f09a2694b59c3a10ee304
title: الخطوة 23
challengeType: 1
dashedName: step-23
---

# --description--

أزل الآن أول سجل في وحدة التحكم وتعيين `rows[rows.length - 1]`. اترك سجل `rows` الثاني لوقت لاحق.

# --hints--

يجب أن تزيل عبارة `console.log(rows[0])`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

يجب أن تزيل إعادة تعيين `rows[rows.length - 1]`.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

لا يجب أن تزيل عبارة `console.log(rows)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
