---
id: 660f4934fb48f63abd5ae371
title: الخطوة 95
challengeType: 1
dashedName: step-95
---

# --description--

استخدام `done` لتتبع عدد الصفوف التي تم توليدها يعمل، لكن يمكنك تحسين المنطق أكثر قليلاً.

للمصفوفات خاصية `length` خاصة تسمح لك برؤية عدد القيم، أو <dfn>العناصر</dfn>، الموجودة في المصفوفة. يمكنك الوصول إلى هذه الخاصية باستخدام صيغة نحوية مثل `myArray.length`.

لاحظ أن `rows.length` في استدعاء `padRow` سيعطيك خطأ بمقدار واحد، لأن `done` يُزاد *قبل* الاستدعاء.

حدّث شرطك ليتحقق مما إذا كانت `rows.length` أقل من `count`.

# --hints--

يجب أن يتحقق حلقة `while` الخاصة بك مما إذا كانت `rows.length` أقل من `count`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
