---
id: 660f44f10ea40f300b896a5e
title: الخطوة 84
challengeType: 1
dashedName: step-84
---

# --description--

الآن بعد أن تدربت على العمل مع عبارات `if...else if...else`، يمكنك إزالتها من شفرتك.

بعد إكمال ذلك، استخدم `let` لإعلان متغير `continueLoop` وتعيين القيمة المنطقية `false` له. ثم استخدم `let` لإعلان متغير `done` وتعيين القيمة `0` له.

# --hints--

يجب ألا تحتوي شفرتك على عبارة `if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

يجب ألا تحتوي شفرتك على عبارة `else if`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

يجب ألا تحتوي شفرتك على عبارة `else`.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

يجب استخدام `let` لإعلان متغير `continueLoop`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

يجب أن تكون قيمة متغير `continueLoop` هي `false`.

```js
assert.isFalse(continueLoop);
```

يجب استخدام `let` لإعلان متغير `done`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

يجب أن تكون قيمة متغير `done` هي `0`.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
