---
id: 660f23b53db70af0f2620e78
title: الخطوة 43
challengeType: 1
dashedName: step-43
---

# --description--

الآن تظهر كل أرقامك في نفس السطر. هذا لن يعمل لإنشاء هرم.

ستحتاج إلى إضافة سطر جديد لكل صف. مع ذلك، الضغط على مفتاح العودة لإدخال فاصل سطر بين علامات الاقتباس في JavaScript سيؤدي إلى خطأ في التحليل. بدلًا من ذلك، تحتاج إلى استخدام <dfn>تخطى التسلسل</dfn> الخاص `\n`، الذي يُفسر كسطر جديد عند تسجيل السلسلة النصية. على سبيل المثال:

```js
lineOne = lineOne + "\n" + lineTwo;
```

استخدم عامل الجمع الثاني لإلحاق سطر جديد بعد قيمتي `result` و `row`.

# --hints--

يجب أن تستخدم تخطى التسلسل `\n`. تذكر أنه يجب أن يكون سلسلة نصية، لذا يجب أن يكون ملفوفًا بين علامات اقتباس.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

يجب أن تدمج متغير `row` مع متغير `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

يجب أن تدمج تخطى التسلسل `\n` مع متغير `row`.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

يجب أن تعيّن كامل عملية الدمج مرة أخرى إلى متغير `result`.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
