---
id: 660f45ccf4ca5c31f253005a
title: الخطوة 86
challengeType: 1
dashedName: step-86
---

# --description--

في الوقت الحالي، إذا غيرت `continueLoop` إلى true، ستعمل حلقة `while` إلى الأبد. يُسمى هذا <dfn>الحلقة اللانهائية</dfn>، ويجب أن تحذر من تجنبها. يمكن أن تؤدي الحلقة اللانهائية إلى تجميد النظام، مما يتطلب إعادة تشغيل كاملة للهروب منها.

لتجنب ذلك، ابدأ باستخدام معامل التزايد لزيادة قيمة المتغير `done` داخل الحلقة.

# --hints--

يجب أن تزيد حلقة `while` قيمة المتغير `done`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
