---
id: 660f1e3f047bf4e403268713
title: الخطوة 39
challengeType: 1
dashedName: step-39
---

# --description--

للأسف، لا يمكنك الآن رؤية ما يفعله التكرار الخاص بك.

استخدم `let` لإعلان متغير `result`، وقم بتعيين سلسلة نصية فارغة له. تمثّل السلسلة النصية الفارغة بعلامات اقتباس لا تحتوي على أي شيء بينها، مثل `""`.

# --hints--

يجب إعلان متغير `result`.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

يجب استخدام `let` لإعلان متغير `result` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

يجب أن يكون متغير `result` سلسلة نصية فارغة.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
