---
id: 660f39b444fd6f16d1e49c1f
title: الخطوة 73
challengeType: 1
dashedName: step-73
---

# --description--

بدلاً من تمرير `i + 1` إلى استدعاء `padRow`، يمكنك بدء الحلقة عند `1`. هذا يسمح لك بإنشاء حلقة تبدأ بفهرس واحد.

حدّث المتكرر ليبدأ عند `1` بدلاً من `0`.

# --hints--

يجب أن تهيئ حلقة `for` المتغير `i` عند `1`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
