---
id: 660f1bf673487ae0bb25b900
title: الخطوة 37
challengeType: 1
dashedName: step-37
---

# --description--

يجب أن يعمل التكرار الآن ثماني مرات. داخل جسم التكرار، اطبع قيمة المتغير `i` ولاحظ ما يحدث.

# --hints--

يجب أن تستخدم `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

يجب أن تسجل قيمة `i`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

يجب أن تسجل قيمة `i` في تكرار `for` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
