---
id: 660f455b044d3230ed971e98
title: الخطوة 85
challengeType: 1
dashedName: step-85
---

# --description--

تُنفذ حلقة <dfn>`while`</dfn> مرارًا وتكرارًا حتى لا يعود `condition` المحدد صحيحًا. لها الصيغة النحوية التالية:

```js
while (condition) {
  logic;
}
```

استخدم تلك الصيغة النحوية لإعلان حلقة `while` مع `continueLoop` كشرط. يجب أن يكون جسم الحلقة فارغًا.

# --hints--

يجب أن تستخدم حلقة `while`.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

يجب أن تستخدم حلقة `while` الخاصة بك `continueLoop` كشرط.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
