---
id: 660f207334fabaeac3269c38
title: الخطوة 41
challengeType: 1
dashedName: step-41
---

# --description--

للتعامل مع سلسلة النص `result`، ستستخدم نوعًا مختلفًا من الحلقات. تحديدًا، حلقة `for...of` التي تُكرر خلال كل عنصر في كائن قابل للتكرار وتُعيّنه مؤقتًا إلى متغير.

صيغة الحلقة `for...of` تبدو كالتالي:

```js
for (const value of iterable) {

}
```

لاحظ أنه يمكنك استخدام `const` لأن المتغير موجود فقط خلال تكرار واحد، وليس طوال الحلقة بأكملها.

أنشئ حلقة `for...of` للتكرار خلال مصفوفة `rows`، مع تعيين كل قيمة إلى متغير `row`.

# --hints--

يجب أن تستخدم كلمة `for` أخرى.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

يجب أن تعلن عن متغير `row`.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

يجب أن تستخدم `const` لإعلان متغير `row` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

يجب أن تعلن حلقة `for...of` عن متغير `row` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

يجب استخراج متغير `row` من `rows` باستخدام كلمة `of`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

يجب أن يكون جسم حلقة `for...of` فارغًا.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
