---
id: 660f383d4c772c12ff59904b
title: الخطوة 70
challengeType: 1
dashedName: step-70
---

# --description--

يمكنك تمرير تعبيرات كاملة كمعلمة. ستستقبل الدالة نتيجة تقييم ذلك التعبير. على سبيل المثال، هذان الاستدعاءان للدالة أدناه سيعطيان نفس النتيجة. (لاحظ أن قاعدة ترتيب العمليات PEMDAS—الأقواس، الأسس، الضرب، القسمة، الجمع، الطرح—تُطبق):

```js
test(2 * 3 + 1);
test(7);
```

النظر إلى النمط مرة أخرى:

```md
··#··
·###·
#####
```

حدّث قيمة `character` لتتكرر `2 * rowNumber - 1` مرة.

افتح وحدة التحكم مرة أخرى لرؤية النتيجة المحدثة.

# --hints--

يجب أن تمرر دالتك `2 * rowNumber - 1` إلى استدعاء `.repeat()`.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
