---
id: 660f4a1472f8e63d76162ce5
title: الخطوة 99
challengeType: 1
dashedName: step-99
---

# --description--

ماذا لو جعلت هرمك مقلوبًا، أو <dfn>معكوسًا</dfn>؟ حان الوقت لتجربته!

ابدأ بإنشاء حلقة `for` جديدة. أعلن عن المتكرر `i` وقم بتعيين قيمته إلى `count`، ثم استخدم القيمة المنطقية `false` كشرط وللتكرار.

# --hints--

يجب أن تحتوي شفرتك على حلقة `for`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

يجب أن تهيئ حلقة `for` المتغير `i` بقيمة `count`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

يجب أن تستخدم حلقة `for` القيمة `false` كشرط.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


يجب أن تستخدم حلقة `for` القيمة `false` للتكرار.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
