---
id: 660f039ff313dbb696b007ca
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

عندما تحتوي المصفوفة على قيم، أو <dfn>عناصر</dfn>، تُفصل هذه القيم بفواصل. هذه مصفوفة تحتوي على سلسلتين نصيتين:

```js
let array = ["first", "second"];
```

غيّر إعلان `rows` ليكون مصفوفة تحتوي على السلاسل النصية `"Naomi"`، و`"Quincy"`، و`"CamperChan"`. ترتيب القيم في المصفوفة مهم، فاتبع هذا الترتيب. تذكر أن السلاسل النصية حساسة لحالة الأحرف.

# --hints--

العنصر الأول في مصفوفة يجب أن يكون السلسلة النصية `"Naomi"`.

```js
assert.equal(rows[0], "Naomi");
```

العنصر الثاني في مصفوفة يجب أن يكون السلسلة النصية `"Quincy"`.

```js
assert.equal(rows[1], "Quincy");
```

العنصر الثالث في مصفوفة يجب أن يكون السلسلة النصية `"CamperChan"`.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
