---
id: 6610c48c4ea0891afa7c4696
title: الخطوة 54
challengeType: 1
dashedName: step-54
---

# --description--

إذا فتحت وحدة التحكم مرة أخرى، سترى أن دالة `padRow` تُرجع `undefined`، رغم أنك عرّفت قيمة إرجاع! هذا لأن المعلمات تحتاج إلى قيمة عند **استدعاء** الدالة.

عندما تمرر قيمة إلى استدعاء دالة، تُسمى هذه القيمة <dfn>معلمة</dfn>. إليك مثالًا على استدعاء دالة `demo` وتمرير `"Naomi"` كمعلمة لـ `name`.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

مرر اسمك كمعلمة لـ `name` في استدعاء `padRow`. تذكر أن اسمك سلسلة نصية، لذا عليك استخدام علامات الاقتباس.

# --hints--

يجب أن تمرر سلسلة نصية إلى استدعاء `padRow()`.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
