---
id: 6610c87eac0f0b256d7b037e
title: الخطوة 62
challengeType: 1
dashedName: step-62
---

# --description--

بما أن دالتك لم تعد تستخدم المعلمة، فإن تغيير المعلمة لم يؤثر عليها.

تابع واحذف إعلان `test` وعبارة `return` من دالة `padRow` الخاصة بك، ليصبح الدالة فارغة مرة أخرى.

# --hints--

يجب ألا تحتوي دالة `padRow` على متغير `test`.

```js
assert.notMatch(padRow.toString(), /test/);
```

يجب ألا تُرجع دالة `padRow` قيمة.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
