---
id: 6610bbed59bc2a0194d85533
title: الخطوة 22
challengeType: 1
dashedName: step-22
---

# --description--

هل لاحظت كيف تغيرت القيمة داخل مصفوفة `rows` مباشرة؟ هذا يسمى <dfn>تحوُّل</dfn>. مع تعلمك المزيد عن المصفوفات، ستتعلم متى تحوِّل المصفوفة ومتى لا يجب عليك ذلك.

قبل المتابعة، هذه فرصة رائعة لتعلم استخدام شائع للمصفوفات. حاليًا، شفرتك تصل إلى العنصر الأخير في المصفوفة باستخدام `rows[2]`. لكن قد لا تعرف عدد العناصر في المصفوفة عندما تريد العنصر الأخير.

يمكنك استخدام خاصية `.length` للمصفوفة - هذه تعيد عدد العناصر في المصفوفة. للحصول على العنصر الأخير في أي مصفوفة، يمكنك استخدام الصيغة النحوية التالية:

```js
array[array.length - 1]
```

`array.length` تعيد عدد العناصر في المصفوفة. بطرح `1`، تحصل على فهرس العنصر الأخير في المصفوفة. يمكنك تطبيق نفس المفهوم على مصفوفة `rows` الخاصة بك.

حدّث `rows[2]` لديك ليصل ديناميكيًا إلى العنصر الأخير في مصفوفة `rows`. ارجع إلى المثال أعلاه لمساعدتك.

لا يجب أن ترى أي تغيير في وحدة التحكم لديك.

# --hints--

يجب أن تستخدم خاصية `.length` لمصفوفة `rows` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

يجب أن تستخدم `rows.length` في ترميز الأقواس الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

يجب أن تطرح `1` من الطول في ترميز الأقواس الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
