---
id: 6610bf6fa14d700beed1b109
title: الخطوة 87
challengeType: 1
dashedName: step-87
---

# --description--

يُستخدم <dfn>معامل المساواة</dfn> `==` للتحقق مما إذا كانت قيمتان متساويتين. للمقارنة بين قيمتين، استخدم عبارة مثل `value == 8`.

أسفل `done++` داخل الحلقة، أضف عبارة `if`. يجب أن تتحقق العبارة مما إذا كانت `done` تساوي `count` باستخدام معامل المساواة.

# --hints--

يجب استخدام عبارة `if` داخل الحلقة. أضفها بعد `done++`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

يجب أن تستخدم عبارة `if` معامل المساواة لمقارنة `done` و `count` في الشرط.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
