---
id: 6610b8f6a98d25f4d485a94d
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

القيمة الافتراضية للمتغير غير المهيأ هي `undefined`. هذا نوع بيانات خاص يُمثل قيمة لا تحتوي على تعريف بعد.

لا يزال بإمكانك تعيين قيمة لمتغير غير مهيأ. إليك مثالًا:

```js
let uninitialized;
uninitialized = "assigned";
```

عيّن السلسلة النصية `"Test"` إلى متغير `secondCharacter` أسفل إعلانك. افتح وحدة التحكم لترى كيف تغير سجلّك.

# --hints--

لا يجب تهيئة `secondCharacter`. تذكر أن التهيئة تعني تعيين قيمة عند إعلان المتغير.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

يجب استخدام معامل التعيين على `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

يجب تعيين السلسلة النصية `"Test"` إلى متغير `secondCharacter`.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
