---
id: 6610b8017d1671f2814e8c77
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

عندما تعلن عن متغير دون تهيئته، يُعتبر <dfn>غير مهيأ</dfn>. حاليًا، متغير `secondCharacter` غير مهيأ.

أضف دالة `console.log()` لترى قيمة المتغير `secondCharacter`.

# --hints--

يجب إضافة دالة `console.log` ثانية.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

يجب تسجيل متغير `secondCharacter` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
