---
id: 6610c21b3ef82015573ffbbe
title: الخطوة 51
challengeType: 1
dashedName: step-51
---

# --description--

المتغير `call` يحمل القيمة `undefined` رغم أنك عرّفته! هذا لأن دالة `padRow` لا تُرجع قيمة حاليًا. افتراضيًا، الدوال تُرجع القيمة `undefined`.

لكي تُرجع شيئًا آخر، عليك استخدام الكلمة المفتاحية `return`. إليك مثالًا على دالة تُرجع السلسلة النصية `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

استخدم الكلمة المفتاحية `return` لجعل دالتك تُرجع السلسلة النصية `"Hello!"`.

# --hints--

يجب أن تستخدم دالة `padRow` الكلمة المفتاحية `return`.

```js
assert.match(padRow.toString(), /return/);
```

يجب أن تُرجع دالة `padRow` السلسلة النصية `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
