---
id: 6610c77d50636722e5b6be17
title: الخطوة 60
challengeType: 1
dashedName: step-60
---

# --description--

أسفل جملة `return`، سجّل السلسلة النصية `"This works!"` في وحدة التحكم.

بعد القيام بذلك، سترى أن السلسلة النصية `"This works!"` لا تظهر في وحدة التحكم، وأن السطر `console.log("This works!")` يظهر باللون الرمادي.

انسخ سجل وحدة التحكم والصقه فوق جملة `return`. الآن يجب أن تظهر السلسلة النصية `"This works!"` في وحدة التحكم.

شيء مهم يجب معرفته عن الكلمة المفتاحية `return` هو أنها لا تحدد فقط قيمة تُعاد من الدالة، بل توقف أيضًا تنفيذ الشفرة داخل الدالة أو جملة الكتلة. هذا يعني أن أي شفرة بعد جملة `return` لن تُنفذ.

# --hints--

يجب أن تُعيد دالة `padRow` المتغير `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

يجب أن يأتي أول `console.log` بعد الكلمة المفتاحية `return`.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

يجب أن يأتي ثاني `console.log` قبل الكلمة المفتاحية `return`.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
