---
id: 6610c424b7119919b62932f4
title: الخطوة 53
challengeType: 1
dashedName: step-53
---

# --description--

لا يجب أن تُعيد الدالة قيمة ثابتة مشفرة. يمكنها إعادة القيمة المخزنة في متغير. المعلمات هي متغيرات خاصة بالدالة، لذا يمكن أيضًا إعادتها.

غيّر دالة `padRow` لتُعيد مباشرةً معلمة `name` باستخدام `return`.

# --hints--

يجب أن تُعيد دالتك `padRow` قيمة المعلمة `name`.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
