---
id: 6610c1d97b1671140f95cfbb
title: الخطوة 50
challengeType: 1
dashedName: step-50
---

# --description--

أضف الآن عبارة تسجيل لطباعة قيمة المتغير `call` الخاص بك.

# --hints--

يجب أن تضيف استدعاء `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

يجب أن تسجل المتغير `call` الخاص بك. لا تنس الفاصلة المنقوطة.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
