---
id: 6610c105bbdacc114d6cdc44
title: الخطوة 48
challengeType: 1
dashedName: step-48
---

# --description--

لكي تستخدم دالة، يجب أن تستدعيها. <dfn>استدعاء الدالة</dfn> يخبر تطبيقك بتنفيذ الشفرة من الدالة في أي مكان تختار استدعاؤها فيه. صيغة استدعاء الدالة هي اسم الدالة متبوعًا بأقواس عادية. على سبيل المثال، هذا الشفرة تعرف وتستدعي دالة `test`.

```js
function test() {

}
test();
```

استدعي دالة `padRow` الخاصة بك.

# --hints--

يجب أن تستدعي دالة `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

يجب أن يكون جسم دالة `padRow` فارغًا.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
