---
id: 6610c8cfe4cf4d278e35c156
title: الخطوة 63
challengeType: 1
dashedName: step-63
---

# --description--

كما هو متوقع، دالتك الآن تُرجع `undefined` مرة أخرى. لم يعد متغير `call` ضروريًا، لذا احذف إعلان `call` و`console.log` الخاص بمتغير `call`.

# --hints--

يجب ألا يكون لديك إعلان `call`.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

يجب ألا تسجل متغير `call` في وحدة التحكم.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
