---
id: 6610b741b54b90f0c0fb3d58
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

عندما تكون أسماء المتغيرات أكثر من كلمة واحدة، توجد اصطلاحات محددة لكيفية كتابة الحروف الكبيرة في الكلمات. في JavaScript، الاصطلاح المستخدم هو <dfn>camel case</dfn>.

تعني camel case أن الكلمة الأولى في الاسم تكون كلها أحرف صغيرة، لكن الكلمات التالية تبدأ بحرف كبير. إليك بعض الأمثلة على camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

استخدم camel case لإعلان متغير جديد باسم `secondCharacter`.

# --hints--

يجب إعلان متغير `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

يجب استخدام `let` لإعلان متغير `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

لا يجب تعيين قيمة لمتغير `secondCharacter`. لا تنس الفاصلة المنقوطة.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
