---
id: 6610c83b52583e245a079217
title: الخطوة 61
challengeType: 1
dashedName: step-61
---

# --description--

الآن المتغير `call` يحمل القيمة `"Testing"`. لكن دالتك لم تعد تستخدم المعلمة `name`.

احذف المعلمة `name` من إعلان دالتك، ثم احذف السلسلة النصية `"CamperChan"` من استدعاء `padRow`.

أيضًا، احذف كلا `console.log` من دالة `padRow`.
# --hints--

يجب ألا تحتوي دالة `padRow` على معلمة `name`.

```js
assert.notMatch(padRow.toString(), /name/);
```

يجب ألا تمرر `"CamperChan"` إلى استدعاء `padRow`.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

يجب أن تستمر في استدعاء دالة `padRow`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

يجب ألا يكون لديك `console.log` قبل كلمة `return`.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

يجب ألا يكون لديك `console.log` بعد كلمة `return`.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
