---
id: 6610c71600966a2191d3a64a
title: الخطوة 59
challengeType: 1
dashedName: step-59
---

# --description--

القيم التي تُعاد من دالة تُستخدم عن طريق استدعاء الدالة. يمكنك استخدام استدعاء الدالة مباشرةً كالقيمة التي تُعيدها، أو تخزين القيمة المعادة في متغير. بهذه الطريقة، يمكنك استخدام القيمة المُعيَّنة لمتغير ذو نطاق محلي خارج الدالة التي أنشئ فيها.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

لاستخدام القيمة `"Testing"` الخاصة بك، أعدها من دالة `padRow` عن طريق تحديث جملة `return` لتُعيد فقط المتغير `test`.


# --hints--

يجب أن تُعيد دالة `padRow` المتغير `test`.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
