---
id: 6610b9f7619764fad5fd516d
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

يمكنك أيضًا تعيين قيمة متغير إلى متغير آخر. على سبيل المثال:

```js
let first = "One";
let second = "Two";
second = first;
```

سيحمل المتغير `second` الآن القيمة `"One"`.

لرؤية ذلك عمليًا، غيّر تعيين `secondCharacter` من `"Test"` إلى متغير `character` الخاص بك.

ثم افتح وحدة التحكم لترى ما يتم تسجيله.

# --hints--

لا يجب أن تعيّن القيمة `"Test"` إلى متغير `secondCharacter`.

```js
assert.notEqual(secondCharacter, "Test");
```

يجب أن تعيّن قيمة متغير `character` إلى متغير `secondCharacter`. لا تنس الفاصلة المنقوطة.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

يجب أن يحمل متغير `secondCharacter` الآن القيمة `"World"`.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
