---
id: 6610c2d8d67563174fcf96dc
title: الخطوة 52
challengeType: 1
dashedName: step-52
---

# --description--

عندما يكون لديك قيمة مكتوبة صراحة في شفرتك، مثل السلسلة النصية `"Hello!"` في دالتك، تُعتبر <dfn>مشفرة بشكل ثابت</dfn>. تشفير قيمة داخل دالة قد يجعلها أقل قابلية لإعادة الاستخدام مما ترغب.

بدلاً من ذلك، يمكنك تعريف <dfn>معلمات</dfn> للدالة. المعلمات هي متغيرات خاصة تُعطى قيمة عند استدعاء الدالة، ويمكن استخدامها داخل الدالة لتغيير نتيجة الشفرة بشكل ديناميكي.

لإضافة معلمة إلى دالتك، تحتاج إلى إضافة اسم متغير داخل الأقواس العادية. على سبيل المثال، هذه الدالة `demo` تحتوي على معلمة `name`:

```js
function demo(name) {

}
```

يبدو أن `name` معلمة مفيدة، فتابع وأضفها إلى دالة `padRow` الخاصة بك.

# --hints--

يجب أن تحتوي دالة `padRow` على معلمة `name`.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
