---
id: 6610c538372aa61cc0f5b122
title: الخطوة 57
challengeType: 1
dashedName: step-57
---

# --description--

المتغيرات في JavaScript متاحة في <dfn>نطاق</dfn> محدد. بعبارة أخرى، مكان إعلان المتغير يحدد أين في شفرتك يمكن استخدامه.

النطاق الأول هو النطاق العام. المتغيرات التي تُعلن خارج أي "كتلة" مثل دالة أو حلقة `for` تكون في <dfn>النطاق العام</dfn>. متغيراتك `character` و`count` و`rows` كلها في النطاق العام.

عندما يكون المتغير في النطاق العام، يمكن للدالة الوصول إليه في تعريفها. إليك مثالًا لدالة تستخدم متغير `title` في النطاق العام:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

هذا المثال سيُرجع `"Professor Naomi"`. حدّث دالتك `padRow` لتُرجع قيمة ربط متغير `character` في بداية معلمة `name`.

# --hints--

يجب أن تربط دالتك `padRow` المتغير `character` في بداية `name`.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

يجب أن تُرجع دالتك `padRow` نتيجة `character + name`.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
