---
id: 6610c16c4fa0df12c0e30675
title: الخطوة 49
challengeType: 1
dashedName: step-49
---

# --description--

أنت تستدعي دالتك `padRow`، لكنك لا تفعل شيئًا مع استدعاء الدالة هذا. كل الدوال في JavaScript <dfn>تعيد</dfn> قيمة، مما يعني أنها توفر النتيجة المحددة لاستدعائها لتستخدمها في مكان آخر.

لرؤية نتيجة استدعاء دالتك `padRow`، أعلن عن متغير `call` وقم بتعيين استدعاء `padRow` الحالي له.

# --hints--

يجب أن تعلن عن متغير `call`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

يجب أن تستخدم `const` لإعلان متغير `call`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

يجب أن تعيّن لـ `call` نتيجة استدعاء دالتك `padRow`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
