---
id: 6610c6541c82551f95e765ab
title: الخطوة 58
challengeType: 1
dashedName: step-58
---

# --description--

يمكن أيضًا إعلان المتغيرات داخل دالة. تُعتبر هذه المتغيرات ضمن <dfn>النطاق المحلي</dfn>، أو <dfn>نطاق الكتلة</dfn>. المتغير المعلن داخل دالة يمكن استخدامه فقط داخل تلك الدالة. إذا حاولت الوصول إليه خارج الدالة، ستحصل على خطأ مرجع.

لرؤية هذا عمليًا، استخدم `const` لإعلان متغير `test` داخل دالتك `padRow`. ابدأه بالقيمة `"Testing"`.

ثم، أسفل دالتك، حاول تسجيل `test` في وحدة التحكم. سترى خطأ لأنه غير معرف خارج النطاق المحلي للدالة. احذف ذلك `console.log` لتجاوز الاختبارات والمتابعة.

# --hints--

يجب أن تعلن دالتك عن متغير `test`.

```js
assert.match(padRow.toString(), /const\s+test/);
```

يجب أن تبدأ `test` بالقيمة `"Testing"`. لا تنس الفاصلة المنقوطة.

```js
assert.match(padRow.toString(), /const\s+test\s*=\s*('|")Testing\1;/)
```

يجب أن يكون متغير `test` قبل كلمة `return`.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
