---
id: 664599653fcd6e97104f9261
title: الخطوة 83
challengeType: 1
dashedName: step-83
---

# --description--

أحيانًا تريد تشغيل شفرة مختلفة عندما تكون كل شروط `if...else if` هي `false`. يمكنك فعل ذلك بإضافة كتلة `else`.

كتلة `else` تُنفذ فقط إذا لم تتحقق الشروط في كتلتي `if` و `else if`.

هنا أُضيفت كتلة `else` إلى كتلة `else if`.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

أضف كتلة `else` إلى كتلة `else if`. داخل كتلة `else`، سجّل السلسلة النصية `"This is the else block"` في وحدة التحكم.

لرؤية النتائج في وحدة التحكم، يمكنك تغيير `<` يدويًا في جملة `else if` إلى `>`. هذا سيجعل الشرط `false` وسيتم تنفيذ كتلة `else`.

# --hints--

يجب أن يكون لديك كتلة `else`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

يجب أن تسجل كتلة `else` السلسلة النصية `"This is the else block"` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

