---
id: 664ed97c55a99e5ffab759dc
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

قبل المتابعة، خذ لحظة لمراجعة المفاهيم التي تعلمتها.

استخدم الكلمة المفتاحية `let` لإعلان متغير `profession` ومتغير `age`. ابدأ تهيئة `profession` بالسلسلة النصية `"teacher"`، لكن لا تهيئ `age` بأي قيمة.

سجل كلا المتغيرين في وحدة التحكم لرؤية النتائج.

# --hints--

يجب إعلان متغير `profession`.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

يجب استخدام `let` لإعلان متغير `profession`.

```js
assert.match(code, /let\s+profession/);
```

يجب تعيين السلسلة النصية `"teacher"` إلى متغير `profession`.

```js
assert.strictEqual(profession, "teacher");
```

يجب إعلان متغير `age`.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

يجب استخدام `let` لإعلان متغير `age`.

```js
assert.match(code, /let\s+age/);
```

يجب عدم تعيين قيمة إلى متغير `age`.

```js
assert.isUndefined(age);
```

يجب تسجيل `age` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

يجب تسجيل `profession` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
