---
id: 664ed9344ba3495e1f67f8ea
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

بما أنك راجعت الإعلان والتهيئة، احذف الشفرة التي كتبتها لمراجعتك.

لا تحذف المتغير `character`.

# --hints--

يجب ألا يكون لديك متغير `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

يجب ألا يكون لديك متغير `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

يجب ألا يكون لديك عبارة `console.log` لـ `age`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

يجب ألا يكون لديك عبارة `console.log` لـ `profession`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

يجب أن يبقى لديك متغير `character`.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
