---
id: 665732da4815b70bb083915e
title: الخطوة 24
challengeType: 1
dashedName: step-24
---


# --description--

في الخطوات القليلة الماضية، تعلمت كل شيء عن العمل مع المصفوفات. خذ لحظة لمراجعة ما تعلمته.

ابدأ بإعلان متغير `cities` وتهيئته كمصفوفة من السلاسل النصية `"London"`، `"New York"`، و`"Mumbai"`. ثم سجّل هذا المتغير في وحدة التحكم.

بعد التسجيل، غيّر العنصر الأخير في `cities` إلى السلسلة النصية `"Mexico City"`، ثم سجّل متغير `cities` مرة أخرى.

عند التنفيذ الصحيح، يجب أن ترى هذا الإخراج في وحدة التحكم.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

يجب استخدام `let` لإعلان متغير `cities`.

```js
assert.match(code, /let\s+cities/);
```

يجب تعيين مصفوفة من السلاسل النصية `"London"`، `"New York"`، و`"Mumbai"` إلى المتغير `cities`.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

يجب استخدام `console.log()` لتسجيل المصفوفة الكاملة `cities` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

يجب تحديث العنصر الأخير في مصفوفة `cities` إلى السلسلة النصية `"Mexico City"`. تذكر أنه يمكنك الوصول إلى العنصر الأخير في المصفوفة باستخدام `array[array.length - 1]`.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

يجب أن يحتوي شفرتك على عبارتي `console.log(cities)`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
