---
id: 66643c9a9e0ad1c787f5fe2b
title: الخطوة 56
challengeType: 1
dashedName: step-56
---

# --description--

بعد إكمال تلك المراجعة السريعة، يجب حذف دالة `addTwoNumbers`، والمتغير `sum`، وعبارة السجل.

# --hints--

يجب ألا تحتوي شفرتك على دالة `addTwoNumbers`.

```js
assert.notMatch(code, /addTwoNumbers/);
```

يجب ألا تحتوي شفرتك على متغير `sum`.

```js
assert.notMatch(code, /sum/);
```

يجب ألا تسجل المتغير `sum`.

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
