---
id: 66643c93e05093c728abdbe9
title: الخطوة 55
challengeType: 1
dashedName: step-55
---

# --description--

قبل المتابعة، خذ لحظة لمراجعة كيفية عمل الدوال.

أعلن عن دالة باسم `addTwoNumbers`. يجب أن تأخذ هذه الدالة معلمتين وتُرجع مجموع هاتين المعلمتين.

يجب ألا تستخدم دالتك قيمًا مبرمجة بشكل ثابت. مثال على دالة مبرمجة بشكل ثابت قد يكون:

```js
function sayName(firstName, lastName) {
  return "John Doe";
}

sayName("Camper", "Cat");
```

ستُرجع هذه الدالة `"John Doe"` بغض النظر عن المعلمات الممررة إلى `firstName` و`lastName`، لذا تُعتبر `"John Doe"` قيمة مبرمجة بشكل ثابت.

أعلن عن متغير `sum` وقم بتعيين قيمة استدعاء دالتك `addTwoNumbers` مع المعلمتين `5` و`10`. سجّل متغير `sum` في وحدة التحكم.

# --hints--

يجب أن يكون لديك دالة تسمى `addTwoNumbers`.

```js
assert.isFunction(addTwoNumbers);
```

يجب أن تحتوي دالتك `addTwoNumbers` على معلمتين.

```js
assert.lengthOf(addTwoNumbers, 2);
```

يجب أن تُرجع دالتك مجموع المعلمتين.

```js
assert.strictEqual(addTwoNumbers(5,10), 15);
```

يجب ألا تُرجع دالتك قيمة مبرمجة بشكل ثابت. أي يجب أن تعمل مع أي معلمتين رقميتين.

```js
assert.strictEqual(addTwoNumbers(3, 5), 8);
```

يجب أن تعلن عن متغير `sum`.

```js
assert.isDefined(sum);
```

يجب أن تكون قيمة متغير `sum` هي `15`.

```js
assert.strictEqual(sum, 15);
```

يجب أن تُعيّن لـ `sum` القيمة الناتجة من استدعاء دالة `addTwoNumbers` مع المعلمتين `5` و`10`.

```js
assert.match(code, /sum\s*=\s*addTwoNumbers\s*\(/);
```

يجب أن تسجل متغير `sum` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
