---
id: 5d5a813321b9e3db6c106a46
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript ist eine leistungsstarke Sprache, die es dir ermöglicht Webseiten zu erstellen, die interaktiv sind.

*Hinweis*: Für alle verbleibenden Projekte in diesem Lehrplan benötigst du Grundwissen in HTML und CSS. Falls du vorher noch nicht mit HTML und CSS gearbeitet hast, bearbeite erst[Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

To get started, create your standard HTML boilerplate with a `DOCTYPE`, `html`, `head`, and `body`, then add a `meta` tag for the `charset`. Füge ein `title`-Element hinzu und verwende dafür den Text `RPG - Dragon Repeller`. Füge einen `link`-Tag für dein Stylesheet hinzu, um die `styles.css`-Datei zu verknüpfen.

Finally, create a `div` element with `id` set to `game` within your `body`.

# --hints--

Dein Code sollte über die `DOCTYPE`-Referenz verfügen.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest ein Leerzeichen nach der `DOCTYPE`-Referenz einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag haben. Vergiss nicht das `lang`-Attribut.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag enthalten.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTMLs stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein abschließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein abschließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Deine `head` und `body`-Elemente sollten Geschwister sein.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Dein Code sollte über ein `meta`-Element verfügen.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Dein `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` enthalten.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Du solltest ein schließendes `title`-Tag haben.

```js
assert(code.match(/<\/title\s*>/i));
```

Dein Code sollte ein `link`-Element haben.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Du solltest über ein `div`-Element verfügen.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your `div` element should have an `id` attribute with the value `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Dein `div`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
