---
id: 62a240c67f3dbb1a1e6d95ee
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Gib deinem `#game` eine maximale Breite von `500px` und eine maximale Höhe von `400px`. Setze die `background-color` auf `#ffffff` und die `color` auf `#ffffff`.

Nutze Margins, um das Ganze zu zentrieren, indem die obere Margin auf `30px`, untere Margin auf `0px` gesetzt wird und die linke und rechte Margin auf `auto`.

Füge zum Schluss jeweils `10px` als Padding bei allen vier Seiten hinzu.

# --hints--

Du solltest einen `#game`-Selektor haben.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

Dein `#game`-Selektor sollte eine `max-width` von `500px` haben.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

Dein `#game`-Selektor sollte eine `max-height` von `400px` haben.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

Dein `#game`-Selektor sollte eine `background-color` von `#ffffff` haben.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

Dein `#game`-Selektor sollte eine `color` von `#ffffff` haben.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

Dein `#game`-Selektor sollte eine `margin-top` auf `30px` gesetzt haben.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

Dein `#game`-Selektor sollte eine `margin-left` auf `auto` gesetzt haben.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

Dein `#game`-Selektor sollte eine `margin-right` auf `auto` gesetzt haben.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


Dein `#game`-Selektor sollte eine `margin-bottom` auf `0px` gesetzt haben.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

Dein `#game`-Selektor sollte `10px`-Padding auf allen Seiten haben.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
