---
id: 62a3b41c9494f937560640ab
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Deklariere zwei weitere Variablen namens `monsterHealth` und `inventory`.

For your `inventory` variable, assign it the value of an array containing the string `"stick"`.

Remember that you worked with arrays in the previous project like this:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Du solltest `let` verwenden, um deine `monsterHealth` Variable zu deklarieren.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Du solltest camelCase verwenden, um deine `monsterHealth`-Variable zu benennen.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` sollte keinen Wert haben.

```js
assert.isUndefined(monsterHealth);
```

Du solltest deiner `monsterHealth`-Variablen keinen Wert zuweisen.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Du solltest `let` verwenden, um deine `inventory`-Variable zu deklarieren.

```js
assert.match(code, /let\s+inventory/i);
```

Deine `inventory`-Variable sollte weiterhin ein Array sein.

```js
assert.isArray(inventory);
```

Deine `inventory`-Variable sollte nur einen Wert haben.

```js
assert.lengthOf(inventory, 1);
```

Deine `inventory`-Variable sollte den String `"stick"` enthalten.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
