---
id: 62a3a0a3c0a4b32915d26a6e
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Entferne deine `console.log("Hello World");` Zeile, um mit dem Schreiben deines Projektcodes zu beginnen.

Im vorherigen Projekt hast du gelernt, wie man mit Variablen arbeitet und das `let` Schlüsselwort verwendet:

```js
let camperbot;
```

Sie haben auch gelernt, eine Variable mit einem Wert wie diesem zu initialisieren:

```js
let age = 32;
```

Use the `let` keyword to declare a variable called `xp` and assign it the value of `0`, a number.

# --hints--

Du solltest keine `console.log("Hello World");`-Zeile in deinem Code haben.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Du solltest mit dem `let`-Schlüsselwort deine Variable deklarieren.

```js
assert.match(code, /let/);
```

Du solltest eine Variable namens `xp` deklarieren.

```js
assert.match(code, /let\s+xp/);
```

`xp` sollte den Wert `0` besitzen.

```js
assert.equal(xp, 0);
```

Du solltest die `xp`-Variable mit `0` initialisieren.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
