---
id: 62a3b3eab50e193608c19fc6
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Deklariere eine Variable namens `fighting`, aber initialisiere sie nicht mit einem Wert.

# --hints--

Du solltest `let` verwenden, um eine Variable `fighting` zu deklarieren.

```js
assert.match(code, /let\s+fighting/);
```

Deine `fighting`-Variable sollte keinen Wert haben.

```js
assert.isUndefined(fighting);
```

Du solltest deiner `fighting`-Variable keinen Wert zuweisen.

```js
assert.match(code, /let\s+fighting\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
--fcc-editable-region--
```
