---
id: 69f0a37c3bf194a7c9c62492
title: 建置裝置借用帳本
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

你是一位軟體開發人員，負責創建一個帳本應用程式來管理公司服務台提供的 IT 硬體裝置。該帳本是一個物件，其鍵是資產標籤，值則是物件：

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該將帳本建模為一個物件，其中每個裝置以其資產標籤作為鍵，並擁有 `type`、`status`、`borrower`（包含 `name` 和 `email`）以及 `dueDate` 屬性。

1. 你應該實作一個 `checkoutDevice(ledger, assetTag, borrower)` 函式，該函式接收一個 ledger 物件，並使用 `serializeLedger` 和 `loadLedger` 複製它，而不會改變原本的物件。在複製的物件上，應設定 `borrower` 物件的 `name` 和 `email`，更新 `status` 為 `"CheckedOut"`，並傳回一個形式為 `{ ledger: updatedLedger, message: confirmationString }` 的物件，其中 `message` 包含資產標籤和 `borrower` 物件的 `name`。

   - 如果 `assetTag` 在帳本中不存在，該函式應傳回未改變的帳本，並附帶一則包含資產標籤且說明未找到的 `message`。
   - 如果該裝置已經是 `"CheckedOut"`，函式應傳回未改變的帳本，保留現有借用者，並附帶一則包含資產標籤且說明裝置已被借出的 `message`。

1. 你應該實作一個 `checkinDevice(ledger, assetTag)` 函式，該函式接收一個 ledger 物件，並使用 `serializeLedger` 和 `loadLedger` 複製它，而不會改變原本的物件。在複製的物件上，應清除 `borrower` 物件的 `name` 和 `email`，將 `dueDate` 重設為 `""`，更新 `status` 為 `"CheckedIn"`，並傳回一個形式為 `{ ledger: updatedLedger, message: confirmationString }` 的物件，其中 `message` 包含資產標籤。

   - 如果 `assetTag` 在帳本中不存在，該函式應傳回未改變的帳本，並附帶一則包含資產標籤且說明未找到的 `message`。

1. 你應該實作一個 `listOverdueDevices(ledger, today)` 函式，該函式接收一個帳本物件並傳回一個逾期裝置物件的陣列，依 `dueDate` 以升冪排序。當裝置的 `status` 是 `"CheckedOut"` 且 `dueDate` 嚴格早於 `today` 時，該裝置即為逾期。`today` 和 `dueDate` 都是以月/日/年格式的字串，月和日可能有或沒有補零（例如，`"9/5/2025"` 和 `"09/05/2025"` 都有效）。

1. 你不應該在程式碼中使用 JavaScript 的 `Date` 物件。

1. 你應該實作一個 `serializeLedger(ledger)` 函式，將帳本物件轉換成 JSON 字串，並實作一個 `loadLedger(json)` 函式，將 JSON 字串轉換回 JavaScript 物件。

# --hints--

你應該創建一個名為 `checkoutDevice` 的函式。

```js
assert.isFunction(checkoutDevice);
```

你應該創建一個名為 `checkinDevice` 的函式。

```js
assert.isFunction(checkinDevice);
```

你應該創建一個名為 `listOverdueDevices` 的函式。

```js
assert.isFunction(listOverdueDevices);
```

你應該創建一個名為 `serializeLedger` 的函式。

```js
assert.isFunction(serializeLedger);
```

你應該創建一個名為 `loadLedger` 的函式。

```js
assert.isFunction(loadLedger);
```

你的 `checkoutDevice` 函式不應該改變原始帳本。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

你的 `checkoutDevice` 函式應該使用 `serializeLedger` 和 `loadLedger` 來複製 ledger。

```js
const source = __helpers.removeJSComments(checkoutDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

你的 `checkoutDevice` 函式應該設定借用者的 `name` 和 `email`，並將 `status` 設為 `"CheckedOut"`。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

你的 `checkoutDevice` 函式應該傳回一則包含資產標籤和借用者 `name` 的確認訊息。

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

你的 `checkoutDevice` 函式應該保留未相關的裝置和未改變的欄位於傳回的帳本中。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

當資產標籤未找到時，你的 `checkoutDevice` 函式應傳回未改變的帳本，並附帶一則包含資產標籤的訊息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

當裝置已被借出時，你的 `checkoutDevice` 函式應傳回未改變的帳本，保留現有借用者，並附帶一則包含資產標籤的訊息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

你的 `checkinDevice` 函式不應該改變原始帳本。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

你的 `checkinDevice` 函式應該使用 `serializeLedger` 和 `loadLedger` 來複製 ledger。

```js
const source = __helpers.removeJSComments(checkinDevice.toString());
assert.match(source, /\bserializeLedger\s*\(/);
assert.match(source, /\bloadLedger\s*\(/);
```

你的 `checkinDevice` 函式應該清除借用者的 `name` 和 `email`，將 `dueDate` 重設為 `""`，並將 `status` 設為 `"CheckedIn"`。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

你的 `checkinDevice` 函式應該傳回一則包含資產標籤的確認訊息。

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

你的 `checkinDevice` 函式應該保留未相關的裝置和未改變的欄位於傳回的帳本中。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

當資產標籤未找到時，你的 `checkinDevice` 函式應傳回未改變的帳本，並附帶一則包含資產標籤的訊息。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

你的 `listOverdueDevices` 函式應該傳回一個陣列。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

你的 `listOverdueDevices` 函式應該只傳回 `status` 為 `"CheckedOut"` 的逾期裝置。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

你的 `listOverdueDevices` 函式應該傳回依 `dueDate` 升冪排序的逾期裝置。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

你的 `listOverdueDevices` 函式即使在月或日未補零的情況下，也應該正確排序。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

你不應該在程式碼中使用 JavaScript 的 `Date` 物件。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

你的 `serializeLedger` 函式應該傳回一個代表所傳帳本的 JSON 字串。

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

你的 `loadLedger` 函式應該傳回 JSON 字串所代表的 JavaScript 物件。

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = loadLedger(serializeLedger(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
