---
id: 6a89c69cbe1bb69845992d8b
title: Tengeneza kizalishaji cha safu ya nambari
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

Katika maabara hii, utaunda kitendakazi cha `range_of_numbers` kinachotumia urudiaji (recursion) kuunda orodha ya nambari ndani ya anuwai iliyobainishwa.

Kwa mfano:

- `range_of_numbers(3, 9)` inapaswa kurudisha `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` inapaswa kurudisha `[5]`.

Kamilisha hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `range_of_numbers` chenye vigezo viwili vinavyoitwa `start_num` na `end_num`.
1. Kitendakazi kinapaswa kurudisha orodha ya nambari kamili mfululizo inayotangulia na `start_num` na kumalizika na `end_num`, ikijumuisha thamani zote mbili.
1. Thamani ya `start_num` itakuwa daima ndogo au sawa na `end_num`.
1. Kitendakazi chako kinapaswa kutumia urudiaji kwa kujiita mwenyewe.
1. Kitendakazi chako hakipaswi kutumia mizunguko ya `for` au `while`, comprehensions, au kitendakazi cha `range()`.
1. Wakati `start_num` ni sawa na `end_num`, kitendakazi kinapaswa kurudisha orodha yenye `start_num`. Hii ni kesi ya msingi.
1. Kwa kesi ya urudiaji, kitendakazi kinapaswa kujiita mwenyewe kwa hoja inayosogea kuelekea kesi ya msingi, kisha kuongeza nambari ya sasa kwenye orodha iliyorejeshwa.
1. Kitendakazi hakipaswi kutumia vigezo vya kimataifa kuhifadhi orodha iliyorejeshwa.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `range_of_numbers` chenye vigezo vinavyoitwa `start_num` na `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` inapaswa kurudisha `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` inapaswa kurudisha `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` inapaswa kurudisha `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` inapaswa kurudisha `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Kitendakazi chako cha `range_of_numbers` kinapaswa kujiita mwenyewe.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Kitendakazi chako cha `range_of_numbers` hakipaswi kutumia mizunguko, comprehensions, au kitendakazi cha `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Kuitwa `range_of_numbers(1, 3)` hakipaswi kuathiri wito baadaye. Baada ya wito huo, `range_of_numbers(6, 9)` bado inapaswa kurudisha `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
