---
id: 6842a6cd9836f0114a5b7a8a
title: Construir um Gerador de Padrão Numérico
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

Neste laboratório, você praticará o básico de Python construindo um pequeno aplicativo que cria um padrão numérico.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve definir uma função chamada `number_pattern` que recebe um único parâmetro `n` (representando um número inteiro positivo).
1. `number_pattern` deve usar um loop `for`.
1. `number_pattern(n)` deve retornar uma string com todos os inteiros começando de 1 até `n` (incluído) separados por um espaço. Por exemplo, `number_pattern(4)` deve retornar a string `1 2 3 4`.
1. Se o argumento passado para a função não for um valor inteiro, a função deve retornar `Argument must be an integer value.`.
1. Se o argumento passado para a função for menor que 1, a função deve retornar `Argument must be an integer greater than 0.`.

# --hints--

Você deve ter uma função `number_patter`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

A função `number_pattern` deve ter um parâmetro chamado `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` deve retornar `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` deve retornar `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` deve retornar uma lista de números separados por espaço para qualquer número inteiro positivo.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` deve retornar `Argumento deve ser um valor inteiro.` quando for passado um valor que não seja um inteiro.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` deve retornar `Argumento deve ser um inteiro maior que 0.` quando receber um inteiro não positivo.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
