---
id: 67c562286b29447da020d407
title: بناء مولد كلمة مرور لمرة واحدة
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

في هذا المختبر، ستُولد كلمة مرور لمرة واحدة (OTP) مكونة من 6 أرقام وتعرضها للمستخدم. تنتهي صلاحية كلمة المرور بعد 5 ثوانٍ، وسيتم توليد كلمة مرور جديدة عند نقر المستخدم على زر `"Generate New OTP"`.

**الهدف:** أتمم قصص المستخدم أدناه واجتز جميع الاختبارات لإكمال المختبر.

**قصص المستخدم:**

1. يجب عليك استخدام الخطاف `useEffect` لإدارة مؤقت العد التنازلي.
2. يجب أن يُرجع مكوِّن `OTPGenerator` عنصر `div` يحمل اسم الصنف `container`.
3. يجب أن يحتوي عنصر `div` الذي يحمل الصنف `container` على العناصر التالية:

- عنصر `h1` يحمل المعرف `otp-title` والنص `"OTP Generator"`.
    - عنصر `h2` يحمل المعرف `otp-display` يعرض إما الرسالة `"Click 'Generate OTP' to get a code"` أو يعرض كلمة المرور المولدة إذا كانت متوفرة.
    - عنصر `p` يحمل المعرف `otp-timer` وسم `aria-live` مضبوط على قيمة صحيحة بحيث:
        - يبدأ فارغًا.
        - يعرض الرسالة `"Expires in: X seconds"` بعد نقر الزر، حيث يمثل `X` الوقت المتبقي قبل انتهاء صلاحية كلمة المرور.
        - يعرض الرسالة `"OTP expired. Click the button to generate a new OTP."` بمجرد وصول العد التنازلي إلى `0`.
    - عنصر `button` يحمل المعرف `generate-otp-button` وعليه المسمى `"Generate OTP"`. عند النقر عليه، يجب أن يولد كلمة مرور جديدة ويبدأ عدًا تنازليًا لمدة 5 ثوانٍ.
    - يجب تعطيل زر `"Generate OTP"` أثناء تشغيل العد التنازلي.

4. يجب أن تتأكد من توقف مؤقت العد التنازلي تلقائيًا بمجرد وصوله إلى `0` ثانية لمنع التحديثات غير الضرورية.
5. يجب أن تكون كلمة المرور المولدة مكونة من 6 أرقام.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

يجب ألا تزيل التعيين الحالي `const { useState, useEffect, useRef } = React;` من الشفرة.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

يجب أن تعرض عنصر `div` يحمل الصنف `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

يجب أن يحتوي `.container` على عنصر `h1` يحمل المعرف `otp-title` والنص `"OTP Generator"`.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

يجب أن يحتوي `.container` على عنصر `h2` يحمل المعرف `otp-display` يعرض كلمة المرور عند توليدها.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

في البداية، يجب أن يعرض عنصر `h2` داخل `.container` الرسالة `"Click 'Generate OTP' to get a code"`.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

يجب أن يحتوي عنصر `.container` على عنصر `p` يحمل المعرف `otp-timer`.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

يجب أن يحمل عنصر `p` سمة `aria-live` مضبوطة على `assertive` أو `polite`.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

في البداية، يجب أن يكون عنصر `p` فارغًا.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

يجب أن يحتوي `.container` على عنصر `button` يحمل المعرف `generate-otp-button` والنص `"Generate OTP"`.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

عند النقر على الزر، يجب أن يعرض كلمة مرور جديدة في عنصر `h2` الذي يحمل المعرف `otp-display`.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

يجب أن تكون كلمة المرور المولدة مكونة من 6 أرقام.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


عند النقر على الزر، يجب أن يعرض عنصر `p` الذي يحمل المعرف `otp-timer` عدًا تنازليًا لمدة 5 ثوانٍ.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

يجب أن تتحدث الرسالة في عنصر `p` الذي يحمل المعرف `otp-timer` كل ثانية لتُظهر الوقت المتبقي.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

يجب تعطيل زر `"Generate OTP"` أثناء تشغيل مؤقت العد التنازلي.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

يجب تمكين زر `"Generate OTP"` بمجرد وصول مؤقت العد التنازلي إلى 0.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

عندما يصل مؤقت العد التنازلي إلى `0`، يجب أن تعرض الرسالة `OTP expired. Click the button to generate a new OTP.`.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

يجب أن تصدر مكوِّن `OTPGenerator`.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
