---
id: afcc8d540bea9ea2669306b6
title: Crea una funzione per ripetere una stringa
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

In questo laboratorio, creerai una funzione che ripete una stringa data un numero specifico di volte. Per questo laboratorio, non usare il metodo integrato `.repeat()`.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione chiamata `repeatStringNumTimes` che prende due parametri: una stringa e un numero.
2. La funzione dovrebbe restituire la stringa ripetuta il numero specificato di volte.
3. Se il numero è minore o uguale a zero, la funzione dovrebbe restituire una stringa vuota.

# --hints--

Dovresti creare una funzione chiamata `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` dovrebbe prendere due parametri.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

La funzione `repeatStringNumTimes` dovrebbe sempre restituire una stringa.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` dovrebbe restituire la stringa `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` dovrebbe restituire la stringa `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` dovrebbe restituire la stringa `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` dovrebbe restituire la stringa `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` dovrebbe restituire la stringa `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` dovrebbe restituire una stringa vuota (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` dovrebbe restituire `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

Non dovrebbe essere usato il metodo integrato `repeat()`.

```js
const repeatSpy = __helpers.spyOn(String.prototype, 'repeat');
try {
  repeatStringNumTimes('abc', 3);
  assert.isEmpty(repeatSpy.calls);
} finally {
  repeatSpy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
