---
id: 68599af28ead38c22590a89d
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Debe haber una manera de conectar los botones con el contenido que representan. Esto se hace usando el atributo `aria-controls` en cada bot贸n. M谩s adelante, agregar谩s el contenido y los atributos `id` correspondientes a los elementos de contenido.

Comienza agregando el atributo `aria-controls` a cada uno de los elementos `button`. El valor del atributo `aria-controls` del primer bot贸n debe establecerse en `"panel-earth"`. El atributo `aria-controls` del segundo bot贸n debe establecerse en `"panel-saturn"`. El atributo `aria-controls` del tercer bot贸n debe establecerse en `"panel-mars"`.

# --hints--

Cada uno de tus elementos `button` debe tener un atributo `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

El primer elemento `button` debe tener un atributo `aria-controls` establecido en `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

El segundo elemento `button` debe tener un atributo `aria-controls` establecido en `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

El tercer elemento `button` debe tener un atributo `aria-controls` establecido en `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab"> Earth</button>
        <button role="tab"> Saturn</button>
        <button role="tab"> Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
