---
id: 685976a92ffab8b2dc744b72
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `.tabs`, agrega un elemento `div` con un atributo `role` establecido en `"tablist"` y un atributo `aria-labelledby` establecido en `"tabs-title"`.  

El atributo `aria-labelledby` es necesario aquí para asociar la lista de pestañas con una etiqueta visible, ayudando a tecnologías de asistencia como los lectores de pantalla a anunciar el propósito de la lista de pestañas a los usuarios.

# --hints--

Debes tener un elemento `div` dentro de tu elemento `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Tu elemento `div` debe tener un atributo `role` establecido en `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Tu elemento `div` debe tener un atributo `aria-labelledby` establecido en `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
