---
id: 685db0bd1280154ae85bcf9f
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, construirás una estructura de pestañas accesible de planetas usando HTML, CSS y JavaScript. El boilerplate de HTML y el CSS han sido proporcionados para ti.

En estos primeros pasos, construirás el HTML restante para el tablist.

Comienza creando un elemento `div` con una clase establecida en `"tabs"`. Dentro de ese elemento `div`, agrega un elemento `h2` con un atributo `id` establecido en `"tabs-title"` y el contenido de texto `Planets`.

# --hints--

Debe tener un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Tu elemento `div` debe tener una clase establecida en `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Tu elemento `div` debe tener un elemento `h2` dentro de él.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Tu elemento `h2` debe tener un atributo `id` establecido en `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Tu elemento `h2` debe tener el contenido de texto `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
