---
id: 68af277560e009597dd861e8
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Тепер видаліть свій рядок `print(movie1)`. Потім помістіть своє присвоєння `movie1` у блок `try` і створіть умову `except`, яка ловить `ValueError as e` і використовує f-рядок, щоб вивести `Validation Error:` з пробілом і повідомленням про помилку.

Не соромтеся змінювати аргументи, передані для створення `movie1`, щоб побачити різні повідомлення про помилки.

# --hints--

У вашому коді не повинно бути `print(movie1)`.

```js
({
  test: () => runPython(`assert not _Node(_code).has_call("print(movie1)")`)
})
```

У вас має бути блок `try` з `movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)` всередині.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_body().has_stmt("movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)")`)
})
```

У вас має бути умова `except`, яка ловить `ValueError as e`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].has_except("ValueError", "e")`)
})
```

Використайте f-рядок, щоб вивести `Validation Error:` з пробілом і повідомленням про помилку всередині блоку `except`.

```js
({
    test: () => runPython(`assert _Node(_code).find_trys()[0].find_except("ValueError", "e").has_call("print(f'Validation Error: {e}')")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--
movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
--fcc-editable-region--
```
