---
id: 68af273960e009597dd861e2
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Ви хочете мати зручний для читання вивід для налагодження та виведення. Створіть метод `__str__`, який повертає рядок у форматі `title (year) - duration min, director` (де `title`, `year`, `duration` і `director` мають бути замінені на значення відповідних атрибутів), щоб при виведенні фільму отримати гарний підсумок.

# --hints--

Ваш клас `Movie` має мати метод `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

Ваш метод `__str__` має повертати рядок у форматі `title (year) - duration min, director`, де `title`, `year`, `duration` і `director` мають бути замінені на значення відповідних атрибутів.

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    
--fcc-editable-region--
```
