---
id: 68af274560e009597dd861e3
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер створіть екземпляр `Movie`, передавши аргументи на свій вибір, і призначте його змінній з ім’ям `movie1`. Потім виведіть `movie1` у консоль.

# --hints--

У вас має бути змінна з ім’ям `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

Ви маєте призначити екземпляр `Movie` змінній `movie1`.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

Ви маєте вивести `movie1` у консоль.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
