---
id: 68af274e60e009597dd861e4
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

У попередніх кроках ви створили базовий об’єкт фільму. Тепер додайте валідацію, щоб переконатися, що фільми не можуть бути створені з недійсними атрибутами.

На початку методу `__init__` створіть інструкцію `if`, яка перевіряє, чи є `title` порожнім або містить лише пробіли (для побудови умови можна використати метод `strip`).

Якщо так, викличте помилку `ValueError` з повідомленням `Title cannot be empty`.

# --hints--

Ви повинні викликати помилку `ValueError` з повідомленням `Title cannot be empty`, коли `title` порожній або містить лише пробіли.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

Ви не повинні викликати помилку `ValueError`, коли `title` є дійсним рядком.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
