---
id: 68af272f60e009597dd861e1
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите каталог медіаконтенту. Почніть із визначення класу з назвою `Movie`.

Додайте до свого класу метод `__init__` з такими параметрами: `self`, `title`, `year`, `director` і `duration`. Усередині методу `__init__` призначте параметри `title`, `year`, `director` і `duration` атрибутам екземпляра з такими ж іменами.

# --hints--

Ви повинні створити клас з назвою `Movie`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Movie")`)
})
```

Ваш клас `Movie` має мати метод `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__init__")`)
})
```

Ваш метод `__init__` повинен мати такі параметри: `self`, `title`, `year`, `director` і `duration`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_args("self, title, year, director, duration")`)
})
```

Ви повинні призначити `title` `self.title` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.title = title")`)
})
```

Ви повинні призначити `year` `self.year` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.year = year")`)
})
```

Ви повинні призначити `director` `self.director` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.director = director")`)
})
```

Ви повинні призначити `duration` `self.duration` всередині методу `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").find_function("__init__").has_stmt("self.duration = duration")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
