---
id: 587d78af367417b2b2512b03
title: Erstellen Sie ein Umfrageformular
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**Ziel:** Erstellen Sie eine App, die funktional ähnlich ist wie <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>. **Kopieren Sie dieses Demo-Projekt nicht**.

**User Stories:**

1. Sie sollten einen Seitentitel in einem `h1`-Element mit einem `id` von `title` haben
1. Sie sollten eine kurze Erklärung in einem `p`-Element mit einem `id` von `description` haben
1. Sie sollten ein `form`-Element mit einem `id` von `survey-form` haben
1. Innerhalb des Formularelements sind Sie **verpflichtet**, Ihren Namen in einem `input`-Feld einzugeben, das ein `id` von `name` und ein `type` von `text` hat
1. Innerhalb des Formularelements sind Sie **verpflichtet**, Ihre E-Mail in einem `input`-Feld einzugeben, das ein `id` von `email` hat
1. Wenn Sie eine E-Mail eingeben, die nicht korrekt formatiert ist, sehen Sie einen HTML5-Validierungsfehler
1. Innerhalb des Formulars können Sie eine Zahl in einem `input`-Feld eingeben, das ein `id` von `number` hat
1. Die Zahleneingabe sollte keine Nicht-Zahlen akzeptieren, entweder indem Sie daran gehindert werden, sie einzugeben, oder indem ein HTML5-Validierungsfehler angezeigt wird (abhängig von Ihrem Browser).
1. Wenn Sie Zahlen außerhalb des Bereichs der Zahleneingabe eingeben, die durch die `min`- und `max`-Attribute definiert sind, sehen Sie einen HTML5-Validierungsfehler
1. Für die Eingabefelder Name, E-Mail und Zahl können Sie entsprechende `label`-Elemente im Formular sehen, die den Zweck jedes Felds mit den folgenden IDs beschreiben: `id="name-label"`, `id="email-label"` und `id="number-label"`
1. Für die Eingabefelder Name, E-Mail und Zahl sehen Sie Platzhaltertext, der eine Beschreibung oder Anweisungen für jedes Feld gibt
1. Innerhalb des Formularelements sollten Sie ein `select`-Dropdown-Element mit einem `id` von `dropdown` und mindestens zwei auswählbaren Optionen haben
1. Innerhalb des Formularelements können Sie eine Option aus einer Gruppe von mindestens zwei Optionsschaltflächen auswählen, die mit dem `name`-Attribut gruppiert sind
1. Innerhalb des Formularelements können Sie mehrere Felder aus einer Reihe von Kontrollkästchen auswählen, von denen jedes ein `value`-Attribut haben muss
1. Innerhalb des Formularelements wird Ihnen ein `textarea` für zusätzliche Kommentare präsentiert
1. Innerhalb des Formularelements wird Ihnen eine Schaltfläche mit `id` von `submit` präsentiert, um alle Eingaben abzusenden

Erfüllen Sie die User Stories und bestehen Sie alle untenstehenden Tests, um dieses Projekt abzuschließen. Verleihen Sie ihm Ihren eigenen persönlichen Stil. Viel Spaß beim Programmieren!

**Hinweis:** Stellen Sie sicher, dass Sie `<link rel="stylesheet" href="styles.css">` in Ihrem HTML hinzufügen, um Ihr Stylesheet zu verlinken und Ihr CSS anzuwenden

# --hints--

Sie sollten ein `h1`-Element mit einem `id` von `title` haben.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Ihr `#title` sollte nicht leer sein.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Sie sollten ein `p`-Element mit einem `id` von `description` haben.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Ihr `#description` sollte nicht leer sein.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Sie sollten ein `form`-Element mit einem `id` von `survey-form` haben.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Sie sollten ein `input`-Element mit einem `id` von `name` haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#name` sollte ein `type` von `text` haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Ihr `#name` sollte eine Eingabe erfordern.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ihr `#name` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Sie sollten ein `input`-Element mit einem `id` von `email` haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#email` sollte ein `type` von `email` haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Ihr `#email` sollte eine Eingabe erfordern.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ihr `#email` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Sie sollten ein `input`-Element mit einem `id` von `number` haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ihr `#number` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Ihr `#number` sollte ein `type` von `number` haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Ihr `#number` sollte ein `min`-Attribut mit einem numerischen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Ihr `#number` sollte ein `max`-Attribut mit einem numerischen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Sie sollten ein `label`-Element mit einem `id` von `name-label` haben.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Sie sollten ein `label`-Element mit einem `id` von `email-label` haben.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Sie sollten ein `label`-Element mit einem `id` von `number-label` haben.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ihr `#name-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#email-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#number-label` sollte Text enthalten, der die Eingabe beschreibt.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ihr `#name-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Ihr `#email-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Ihr `#number-label` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Ihr `#name` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ihr `#email` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ihr `#number` sollte ein `placeholder`-Attribut und einen Wert haben.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Sie sollten ein `select`-Feld mit einem `id` von `dropdown` haben.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Ihr `#dropdown` sollte mindestens zwei auswählbare (nicht deaktivierte) `option`-Elemente haben.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Ihr `#dropdown` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Sie sollten mindestens zwei `input`-Elemente mit einem `type` von `radio` (Optionsschaltflächen) haben.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Sie sollten mindestens zwei Optionsschaltflächen haben, die Nachfahren von `#survey-form` sind.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Alle Ihre Optionsschaltflächen sollten ein `value`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Alle Ihre Optionsschaltflächen sollten ein `name`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Jede Gruppe von Optionsschaltflächen sollte mindestens 2 Optionsschaltflächen haben.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

Sie sollten mindestens zwei `input`-Elemente mit einem `type` von `checkbox` (Kontrollkästchen) haben, die Nachfahren von `#survey-form` sind.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Alle Ihre Kontrollkästchen innerhalb von `#survey-form` sollten ein `value`-Attribut und einen Wert haben.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Sie sollten mindestens ein `textarea`-Element haben, das ein Nachfahre von `#survey-form` ist.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Sie sollten ein `input` oder `button`-Element mit einem `id` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Ihr `#submit` sollte ein `type` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Ihr `#submit` sollte ein Nachfahre von `#survey-form` sein.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
