---
id: 68e6a99a0fd7386fb66e156f
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, eserciterai lo stile delle caselle di spunta creando un pannello delle impostazioni di gioco.

Tutto il codice boilerplate HTML (`DOCTYPE`, `html`, `head` e `body`) è già fornito.

All’interno del `body`, crea un elemento `div` con un elemento `h1` annidato all’interno. Dai all’elemento `h1` il testo `Game Settings`

# --hints--

Dovresti avere un elemento `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Dovresti avere un elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Il tuo elemento `h1` dovrebbe essere dentro il tuo elemento `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Il tuo elemento `h1` dovrebbe avere il testo `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
