---
id: 68e9a6bf1a74721a4e43e217
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Ora rimuoverai la casella di spunta predefinita applicata dai browser. All’interno del tuo selettore `input[type='checkbox']`, aggiungi `appearance` con un valore di `none`.

Impostare la proprietà `appearance` su `none` cancellerà l’aspetto che il browser applica alle caselle di spunta, permettendoti di mostrare lo stile che desideri.

Dopo averlo fatto, dato che la casella di spunta non sarà più visibile, imposta un `border` con uno spessore `2px`, uno stile `solid` e un codice esadecimale `#f1be32`.

# --hints--

Dovresti aggiungere una proprietà `appearance` al tuo selettore `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Dovresti avere una proprietà `appearance` con valore `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Dovresti aggiungere una proprietà `border` al tuo selettore `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Dovresti avere una proprietà `border` con valore `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
