---
id: 68e9a3194160f8f5620ddf21
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Immediatamente dopo il tuo elemento `h1`, crea quattro elementi `label`.

All'interno di ciascuno degli elementi `label`, crea un elemento `input` con l'attributo `type` impostato su `"checkbox"`.

Dopo ogni elemento `input`, assegna ai tuoi elementi `label` i seguenti testi, in ordine: `Sound Effects`, `Background Music`, `Hard Mode` e `Haptic Feedback`.

# --hints--

Dovresti avere quattro elementi `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Il tuo primo elemento `label` dovrebbe avere un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Il tuo primo elemento `input` dovrebbe avere l'attributo `type` impostato su `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Il tuo primo elemento `label` dovrebbe avere il testo `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Il tuo testo `Sound Effects` dovrebbe venire dopo l'elemento `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Il tuo secondo elemento `label` dovrebbe avere un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Il tuo secondo elemento `input` dovrebbe avere l'attributo `type` impostato su `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Il tuo secondo elemento `label` dovrebbe avere il testo `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Il tuo testo `Background Music` dovrebbe venire dopo l'elemento `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Il tuo terzo elemento `label` dovrebbe avere un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Il tuo terzo elemento `input` dovrebbe avere l'attributo `type` impostato su `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Il tuo terzo elemento `label` dovrebbe avere il testo `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Il tuo testo `Hard Mode` dovrebbe venire dopo l'elemento `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Il tuo quarto elemento `label` dovrebbe avere un elemento `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Il tuo quarto elemento `input` dovrebbe avere l'attributo `type` impostato su `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Il tuo quarto elemento `label` dovrebbe avere il testo `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Il tuo testo `Haptic Feedback` dovrebbe venire dopo l'elemento `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
