---
id: 68e9a6201e2efa142d62bfe1
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Adesso ingrandirai le caselle di spunta per migliorarne la visibilità.

Inizia creando un selettore per `input`, ma che punti specificamente al tuo `[type="checkbox"]`.

All’interno del selettore, imposta la `width` e la `height` su `20px`. Questo le rende più grandi rispetto a prima.

E per conformarti all’impostazione `cursor` che hai applicato alle etichette, aggiungi `cursor` e assegna il valore `pointer`. Dopo di che, quando passi il mouse sopra le caselle di spunta, il cursore diventerà una mano puntatrice.

# --hints--

Dovresti avere un selettore nel tuo CSS di `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Dovresti aggiungere una proprietà `width` al selettore `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Dovresti avere una proprietà `width` con valore `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Dovresti aggiungere una proprietà `height` al selettore `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Dovresti avere una proprietà `height` con valore `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Dovresti aggiungere una proprietà `cursor` al selettore `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Dovresti avere una proprietà `cursor` con valore `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
