---
id: 68ee86e3ee07b01626a392cc
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

All'interno del tuo selettore `.settings-card` imposta la proprietà `margin` su `auto`.

Impostare la proprietà `margin` su `auto` regola automaticamente i margini di un elemento per distribuire in modo uniforme lo spazio rimanente nel suo contenitore, comunemente usato per centrare orizzontalmente elementi di tipo blocco.

Infine, imposta una proprietà `text-align` con il valore `left`. Questo allineerà il contenuto in linea, come il testo, sul lato sinistro del suo elemento contenitore.

# --hints--

Dovresti aggiungere una proprietà `margin` al selettore di classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Dovresti avere una proprietà `margin` con valore `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Dovresti aggiungere una proprietà `text-align` al selettore di classe `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Dovresti avere una proprietà `text-align` con valore `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
