---
id: 68f15cd397503429f0345262
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Nella dichiarazione per il selettore `input[type="checkbox"]:checked::after`, imposta la proprietà `display` su `block`. Impostare la proprietà `display` su `block` rende l'elemento un elemento a blocco, il che significa che occupa tutta la larghezza del suo contenitore e inizia su una nuova riga, permettendoti di controllarne più facilmente larghezza, altezza e spaziatura.

Poi, imposta `text-align` su `center`. Questo centrerà orizzontalmente il contenuto in linea (come testo o elementi in linea) all'interno del blocco. Il blocco continuerà a occupare tutta la larghezza del suo contenitore ma tutto ciò che contiene sarà allineato al centro.

# --hints--

Dovresti aggiungere una proprietà `display` al tuo selettore combinato di tipo con un selettore di pseudo-classe e pseudo-elemento `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Dovresti avere una proprietà `display` con valore `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Dovresti aggiungere una proprietà `text-align` al tuo selettore combinato di tipo con un selettore di pseudo-classe e pseudo-elemento `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Dovresti avere una proprietà `text-align` con valore `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
