---
id: 68f15ef0446464325ce1aa45
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Hai bisogno di un po' di spazio tra le caselle di spunta e le etichette. Inizia usando il selettore per `label` nel tuo CSS e imposta una proprietà `display` con valore `block`.

Poi, imposta una proprietà `margin` con valore `8px auto` per aggiungere spazio verticale tra gli elementi. Questo creerà uno spazio coerente sopra e sotto ogni elemento, aiutando a separare il contenuto e migliorare la leggibilità.

Infine, imposta la proprietà `cursor` su `pointer` sugli elementi `label`. Questo cambierà il cursore in un'icona a forma di mano quando l'utente passa sopra un `label`, segnalando che l'elemento è cliccabile e migliorando l'esperienza utente complessiva.

# --hints--

Dovresti avere un selettore di tipo nel tuo CSS per `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Dovresti aggiungere una proprietà `display` al tuo selettore `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Dovresti avere una proprietà `display` con valore `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Dovresti aggiungere la proprietà `margin` al tuo selettore `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Dovresti avere una proprietà `margin` con valore `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Dovresti aggiungere una proprietà `cursor` al tuo selettore `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Dovresti avere una proprietà `cursor` con valore `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
