---
id: 68f1560aff1214103ab6b59b
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Ora inizierai ad abbellire la pagina con un po' di stile CSS. Inizia creando il selettore `body`.

Imposta il `body` per avere una proprietà `height` con un valore di `100vh` e una proprietà `background-color` con un valore di `#f0f0f0`.

La proprietà `height` di `100vh` fa sì che `body` occupi tutta l'altezza della viewport del browser, mentre il colore di sfondo grigio chiaro fornisce una base sottile per la pagina.

Infine, imposta una proprietà `text-align` con il valore `center`. Questo centrerà tutto il contenuto in linea contenuto nella pagina, a meno che un elemento figlio non la sovrascriva con il proprio allineamento.

# --hints--

Dovresti usare il selettore di tipo `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Dovresti aggiungere la proprietà `height` all'interno del selettore di tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Dovresti avere una proprietà `height` con il valore `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Dovresti aggiungere la proprietà `background-color` all'interno del selettore di tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Dovresti avere una proprietà `background-color` con un valore `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Dovresti aggiungere la proprietà `text-align` all'interno del selettore di tipo `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Dovresti avere una proprietà `text-align` con il valore `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
