---
id: 68f29dbf9125476b4808d6fb
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Ora è il momento di definire lo stile del contenitore `settings-card`. Crea un selettore di classe per `settings-card`.

Qui andrà tutto il tuo formato per il contenitore. Imposta la proprietà `max-width` su `250px` per definire la dimensione complessiva del contenitore.

Poi, imposta il tuo `padding` su `20px` in modo che il contenuto abbia spazio tra sé e il bordo del contenitore.

Dopodiché, crea un bordo arrotondato impostando il tuo `border-radius` su `10px`.

Quindi imposta un `box-shadow` con i valori di `0 2px 6px rgba(0,0,0,0.2)`. Questo creerà un effetto "sollevato" sottile che darà profondità al contenitore.

# --hints--

Dovresti avere un selettore di classe nel tuo css per `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Dovresti aggiungere la proprietà `max-width` al selettore di classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Dovresti avere una proprietà `max-width` con valore `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Dovresti aggiungere una proprietà `padding` al selettore di classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Dovresti avere una proprietà `padding` con valore `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Dovresti aggiungere una proprietà `border-radius` al selettore di classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Dovresti avere una proprietà `border-radius` con valore `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Dovresti aggiungere una proprietà `box-shadow` al selettore di classe `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Dovresti avere un `box-shadow` con valore `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
