---
id: 66bb3e20d3dc5b6d0a21f5dd
title: 建置一個即時計數器
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

在這個實驗中，你將使用 JavaScript 建置一個即時字元計數器。字元計數器會顯示在 `textarea` 元素中輸入的字元數量。當使用者在 `textarea` 中輸入時，計數器會即時更新。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `textarea` 元素，其 `id` 為 `text-input`。
2. 應該有一個 `p` 元素，其 `id` 為 `char-count`。它最初應包含文字 `Character Count: 0/50`。這個佔位文字應該被 `textarea` 元素中目前的字元數取代。
3. 當 `#text-input` 元素包含文字 `hello coder` 時，`#char-count` 元素應該包含文字 `"Character Count: 11/50"`
4. 當 `#text-input` 元素包含文字 `i am a programmer` 時，`#char-count` 元素應該包含文字 `"Character Count: 17/50"`
5. 當 `#text-input` 元素包含文字 `hello world` 時，`#char-count` 元素應該包含文字 `"Character Count: 11/50"`
6. 當 `#text-input` 元素包含文字 `I am learning a new language and it is called c++.` 時，`#char-count` 元素應該包含文字 `"Character Count: 50/50"`。**注意**：雖然 `maxlength` 屬性可以達到正確的功能，但本實驗室要求僅使用 JavaScript 解法。
7. 使用者不應該能在 `textarea` 元素中輸入超過 `50` 個字元。當字元數達到 `50` 時，任何多餘的輸入應自動被修剪，並且應以 `red` 顯示文字 `Character Count: 50/50`。

**注意：** 請務必在你的 HTML 中連結你的樣式表和 JavaScript 檔案。

# --hints--

你應該有一個 `textarea` 元素，其 `id` 為 `text-input`。 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

你應該有一個 `p` 元素，其 `id` 為 `char-count`。

```js
assert.exists("p#char-count");
```

`#char-count` 元素應該最初包含文字 `Character Count: 0/50`。

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

當 `#text-input` 元素包含文字 `hello coder` 時，`#char-count` 元素應該包含文字 `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

當字元數為 `50` 時，文字應該以 `red` 顯示。

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

如果字元數大於或相等於 `50`，多餘的輸入應該被修剪。

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
