---
id: 669fcb06c3034a39f5431a38
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `<body>`.

```js
assert.match(code, /<body>/i);
```

Unapaswa kuwa na tagi ya kufunga ya `</body>`.

```js
assert.match(code, /<\/body>/i);
```

Hupaswi kubadilisha kipengele chako cha `head` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Kipengele chako cha `body` kinapaswa kuja baada ya kipengele chako cha `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
