---
id: 0095d0567b6817114989a1dd
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaweka thamani algoriti ya njia fupi zaidi kwa kutumia JavaScript. Utaandika kitendakazi kinachohesabu njia fupi zaidi kati ya nodi katika grafu yenye uzito na kurudisha umbali pamoja na njia zilizochukuliwa.

Kwa mfano, ukizingatia grafu ambapo miji imeunganishwa na barabara zenye umbali tofauti, algoriti itapata njia fupi zaidi kutoka mji mmoja hadi mwingine. Ikiwa unataka kusafiri kutoka Mji A hadi Mji D, algoriti inaweza kugundua kuwa kwenda A ⇨ B ⇨ C ⇨ D (jumla: km 15) ni fupi zaidi kuliko kwenda moja kwa moja A ⇨ D (km 20).

Ili kuanza, fafanua kigezo kinachoitwa `INF` na uweke thamani `Infinity`. Baadaye, utakitumia kuonyesha kuwa hakuna muunganisho wa moja kwa moja kati ya nodi mbili.

# --hints--

Unapaswa kufafanua kigezo kinachoitwa `INF`.

```js
assert.isDefined(INF);
```

Kigezo chako cha `INF` kinapaswa kuwa na thamani `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
