---
id: 69c937c73f23b1140bcd297a
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Kigezo cha `targetNode` ni hiari. Ukikosa kupeanwa, kitendakazi kitaweka hesabu njia fupi kutoka `startNode` hadi nodi zote nyingine katika grafu.

Toa thamani ya msingi ya `targetNode` kwa kigezo cha `null`.

# --hints--

Kigezo cha `targetNode` cha kitendakazi chako cha `shortestPath` kinapaswa kuwa na thamani ya msingi `null`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {

}
--fcc-editable-region--
```
