---
id: 69c918841987bf16c1d1baba
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Utahitaji safu ya data ya pande mbili ili wakilisha matriki ya ujirani ya grafu. Kila thamani inawakilisha uzito (umbali) wa kivuko kati ya nodi mbili. Thamani ya `INF` ina maana hakuna kivuko cha moja kwa moja.

Unda kigezo kingine kinachoitwa `adjMatrix` na weka thamani safu ya data ya pande mbili yenye thamani zifuatazo:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Unapaswa kuunda kigezo kinachoitwa `adjMatrix`.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` inapaswa kuwa safu ya data ya pande mbili yenye mistari 6 na nguzo 6.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` inapaswa kuwa na thamani sahihi.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
