---
id: 69c918841987bf16c1d1bacb
title: Hatua 20
challengeType: 1
dashedName: step-20
---

# --description--

Mara mzunguko mkuu unapomalizika, unahitaji kuamua nodi gani au nodi zipi kuonyesha matokeo yake.

Ikiwa `targetNode` maalum imetolewa, onyesha matokeo kwa nodi hiyo tu. Vinginevyo, onyesha matokeo kwa nodi zote.

Katika JavaScript, unaweza kupata viashiria vyote vya safu ya data kama inayoweza kurudiwa ukitumia `[...Array(n).keys()]`.

Baada ya mzunguko wa nje wa `for`, tengeneza kigezo kinachoitwa `targets`. Ukitumia kikokotoo cha masharti cha pande tatu, ikiwa `targetNode` si sawa kabisa na `null`, weka thamani `[targetNode]`, vinginevyo weka `[...Array(n).keys()]`.

# --hints--

Unapaswa kutengeneza kigezo `targets` baada ya mzunguko mkuu.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

Kigezo chako cha `targets` kinapaswa kutumia kikokotoo cha masharti cha pande tatu kinachokagua `targetNode !== null`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

Wakati lengo linapotolewa, `targets` inapaswa kuwa `[targetNode]`, vinginevyo inapaswa kuwa `[...Array(n).keys()]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
