---
id: 69c918841987bf16c1d1bac3
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Ndani ya mzunguko wa ndani wa `for`, unahitaji kuangalia kama nodi ya sasa haijatembelewa na iko karibu zaidi kuliko bora uliyoipata hadi sasa.

Ongeza tamko la `if` linaloangalia kama `nodeNo` haijatembelewa na `distances[nodeNo]` ni ndogo kuliko `minDistance`.

Acha mwili ukiwa tupu kwa sasa.

# --hints--

Unapaswa kuwa na tamko la `if` lenye sharti la `!visited[nodeNo] && distances[nodeNo] < minDistance` ndani ya mzunguko wa ndani wa `for`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
