---
id: 69c918841987bf16c1d1babd
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Unahitaji kufuatilia umbali mfupi unaojulikana kutoka nodi ya mwanzo hadi kila nodi nyingine. Kuanza, utadhani kila nodi iko umbali usio na kikomo.

Katika JavaScript, unaweza kuunda safu ya data iliyojaa thamani kwa kutumia `new Array(n).fill(value)`. Kwa mfano, `new Array(3).fill(0)` huunda `[0, 0, 0]`.

Ndani ya kitendakazi cha `shortestPath`, unda kigezo kinachoitwa `distances` na weka thamani `new Array(n).fill(INF)`.

# --hints--

Unapaswa kuunda kigezo kinachoitwa `distances` ndani ya `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Kigezo chako cha `distances` kinapaswa kuwekwa kuwa `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
