---
id: 69c918841987bf16c1d1babb
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa utaunda kitendakazi kikuu kinachofanya kazi zote.

Unda kitendakazi kinachoitwa `shortestPath` kinachopokea vigezo vitatu: `matrix`, `startNode`, na `targetNode`.

Katika hatua chache zijazo, utajenga mantiki ya kitendakazi hiki.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Kitendakazi chako cha `shortestPath` kinapaswa kuwa na vigezo `matrix`, `startNode`, na `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
