---
id: 69c918841987bf16c1d1baca
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Unapopata njia fupi zaidi ya kufikia jirani, pia unahitaji kusahihisha njia iliyorekodiwa ya kufikia jirani huyo.

Katika JavaScript, unaweza kuunda safu mpya ya data kwa kuchanganya safu iliyopo na kipengele kipya kwa kutumia opereta ya spread:

```js
const newPath = [...existingArray, newElement];
```

Ndani ya blokki iliyopangwa ndani ya `if`, sasisha `paths[nodeNo]` ili iwe njia ya kufikia nodi ya sasa na `nodeNo` ikiongezwa mwishoni kwa kutumia sintaksia ya spread kama ifuatavyo: `[...paths[current], nodeNo]`.

# --hints--

Unapaswa kusahihisha `paths[nodeNo]` ndani ya blokki iliyopangwa ndani ya `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` inapaswa kuwekwa kuwa `[...paths[current], nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
