---
id: 69c918841987bf16c1d1babc
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Unahitaji kuhifadhi nambari ya nodi katika grafu.

Ndani ya kitendakazi cha `shortestPath`, tengeneza kigezo kinachoitwa `n` na umweke thamani `matrix.length`.

# --hints--

Unapaswa kutengeneza kigezo `n` ndani ya kitendakazi cha `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

Kigezo chako cha `n` kinapaswa kuwa na thamani `matrix.length`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
