---
id: 69c918841987bf16c1d1bac1
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Sasa utaongeza mzunguko mkuu unaoendesha algoriti. Unakimbia mara moja kwa kila nodi katika grafu, ukichagua nodi isiyotembelewa iliyo karibu zaidi kila mara.

Ndani ya kitendakazi cha `shortestPath`, ongeza mzunguko wa `for` unaokimbia mara `n` kwa kutumia `i` kama kigeuzi cha mzunguko.

Mwanzoni mwa mwili wa mzunguko, kwa kutumia `let`, tangaza vigezo viwili `minDistance` na `current` vilivyowekwa kwa `INF` na `-1` mtawaliwa. Hivi vitafuatilia nodi isiyotembelewa iliyo karibu zaidi inayopatikana kila mara ya mzunguko.

# --hints--

Unapaswa kuwa na mzunguko wa `for` ndani ya `shortestPath` unaokimbia mara `n`. Usisahau kuweka `i` na `let` ndani yake.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Unapaswa kuwa na kigezo cha `minDistance` kilichowekwa kwa `INF` ndani ya mzunguko.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Unapaswa kuwa na kigezo cha `current` kilichowekwa kwa `-1` ndani ya mzunguko.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
