---
id: 69c918841987bf16c1d1bacd
title: Hatua 22
challengeType: 1
dashedName: step-22
---

# --description--

Unataka kuonyesha matokeo tu kwa nodi ambazo zinaweza kufikiwa na tofauti na nodi ya kuanzia.

Ndani ya mzunguko wa `for...of`, ongeza tamko la `if` ambalo linakagua kama `nodeNo` ni sawa kabisa na `startNode` **au** `distances[nodeNo]` ni sawa kabisa na `INF`. Ikiwa mojawapo ni kweli, tumia `continue` ndani ya blokki ya `if` kuruka kwa mzunguko unaofuata.

# --hints--

Unapaswa kuwa na tamko la `if` lenye sharti la `nodeNo === startNode || distances[nodeNo] === INF` ndani ya mzunguko wako wa `for...of`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*\w+\s*===\s*startNode\s*\|\|\s*distances\s*\[\s*\w+\s*\]\s*===\s*INF\s*\)/
);
```

Unapaswa kutumia `continue` katika blokki ya `if`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*\w+\s*===\s*startNode\s*\|\|\s*distances\s*\[\s*\w+\s*\]\s*===\s*INF\s*\)[\s\S]*?continue/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
  
  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
