---
id: 67fcd561ae787aba05a9d917
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`form` 요소 안에 텍스트 `label`를 가진 `Search for an item:` 요소를 추가하세요. 다음으로, `label`에 `htmlFor` 속성을 `search`로 설정하세요. 이렇게 하면 레이블이 다음에 추가할 `input` 요소와 연결됩니다.

`label` 아래에 `input` 요소를 추가하고 `type`와 `id` 속성을 `search`로 설정하세요. 또한 `input`에 자리 표시자 `Search...`과 `aria-describedby` 속성을 `search-description`로 설정하세요.

마지막으로 `p` 아래에 `input` 요소를 추가하고 텍스트 `Type to filter the list below:`과 `id`를 `search-description`로 설정하세요. 이렇게 하면 화면 읽기 프로그램에 추가적인 문맥을 제공합니다.

# --hints--

`form` 요소는 텍스트 `label`를 가진 `Search for an item:` 요소를 포함해야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const label = testElem.querySelector('form label');
  assert.exists(label);
  assert.equal(label.textContent.toLowerCase().trim(), 'search for an item:');
```

`label`에는 `htmlFor` 속성이 `search`로 설정되어야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const label = testElem.querySelector('form label');
  assert.equal(label.getAttribute('for'), 'search');
```

`form` 요소는 `input`와 `type` 속성이 `id`로 설정된 `search` 요소를 포함해야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const input = testElem.querySelector('form input');
  assert.exists(input);
  assert.equal(input.getAttribute('type'), 'search');
  assert.equal(input.getAttribute('id'), 'search');
```

`input`에는 자리 표시자 `Search...`와 `aria-describedby` 속성이 `search-description`로 설정되어야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const input = testElem.querySelector('form input');
  assert.equal(input.getAttribute('placeholder').toLowerCase().trim(), 'search...');
  assert.equal(input.getAttribute('aria-describedby'), 'search-description');
```

`form` 요소는 텍스트 `p`를 가진 `Type to filter the list below:` 요소를 포함해야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const p = testElem.querySelector('form p');
  assert.exists(p);
  assert.equal(p.textContent.toLowerCase().trim(), 'type to filter the list below:');
```

`p` 요소에는 `id`가 `search-description`로 설정되어야 합니다.

```js
  const testElem = await __helpers.prepTestComponent(window.index.ShoppingList);
  const p = testElem.querySelector('form p');
  assert.equal(p.getAttribute('id'), 'search-description');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Shopping List</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main id="app-container"></main>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { ShoppingList } from "./index.jsx";
      const appContainer = document.getElementById("app-container");
      const root = ReactDOM.createRoot(appContainer);
      root.render(<ShoppingList />);
    </script>
  </body>
</html>

```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
  line-height: 1.5;
  color: var(--dark-grey);
  background-color: var(--dark-grey);
}

main,
.container {
  display: flex;
  flex-direction: column;
  align-items: center;
}

.container {
  background-color: var(--light-grey);
  width: 90%;
  margin: 20px;
  padding: 10px;
}

h1 {
  color: var(--dark-grey);
}

form {
  text-align: center;
}

button {
  cursor: pointer;
}

button {
  cursor: pointer;
  width: 100px;
  margin: 3px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

input {
  color: var(--dark-grey);
  margin-left: 5px;
  padding: 3px;
}

li {
  text-align: left;
  margin: 10px 0;
}

.selected-item {
  font-weight: bold;
}

@media (min-width: 425px) {
  .container {
    width: 400px;
  }
}

```

```jsx
const { useState } = React;

export const ShoppingList = () => {
  return (
    <div className="container">
      <h1>Shopping List</h1>
      <form>
      --fcc-editable-region--
      
      --fcc-editable-region--
      </form>
    </div>
  );
};

```
