---
id: 6619240f46cec8e04d77e03a
title: Zoezi la msingi la kitendakazi A 
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

Tengeneza kitendakazi kinachopokea nambari kamili. Kitendakazi hiki kinapaswa kurudisha `integer + 7` iliyotolewa ikiwa nambari kamili ni ndogo kuliko `10`. Ikiwa nambari kamili ni kubwa au sawa na `10`, kinapaswa kurudisha `integer - 3` iliyotolewa.

Jina la kitendakazi linapaswa kuwa `addOrSubtract`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `addOrSubtract`.

```js
assert.isFunction(addOrSubtract);
```

Kitendakazi chako kinapaswa kupokea nambari kamili kama hoja.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

Unapaswa kurudisha nambari kamili iliyotolewa + 7 ikiwa nambari kamili ni ndogo kuliko 10.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

Unapaswa kurudisha nambari kamili iliyotolewa - 3 ikiwa nambari kamili ni kubwa au sawa na 10.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
