---
id: 661e17c6068359c3ccf2f4d8
title: Zoezi la vitendakazi vya msingi D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

Andika kitendakazi, chenye jina `lastLetter`, kinachopokea mfuatano wa herufi kama kigezo na kurudisha herufi ya mwisho ya mfuatano huo.

# --hints--

Unapaswa kuwa na kitendakazi chenye jina `lastLetter`.

```js
assert.isFunction(lastLetter);
```

Kitendakazi chako kinapaswa kupokea mfuatano wa herufi kama kigezo.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

Unapaswa kurudisha herufi ya mwisho ya mfuatano huo.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
