---
id: 661e131f068359c3ccf2f4d6
title: Zoezi la msingi la vitendakazi B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Andika kitendakazi, chenye jina `multiply`, kinachopokea vigezo viwili na kurudisha bidhaa yao.

# --hints--

Unapaswa kuwa na kitendakazi chenye jina `multiply`.

```js
assert.isFunction(multiply);
```

Kitendakazi chako kinapaswa kupokea nambari kamili mbili kama hoja.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

Unapaswa kurudisha bidhaa ya nambari kamili hizo mbili.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
