---
id: 661e151f068359c3ccf2f4d7
title: Zoezi la msingi la kitendakazi C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

Andika kitendakazi, chenye jina `capitalize`, kinachopokea mfuatano wa herufi kama kigezo na kurudisha mfuatano mpya wa herufi wenye herufi ya kwanza kubwa.

# --hints--

Unapaswa kuwa na kitendakazi chenye jina `capitalize`.

```js
assert.isFunction(capitalize);
```

Kitendakazi chako kinapaswa kupokea mfuatano wa herufi kama kigezo.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

Kitendakazi chako kinapaswa kurudisha mfuatano mpya wa herufi wenye herufi ya kwanza kubwa.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
