---
id: 655774d01daeeaa1978b99d5
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

처음에는 그래프의 다른 모든 노드가 출발 노드로부터 무한대 거리에 있다고 간주합니다. 거리가 아직 결정되지 않았기 때문입니다.

`else` 절을 만들고 `distances` 사전의 노드에 무한대 값을 할당하세요. 이를 위해 `float()` 함수에 문자열 `'inf'`를 인수로 사용하여 양의 무한대를 나타내는 부동소수점 숫자를 생성하세요.

# --hints--

`else` 절이 있어야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

새로운 `float('inf')` 절 안에서 `distances[node]`에 `else`를 할당해야 합니다.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
