---
id: 655777060d8ddea6741be1b1
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

`distances`의 모든 거리는 시작 노드를 제외하고 무한대로 설정되어 있습니다. `unvisited` 목록에는 그래프의 모든 노드가 포함되어 있습니다. 하지만 실제로 이 결과를 얻기 위해 `for` 반복문이 필요하지 않습니다.

`for` 반복문과 그 전체 본문을 제거하세요.

# --hints--

`for` 반복문과 그 안에 중첩된 모든 코드를 제거해야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)distances\s*=\s*\{\s*\}\s*\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
