---
id: 65578d4fc3afc3b8f554c882
title: 41단계
challengeType: 20
dashedName: step-41
---

# --description--

`min()`는 키워드 전용 인수도 받습니다. `min()`에 추가 인수로 함수를 전달하면 리스트 항목을 비교하는 방식을 수정할 수 있습니다.

이전 단계에서 작성한 줄의 결과는 알파벳 순서로 가장 먼저 오는 노드입니다. 대신 시작 노드에서 가장 거리가 짧은 방문하지 않은 노드를 선택하려고 합니다.

`key=distances.get`를 `min()` 호출의 두 번째 인수로 전달하세요. 이렇게 하면 비교가 `unvisited` 사전 안에 각 `distances` 리스트 항목이 가진 값에 따라 이루어집니다.

# --hints--

`key=distances.get`를 `min()` 호출의 두 번째 인수로 전달해야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
