---
id: 6557709b0aee699a6a00528c
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

예를 들어, 그래프는 공간 내 두 점 `A`와 `B`가 경로로 연결된 모습을 나타낼 수 있습니다. 이런 그래프는 두 개의 노드가 간선으로 연결된 형태입니다.

기존 `'species'` 키를 문자열 `'A'`로 바꾸세요. 그런 다음, `'B'`와 `'A'` 노드 간 연결을 나타내기 위해 해당 값도 문자열 `'B'`로 바꾸세요.

# --hints--

사전에는 `'A'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

`'A'` 키의 값은 `'B'`여야 합니다.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

사전에는 두 개의 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
