---
id: 655771d889132f9ccd341060
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

그래프의 간선에 거리, 시간 또는 기타 정량적 값을 나타내는 가중치가 연결되어 있으면 이를 *가중 그래프*라고 합니다.

여기서는 이 가중치가 각 노드 또는 공간상의 점 사이의 거리가 됩니다. 가중 그래프를 나타내려면 사전을 수정하여 각 값에 대해 튜플 목록을 사용할 수 있습니다.

튜플의 첫 번째 요소는 연결된 노드이고, 두 번째 요소는 거리를 나타내는 정수입니다.

`my_graph['A']`를 튜플 목록으로 수정하세요. `A-B` 거리 값이 `3`이고 `A-D` 거리 값이 `1`임을 고려하세요.

# --hints--

`my_graph["A"]`는 `('B', 3)`와 `('D', 1)` 튜플을 포함하는 목록이어야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph`는 `'A'`, `'B'`, `'C'`, `'D'`라는 4개의 키를 가져야 합니다.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
