---
id: 65578cee7f2cb8b80127cce2
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

`print` 호출 전에, `while`가 비어 있지 않은 동안 실행되는 `unvisited` 반복문을 만드세요. 반복문 본문을 채우려면 `pass` 키워드를 사용하세요.

# --hints--

`while`가 비어 있지 않은 동안 실행되는 `unvisited` 반복문이 있어야 합니다. `pass` 키워드를 잊지 마세요.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "shortest_path");     
    assert(function_body.match(/^\s+while\s+(unvisited|unvisited\s*!=\s*\[\s*\]|len\s*\(\s*unvisited\s*\)\s*(>|!=)\s*0)\s*:/m));
    const {block_body} = __helpers.python.getBlock(commentless_code, /while\s+(unvisited|unvisited\s*!=\s*\[\s*\]|len\s*\(\s*unvisited\s*\)\s*(>|!=)\s*0)\s*/);
    assert(block_body.match(/\s+pass/))
    
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
--fcc-editable-region--
```
