---
id: 6557712d77ce2d9bd7e63afd
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

그래프에 `B`에 연결된 또 다른 노드를 추가하고 `C`이라고 이름 붙이세요.

기존 사전을 이 배열을 나타내도록 수정하세요: `'C'`에 또 다른 키 `my_graph`를 추가하고 문자열 `'B'`를 값으로 지정하세요.

또한, 기존 `'B'` 키의 값을 리스트 `['A', 'C']`로 변경하여 `'B'` 노드의 여러 연결을 나타내세요.

# --hints--

사전에는 `'A'`, `'B'`, `'C'` 세 개의 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']`의 값은 문자열 `'B'`여야 합니다.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']`은 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']`의 값은 `'A'`와 `'C'`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']`의 값은 문자열 `'B'`이어야 합니다.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
