---
id: 65577a17564ce8a8e06c1460
title: 35단계
challengeType: 20
dashedName: step-35
---

# --description--

사전 내포는 조건부 `if`/`else` 구문도 지원합니다:

```py
{key: val_1 if condition else val_2 for key in dict}
```

위 예제에서 `dict`는 기존 사전입니다. `condition`가 `True`로 평가되면 `key`는 `val_1` 값을 가지며, 그렇지 않으면 `val_2`가 됩니다.

사전 내포를 사용해 `graph`를 기반으로 사전을 만들고 `distances` 변수에 할당하세요. 노드가 시작 노드와 같으면 키에 0 값을 주고, 그렇지 않으면 무한대를 주세요. 후자를 위해 `float('inf')`를 사용하세요.

# --hints--

`distances` 변수에 값을 주기 위해 사전 내포 구문을 사용하세요.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
