---
id: 65578d0f6c78a0b868a43b9c
title: 40단계
challengeType: 20
dashedName: step-40
---

# --description--

`while` 반복문 안에서 가장 먼저 해야 할 일은 방문할 현재 노드를 정의하는 것입니다. 이를 위해 `min()` 함수를 사용할 수 있습니다. 이 함수는 인수로 전달된 반복 가능한 객체에서 가장 작은 항목을 반환합니다.

`pass`를 삭제한 다음, `current`라는 변수를 만들고 `min(unvisited)`를 할당하세요.

# --hints--

`current` 반복문 안에 `while` 변수를 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

`min(unvisited)` 변수에 `current`를 할당해야 합니다. `pass`를 삭제하는 것도 잊지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
