---
id: 655774955b097ea14897db12
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

시작 노드에서의 거리는 0입니다. 알고리즘이 바로 그 지점에서 평가를 시작하기 때문입니다.

루프에서 `node`를 `unvisited`에 추가한 후, 노드가 시작 노드와 같을 때 실행되는 `if` 문을 만드세요. 그런 다음 `0` 사전 안에서 그 노드에 `distances`를 할당하세요.

# --hints--

`if`가 `node`와 같을 때 실행되는 `start` 문을 만드세요.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

새로운 `if` 문 안에서 `0` 사전의 노드에 `distances`를 할당해야 합니다.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
