---
id: 655775221059f5a20493d5d7
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

`for` 반복문 다음에 `print()` 호출을 추가하고, 생성한 변수들의 값을 확인하려면 다음 문자열 `f'Unvisited: {unvisited}\nDistances: {distances}'`를 전달하세요.

# --hints--

`f'Unvisited: {unvisited}\nDistances: {distances}'` 반복문 다음에 `for`를 출력하세요.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(^\s*)for.*:.*^\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
--fcc-editable-region--
```
