---
id: 6557743527cb92a06417ea97
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

그래프를 순회하기 위해 `for` 반복문을 만들고, `.append()` 메서드를 사용해 각 노드를 `unvisited` 리스트의 끝에 추가하세요.

# --hints--

`for` 함수 안에서 `graph`를 순회하기 위해 `shortest_path` 반복문을 만들어야 합니다.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "shortest_path");    
    assert(function_body.match(/^\s*for\s+\w+\s+in\s+graph\s*:/m));
  }
})
```

각 노드를 `unvisited` 반복문 안에서 `for`에 추가해야 합니다.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const block_regex = /for\s+(\w+)\s+in\s+graph\s*/;
    const {block_body} = __helpers.python.getBlock(commentless_code, block_regex);
    const loop_condition = commentless_code.match(block_regex);
    const regex = new RegExp(`^\\s+unvisited\\.append\\s*\\(\\s*${loop_condition[1]}\\s*\\)`, "m");
    assert(block_body.match(regex));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
--fcc-editable-region--
```
