---
id: 655773f8b8b5db9fc6d0ae76
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

방문한 노드를 추적하려면 그래프의 모든 노드 목록이 필요합니다. 노드를 방문하면 그 목록에서 삭제됩니다.

이제 `pass` 키워드를 `unvisited`라는 변수로 바꾸고 빈 리스트를 할당하세요.

# --hints--

`unvisited` 함수 안에 `shortest_path`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

`unvisited` 변수에 빈 리스트를 할당해야 합니다. `pass`를 삭제하는 것도 잊지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
