---
id: 655790d113d14dbb727eaf41
title: 45단계
challengeType: 20
dashedName: step-45
---

# --description--

`distances` 사전 안에 노드까지의 거리가 설정되면, 그 노드까지의 경로도 추적해야 합니다. 처리된 튜플에서 노드의 거리가 업데이트되었다면, 경로의 마지막 항목은 노드 자체입니다.

조건문 안에 `if` 문을 중첩하여 `paths[node]`의 마지막 요소가 `node`와 같을 때 실행되도록 하세요. `pass`를 사용해 `if` 문 본문을 채우세요.

# --hints--

`if` 문을 중첩하여 `paths[node][-1]`이 `node`와 같은지 확인하세요. `pass`를 사용하는 것을 잊지 마세요.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1(\s{4})if\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\2\2pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
```
