---
id: 6557746aad2844a0cd864e12
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

알고리즘이 그래프를 탐색하는 동안, 시작 노드와 다른 노드들 사이의 현재 알려진 최단 거리를 추적해야 합니다.

`for` 반복문 전에 `distances`이라는 새 변수를 만들고 빈 딕셔너리를 할당하세요.

# --hints--

`distances`라는 변수가 있어야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

`distances` 변수는 빈 딕셔너리여야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
