---
id: 655776db1eeae0a620e42a0d
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

이제 `my_graph`와 `'A'`을 인수로 전달하여 함수를 호출하세요.

# --hints--

`shortest_path`를 호출할 때 `my_graph`와 `'A'`를 인수로 전달해야 합니다.

```js
({ test: () => assert.match(code, /^shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
--fcc-editable-region--
```
