---
id: 65578c17d54dfab65cd54b95
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

알고리즘이 시작 노드부터 평가를 시작하므로, `paths` 딕셔너리를 만든 후에는 시작 노드를 `paths` 딕셔너리 내 자신의 리스트에 추가해야 합니다.

`.append()` 메서드를 사용해 `start`를 `paths[start]` 리스트에 추가하세요.


# --hints--

`.append()` 메서드를 사용해 `start`를 `paths[start]`에 추가해야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}paths\s*\[\s*start\s*\]\s*\.append\s*\(\s*start\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
