---
id: 6557716aadbd2d9c42c0e69a
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

마지막 노드 `'D'`를 추가하세요. 이 노드는 `'A'`과 `'C'`에 연결되어 있습니다.

사전을 수정하여 이 구조를 나타내세요. 여러 연결을 나타낼 때는 리스트를 사용하세요.

# --hints--

사전에는 `'A'`, `'B'`, `'C'`, `'D'`라는 4개의 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']`는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph["A"]) is list
  `))
})
```

`my_graph['A']`는 `'B'`와 `'D'`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph["A"]) == 2 and "B" in my_graph["A"] and "D" in my_graph["A"]
  `))
})
```

`my_graph['B']`는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

`my_graph['B']`는 `'A'`와 `'C'`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

`my_graph['C']`는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list
  `))
})
```

`my_graph['C']`는 `'B'`와 `'D'`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph["C"]) == 2 and "B" in my_graph["C"] and "D" in my_graph["C"]
  `))
})
```

`my_graph['D']`는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list
  `))
})
```

`my_graph['D']`는 `'A'`와 `'C'`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph["D"]) == 2 and "A" in my_graph["D"] and "C" in my_graph["D"]
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': ['A', 'C'],
    'C': 'B'
}
--fcc-editable-region--
```
