---
id: 65578cb031cd93b77a285db2
title: 38단계
challengeType: 20
dashedName: step-38
---

# --description--

함수는 시작 노드와 연결된 모든 노드를 탐색할 것입니다. 모든 노드에 대해 최단 경로를 계산합니다. 그런 다음 시작 노드를 방문하지 않은 노드 목록에서 제거합니다.

다음으로 가장 가까운 이웃 노드를 방문하고 이 과정을 모든 노드를 방문할 때까지 반복합니다.

이제 그래프에서 노드를 탐색하는 메인 루프 작업을 진행할 것입니다. 알고리즘 개발 중 무한 루프 실행 문제를 피하려면 함수 호출을 주석으로 바꾸세요.

# --hints--

함수 호출을 주석으로 바꾸세요.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
