---
id: 65577739f57ecca6c39bb4e9
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

`list()` 타입 생성자를 사용하면 반복 가능한 객체에서 리스트를 만들 수 있습니다.

`unvisited` 변수의 할당을 `list()`를 사용하도록 수정하고, 반복 가능한 객체로 `graph`을 전달하세요.

# --hints--

`list()` 딕셔너리에서 리스트를 생성하려면 `graph`를 사용하세요.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

`list(graph)`에 `unvisited`를 할당해야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
