---
id: 65579228c669fcbebffd01d5
title: 52단계
challengeType: 20
dashedName: step-52
---

# --description--

다른 버그는 미묘합니다. 이웃 노드에 대해 더 짧은 거리가 발견되면 `paths[current]`가 이웃 노드 경로인 `paths[node]`에 할당됩니다.

이것은 두 변수가 같은 리스트를 가리킨다는 뜻입니다. 리스트는 변경 가능하므로, 이웃 노드를 경로에 추가하면 `paths[node]`와 `paths[current]` 모두가 수정되는데, 이는 같은 리스트이기 때문입니다. 이로 인해 거리는 올바르지만 경로가 잘못됩니다.

이 버그는 `paths[current]`의 복사본을 이웃 노드 경로에 할당하여 고칠 수 있습니다. 이를 위해 슬라이스 구문을 사용할 수 있습니다:

```py
my_list[:]
```

여기서 `my_list`는 복사하려는 리스트입니다. `paths[node] = paths[current]` 블록 안에 있는 기존 `if` 할당을 `paths[current]`을 슬라이싱하는 방식으로 수정하세요.

# --hints--

이웃 노드 경로에 `paths[current][:]`를 할당해야 합니다.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
