---
id: 6557924d47c325bf27afbe51
title: 53단계
challengeType: 20
dashedName: step-53
---

# --description--

알고리즘은 완성되었지만 출력 결과를 개선할 수 있습니다. 또한, 함수에 추가 인수를 제공하여 두 노드 사이의 경로만 반환하도록 할 수 있습니다.

함수 선언에 `target`를 세 번째 매개변수로 추가하고 기본값을 빈 문자열로 지정하세요.

# --hints--

함수는 세 개의 매개변수를 받아야 합니다: `graph`, `start`, 그리고 `target`. 순서가 중요합니다.

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

`target` 매개변수는 기본값으로 빈 문자열을 가져야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
