---
id: 65576ff7888f9e96f52a4be1
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

사전은 *키*-*값* 쌍 형태로 데이터를 저장합니다. 키는 콜론으로 해당 값과 구분되고, 각 키-값 쌍은 쉼표로 다음 쌍과 구분됩니다:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

사전에 새 키-값 쌍을 추가하세요. 문자열 `'species'`를 키로, 문자열 `'guinea pig'`을 값으로 사용하세요.

# --hints--

`'species'` 사전 안에 `'guinea pig'` 키와 `copper` 값을 포함해야 합니다.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

`copper` 사전에는 단일 키-값 쌍만 있어야 합니다.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
