---
id: 65577791ad8c26a7705e2919
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

딕셔너리 컴프리헨션을 사용하면 기존 딕셔너리에서 시작해 새 딕셔너리를 만들 수 있습니다:

```py
{key: val for key in dict}
```

위 예시에서 `val`는 새 딕셔너리에서 `key`가 가질 값이고, `dict`는 기존 딕셔너리입니다.

시작 노드와 다른 각 노드 사이의 경로를 추적하려고 합니다.

`distances` 변수 다음에 `paths` 변수를 만들고 `graph`의 모든 키를 가진 딕셔너리를 할당하세요. 각 키에 빈 리스트를 할당하고 딕셔너리 컴프리헨션을 사용해 딕셔너리를 만드세요.

# --hints--

`paths`라는 변수가 있어야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

`paths` 변수는 딕셔너리 컴프리헨션 구문을 사용해 graph의 각 노드에 빈 리스트를 할당해야 합니다.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
