---
id: 655773b0591c5f9f4045883e
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

알고리즘은 지정된 노드에서 시작합니다. 그런 다음 그래프를 탐색하여 시작 노드, 즉 *source*와 다른 모든 노드 사이의 최단 경로를 찾습니다.

이를 위해 함수에는 `graph`와 `start` 두 개의 매개변수가 필요합니다. 함수 선언에 이들을 추가하세요.

# --hints--

함수는 `graph`와 `start`를 이 순서대로 매개변수로 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
