---
id: 655791e6cf5e03be3de73451
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

함수 호출의 주석을 해제하려 하면 작동하지 않습니다. 고쳐야 할 버그가 몇 가지 있습니다. 첫 번째는 중첩된 `if`에서 `paths[node]` 목록에 없을 수도 있는 요소에 접근하려 하기 때문에 발생합니다. 따라서 `paths[node]`에 접근하기 전에 `paths[node][-1]`가 비어 있지 않은지 확인해야 합니다.

중첩된 `if` 문에 추가 조건을 넣어 `paths[node]`에 접근하기 전에 `paths[node][-1]`가 비어 있지 않은지 확인하세요.

# --hints--

중첩된 `paths[node]` 문에 `if`을 첫 번째 조건으로 추가하세요. 조건을 결합하려면 `and` 연산자를 사용하세요.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
