---
id: 65577236b056379d5dbc7000
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

이제 새 그래프의 각 노드 사이 최단 경로를 계산하는 알고리즘을 개발하기 시작할 것입니다.

`shortest_path`라는 빈 함수를 선언하세요. 함수 본문을 채우려면 `pass` 키워드를 사용하세요.

# --hints--

`shortest_path`라는 이름의 함수를 정의해야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
