---
id: 6559d70c5161b16ff1d6530d
title: 54단계
challengeType: 20
dashedName: step-54
---

# --description--

Python은 삼항 구문을 사용해 `if`/`else` 조건문을 간결하게 작성하는 방법을 제공합니다:

```py
val_1 if condition else val_2
```

위 표현식은 `val_1`이 참이면 `condition`로 평가되고, 그렇지 않으면 `val_2`로 평가됩니다.

`print` 호출을 삭제하고 `targets_to_print` 반복문 뒤에 `while`라는 변수를 만드세요. 삼항 구문을 사용해 `[target]`이 참일 때 `target`을 할당하고, 그렇지 않으면 `graph`을 할당하세요.


# --hints--

`print` 호출을 삭제해야 합니다.

```js
({ test: () => assert.isFalse( /print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\\nPaths:\s\{\s*paths\s*\}\1\s*\)/.test(code)) })
```

`targets_to_print` 반복문 뒤에 `while`라는 변수를 만들어야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=/ms));
  }
})
```

삼항 구문을 사용해 `[target]`이 참일 때 `target`를, 그렇지 않으면 `graph`을 `targets_to_print` 변수에 할당해야 합니다.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=\s*\[\s*target\s*\]\s+if\s+target\s+else\s+graph/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start, target = ''):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
