---
id: 65661b72d6745ebec6a96923
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

마찬가지로, `C-D` 거리가 `7`임을 고려하여 나머지 두 리스트를 수정하세요.

# --hints--

`my_graph['C']`는 튜플 `('B', 4)`과 `('D', 7)`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']`는 튜플 `('A', 1)`과 `('C', 7)`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph`는 `'A'`, `'B'`, `'C'`, `'D'`라는 4개의 키를 가져야 합니다.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']`는 튜플 `('B', 3)`와 `('D', 1)`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']`는 튜플 `('A', 3)`과 `('C', 4)`를 포함하는 리스트여야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
