---
id: 6566195b0a021bb660b2b4b1
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

이제 `my_graph['B']`를 튜플 목록으로 수정하세요. 튜플의 첫 번째 요소는 연결된 노드이고, 두 번째 요소는 거리입니다. `B-C` 거리는 `4`입니다.

# --hints--

`my_graph['B']`는 `('A', 3)`와 `('C', 4)` 튜플을 포함하는 목록이어야 합니다.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph`는 `'A'`, `'B'`, `'C'`, `'D'`라는 이름의 4개 키를 가져야 합니다.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
