---
id: 65789506b30453080f77470c
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

지금까지 다양한 데이터 유형을 만났습니다:

- 정수, 문자열, 튜플, 불리언과 같은 불변 데이터 유형
- 리스트와 딕셔너리 같은 변경 가능한 데이터 유형

딕셔너리는 중괄호 쌍 `{}`로 구분됩니다.

먼저 `copper`이라는 변수를 만들고, 빈 리스트를 대괄호 쌍으로 만드는 것처럼 중괄호 쌍을 사용해 빈 딕셔너리를 할당하세요.

# --hints--

`copper`라는 변수가 있어야 합니다.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

`copper` 변수에는 빈 딕셔너리 값이 있어야 합니다. 중괄호 쌍을 사용하세요.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
