---
id: 6578b13757611e2825beb8a5
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

키는 딕셔너리 내에서 고유해야 하며, 변경 불가능한 데이터 타입만 사용할 수 있습니다. 즉, 리스트나 다른 딕셔너리를 키로 사용할 수 없습니다.

딕셔너리에 또 다른 키 `'age'`를 추가하고 값으로 정수 `2`를 지정하세요.

# --hints--

`'age'` 딕셔너리 안에서 `2` 바로 뒤에 값이 `'species': 'guinea pig'`인 `copper` 키가 있어야 합니다. 쉼표를 잊지 마세요.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

`copper` 딕셔너리에 키-값 쌍이 두 개 있어야 합니다.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
