---
id: 6578b57361f2f132a02e2a18
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

사전의 키를 통해 저장된 데이터에 접근할 수 있습니다:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

사전 뒤에 위 예시를 따라 `'species'`의 `copper` 키에 접근하고 결과를 출력하세요.

# --hints--

사전의 할당은 수정하지 마세요.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

`copper['species']` 키의 값을 접근하려면 `'species'`를 사용하세요.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

`print()`를 인수로 전달하여 `copper['species']`를 호출하세요.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
