---
id: 6579717f0920131304286804
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

사전을 선언한 후 새 키-값 쌍을 추가하려면, 기존 키에 접근하는 것과 같은 방식으로 키를 지정하고 대입 연산자를 사용해 새 키의 값을 설정할 수 있습니다.

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

`print()` 호출을 삭제하세요. 그런 다음 `copper` 선언 후에 사전에 `'food'` 키를 추가하고 그 값을 `'hay'`으로 설정하세요.

# --hints--

코드에 `print(copper['age'])`가 없어야 합니다.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

사전을 선언한 후 `'food'`에 `copper` 키를 추가해야 합니다.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

사전을 선언한 후 `copper['food']`을 `'hay'`으로 설정해야 합니다.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
