---
id: 6579dd49fa8a8e1fd06b85a9
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

그래프는 요소 쌍 간의 관계를 나타내는 자료 구조입니다.
이 요소들은 *노드*라고 하며, 실제 사물, 개체, 공간상의 점 또는 기타가 될 수 있습니다. 노드 간의 연결은 *간선*이라고 합니다.

다음은 그래프의 시각적 표현입니다:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


`copper` 사전을 `my_graph`으로 이름을 바꾸세요. 이것이 알고리즘을 테스트할 그래프를 나타냅니다.

# --hints--

`copper` 사전의 이름을 `my_graph`으로 바꾸세요.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

`my_graph` 변수는 사전이어야 합니다.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
