---
id: 65797670e0c0d016f17e7660
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 코드 맨 아래에 `copper`를 출력하세요.

# --hints--

코드 맨 아래에 `copper`를 출력해야 합니다.

```js
({ test: () =>
  {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /print\s*\(\s*copper\s*\)\s*$/)
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

--fcc-editable-region--
```
