---
id: 6579cbab9825b8170974c69a
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`del` 키워드를 사용해 딕셔너리에서 키-값 쌍을 삭제할 수 있습니다.

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

`for` 반복문 바로 전에 `del` 키워드를 사용해 `'age'` 키와 그 값을 `copper`에서 삭제하세요.

# --hints--

`del` 반복문 전에 `copper['age']`를 삭제하려면 `for` 키워드를 사용하세요.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
