---
id: 65796fac81f983127558f3f4
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이제 기존 `print()` 호출을 수정하여 `age` 키의 값을 출력하세요.


# --hints--

`copper['age']` 키의 값을 접근하려면 `age`를 사용하세요.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

`print()`를 인수로 전달하여 `copper['age']`를 호출해야 합니다.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
