---
id: 6579ca0923cfa7162089d2f0
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

기존 키의 값을 변경하는 데에도 같은 구문을 사용할 수 있습니다.

`print()` 호출 직전에 `'species'` 키에 접근하여 그 값을 `'Cavia porcellus'`로 다시 할당하세요.

# --hints--

`copper['species']` 호출 전에 `'Cavia porcellus'`에 문자열 `print()`를 다시 할당해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
