---
id: 6579cd5f6dd62c189e53ddbb
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

사전의 기본적인 측면을 검토했으니, 이제 최단 경로 알고리즘을 작성할 수 있습니다.

`copper` 사전 선언 이후의 모든 코드를 삭제하세요.

# --hints--

사전 선언 이후의 모든 코드를 삭제해야 합니다.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
