---
id: 65affe602a2a828b0382a058
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`'age'` 키를 문자열 `'B'`로 바꾸고, 그 값을 문자열 `'A'`로 설정하여 노드 간 양방향 연결을 나타내세요.

# --hints--

사전에는 `'B'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

`'B'` 키의 값은 `'A'`여야 합니다.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

사전에는 `'A'` 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

`'A'` 키의 값은 `'B'`여야 합니다.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

사전에는 두 개의 키가 있어야 합니다.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
