---
id: 65b7d3319ebec69b983fb91b
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

출력 결과에서 볼 수 있듯이, `.items()`는 각 키-값 쌍을 별도의 튜플에 저장하는 자료 구조를 만듭니다. 이 튜플의 요소를 순회하려면 두 번째 반복 변수(loop variable)를 추가할 수 있습니다.

```py
for i, j in dict.items():
    print(i, j)
```

`for` 반복문을 수정하여 두 개의 반복 변수를 받고, 반복문 본문 안에서 두 변수 모두를 출력하세요.

# --hints--

`for` 반복문을 두 개의 변수로 사용하도록 수정하세요.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

`for` 반복문 안에서 두 반복 변수를 모두 출력하세요.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
