---
id: 65b7cf140d34058d7ea8935f
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

사전 키의 값들을 순회하려면, 한 가지 방법은 `.values()` 메서드를 사용하는 것입니다.

`for` 반복문을 `copper.values()` 대신 `copper`를 순회하도록 수정하고 출력 결과를 확인하세요.

# --hints--

`for` 반복문을 `copper.values()` 대신 `copper`를 순회하도록 수정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
