---
id: 65b7cd2b7bd9a684ccf1dc16
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

사전의 키를 순회하려면 사전을 `for` 반복문에 넣으면 됩니다. 아래 코드는 사전 `dict`의 각 키를 출력할 것입니다.

```py
for i in dict:
   print(i)
```

`print()` 호출을 `for` 반복문으로 바꾸어 `copper`를 순회하며 각 키를 출력하세요.

# --hints--

코드에 `print(copper)`가 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

`for` 사전을 순회하기 위해 `copper` 반복문을 만들어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

`copper` 반복문 안에서 `for`의 각 키를 출력해야 합니다.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
