---
id: 65b7d25992879698180e6a71
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

마지막으로, 키-값 쌍을 순회할 수 있으려면 `.items()` 메서드를 사용할 수 있습니다.

`for` 반복문을 `copper.items()` 대신 `copper.values()`를 순회하도록 수정하세요.

# --hints--

`for` 반복문을 `copper.items()` 대신 `copper.values()`를 순회하도록 수정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
