---
id: 614385513d91ae5c251c2052
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop creerai una pagina di una rivista usando un comunicato stampa di una versione precedente del curriculum di freeCodeCamp.

Inizia aggiungendo un elemento `link` per il foglio di stile del carattere `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, un `link` per il foglio di stile `https://use.fontawesome.com/releases/v5.8.2/css/all.css` di FontAwesome e un `link` per il tuo foglio di stile `./styles.css`.

Il tuo foglio di stile per i caratteri caricherà tre caratteri tipografici separati: `Anton`, `Baskervville` e `Raleway`.

# --hints--

Il tuo codice dovrebbe avere tre elementi `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

I tuoi elementi `link` dovrebbero essere all’interno del tuo elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

I tuoi tre elementi `link` dovrebbero avere un attributo `rel` con valore `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno dei tuoi elementi link dovrebbe avere l’attributo `href` impostato su `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Uno dei tuoi elementi link dovrebbe avere l’attributo `href` impostato su `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Uno dei tuoi elementi `link` dovrebbe avere l’attributo `href` impostato su `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
