---
id: 6143869bb45bd85e3b1926aa
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

All'interno del tuo elemento `.heading`, crea un elemento `header` con l'insieme `class` impostato su `hero`.

In quell'elemento, crea un elemento `img` con l'insieme `src` impostato su `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, l'insieme `alt` impostato su `freecodecamp logo` e l'insieme `class` impostato su `hero-img`.

L'attributo `loading` su un elemento `img` può essere impostato su `lazy` per dire al browser di non caricare la risorsa immagine finché non è necessaria (cioè, quando l'utente scorre l'immagine in vista). Come beneficio aggiuntivo, gli elementi caricati in modo lazy non si caricheranno finché gli elementi non lazy non sono caricati — questo significa che gli utenti con connessioni internet lente possono visualizzare il contenuto della tua pagina senza dover aspettare il caricamento delle immagini.

Dai al tuo nuovo elemento `img` un attributo `loading` impostato su `lazy`.

# --hints--

Dovresti creare un elemento `header`.

```js
assert.exists(document.querySelector('header'));
```

Il tuo elemento `header` dovrebbe essere all'interno del tuo elemento `.heading`.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Il tuo elemento `header` dovrebbe avere l'insieme `class` impostato su `hero`.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

Dovresti creare un elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

Il tuo elemento `img` dovrebbe essere all'interno del tuo elemento `header`.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Il tuo elemento `img` dovrebbe avere l'insieme `src` impostato su `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Il tuo elemento `img` dovrebbe avere l'insieme `alt` impostato su `freecodecamp logo`.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Il tuo elemento `img` dovrebbe avere l'attributo `loading` impostato su `lazy`.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Il tuo elemento `img` dovrebbe avere l'insieme `class` impostato su `hero-img`.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
