---
id: 67fccc2463253b213a000142
title: Costruisci un lettore multimediale
challengeType: 25
dashedName: lab-multimedia-player
demoType: onClick
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con gli elementi `audio` e `video`. In questo laboratorio, costruirai un lettore multimediale che mostrerà una traccia `audio` e `video` con una trascrizione.

Per l'elemento `audio`, dovrai includere un elemento `source` che serve a specificare il media utilizzato.

Ecco un esempio:

```html
<audio controls aria-label="descriptive label goes here">
  <source src="url-to-audio-goes-here" type="audio/mpeg">
</audio>
```

L'elemento `source` può essere usato anche nell'elemento `video` in questo modo:

```html
<video controls width="600" aria-label="descriptive label goes here">
  <source src="link-to-mp4-goes-here" type="video/mp4">
  <!-- Remaining code goes here -->  
</video>
```

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `h1` per il titolo principale della pagina.
2. Dovresti avere tre elementi `section`.
3. All'interno del primo elemento `section`, dovresti avere un elemento `h2` per il titolo della canzone in riproduzione.
4. Sotto l'elemento `h2`, dovresti avere un elemento `audio` con l'attributo `controls` e un attributo `aria-label`.
5. All'interno dell'elemento `audio`, dovresti avere un elemento `source` con un attributo `src` che punta a un file audio e un attributo `type`. Puoi usare questo URL audio se vuoi: `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`
6. All'interno del secondo elemento `section`, dovresti avere un elemento `h2` per il titolo del video in riproduzione.
7. Sotto l'elemento `h2`, dovresti avere un elemento `video` con gli attributi `controls`, `width` e un attributo `aria-label`.
8. All'interno dell'elemento `video`, dovresti avere un elemento `source` con un attributo `src` che punta a un file video e un attributo `type`. Puoi usare questo URL video se vuoi: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`
9. Sotto l'elemento `source`, dovresti avere un elemento `track` con un attributo `src` che punta a un file di sottotitoli e gli attributi `kind`, `srclang` e `label`. Puoi usare questo URL per i sottotitoli se vuoi: `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.vtt`
10. All'interno del terzo elemento `section`, dovresti avere un elemento `h2` per il titolo della sezione, ad esempio "Trascrizione".
11. Sotto l'elemento `h2`, dovresti avere un elemento `p` con la trascrizione del video.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.isNotNull(document.querySelector("h1"));
```

Il tuo `h1` dovrebbe contenere del testo che rappresenta il titolo principale della pagina.

```js  
assert.isNotEmpty(document.querySelector("h1")?.textContent?.trim());  
```

Dovresti avere tre elementi `section`.

```js
const sectionElements = document.querySelectorAll("section");
assert.lengthOf(sectionElements, 3);
```

Il primo `section` dovrebbe contenere un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const firstSection = sectionElements[0];
assert.isNotNull(firstSection?.querySelector("h2"));
```

Il tuo `h2` dovrebbe contenere del testo che rappresenta il titolo della canzone in riproduzione.

```js
const firstSection = document.querySelectorAll("section")[0];
assert.isNotEmpty(firstSection?.querySelector("h2")?.textContent?.trim());
```

Il primo `section` dovrebbe contenere un elemento `audio`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
assert.isNotNull(audio);
```

Il tuo elemento `audio` dovrebbe avere un attributo `controls`.

```js  
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("controls"));
```

Il tuo elemento `audio` dovrebbe avere un attributo `aria-label` con un testo che descrive l'audio.

```js
const firstSection = document.querySelectorAll("section")[0];  
const audio = firstSection?.querySelector("audio");  
assert.isTrue(audio?.hasAttribute("aria-label"));  

const ariaLabel = audio?.getAttribute("aria-label");  
assert.isString(ariaLabel);  
assert.isNotEmpty(ariaLabel.trim());
```

Il tuo elemento `audio` dovrebbe avere un elemento `source`.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isNotNull(source);
```

Il tuo elemento `source` dovrebbe avere un attributo `src` che punta a un file audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source?.getAttribute("src")?.trim().length > 0);
```

Il tuo elemento `source` dovrebbe avere un attributo `type` con un valore che specifica il tipo di audio.

```js
const firstSection = document.querySelectorAll("section")[0];
const audio = firstSection?.querySelector("audio");
const source = audio?.querySelector("source");

assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source?.getAttribute("type")?.trim().length > 0);
```

Il secondo `section` dovrebbe contenere un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const h2 = sectionElements[1]?.querySelector("h2");
assert.isNotNull(h2);
```

Il tuo `h2` dovrebbe contenere del testo che rappresenta il titolo del video in riproduzione.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Il secondo `section` dovrebbe contenere un elemento `video`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isNotNull(video);
```

Il tuo elemento `video` dovrebbe avere un attributo `controls`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("controls"));
```

Il tuo elemento `video` dovrebbe avere un attributo `width` con un valore che specifica la larghezza del video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
assert.isTrue(video?.hasAttribute("width"));

const width = video?.getAttribute("width");
assert.isTrue(width.trim().length > 0);
```

L'elemento `video` dovrebbe avere un elemento `source`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isNotNull(source);
```

L'elemento `source` dovrebbe avere un attributo `src` che punta a un file video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("src"));
assert.isTrue(source.getAttribute("src").trim().length > 0);
```

L'elemento `source` dovrebbe avere un attributo `type` con un valore che specifica il tipo di video.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const source = video?.querySelector("source");
assert.isTrue(source?.hasAttribute("type"));
assert.isTrue(source.getAttribute("type").trim().length > 0);
```

Il tuo elemento `video` dovrebbe avere un elemento `track`.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isNotNull(track);
```

Il tuo elemento `track` dovrebbe avere un attributo `src` che punta a un file di sottotitoli.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("src"));
assert.isTrue(track?.getAttribute("src")?.trim().length > 0);
```

Il tuo elemento `track` dovrebbe avere un attributo `kind` con un testo che descrive il tipo di file.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("kind"));
assert.isTrue(track?.getAttribute("kind")?.trim().length > 0);
```

Il tuo elemento `track` dovrebbe avere un attributo `srclang` con un testo che descrive la lingua dei sottotitoli.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("srclang"));
assert.isTrue(track?.getAttribute("srclang")?.trim().length > 0);
```

Il tuo elemento `track` dovrebbe avere un attributo `label` con un testo che descrive la lingua dei sottotitoli.

```js
const secondSection = document.querySelectorAll("section")[1];
const video = secondSection?.querySelector("video");
const track = video?.querySelector("track");
assert.isTrue(track?.hasAttribute("label"));
assert.isTrue(track?.getAttribute("label")?.trim().length > 0);
```

Il terzo `section` dovrebbe contenere un elemento `h2`.

```js
const sectionElements = document.querySelectorAll("section");
const thirdSection = sectionElements[2];
assert.isNotNull(thirdSection.querySelector("h2"));
```

Il tuo `h2` dovrebbe contenere del testo che rappresenta il titolo della sezione.

```js
assert.isNotEmpty(document.querySelector("h2")?.textContent?.trim());
```

Il terzo `section` dovrebbe contenere un elemento `p` per la trascrizione.

```js
const thirdSection = document.querySelectorAll("section")[2];
const transcript = thirdSection?.querySelector("p");
assert.isNotNull(transcript);
assert.isAbove(transcript.textContent.trim().length, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
    <h2>Now Playing: Sailing Away</h2>
    <audio controls aria-label="Audio Player for Sailing Away">
        <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
    </audio>
  </section>

  <section>
    <h2>What is a map method and how does it work?</h2>
    <video controls width="600" aria-label="What is a map method and how does it work?">
      <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4">
      
      <track 
      src="what-is-a-map-method.vtt" 
      kind="captions" 
      srclang="en" 
      label="English" 
      default>

    </video>
  </section>

  <section>
    <h2>Transcript</h2>
    <p>What is a map method and how does it work? The map method is a powerful and widely used function in JavaScript that operates on arrays. It's designed to create a new array by applying a given function to each element of the original array. This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
    </p>
  </section>

</body>
</html> 
```
