---
id: 6a89c69c111694edc6d1c628
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Prueba tu función `countup` pasando `5` dentro de una llamada a `print()`. La salida debe ser `[1, 2, 3, 4, 5]`.

Con eso, el taller de la función countup está completo.

# --hints--

Debes llamar a `countup(5)` dentro de la función `print()`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
