---
id: 6a89c69c6af682856558d78b
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

La llamada recursiva devuelve una lista que contiene los números más pequeños. Usa `append()` para agregar `number` a `count_list`, luego devuelve `count_list`.

# --hints--

Debes agregar `number` a `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Debes devolver `count_list` desde tu función `countup`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Tu función `countup` debe devolver `[1, 2, 3]` cuando se llame con `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
