---
id: 6a89c69c8dbfda74151fa0b6
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, crearás una función `countup` que usa recursión para crear una lista de números desde `1` hasta un número dado.

Comienza definiendo una función llamada `countup` con un parámetro llamado `number`. Agrega la palabra clave `pass` al cuerpo de la función para que sea válida mientras está vacía.

# --hints--

Debes tener una función llamada `countup` con un parámetro llamado `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Tu función `countup` debe contener la palabra clave `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
