---
id: 6a89c69c99298438d6fbe60b
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Ahora agrega el caso recursivo. Después de la sentencia `if`, crea una variable llamada `count_list`. Asígnale el resultado de llamar a `countup` con `number - 1`.

Cada llamada usa un número más pequeño, por lo que la recursión avanza hacia el caso base.

# --hints--

Debes asignar `countup(number - 1)` a una variable llamada `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
