---
id: 6a89c69c91bc7365fbd4b733
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Agrega el caso base a tu función `countup`. Reemplaza `pass` con una sentencia `if` que verifique si `number` es menor que `1`. Dentro de la sentencia `if`, devuelve una lista vacía.

# --hints--

Tu función `countup` debe devolver una lista vacía cuando `number` sea menor que `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Tu función `countup` no debe devolver una lista vacía cuando `number` sea mayor o igual a `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
