---
id: 66a37f37ef5823a313de8c26
title: 建造月球軌道
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

在本實驗中，你將使用 HTML 和 CSS 創建月球繞地球軌道的簡單動畫。地球將位於系統的中心，而月球將繞著它運行。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該透過將高度設定為視窗的 `100%` 並在 `body` 元素中設定 flexbox 佈局，將所有元素置中對齊。
   
2. 你應該有一個帶有類別 `space` 的 `div`。

3. 在 `.space` 元素內，應該依序有兩個帶有類別 `earth` 和 `orbit` 的 `div` 元素。

4. 在 `.orbit` 元素內，應該有另一個帶有 `moon` 類別的 `div`。
   
5. 具有 `class` 為 `space` 的 `div` 元素應置中於頁面，並且寬度和高度皆為 `200px`。

6. 具有 `space` `class` 的 `div` 元素應該使用 `relative` 定位。

7. `.earth` 元素應該使用 `absolute` 定位。將 `.earth` 元素的中心定位在其父元素的垂直（top）和水平（left）軸的中點。之後，將 `.earth` 元素向後移動自身寬度和高度的一半，以使其在父元素 `.space` 中置中。

8. `.earth` 元素應該有 `100px` 的寬度和高度。

9.  `.orbit` 元素應該有 `200px` 的寬度和高度。

10.  `.orbit` 元素應該使用 `absolute` 定位。它的右下角應該位於 `.space` 元素的中心，並使用 `transform` 屬性在垂直和水平方向上各移動 `-50%`。

11. 月球繞地球的軌道路徑應該是圓形。

12. `.moon` 元素應該使用 `absolute` 定位，並且寬度和高度為 `30px`。`.moon` 元素應該定位在 `.orbit` 元素的頂部，並且水平置中。

13. 你應該進一步調整 `.moon` 元素的水平位置，將該元素向左移動其寬度的一半。

14. `.earth` 選擇器應該有背景顏色和 `border-radius` 為 `50%`。

15. `.moon` 選擇器應該有背景顏色和 `border-radius` 為 `50%`。

16. 你應該定義一個 `@keyframes orbit` 動畫，使 `.orbit` 元素繞其中心旋轉 360 度。你應該將此動畫以持續時間為 `5` 秒、線性時間函式以及無限次迭代的方式套用到 `.orbit` 元素。

**注意：** 請務必在你的 HTML 中聯結你的樣式表並套用你的 CSS。

# --hints--

`body` 元素的內容應該透過將高度設定為視窗高度的 `100%`、`display` 設為 `flex`、`justify-content` 和 `align-items` 設為 `center`，使其置中於頁面。

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

你應該有一個主要的 `div`，其類別為 `space`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

`.space` 元素應該有寬度和高度為 `200px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

`.space` 元素應該具有 `position` 屬性且值為 `relative`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

你應該在具有 `class` 為 `space` 的 `div` 裡面放一個 `class` 為 `earth` 的 `div`。

```js
assert.exists(document.querySelector('div.space div.earth'));
```

`.earth` 元素應該是 `.space` 元素的第一個子元素。

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

`.earth` 元素應該有 `100px` 的寬度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

`.earth` 元素應該使用 `absolute` 定位。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

`.earth` 元素的 `top` 和 `left` 位置應為 `50%`。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

應將 `.earth` 元素向後和向上移動 `50%`，方法是將 `transform` 屬性設為 `translate(-50%, -50%)`。


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

你應該在 `.space` 元素內有一個帶有 `orbit` 類別的 `div`。

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

`.orbit` 元素應該位於 `.earth` 元素之後。

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

`.orbit` 元素應該有 `200px` 的寬度和高度。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

`.orbit` 元素應該使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

`.orbit` 元素應該使用 `transform` 屬性，將其在垂直和水平方向上各移動 `-50%`，以置中於 `.space` 元素。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

你應該在 `.orbit` 元素內有一個帶有 `moon` 類別的 `div`。

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

`.moon` 元素的寬度和高度應為 `30px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

`.moon` 元素應該使用 `absolute` 定位。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

`.moon` 元素的頂邊應與其包含區塊的頂邊對齊。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

你應該將 `.moon` 元素的左邊緣移動到包含區塊的中間（從左側算起 `50%`）。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

應該透過沿 x 軸平移 `-50%` 來水平調整 `.moon` 元素。

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

`.earth` 選擇器應該有背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

`.earth` 選擇器應該有 `border-radius` 為 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

`.moon` 選擇器應該有背景顏色。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

`.moon` 選擇器應該有 `border-radius` 為 `50%`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

你應該有一個 `@keyframes` 規則。

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

你的新 `@keyframes` 規則應命名為 `orbit`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

你的 `@keyframes orbit` 規則應該有一個 `0%` 選擇器。

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

你的 `@keyframes orbit` 規則應該有一個 `100%` 選擇器。

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

你的 `0%` 選擇器應該有一個設定為 `rotate(0deg) translate(-50%, -50%)` 的 `transform` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

你的 `100%` 選擇器應該在你的 `0%` 選擇器之後。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

你的 `100%` 選擇器應該有一個設定為 `rotate(360deg) translate(-50%, -50%)` 的 `transform` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

你的 `.orbit` 選擇器應該有一個設定為 `orbit 5s linear infinite` 的 `animation` 屬性。

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

