---
id: 579e2a2c335b9d72dd32e05c
title: 實作 Slice 和 Splice 演算法
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

在這個實驗中，你需要創建一個演算法來合併兩個陣列。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 創建一個 `frankenSplice` 函式（程式），該函式接受兩個陣列和一個索引。
1. 將第一個陣列的每個元素依序複製到第二個陣列中，從指定的索引開始，並傳回結果陣列。
1. 在函式執行後，輸入的陣列應保持不變。

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` 應該傳回 `[4, 1, 2, 3, 5]`。

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` 應該傳回 `["a", 1, 2, "b"]`。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` 應該傳回 `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`。

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

第一個陣列中的所有元素應該以原始順序添加到第二個陣列中。`frankenSplice([1, 2, 3, 4], [], 0)` 應該傳回 `[1, 2, 3, 4]`。

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

第一個陣列在函式執行後應保持不變。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

第二個陣列在函式執行後應保持不變。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
