---
id: 668f08ea07b99b1f4a91acab
title: Створіть сторінку рецепта
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати оголошення `!DOCTYPE html`.
1. Ви повинні мати елемент `html` з атрибутом `lang`, встановленим у `en`.
1. Ви повинні мати елемент `head`, який містить елемент `title` з назвою вашого рецепта, а також елемент `meta` з атрибутом `charset`, встановленим у `UTF-8`.
1. Ви повинні мати елемент `body`.
1. Ви повинні мати елемент `h1` з назвою вашого рецепта.
1. Ви повинні мати елемент `p`, який вводить рецепт під `h1`.
1. Ви повинні мати один елемент `h2` з текстом `Ingredients` для розділу інгредієнтів.
1. Ви повинні мати невпорядкований список (елемент `ul`) з принаймні чотирма елементами списку (елементи `li`), які перелічують ваші інгредієнти під першим елементом `h2`.
1. Ви повинні мати другий елемент `h2` з текстом `Instructions` для розділу інструкцій.
1. Ви повинні мати впорядкований список (елемент `ol`) з принаймні чотирма елементами списку, які перелічують кроки рецепта у правильному порядку, під другим елементом `h2`.
1. Ви повинні мати один елемент `img` з атрибутом `src`, встановленим на дійсне зображення, ви можете використати `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`, якщо хочете, а також атрибут `alt`, який описує зображення.

# --hints--

Ваша сторінка рецепта повинна мати оголошення `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Ви повинні мати елемент `html` з атрибутом `lang`, встановленим у `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Ви повинні мати елемент `head` всередині елемента `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `title` всередині вашого елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Ваш елемент `title` повинен містити назву вашого рецепта.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Ви повинні мати елемент `meta` всередині вашого елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Ваш елемент `meta` повинен мати атрибут `charset`, встановлений у `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Ви повинні мати елемент `body` всередині вашого елемента `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `h1` з назвою вашого рецепта.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

У вас повинен бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ви повинні мати елемент `p` під вашим елементом `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Ваш перший елемент `p` повинен описувати ваш рецепт.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Ваш перший елемент `h2` повинен містити текст `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Ви повинні мати невпорядкований список під вашим першим елементом `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Ви повинні мати принаймні чотири елементи списку у вашому невпорядкованому списку з інгредієнтами.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Ваш другий елемент `h2` повинен містити текст `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Ви повинні мати впорядкований список під вашим другим елементом `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Ви повинні мати принаймні чотири елементи списку у вашому впорядкованому списку з інструкціями.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Ви повинні мати принаймні один елемент `img`.

```js
assert.exists(document.querySelector('img'));
```

Усі ваші елементи `img` повинні мати дійсний атрибут `src` та значення.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Усі ваші елементи `img` повинні мати атрибут `alt`, який описує зображення.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
