---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Побудуйте лічильник символів у реальному часі
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

У цій лабораторній роботі ви створите лічильник символів у реальному часі за допомогою JavaScript. Лічильник символів відображатиме кількість символів, введених у `textarea` елемент. Лічильник оновлюватиметься в реальному часі під час введення користувачем у `textarea`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати `textarea` елемент з `id` `text-input`.
2. Має бути `p` елемент з `id` `char-count`. Спочатку він має містити текст `Character Count: 0/50`. Цей заповнювач має бути замінений на поточну кількість символів у `textarea` елементі.
3. Коли `#text-input` елемент містить текст `hello coder`, `#char-count` елемент має містити текст `"Character Count: 11/50"`.
4. Коли `#text-input` елемент містить текст `i am a programmer`, `#char-count` елемент має містити текст `"Character Count: 17/50"`.
5. Коли `#text-input` елемент містить текст `hello world`, `#char-count` елемент має містити текст `"Character Count: 11/50"`.
6. Коли `#text-input` елемент містить текст `I am learning a new language and it is called c++.`, `#char-count` елемент має містити текст `"Character Count: 50/50"`. **ПРИМІТКА**: Хоча атрибут `maxlength` забезпечив би правильну функціональність, ця лабораторна робота вимагає рішення виключно на JavaScript.  
7. Користувач не повинен мати змогу вводити більше ніж `50` символів у `textarea` елемент. Коли кількість символів досягає `50`, будь-який зайвий ввід має автоматично обрізатися, а текст `Character Count: 50/50` має відображатися у `red`.

**Примітка:** Обов’язково підключіть свій файл стилів і JavaScript у ваш HTML.

# --hints--

Ви повинні мати `textarea` елемент з `id` `text-input`. 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Ви повинні мати `p` елемент з `id` `char-count`.

```js
assert.exists("p#char-count");
```

`#char-count` елемент спочатку має містити текст `Character Count: 0/50`.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Коли `#text-input` елемент містить текст `hello coder`, `#char-count` елемент має містити текст `"Character Count: 11/50"`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Коли кількість символів дорівнює `50`, текст має відображатися у `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Якщо кількість символів більша або дорівнює `50`, зайвий ввід має бути обрізаний.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
