---
id: 6998b99a8e195cd0c1f211df
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript 객체를 사용해 간단한 야생동물 추적기를 만드실 거예요.

다음은 객체의 예시입니다:

```js
const animal = {}; // empty object
```

이 코드는 빈 객체를 만듭니다.

이제 `tiger`라는 변수를 만들고 빈 객체를 할당하세요.

# --hints--

`tiger`라는 변수를 만들어야 합니다.

```js
assert.exists(tiger);
```

`tiger` 변수에는 빈 객체가 할당되어야 합니다.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
