---
id: 6999a0c1698b93d5d41156ed
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 `tiger` 객체에 `isEndangered`이라는 또 다른 속성을 추가하세요.

`isEndangered` 속성을 `true`으로 설정하세요.

# --hints--

`tiger`에는 `isEndangered` 속성이 있어야 합니다.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered`는 `true`와 같아야 합니다.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
