---
id: 6999b0d79ea938edcb032237
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

Wildlife Tracker 워크숍의 마지막 단계에서는 대괄호 표기법을 사용해 객체의 속성에 접근하는 함수를 만드세요.

다음은 대괄호 표기법으로 속성에 접근하는 예시입니다:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

이 예시에서 대괄호 표기법은 변수를 사용해 속성에 접근할 수 있게 합니다.

이제 `getProperty`라는 함수를 만드세요.

이 함수는 두 개의 매개변수 `animal`과 `propertyName`를 받아야 합니다.

대괄호 표기법을 사용해 속성 값을 반환하세요.

함수를 만든 후 `console.log`를 사용해 `getProperty(tiger, "species")`와 `getProperty(elephant, "age")`를 호출하여 반환된 값을 콘솔에서 확인하세요.

이로써 Wildlife Tracker 워크숍이 완료됩니다!

# --hints--

`getProperty`라는 이름의 함수를 만드세요.

```js
assert.isFunction(getProperty);
```

`getProperty` 함수는 두 개의 매개변수 `animal`과 `propertyName`를 가져야 합니다.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty`은 대괄호 표기법을 사용해 속성에 접근해야 합니다.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

`getProperty(tiger, "species")`을 콘솔에 출력하세요.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

`getProperty(elephant, "age")`을 콘솔에 출력하세요.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

`getProperty(tiger, "species")`을 호출하면 `"Tiger"`가 반환되어야 합니다.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

`getProperty(elephant, "age")`을 호출하면 `12`가 반환되어야 합니다.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty`은 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
