---
id: 6999faf886634b72be7a819b
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이 야생동물 추적기에서는 객체에 새 속성을 추가해 업데이트할 것입니다.

다음은 속성을 가진 객체를 만드는 예시입니다:

```js
const animal = {
  name: "Lion" // property: value
};
```

이제 `tiger` 객체를 업데이트하여 `species`라는 속성을 포함하도록 하세요.

그 값은 `"Tiger"`로 설정하세요.

# --hints--

`tiger`는 객체여야 합니다.

```js
assert.isObject(tiger);
```

`tiger`는 `species` 속성을 가져야 합니다.

```js
assert.property(tiger, "species");
```

`tiger.species`는 `"Tiger"`와 같아야 합니다.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
