---
id: 6999adc72c02812a28c01e31
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이번 단계에서는 객체의 속성을 업데이트하는 함수를 만드세요.

다음은 함수 안에서 속성을 업데이트하는 예시입니다:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

이 예시에서는 `age` 속성이 새 값으로 업데이트됩니다.

이제 `updateAge`라는 함수를 만드세요. 이 함수는 두 개의 매개변수 `animal`와 `newAge`를 받아야 합니다.

함수 안에서 `age` 객체의 `animal` 속성을 `newAge`로 업데이트하세요. 업데이트된 `animal` 객체를 반환하세요.

함수를 만든 후 `console.log`를 사용해 `updateAge(elephant, 12)`를 호출하여 콘솔에서 업데이트된 `elephant` 객체를 확인하세요.

# --hints--

`updateAge`라는 이름의 함수를 만드세요.

```js
assert.isFunction(updateAge);
```

`updateAge` 함수는 두 개의 매개변수 `animal`와 `newAge`를 가져야 합니다.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge`는 점 표기법을 사용해 `age` 속성을 업데이트해야 합니다.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

`updateAge` 함수는 업데이트된 `animal` 객체를 반환해야 합니다.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

`updateAge(elephant, 12)`를 콘솔에 출력하세요.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

`updateAge(elephant, 12)`를 호출하면 age 속성이 업데이트됩니다.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge`는 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
