---
id: 69999ff7f7862fdfb885ea5d
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이제 `tiger` 객체에 `age`이라는 새 속성을 추가해 업데이트하세요.

`age` 속성의 값을 `5`으로 설정하세요.

# --hints--

`tiger`에는 `age` 속성이 있어야 합니다.

```js
assert.property(tiger, "age");
```

`tiger.age`는 `5`와 같아야 합니다.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
