---
id: 6999ac171928136168654774
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이번 단계에서는 객체의 `species` 속성에 접근하는 함수를 만드세요.

객체에서 속성에 접근하는 간단한 예는 다음과 같습니다:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

이 예제에서 함수는 객체를 매개변수로 받아 그 객체의 `species` 속성을 반환합니다.

이제 `getSpecies`라는 함수를 만드세요.

함수는 `animal`라는 이름의 매개변수를 하나 받아야 합니다.

점 표기법을 사용해 `species`의 `animal` 속성을 반환하세요.

함수를 만든 후 `console.log`를 사용해 `getSpecies(tiger)`를 호출하여 콘솔에서 결과를 확인하세요.

# --hints--

`getSpecies`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(getSpecies);
```

`getSpecies` 함수는 `animal`라는 단일 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getSpecies(tiger)`를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies`는 점 표기법을 사용해 species 속성에 접근해야 합니다.

```js
assert.match(code, /animal\.species/);
```

`getSpecies(tiger)`를 호출하면 `"Tiger"`가 반환되어야 합니다.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies`는 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
