---
id: 6999b0198d18d402c36cc5fe
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이번 단계에서는 객체에서 속성을 제거하는 함수를 만드세요.

다음은 `delete`를 사용해 속성을 제거하는 예시입니다.

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

이 예시에서는 `canFly` 객체에서 `bird` 속성이 제거되었습니다.

이제 `removeEndangeredStatus`라는 함수를 만드세요. 이 함수는 `animal`라는 매개변수를 하나 받습니다.

함수 내부에서 `isEndangered` 키워드를 사용해 `animal` 객체에서 `delete` 속성을 제거하세요. 그리고 업데이트된 `animal` 객체를 반환하세요.

함수를 만든 후에는 `console.log`를 사용해 `removeEndangeredStatus(tiger)`을 호출하여 콘솔에서 업데이트된 객체를 확인하세요.

# --hints--

`removeEndangeredStatus`라는 함수를 만드세요.

```js
assert.isFunction(removeEndangeredStatus);
```

`removeEndangeredStatus` 함수는 `animal`라는 매개변수 하나를 가져야 합니다.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`removeEndangeredStatus` 함수는 업데이트된 `animal` 객체를 반환해야 합니다.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus`는 `delete` 키워드를 사용해 `isEndangered` 속성을 제거해야 합니다.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

`removeEndangeredStatus(tiger)`를 콘솔에 출력하세요.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

`removeEndangeredStatus(tiger)`를 호출하면 `isEndangered` 속성이 제거되어야 합니다.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus`는 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
