---
id: 6999b07c34c2cde21ee6e1d8
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이번 단계에서는 객체에 특정 속성이 있는지 확인하는 함수를 만드세요.

다음은 `hasOwnProperty`를 사용하는 예시입니다:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

이 예시에서 `hasColor` 함수는 `hasOwnProperty`를 사용해 객체에 `color` 속성이 존재하는지 확인합니다.

이제 `hasHabitat`라는 함수를 만드세요. 이 함수는 `animal`라는 매개변수를 하나 받습니다.

`hasOwnProperty` 메서드를 사용해 `true` 객체에 `animal`라는 속성이 있으면 `"habitat"`를 반환하고, 그렇지 않으면 `false`를 반환하세요.

함수를 만든 후 `console.log`를 사용해 `hasHabitat(tiger)`와 `hasHabitat(elephant)`를 호출하여 두 결과를 콘솔에서 확인하세요.

# --hints--

`hasHabitat`라는 함수를 만드세요.

```js
assert.isFunction(hasHabitat);
```

`hasHabitat` 함수는 `animal`라는 매개변수 하나를 가져야 합니다.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat`는 `hasOwnProperty` 메서드를 사용해 `habitat` 속성을 확인해야 합니다.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

`hasHabitat(tiger)`를 콘솔에 출력하세요.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

`hasHabitat(elephant)`를 콘솔에 출력하세요.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

`hasHabitat(tiger)`를 호출하면 서식지를 추가한 후 `true`를 반환해야 합니다.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

서식지가 없으면 `hasHabitat(elephant)`를 호출했을 때 `false`를 반환해야 합니다.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat`는 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
