---
id: 699a00de9f564bb0effaa14d
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이번 단계에서는 객체의 `age` 속성에 접근하는 함수를 만드세요.

예시는 다음과 같습니다:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

이 예시에서 함수는 객체를 매개변수로 받아 그 객체의 `age` 속성을 반환합니다.

이제 `getAge`라는 이름의 함수를 만드세요.

함수는 `animal`라는 이름의 매개변수를 하나 받아야 합니다.

점 표기법을 사용해 `age`의 `animal` 속성을 반환하세요.

함수를 만든 후에는 `console.log`를 사용해 `getAge(tiger)`를 호출하여 콘솔에서 결과를 확인하세요.

# --hints--

`getAge`라는 이름의 함수를 만드세요.

```js
assert.isFunction(getAge);
```

`getAge` 함수는 `animal`이라는 단일 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge`는 점 표기법을 사용해 `age` 속성에 접근해야 합니다.

```js
assert.match(code, /animal\.age/);
```

`getAge(tiger)`를 콘솔에 출력하세요.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

`getAge(tiger)`를 호출하면 `5`를 반환해야 합니다.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge`는 함수 매개변수를 사용해 어떤 객체와도 작동해야 합니다.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
