---
id: 69ba8da6b296391d913096a0
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 `elephant`라는 두 번째 객체를 만드세요.

다음 속성을 추가하세요:

- 값이 `species`인 `"Elephant"`
- 값이 `age`인 `10`
- 값이 `isEndangered`인 `true`

# --hints--

`elephant`라는 변수를 만들어야 합니다.

```js
assert.exists(elephant);
```

`elephant`는 `species`과 같은 `"Elephant"` 속성을 가져야 합니다.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant`는 `age`과 같은 `10` 속성을 가져야 합니다.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant`는 `isEndangered`과 같은 `true` 속성을 가져야 합니다.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
