---
id: 68773ee26f332a80bc0295db
title: реалізувати алгоритм вежі ханоя
challengeType: 27
saveSubmissionToDB: true
dashedName: implement-the-tower-of-hanoi-algorithm
---

# --description--

У цій лабораторній роботі ви розв’яжете математичну головоломку, відому як вежа Ханоя. Головоломка складається з трьох стрижнів і кількох дисків різного діаметра.

<img alt="послідовність ходів, необхідних для розв’язання головоломки вежі Ханоя з 3 дисками" src="https://cdn.freecodecamp.org/curriculum/python/tower-of-hanoi.gif" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Головоломка починається з того, що диски складені на першому стрижні за спаданням розміру, найменший диск зверху, а найбільший — знизу.

Мета головоломки вежі Ханоя — перемістити всі диски на останній стрижень. Для цього потрібно дотримуватися трьох простих правил:

- Можна рухати лише верхні диски.
- Можна рухати лише один диск за раз.
- Не можна класти більші диски на менші.

**Мета:** Виконайте наведені нижче історії користувача і пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати функцію з назвою `hanoi_solver`, яка приймає ціле число, що представляє загальну кількість дисків головоломки, як аргумент.
1. Функція `hanoi_solver` повинна розв’язувати головоломку, дотримуючись заданих правил, за 2<sup>n</sup> - 1 ходів, де `n` — загальна кількість дисків.
1. Функція `hanoi_solver` повинна повертати рядок зі всіма ходами, зробленими для розв’язання головоломки, включно з початковим розташуванням, причому кожен хід на новому рядку. Стрижні мають бути представлені як списки цілих чисел (найменший диск представлений числом `1`), причому кожен стрижень відокремлений пробілом. Наприклад, `hanoi_solver(3)` має повертати таке:

```md
[3, 2, 1] [] []
[3, 2] [] [1]
[3] [2] [1]
[3] [2, 1] []
[] [2, 1] [3]
[1] [2] [3]
[1] [] [3, 2]
[] [] [3, 2, 1]
```

# --hints--

Ви повинні мати функцію з назвою `hanoi_solver`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function('hanoi_solver')`)) })
```

Ваша функція `hanoi_solver` має приймати один аргумент.

```js
({ test: () => runPython(`
import inspect
sig = inspect.signature(hanoi_solver)
assert len(sig.parameters) == 1
`) })
```

Ваша функція `hanoi_solver` має повертати рядок.

```js
({ test: () => runPython(`
assert isinstance(hanoi_solver(2), str)
assert isinstance(hanoi_solver(6), str)
assert isinstance(hanoi_solver(11), str)
`) })
```

`hanoi_solver(2)` має повертати `[2, 1] [] []\n[2] [1] []\n[] [1] [2]\n[] [] [2, 1]`.

```js
({ test: () => runPython(`
expected = """[2, 1] [] []
[2] [1] []
[] [1] [2]
[] [] [2, 1]"""
assert hanoi_solver(2) == expected
`) })
```

`hanoi_solver(3)` має повертати `[3, 2, 1] [] []\n[3, 2] [] [1]\n[3] [2] [1]\n[3] [2, 1] []\n[] [2, 1] [3]\n[1] [2] [3]\n[1] [] [3, 2]\n[] [] [3, 2, 1]`.

```js
({ test: () => runPython(`
expected = """[3, 2, 1] [] []
[3, 2] [] [1]
[3] [2] [1]
[3] [2, 1] []
[] [2, 1] [3]
[1] [2] [3]
[1] [] [3, 2]
[] [] [3, 2, 1]"""
assert hanoi_solver(3) == expected
`) })
```

`hanoi_solver(4)` має повертати `[4, 3, 2, 1] [] []\n[4, 3, 2] [1] []\n[4, 3] [1] [2]\n[4, 3] [] [2, 1]\n[4] [3] [2, 1]\n[4, 1] [3] [2]\n[4, 1] [3, 2] []\n[4] [3, 2, 1] []\n[] [3, 2, 1] [4]\n[] [3, 2] [4, 1]\n[2] [3] [4, 1]\n[2, 1] [3] [4]\n[2, 1] [] [4, 3]\n[2] [1] [4, 3]\n[] [1] [4, 3, 2]\n[] [] [4, 3, 2, 1]`.

```js
({ test: () => runPython(`
expected = """[4, 3, 2, 1] [] []
[4, 3, 2] [1] []
[4, 3] [1] [2]
[4, 3] [] [2, 1]
[4] [3] [2, 1]
[4, 1] [3] [2]
[4, 1] [3, 2] []
[4] [3, 2, 1] []
[] [3, 2, 1] [4]
[] [3, 2] [4, 1]
[2] [3] [4, 1]
[2, 1] [3] [4]
[2, 1] [] [4, 3]
[2] [1] [4, 3]
[] [1] [4, 3, 2]
[] [] [4, 3, 2, 1]"""
assert hanoi_solver(4) == expected
`) })
```

`hanoi_solver(5)` має повертати `[5, 4, 3, 2, 1] [] []\n[5, 4, 3, 2] [] [1]\n[5, 4, 3] [2] [1]\n[5, 4, 3] [2, 1] []\n[5, 4] [2, 1] [3]\n[5, 4, 1] [2] [3]\n[5, 4, 1] [] [3, 2]\n[5, 4] [] [3, 2, 1]\n[5] [4] [3, 2, 1]\n[5] [4, 1] [3, 2]\n[5, 2] [4, 1] [3]\n[5, 2, 1] [4] [3]\n[5, 2, 1] [4, 3] []\n[5, 2] [4, 3] [1]\n[5] [4, 3, 2] [1]\n[5] [4, 3, 2, 1] []\n[] [4, 3, 2, 1] [5]\n[1] [4, 3, 2] [5]\n[1] [4, 3] [5, 2]\n[] [4, 3] [5, 2, 1]\n[3] [4] [5, 2, 1]\n[3] [4, 1] [5, 2]\n[3, 2] [4, 1] [5]\n[3, 2, 1] [4] [5]\n[3, 2, 1] [] [5, 4]\n[3, 2] [] [5, 4, 1]\n[3] [2] [5, 4, 1]\n[3] [2, 1] [5, 4]\n[] [2, 1] [5, 4, 3]\n[1] [2] [5, 4, 3]\n[1] [] [5, 4, 3, 2]\n[] [] [5, 4, 3, 2, 1]`.

```js
({ test: () => runPython(`
expected = """[5, 4, 3, 2, 1] [] []
[5, 4, 3, 2] [] [1]
[5, 4, 3] [2] [1]
[5, 4, 3] [2, 1] []
[5, 4] [2, 1] [3]
[5, 4, 1] [2] [3]
[5, 4, 1] [] [3, 2]
[5, 4] [] [3, 2, 1]
[5] [4] [3, 2, 1]
[5] [4, 1] [3, 2]
[5, 2] [4, 1] [3]
[5, 2, 1] [4] [3]
[5, 2, 1] [4, 3] []
[5, 2] [4, 3] [1]
[5] [4, 3, 2] [1]
[5] [4, 3, 2, 1] []
[] [4, 3, 2, 1] [5]
[1] [4, 3, 2] [5]
[1] [4, 3] [5, 2]
[] [4, 3] [5, 2, 1]
[3] [4] [5, 2, 1]
[3] [4, 1] [5, 2]
[3, 2] [4, 1] [5]
[3, 2, 1] [4] [5]
[3, 2, 1] [] [5, 4]
[3, 2] [] [5, 4, 1]
[3] [2] [5, 4, 1]
[3] [2, 1] [5, 4]
[] [2, 1] [5, 4, 3]
[1] [2] [5, 4, 3]
[1] [] [5, 4, 3, 2]
[] [] [5, 4, 3, 2, 1]"""
assert hanoi_solver(5) == expected
`) })
```

`hanoi_solver(n)` має розв’язувати головоломку вежі Ханоя для будь-якого додатного значення `n`.

```js
({ test: () => runPython(`
expected = """[6, 5, 4, 3, 2, 1] [] []
[6, 5, 4, 3, 2] [1] []
[6, 5, 4, 3] [1] [2]
[6, 5, 4, 3] [] [2, 1]
[6, 5, 4] [3] [2, 1]
[6, 5, 4, 1] [3] [2]
[6, 5, 4, 1] [3, 2] []
[6, 5, 4] [3, 2, 1] []
[6, 5] [3, 2, 1] [4]
[6, 5] [3, 2] [4, 1]
[6, 5, 2] [3] [4, 1]
[6, 5, 2, 1] [3] [4]
[6, 5, 2, 1] [] [4, 3]
[6, 5, 2] [1] [4, 3]
[6, 5] [1] [4, 3, 2]
[6, 5] [] [4, 3, 2, 1]
[6] [5] [4, 3, 2, 1]
[6, 1] [5] [4, 3, 2]
[6, 1] [5, 2] [4, 3]
[6] [5, 2, 1] [4, 3]
[6, 3] [5, 2, 1] [4]
[6, 3] [5, 2] [4, 1]
[6, 3, 2] [5] [4, 1]
[6, 3, 2, 1] [5] [4]
[6, 3, 2, 1] [5, 4] []
[6, 3, 2] [5, 4, 1] []
[6, 3] [5, 4, 1] [2]
[6, 3] [5, 4] [2, 1]
[6] [5, 4, 3] [2, 1]
[6, 1] [5, 4, 3] [2]
[6, 1] [5, 4, 3, 2] []
[6] [5, 4, 3, 2, 1] []
[] [5, 4, 3, 2, 1] [6]
[] [5, 4, 3, 2] [6, 1]
[2] [5, 4, 3] [6, 1]
[2, 1] [5, 4, 3] [6]
[2, 1] [5, 4] [6, 3]
[2] [5, 4, 1] [6, 3]
[] [5, 4, 1] [6, 3, 2]
[] [5, 4] [6, 3, 2, 1]
[4] [5] [6, 3, 2, 1]
[4, 1] [5] [6, 3, 2]
[4, 1] [5, 2] [6, 3]
[4] [5, 2, 1] [6, 3]
[4, 3] [5, 2, 1] [6]
[4, 3] [5, 2] [6, 1]
[4, 3, 2] [5] [6, 1]
[4, 3, 2, 1] [5] [6]
[4, 3, 2, 1] [] [6, 5]
[4, 3, 2] [1] [6, 5]
[4, 3] [1] [6, 5, 2]
[4, 3] [] [6, 5, 2, 1]
[4] [3] [6, 5, 2, 1]
[4, 1] [3] [6, 5, 2]
[4, 1] [3, 2] [6, 5]
[4] [3, 2, 1] [6, 5]
[] [3, 2, 1] [6, 5, 4]
[] [3, 2] [6, 5, 4, 1]
[2] [3] [6, 5, 4, 1]
[2, 1] [3] [6, 5, 4]
[2, 1] [] [6, 5, 4, 3]
[2] [1] [6, 5, 4, 3]
[] [1] [6, 5, 4, 3, 2]
[] [] [6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(6) == expected

expected = """[7, 6, 5, 4, 3, 2, 1] [] []
[7, 6, 5, 4, 3, 2] [] [1]
[7, 6, 5, 4, 3] [2] [1]
[7, 6, 5, 4, 3] [2, 1] []
[7, 6, 5, 4] [2, 1] [3]
[7, 6, 5, 4, 1] [2] [3]
[7, 6, 5, 4, 1] [] [3, 2]
[7, 6, 5, 4] [] [3, 2, 1]
[7, 6, 5] [4] [3, 2, 1]
[7, 6, 5] [4, 1] [3, 2]
[7, 6, 5, 2] [4, 1] [3]
[7, 6, 5, 2, 1] [4] [3]
[7, 6, 5, 2, 1] [4, 3] []
[7, 6, 5, 2] [4, 3] [1]
[7, 6, 5] [4, 3, 2] [1]
[7, 6, 5] [4, 3, 2, 1] []
[7, 6] [4, 3, 2, 1] [5]
[7, 6, 1] [4, 3, 2] [5]
[7, 6, 1] [4, 3] [5, 2]
[7, 6] [4, 3] [5, 2, 1]
[7, 6, 3] [4] [5, 2, 1]
[7, 6, 3] [4, 1] [5, 2]
[7, 6, 3, 2] [4, 1] [5]
[7, 6, 3, 2, 1] [4] [5]
[7, 6, 3, 2, 1] [] [5, 4]
[7, 6, 3, 2] [] [5, 4, 1]
[7, 6, 3] [2] [5, 4, 1]
[7, 6, 3] [2, 1] [5, 4]
[7, 6] [2, 1] [5, 4, 3]
[7, 6, 1] [2] [5, 4, 3]
[7, 6, 1] [] [5, 4, 3, 2]
[7, 6] [] [5, 4, 3, 2, 1]
[7] [6] [5, 4, 3, 2, 1]
[7] [6, 1] [5, 4, 3, 2]
[7, 2] [6, 1] [5, 4, 3]
[7, 2, 1] [6] [5, 4, 3]
[7, 2, 1] [6, 3] [5, 4]
[7, 2] [6, 3] [5, 4, 1]
[7] [6, 3, 2] [5, 4, 1]
[7] [6, 3, 2, 1] [5, 4]
[7, 4] [6, 3, 2, 1] [5]
[7, 4, 1] [6, 3, 2] [5]
[7, 4, 1] [6, 3] [5, 2]
[7, 4] [6, 3] [5, 2, 1]
[7, 4, 3] [6] [5, 2, 1]
[7, 4, 3] [6, 1] [5, 2]
[7, 4, 3, 2] [6, 1] [5]
[7, 4, 3, 2, 1] [6] [5]
[7, 4, 3, 2, 1] [6, 5] []
[7, 4, 3, 2] [6, 5] [1]
[7, 4, 3] [6, 5, 2] [1]
[7, 4, 3] [6, 5, 2, 1] []
[7, 4] [6, 5, 2, 1] [3]
[7, 4, 1] [6, 5, 2] [3]
[7, 4, 1] [6, 5] [3, 2]
[7, 4] [6, 5] [3, 2, 1]
[7] [6, 5, 4] [3, 2, 1]
[7] [6, 5, 4, 1] [3, 2]
[7, 2] [6, 5, 4, 1] [3]
[7, 2, 1] [6, 5, 4] [3]
[7, 2, 1] [6, 5, 4, 3] []
[7, 2] [6, 5, 4, 3] [1]
[7] [6, 5, 4, 3, 2] [1]
[7] [6, 5, 4, 3, 2, 1] []
[] [6, 5, 4, 3, 2, 1] [7]
[1] [6, 5, 4, 3, 2] [7]
[1] [6, 5, 4, 3] [7, 2]
[] [6, 5, 4, 3] [7, 2, 1]
[3] [6, 5, 4] [7, 2, 1]
[3] [6, 5, 4, 1] [7, 2]
[3, 2] [6, 5, 4, 1] [7]
[3, 2, 1] [6, 5, 4] [7]
[3, 2, 1] [6, 5] [7, 4]
[3, 2] [6, 5] [7, 4, 1]
[3] [6, 5, 2] [7, 4, 1]
[3] [6, 5, 2, 1] [7, 4]
[] [6, 5, 2, 1] [7, 4, 3]
[1] [6, 5, 2] [7, 4, 3]
[1] [6, 5] [7, 4, 3, 2]
[] [6, 5] [7, 4, 3, 2, 1]
[5] [6] [7, 4, 3, 2, 1]
[5] [6, 1] [7, 4, 3, 2]
[5, 2] [6, 1] [7, 4, 3]
[5, 2, 1] [6] [7, 4, 3]
[5, 2, 1] [6, 3] [7, 4]
[5, 2] [6, 3] [7, 4, 1]
[5] [6, 3, 2] [7, 4, 1]
[5] [6, 3, 2, 1] [7, 4]
[5, 4] [6, 3, 2, 1] [7]
[5, 4, 1] [6, 3, 2] [7]
[5, 4, 1] [6, 3] [7, 2]
[5, 4] [6, 3] [7, 2, 1]
[5, 4, 3] [6] [7, 2, 1]
[5, 4, 3] [6, 1] [7, 2]
[5, 4, 3, 2] [6, 1] [7]
[5, 4, 3, 2, 1] [6] [7]
[5, 4, 3, 2, 1] [] [7, 6]
[5, 4, 3, 2] [] [7, 6, 1]
[5, 4, 3] [2] [7, 6, 1]
[5, 4, 3] [2, 1] [7, 6]
[5, 4] [2, 1] [7, 6, 3]
[5, 4, 1] [2] [7, 6, 3]
[5, 4, 1] [] [7, 6, 3, 2]
[5, 4] [] [7, 6, 3, 2, 1]
[5] [4] [7, 6, 3, 2, 1]
[5] [4, 1] [7, 6, 3, 2]
[5, 2] [4, 1] [7, 6, 3]
[5, 2, 1] [4] [7, 6, 3]
[5, 2, 1] [4, 3] [7, 6]
[5, 2] [4, 3] [7, 6, 1]
[5] [4, 3, 2] [7, 6, 1]
[5] [4, 3, 2, 1] [7, 6]
[] [4, 3, 2, 1] [7, 6, 5]
[1] [4, 3, 2] [7, 6, 5]
[1] [4, 3] [7, 6, 5, 2]
[] [4, 3] [7, 6, 5, 2, 1]
[3] [4] [7, 6, 5, 2, 1]
[3] [4, 1] [7, 6, 5, 2]
[3, 2] [4, 1] [7, 6, 5]
[3, 2, 1] [4] [7, 6, 5]
[3, 2, 1] [] [7, 6, 5, 4]
[3, 2] [] [7, 6, 5, 4, 1]
[3] [2] [7, 6, 5, 4, 1]
[3] [2, 1] [7, 6, 5, 4]
[] [2, 1] [7, 6, 5, 4, 3]
[1] [2] [7, 6, 5, 4, 3]
[1] [] [7, 6, 5, 4, 3, 2]
[] [] [7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(7) == expected

expected = """[8, 7, 6, 5, 4, 3, 2, 1] [] []
[8, 7, 6, 5, 4, 3, 2] [1] []
[8, 7, 6, 5, 4, 3] [1] [2]
[8, 7, 6, 5, 4, 3] [] [2, 1]
[8, 7, 6, 5, 4] [3] [2, 1]
[8, 7, 6, 5, 4, 1] [3] [2]
[8, 7, 6, 5, 4, 1] [3, 2] []
[8, 7, 6, 5, 4] [3, 2, 1] []
[8, 7, 6, 5] [3, 2, 1] [4]
[8, 7, 6, 5] [3, 2] [4, 1]
[8, 7, 6, 5, 2] [3] [4, 1]
[8, 7, 6, 5, 2, 1] [3] [4]
[8, 7, 6, 5, 2, 1] [] [4, 3]
[8, 7, 6, 5, 2] [1] [4, 3]
[8, 7, 6, 5] [1] [4, 3, 2]
[8, 7, 6, 5] [] [4, 3, 2, 1]
[8, 7, 6] [5] [4, 3, 2, 1]
[8, 7, 6, 1] [5] [4, 3, 2]
[8, 7, 6, 1] [5, 2] [4, 3]
[8, 7, 6] [5, 2, 1] [4, 3]
[8, 7, 6, 3] [5, 2, 1] [4]
[8, 7, 6, 3] [5, 2] [4, 1]
[8, 7, 6, 3, 2] [5] [4, 1]
[8, 7, 6, 3, 2, 1] [5] [4]
[8, 7, 6, 3, 2, 1] [5, 4] []
[8, 7, 6, 3, 2] [5, 4, 1] []
[8, 7, 6, 3] [5, 4, 1] [2]
[8, 7, 6, 3] [5, 4] [2, 1]
[8, 7, 6] [5, 4, 3] [2, 1]
[8, 7, 6, 1] [5, 4, 3] [2]
[8, 7, 6, 1] [5, 4, 3, 2] []
[8, 7, 6] [5, 4, 3, 2, 1] []
[8, 7] [5, 4, 3, 2, 1] [6]
[8, 7] [5, 4, 3, 2] [6, 1]
[8, 7, 2] [5, 4, 3] [6, 1]
[8, 7, 2, 1] [5, 4, 3] [6]
[8, 7, 2, 1] [5, 4] [6, 3]
[8, 7, 2] [5, 4, 1] [6, 3]
[8, 7] [5, 4, 1] [6, 3, 2]
[8, 7] [5, 4] [6, 3, 2, 1]
[8, 7, 4] [5] [6, 3, 2, 1]
[8, 7, 4, 1] [5] [6, 3, 2]
[8, 7, 4, 1] [5, 2] [6, 3]
[8, 7, 4] [5, 2, 1] [6, 3]
[8, 7, 4, 3] [5, 2, 1] [6]
[8, 7, 4, 3] [5, 2] [6, 1]
[8, 7, 4, 3, 2] [5] [6, 1]
[8, 7, 4, 3, 2, 1] [5] [6]
[8, 7, 4, 3, 2, 1] [] [6, 5]
[8, 7, 4, 3, 2] [1] [6, 5]
[8, 7, 4, 3] [1] [6, 5, 2]
[8, 7, 4, 3] [] [6, 5, 2, 1]
[8, 7, 4] [3] [6, 5, 2, 1]
[8, 7, 4, 1] [3] [6, 5, 2]
[8, 7, 4, 1] [3, 2] [6, 5]
[8, 7, 4] [3, 2, 1] [6, 5]
[8, 7] [3, 2, 1] [6, 5, 4]
[8, 7] [3, 2] [6, 5, 4, 1]
[8, 7, 2] [3] [6, 5, 4, 1]
[8, 7, 2, 1] [3] [6, 5, 4]
[8, 7, 2, 1] [] [6, 5, 4, 3]
[8, 7, 2] [1] [6, 5, 4, 3]
[8, 7] [1] [6, 5, 4, 3, 2]
[8, 7] [] [6, 5, 4, 3, 2, 1]
[8] [7] [6, 5, 4, 3, 2, 1]
[8, 1] [7] [6, 5, 4, 3, 2]
[8, 1] [7, 2] [6, 5, 4, 3]
[8] [7, 2, 1] [6, 5, 4, 3]
[8, 3] [7, 2, 1] [6, 5, 4]
[8, 3] [7, 2] [6, 5, 4, 1]
[8, 3, 2] [7] [6, 5, 4, 1]
[8, 3, 2, 1] [7] [6, 5, 4]
[8, 3, 2, 1] [7, 4] [6, 5]
[8, 3, 2] [7, 4, 1] [6, 5]
[8, 3] [7, 4, 1] [6, 5, 2]
[8, 3] [7, 4] [6, 5, 2, 1]
[8] [7, 4, 3] [6, 5, 2, 1]
[8, 1] [7, 4, 3] [6, 5, 2]
[8, 1] [7, 4, 3, 2] [6, 5]
[8] [7, 4, 3, 2, 1] [6, 5]
[8, 5] [7, 4, 3, 2, 1] [6]
[8, 5] [7, 4, 3, 2] [6, 1]
[8, 5, 2] [7, 4, 3] [6, 1]
[8, 5, 2, 1] [7, 4, 3] [6]
[8, 5, 2, 1] [7, 4] [6, 3]
[8, 5, 2] [7, 4, 1] [6, 3]
[8, 5] [7, 4, 1] [6, 3, 2]
[8, 5] [7, 4] [6, 3, 2, 1]
[8, 5, 4] [7] [6, 3, 2, 1]
[8, 5, 4, 1] [7] [6, 3, 2]
[8, 5, 4, 1] [7, 2] [6, 3]
[8, 5, 4] [7, 2, 1] [6, 3]
[8, 5, 4, 3] [7, 2, 1] [6]
[8, 5, 4, 3] [7, 2] [6, 1]
[8, 5, 4, 3, 2] [7] [6, 1]
[8, 5, 4, 3, 2, 1] [7] [6]
[8, 5, 4, 3, 2, 1] [7, 6] []
[8, 5, 4, 3, 2] [7, 6, 1] []
[8, 5, 4, 3] [7, 6, 1] [2]
[8, 5, 4, 3] [7, 6] [2, 1]
[8, 5, 4] [7, 6, 3] [2, 1]
[8, 5, 4, 1] [7, 6, 3] [2]
[8, 5, 4, 1] [7, 6, 3, 2] []
[8, 5, 4] [7, 6, 3, 2, 1] []
[8, 5] [7, 6, 3, 2, 1] [4]
[8, 5] [7, 6, 3, 2] [4, 1]
[8, 5, 2] [7, 6, 3] [4, 1]
[8, 5, 2, 1] [7, 6, 3] [4]
[8, 5, 2, 1] [7, 6] [4, 3]
[8, 5, 2] [7, 6, 1] [4, 3]
[8, 5] [7, 6, 1] [4, 3, 2]
[8, 5] [7, 6] [4, 3, 2, 1]
[8] [7, 6, 5] [4, 3, 2, 1]
[8, 1] [7, 6, 5] [4, 3, 2]
[8, 1] [7, 6, 5, 2] [4, 3]
[8] [7, 6, 5, 2, 1] [4, 3]
[8, 3] [7, 6, 5, 2, 1] [4]
[8, 3] [7, 6, 5, 2] [4, 1]
[8, 3, 2] [7, 6, 5] [4, 1]
[8, 3, 2, 1] [7, 6, 5] [4]
[8, 3, 2, 1] [7, 6, 5, 4] []
[8, 3, 2] [7, 6, 5, 4, 1] []
[8, 3] [7, 6, 5, 4, 1] [2]
[8, 3] [7, 6, 5, 4] [2, 1]
[8] [7, 6, 5, 4, 3] [2, 1]
[8, 1] [7, 6, 5, 4, 3] [2]
[8, 1] [7, 6, 5, 4, 3, 2] []
[8] [7, 6, 5, 4, 3, 2, 1] []
[] [7, 6, 5, 4, 3, 2, 1] [8]
[] [7, 6, 5, 4, 3, 2] [8, 1]
[2] [7, 6, 5, 4, 3] [8, 1]
[2, 1] [7, 6, 5, 4, 3] [8]
[2, 1] [7, 6, 5, 4] [8, 3]
[2] [7, 6, 5, 4, 1] [8, 3]
[] [7, 6, 5, 4, 1] [8, 3, 2]
[] [7, 6, 5, 4] [8, 3, 2, 1]
[4] [7, 6, 5] [8, 3, 2, 1]
[4, 1] [7, 6, 5] [8, 3, 2]
[4, 1] [7, 6, 5, 2] [8, 3]
[4] [7, 6, 5, 2, 1] [8, 3]
[4, 3] [7, 6, 5, 2, 1] [8]
[4, 3] [7, 6, 5, 2] [8, 1]
[4, 3, 2] [7, 6, 5] [8, 1]
[4, 3, 2, 1] [7, 6, 5] [8]
[4, 3, 2, 1] [7, 6] [8, 5]
[4, 3, 2] [7, 6, 1] [8, 5]
[4, 3] [7, 6, 1] [8, 5, 2]
[4, 3] [7, 6] [8, 5, 2, 1]
[4] [7, 6, 3] [8, 5, 2, 1]
[4, 1] [7, 6, 3] [8, 5, 2]
[4, 1] [7, 6, 3, 2] [8, 5]
[4] [7, 6, 3, 2, 1] [8, 5]
[] [7, 6, 3, 2, 1] [8, 5, 4]
[] [7, 6, 3, 2] [8, 5, 4, 1]
[2] [7, 6, 3] [8, 5, 4, 1]
[2, 1] [7, 6, 3] [8, 5, 4]
[2, 1] [7, 6] [8, 5, 4, 3]
[2] [7, 6, 1] [8, 5, 4, 3]
[] [7, 6, 1] [8, 5, 4, 3, 2]
[] [7, 6] [8, 5, 4, 3, 2, 1]
[6] [7] [8, 5, 4, 3, 2, 1]
[6, 1] [7] [8, 5, 4, 3, 2]
[6, 1] [7, 2] [8, 5, 4, 3]
[6] [7, 2, 1] [8, 5, 4, 3]
[6, 3] [7, 2, 1] [8, 5, 4]
[6, 3] [7, 2] [8, 5, 4, 1]
[6, 3, 2] [7] [8, 5, 4, 1]
[6, 3, 2, 1] [7] [8, 5, 4]
[6, 3, 2, 1] [7, 4] [8, 5]
[6, 3, 2] [7, 4, 1] [8, 5]
[6, 3] [7, 4, 1] [8, 5, 2]
[6, 3] [7, 4] [8, 5, 2, 1]
[6] [7, 4, 3] [8, 5, 2, 1]
[6, 1] [7, 4, 3] [8, 5, 2]
[6, 1] [7, 4, 3, 2] [8, 5]
[6] [7, 4, 3, 2, 1] [8, 5]
[6, 5] [7, 4, 3, 2, 1] [8]
[6, 5] [7, 4, 3, 2] [8, 1]
[6, 5, 2] [7, 4, 3] [8, 1]
[6, 5, 2, 1] [7, 4, 3] [8]
[6, 5, 2, 1] [7, 4] [8, 3]
[6, 5, 2] [7, 4, 1] [8, 3]
[6, 5] [7, 4, 1] [8, 3, 2]
[6, 5] [7, 4] [8, 3, 2, 1]
[6, 5, 4] [7] [8, 3, 2, 1]
[6, 5, 4, 1] [7] [8, 3, 2]
[6, 5, 4, 1] [7, 2] [8, 3]
[6, 5, 4] [7, 2, 1] [8, 3]
[6, 5, 4, 3] [7, 2, 1] [8]
[6, 5, 4, 3] [7, 2] [8, 1]
[6, 5, 4, 3, 2] [7] [8, 1]
[6, 5, 4, 3, 2, 1] [7] [8]
[6, 5, 4, 3, 2, 1] [] [8, 7]
[6, 5, 4, 3, 2] [1] [8, 7]
[6, 5, 4, 3] [1] [8, 7, 2]
[6, 5, 4, 3] [] [8, 7, 2, 1]
[6, 5, 4] [3] [8, 7, 2, 1]
[6, 5, 4, 1] [3] [8, 7, 2]
[6, 5, 4, 1] [3, 2] [8, 7]
[6, 5, 4] [3, 2, 1] [8, 7]
[6, 5] [3, 2, 1] [8, 7, 4]
[6, 5] [3, 2] [8, 7, 4, 1]
[6, 5, 2] [3] [8, 7, 4, 1]
[6, 5, 2, 1] [3] [8, 7, 4]
[6, 5, 2, 1] [] [8, 7, 4, 3]
[6, 5, 2] [1] [8, 7, 4, 3]
[6, 5] [1] [8, 7, 4, 3, 2]
[6, 5] [] [8, 7, 4, 3, 2, 1]
[6] [5] [8, 7, 4, 3, 2, 1]
[6, 1] [5] [8, 7, 4, 3, 2]
[6, 1] [5, 2] [8, 7, 4, 3]
[6] [5, 2, 1] [8, 7, 4, 3]
[6, 3] [5, 2, 1] [8, 7, 4]
[6, 3] [5, 2] [8, 7, 4, 1]
[6, 3, 2] [5] [8, 7, 4, 1]
[6, 3, 2, 1] [5] [8, 7, 4]
[6, 3, 2, 1] [5, 4] [8, 7]
[6, 3, 2] [5, 4, 1] [8, 7]
[6, 3] [5, 4, 1] [8, 7, 2]
[6, 3] [5, 4] [8, 7, 2, 1]
[6] [5, 4, 3] [8, 7, 2, 1]
[6, 1] [5, 4, 3] [8, 7, 2]
[6, 1] [5, 4, 3, 2] [8, 7]
[6] [5, 4, 3, 2, 1] [8, 7]
[] [5, 4, 3, 2, 1] [8, 7, 6]
[] [5, 4, 3, 2] [8, 7, 6, 1]
[2] [5, 4, 3] [8, 7, 6, 1]
[2, 1] [5, 4, 3] [8, 7, 6]
[2, 1] [5, 4] [8, 7, 6, 3]
[2] [5, 4, 1] [8, 7, 6, 3]
[] [5, 4, 1] [8, 7, 6, 3, 2]
[] [5, 4] [8, 7, 6, 3, 2, 1]
[4] [5] [8, 7, 6, 3, 2, 1]
[4, 1] [5] [8, 7, 6, 3, 2]
[4, 1] [5, 2] [8, 7, 6, 3]
[4] [5, 2, 1] [8, 7, 6, 3]
[4, 3] [5, 2, 1] [8, 7, 6]
[4, 3] [5, 2] [8, 7, 6, 1]
[4, 3, 2] [5] [8, 7, 6, 1]
[4, 3, 2, 1] [5] [8, 7, 6]
[4, 3, 2, 1] [] [8, 7, 6, 5]
[4, 3, 2] [1] [8, 7, 6, 5]
[4, 3] [1] [8, 7, 6, 5, 2]
[4, 3] [] [8, 7, 6, 5, 2, 1]
[4] [3] [8, 7, 6, 5, 2, 1]
[4, 1] [3] [8, 7, 6, 5, 2]
[4, 1] [3, 2] [8, 7, 6, 5]
[4] [3, 2, 1] [8, 7, 6, 5]
[] [3, 2, 1] [8, 7, 6, 5, 4]
[] [3, 2] [8, 7, 6, 5, 4, 1]
[2] [3] [8, 7, 6, 5, 4, 1]
[2, 1] [3] [8, 7, 6, 5, 4]
[2, 1] [] [8, 7, 6, 5, 4, 3]
[2] [1] [8, 7, 6, 5, 4, 3]
[] [1] [8, 7, 6, 5, 4, 3, 2]
[] [] [8, 7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(8) == expected

expected = """[9, 8, 7, 6, 5, 4, 3, 2, 1] [] []
[9, 8, 7, 6, 5, 4, 3, 2] [] [1]
[9, 8, 7, 6, 5, 4, 3] [2] [1]
[9, 8, 7, 6, 5, 4, 3] [2, 1] []
[9, 8, 7, 6, 5, 4] [2, 1] [3]
[9, 8, 7, 6, 5, 4, 1] [2] [3]
[9, 8, 7, 6, 5, 4, 1] [] [3, 2]
[9, 8, 7, 6, 5, 4] [] [3, 2, 1]
[9, 8, 7, 6, 5] [4] [3, 2, 1]
[9, 8, 7, 6, 5] [4, 1] [3, 2]
[9, 8, 7, 6, 5, 2] [4, 1] [3]
[9, 8, 7, 6, 5, 2, 1] [4] [3]
[9, 8, 7, 6, 5, 2, 1] [4, 3] []
[9, 8, 7, 6, 5, 2] [4, 3] [1]
[9, 8, 7, 6, 5] [4, 3, 2] [1]
[9, 8, 7, 6, 5] [4, 3, 2, 1] []
[9, 8, 7, 6] [4, 3, 2, 1] [5]
[9, 8, 7, 6, 1] [4, 3, 2] [5]
[9, 8, 7, 6, 1] [4, 3] [5, 2]
[9, 8, 7, 6] [4, 3] [5, 2, 1]
[9, 8, 7, 6, 3] [4] [5, 2, 1]
[9, 8, 7, 6, 3] [4, 1] [5, 2]
[9, 8, 7, 6, 3, 2] [4, 1] [5]
[9, 8, 7, 6, 3, 2, 1] [4] [5]
[9, 8, 7, 6, 3, 2, 1] [] [5, 4]
[9, 8, 7, 6, 3, 2] [] [5, 4, 1]
[9, 8, 7, 6, 3] [2] [5, 4, 1]
[9, 8, 7, 6, 3] [2, 1] [5, 4]
[9, 8, 7, 6] [2, 1] [5, 4, 3]
[9, 8, 7, 6, 1] [2] [5, 4, 3]
[9, 8, 7, 6, 1] [] [5, 4, 3, 2]
[9, 8, 7, 6] [] [5, 4, 3, 2, 1]
[9, 8, 7] [6] [5, 4, 3, 2, 1]
[9, 8, 7] [6, 1] [5, 4, 3, 2]
[9, 8, 7, 2] [6, 1] [5, 4, 3]
[9, 8, 7, 2, 1] [6] [5, 4, 3]
[9, 8, 7, 2, 1] [6, 3] [5, 4]
[9, 8, 7, 2] [6, 3] [5, 4, 1]
[9, 8, 7] [6, 3, 2] [5, 4, 1]
[9, 8, 7] [6, 3, 2, 1] [5, 4]
[9, 8, 7, 4] [6, 3, 2, 1] [5]
[9, 8, 7, 4, 1] [6, 3, 2] [5]
[9, 8, 7, 4, 1] [6, 3] [5, 2]
[9, 8, 7, 4] [6, 3] [5, 2, 1]
[9, 8, 7, 4, 3] [6] [5, 2, 1]
[9, 8, 7, 4, 3] [6, 1] [5, 2]
[9, 8, 7, 4, 3, 2] [6, 1] [5]
[9, 8, 7, 4, 3, 2, 1] [6] [5]
[9, 8, 7, 4, 3, 2, 1] [6, 5] []
[9, 8, 7, 4, 3, 2] [6, 5] [1]
[9, 8, 7, 4, 3] [6, 5, 2] [1]
[9, 8, 7, 4, 3] [6, 5, 2, 1] []
[9, 8, 7, 4] [6, 5, 2, 1] [3]
[9, 8, 7, 4, 1] [6, 5, 2] [3]
[9, 8, 7, 4, 1] [6, 5] [3, 2]
[9, 8, 7, 4] [6, 5] [3, 2, 1]
[9, 8, 7] [6, 5, 4] [3, 2, 1]
[9, 8, 7] [6, 5, 4, 1] [3, 2]
[9, 8, 7, 2] [6, 5, 4, 1] [3]
[9, 8, 7, 2, 1] [6, 5, 4] [3]
[9, 8, 7, 2, 1] [6, 5, 4, 3] []
[9, 8, 7, 2] [6, 5, 4, 3] [1]
[9, 8, 7] [6, 5, 4, 3, 2] [1]
[9, 8, 7] [6, 5, 4, 3, 2, 1] []
[9, 8] [6, 5, 4, 3, 2, 1] [7]
[9, 8, 1] [6, 5, 4, 3, 2] [7]
[9, 8, 1] [6, 5, 4, 3] [7, 2]
[9, 8] [6, 5, 4, 3] [7, 2, 1]
[9, 8, 3] [6, 5, 4] [7, 2, 1]
[9, 8, 3] [6, 5, 4, 1] [7, 2]
[9, 8, 3, 2] [6, 5, 4, 1] [7]
[9, 8, 3, 2, 1] [6, 5, 4] [7]
[9, 8, 3, 2, 1] [6, 5] [7, 4]
[9, 8, 3, 2] [6, 5] [7, 4, 1]
[9, 8, 3] [6, 5, 2] [7, 4, 1]
[9, 8, 3] [6, 5, 2, 1] [7, 4]
[9, 8] [6, 5, 2, 1] [7, 4, 3]
[9, 8, 1] [6, 5, 2] [7, 4, 3]
[9, 8, 1] [6, 5] [7, 4, 3, 2]
[9, 8] [6, 5] [7, 4, 3, 2, 1]
[9, 8, 5] [6] [7, 4, 3, 2, 1]
[9, 8, 5] [6, 1] [7, 4, 3, 2]
[9, 8, 5, 2] [6, 1] [7, 4, 3]
[9, 8, 5, 2, 1] [6] [7, 4, 3]
[9, 8, 5, 2, 1] [6, 3] [7, 4]
[9, 8, 5, 2] [6, 3] [7, 4, 1]
[9, 8, 5] [6, 3, 2] [7, 4, 1]
[9, 8, 5] [6, 3, 2, 1] [7, 4]
[9, 8, 5, 4] [6, 3, 2, 1] [7]
[9, 8, 5, 4, 1] [6, 3, 2] [7]
[9, 8, 5, 4, 1] [6, 3] [7, 2]
[9, 8, 5, 4] [6, 3] [7, 2, 1]
[9, 8, 5, 4, 3] [6] [7, 2, 1]
[9, 8, 5, 4, 3] [6, 1] [7, 2]
[9, 8, 5, 4, 3, 2] [6, 1] [7]
[9, 8, 5, 4, 3, 2, 1] [6] [7]
[9, 8, 5, 4, 3, 2, 1] [] [7, 6]
[9, 8, 5, 4, 3, 2] [] [7, 6, 1]
[9, 8, 5, 4, 3] [2] [7, 6, 1]
[9, 8, 5, 4, 3] [2, 1] [7, 6]
[9, 8, 5, 4] [2, 1] [7, 6, 3]
[9, 8, 5, 4, 1] [2] [7, 6, 3]
[9, 8, 5, 4, 1] [] [7, 6, 3, 2]
[9, 8, 5, 4] [] [7, 6, 3, 2, 1]
[9, 8, 5] [4] [7, 6, 3, 2, 1]
[9, 8, 5] [4, 1] [7, 6, 3, 2]
[9, 8, 5, 2] [4, 1] [7, 6, 3]
[9, 8, 5, 2, 1] [4] [7, 6, 3]
[9, 8, 5, 2, 1] [4, 3] [7, 6]
[9, 8, 5, 2] [4, 3] [7, 6, 1]
[9, 8, 5] [4, 3, 2] [7, 6, 1]
[9, 8, 5] [4, 3, 2, 1] [7, 6]
[9, 8] [4, 3, 2, 1] [7, 6, 5]
[9, 8, 1] [4, 3, 2] [7, 6, 5]
[9, 8, 1] [4, 3] [7, 6, 5, 2]
[9, 8] [4, 3] [7, 6, 5, 2, 1]
[9, 8, 3] [4] [7, 6, 5, 2, 1]
[9, 8, 3] [4, 1] [7, 6, 5, 2]
[9, 8, 3, 2] [4, 1] [7, 6, 5]
[9, 8, 3, 2, 1] [4] [7, 6, 5]
[9, 8, 3, 2, 1] [] [7, 6, 5, 4]
[9, 8, 3, 2] [] [7, 6, 5, 4, 1]
[9, 8, 3] [2] [7, 6, 5, 4, 1]
[9, 8, 3] [2, 1] [7, 6, 5, 4]
[9, 8] [2, 1] [7, 6, 5, 4, 3]
[9, 8, 1] [2] [7, 6, 5, 4, 3]
[9, 8, 1] [] [7, 6, 5, 4, 3, 2]
[9, 8] [] [7, 6, 5, 4, 3, 2, 1]
[9] [8] [7, 6, 5, 4, 3, 2, 1]
[9] [8, 1] [7, 6, 5, 4, 3, 2]
[9, 2] [8, 1] [7, 6, 5, 4, 3]
[9, 2, 1] [8] [7, 6, 5, 4, 3]
[9, 2, 1] [8, 3] [7, 6, 5, 4]
[9, 2] [8, 3] [7, 6, 5, 4, 1]
[9] [8, 3, 2] [7, 6, 5, 4, 1]
[9] [8, 3, 2, 1] [7, 6, 5, 4]
[9, 4] [8, 3, 2, 1] [7, 6, 5]
[9, 4, 1] [8, 3, 2] [7, 6, 5]
[9, 4, 1] [8, 3] [7, 6, 5, 2]
[9, 4] [8, 3] [7, 6, 5, 2, 1]
[9, 4, 3] [8] [7, 6, 5, 2, 1]
[9, 4, 3] [8, 1] [7, 6, 5, 2]
[9, 4, 3, 2] [8, 1] [7, 6, 5]
[9, 4, 3, 2, 1] [8] [7, 6, 5]
[9, 4, 3, 2, 1] [8, 5] [7, 6]
[9, 4, 3, 2] [8, 5] [7, 6, 1]
[9, 4, 3] [8, 5, 2] [7, 6, 1]
[9, 4, 3] [8, 5, 2, 1] [7, 6]
[9, 4] [8, 5, 2, 1] [7, 6, 3]
[9, 4, 1] [8, 5, 2] [7, 6, 3]
[9, 4, 1] [8, 5] [7, 6, 3, 2]
[9, 4] [8, 5] [7, 6, 3, 2, 1]
[9] [8, 5, 4] [7, 6, 3, 2, 1]
[9] [8, 5, 4, 1] [7, 6, 3, 2]
[9, 2] [8, 5, 4, 1] [7, 6, 3]
[9, 2, 1] [8, 5, 4] [7, 6, 3]
[9, 2, 1] [8, 5, 4, 3] [7, 6]
[9, 2] [8, 5, 4, 3] [7, 6, 1]
[9] [8, 5, 4, 3, 2] [7, 6, 1]
[9] [8, 5, 4, 3, 2, 1] [7, 6]
[9, 6] [8, 5, 4, 3, 2, 1] [7]
[9, 6, 1] [8, 5, 4, 3, 2] [7]
[9, 6, 1] [8, 5, 4, 3] [7, 2]
[9, 6] [8, 5, 4, 3] [7, 2, 1]
[9, 6, 3] [8, 5, 4] [7, 2, 1]
[9, 6, 3] [8, 5, 4, 1] [7, 2]
[9, 6, 3, 2] [8, 5, 4, 1] [7]
[9, 6, 3, 2, 1] [8, 5, 4] [7]
[9, 6, 3, 2, 1] [8, 5] [7, 4]
[9, 6, 3, 2] [8, 5] [7, 4, 1]
[9, 6, 3] [8, 5, 2] [7, 4, 1]
[9, 6, 3] [8, 5, 2, 1] [7, 4]
[9, 6] [8, 5, 2, 1] [7, 4, 3]
[9, 6, 1] [8, 5, 2] [7, 4, 3]
[9, 6, 1] [8, 5] [7, 4, 3, 2]
[9, 6] [8, 5] [7, 4, 3, 2, 1]
[9, 6, 5] [8] [7, 4, 3, 2, 1]
[9, 6, 5] [8, 1] [7, 4, 3, 2]
[9, 6, 5, 2] [8, 1] [7, 4, 3]
[9, 6, 5, 2, 1] [8] [7, 4, 3]
[9, 6, 5, 2, 1] [8, 3] [7, 4]
[9, 6, 5, 2] [8, 3] [7, 4, 1]
[9, 6, 5] [8, 3, 2] [7, 4, 1]
[9, 6, 5] [8, 3, 2, 1] [7, 4]
[9, 6, 5, 4] [8, 3, 2, 1] [7]
[9, 6, 5, 4, 1] [8, 3, 2] [7]
[9, 6, 5, 4, 1] [8, 3] [7, 2]
[9, 6, 5, 4] [8, 3] [7, 2, 1]
[9, 6, 5, 4, 3] [8] [7, 2, 1]
[9, 6, 5, 4, 3] [8, 1] [7, 2]
[9, 6, 5, 4, 3, 2] [8, 1] [7]
[9, 6, 5, 4, 3, 2, 1] [8] [7]
[9, 6, 5, 4, 3, 2, 1] [8, 7] []
[9, 6, 5, 4, 3, 2] [8, 7] [1]
[9, 6, 5, 4, 3] [8, 7, 2] [1]
[9, 6, 5, 4, 3] [8, 7, 2, 1] []
[9, 6, 5, 4] [8, 7, 2, 1] [3]
[9, 6, 5, 4, 1] [8, 7, 2] [3]
[9, 6, 5, 4, 1] [8, 7] [3, 2]
[9, 6, 5, 4] [8, 7] [3, 2, 1]
[9, 6, 5] [8, 7, 4] [3, 2, 1]
[9, 6, 5] [8, 7, 4, 1] [3, 2]
[9, 6, 5, 2] [8, 7, 4, 1] [3]
[9, 6, 5, 2, 1] [8, 7, 4] [3]
[9, 6, 5, 2, 1] [8, 7, 4, 3] []
[9, 6, 5, 2] [8, 7, 4, 3] [1]
[9, 6, 5] [8, 7, 4, 3, 2] [1]
[9, 6, 5] [8, 7, 4, 3, 2, 1] []
[9, 6] [8, 7, 4, 3, 2, 1] [5]
[9, 6, 1] [8, 7, 4, 3, 2] [5]
[9, 6, 1] [8, 7, 4, 3] [5, 2]
[9, 6] [8, 7, 4, 3] [5, 2, 1]
[9, 6, 3] [8, 7, 4] [5, 2, 1]
[9, 6, 3] [8, 7, 4, 1] [5, 2]
[9, 6, 3, 2] [8, 7, 4, 1] [5]
[9, 6, 3, 2, 1] [8, 7, 4] [5]
[9, 6, 3, 2, 1] [8, 7] [5, 4]
[9, 6, 3, 2] [8, 7] [5, 4, 1]
[9, 6, 3] [8, 7, 2] [5, 4, 1]
[9, 6, 3] [8, 7, 2, 1] [5, 4]
[9, 6] [8, 7, 2, 1] [5, 4, 3]
[9, 6, 1] [8, 7, 2] [5, 4, 3]
[9, 6, 1] [8, 7] [5, 4, 3, 2]
[9, 6] [8, 7] [5, 4, 3, 2, 1]
[9] [8, 7, 6] [5, 4, 3, 2, 1]
[9] [8, 7, 6, 1] [5, 4, 3, 2]
[9, 2] [8, 7, 6, 1] [5, 4, 3]
[9, 2, 1] [8, 7, 6] [5, 4, 3]
[9, 2, 1] [8, 7, 6, 3] [5, 4]
[9, 2] [8, 7, 6, 3] [5, 4, 1]
[9] [8, 7, 6, 3, 2] [5, 4, 1]
[9] [8, 7, 6, 3, 2, 1] [5, 4]
[9, 4] [8, 7, 6, 3, 2, 1] [5]
[9, 4, 1] [8, 7, 6, 3, 2] [5]
[9, 4, 1] [8, 7, 6, 3] [5, 2]
[9, 4] [8, 7, 6, 3] [5, 2, 1]
[9, 4, 3] [8, 7, 6] [5, 2, 1]
[9, 4, 3] [8, 7, 6, 1] [5, 2]
[9, 4, 3, 2] [8, 7, 6, 1] [5]
[9, 4, 3, 2, 1] [8, 7, 6] [5]
[9, 4, 3, 2, 1] [8, 7, 6, 5] []
[9, 4, 3, 2] [8, 7, 6, 5] [1]
[9, 4, 3] [8, 7, 6, 5, 2] [1]
[9, 4, 3] [8, 7, 6, 5, 2, 1] []
[9, 4] [8, 7, 6, 5, 2, 1] [3]
[9, 4, 1] [8, 7, 6, 5, 2] [3]
[9, 4, 1] [8, 7, 6, 5] [3, 2]
[9, 4] [8, 7, 6, 5] [3, 2, 1]
[9] [8, 7, 6, 5, 4] [3, 2, 1]
[9] [8, 7, 6, 5, 4, 1] [3, 2]
[9, 2] [8, 7, 6, 5, 4, 1] [3]
[9, 2, 1] [8, 7, 6, 5, 4] [3]
[9, 2, 1] [8, 7, 6, 5, 4, 3] []
[9, 2] [8, 7, 6, 5, 4, 3] [1]
[9] [8, 7, 6, 5, 4, 3, 2] [1]
[9] [8, 7, 6, 5, 4, 3, 2, 1] []
[] [8, 7, 6, 5, 4, 3, 2, 1] [9]
[1] [8, 7, 6, 5, 4, 3, 2] [9]
[1] [8, 7, 6, 5, 4, 3] [9, 2]
[] [8, 7, 6, 5, 4, 3] [9, 2, 1]
[3] [8, 7, 6, 5, 4] [9, 2, 1]
[3] [8, 7, 6, 5, 4, 1] [9, 2]
[3, 2] [8, 7, 6, 5, 4, 1] [9]
[3, 2, 1] [8, 7, 6, 5, 4] [9]
[3, 2, 1] [8, 7, 6, 5] [9, 4]
[3, 2] [8, 7, 6, 5] [9, 4, 1]
[3] [8, 7, 6, 5, 2] [9, 4, 1]
[3] [8, 7, 6, 5, 2, 1] [9, 4]
[] [8, 7, 6, 5, 2, 1] [9, 4, 3]
[1] [8, 7, 6, 5, 2] [9, 4, 3]
[1] [8, 7, 6, 5] [9, 4, 3, 2]
[] [8, 7, 6, 5] [9, 4, 3, 2, 1]
[5] [8, 7, 6] [9, 4, 3, 2, 1]
[5] [8, 7, 6, 1] [9, 4, 3, 2]
[5, 2] [8, 7, 6, 1] [9, 4, 3]
[5, 2, 1] [8, 7, 6] [9, 4, 3]
[5, 2, 1] [8, 7, 6, 3] [9, 4]
[5, 2] [8, 7, 6, 3] [9, 4, 1]
[5] [8, 7, 6, 3, 2] [9, 4, 1]
[5] [8, 7, 6, 3, 2, 1] [9, 4]
[5, 4] [8, 7, 6, 3, 2, 1] [9]
[5, 4, 1] [8, 7, 6, 3, 2] [9]
[5, 4, 1] [8, 7, 6, 3] [9, 2]
[5, 4] [8, 7, 6, 3] [9, 2, 1]
[5, 4, 3] [8, 7, 6] [9, 2, 1]
[5, 4, 3] [8, 7, 6, 1] [9, 2]
[5, 4, 3, 2] [8, 7, 6, 1] [9]
[5, 4, 3, 2, 1] [8, 7, 6] [9]
[5, 4, 3, 2, 1] [8, 7] [9, 6]
[5, 4, 3, 2] [8, 7] [9, 6, 1]
[5, 4, 3] [8, 7, 2] [9, 6, 1]
[5, 4, 3] [8, 7, 2, 1] [9, 6]
[5, 4] [8, 7, 2, 1] [9, 6, 3]
[5, 4, 1] [8, 7, 2] [9, 6, 3]
[5, 4, 1] [8, 7] [9, 6, 3, 2]
[5, 4] [8, 7] [9, 6, 3, 2, 1]
[5] [8, 7, 4] [9, 6, 3, 2, 1]
[5] [8, 7, 4, 1] [9, 6, 3, 2]
[5, 2] [8, 7, 4, 1] [9, 6, 3]
[5, 2, 1] [8, 7, 4] [9, 6, 3]
[5, 2, 1] [8, 7, 4, 3] [9, 6]
[5, 2] [8, 7, 4, 3] [9, 6, 1]
[5] [8, 7, 4, 3, 2] [9, 6, 1]
[5] [8, 7, 4, 3, 2, 1] [9, 6]
[] [8, 7, 4, 3, 2, 1] [9, 6, 5]
[1] [8, 7, 4, 3, 2] [9, 6, 5]
[1] [8, 7, 4, 3] [9, 6, 5, 2]
[] [8, 7, 4, 3] [9, 6, 5, 2, 1]
[3] [8, 7, 4] [9, 6, 5, 2, 1]
[3] [8, 7, 4, 1] [9, 6, 5, 2]
[3, 2] [8, 7, 4, 1] [9, 6, 5]
[3, 2, 1] [8, 7, 4] [9, 6, 5]
[3, 2, 1] [8, 7] [9, 6, 5, 4]
[3, 2] [8, 7] [9, 6, 5, 4, 1]
[3] [8, 7, 2] [9, 6, 5, 4, 1]
[3] [8, 7, 2, 1] [9, 6, 5, 4]
[] [8, 7, 2, 1] [9, 6, 5, 4, 3]
[1] [8, 7, 2] [9, 6, 5, 4, 3]
[1] [8, 7] [9, 6, 5, 4, 3, 2]
[] [8, 7] [9, 6, 5, 4, 3, 2, 1]
[7] [8] [9, 6, 5, 4, 3, 2, 1]
[7] [8, 1] [9, 6, 5, 4, 3, 2]
[7, 2] [8, 1] [9, 6, 5, 4, 3]
[7, 2, 1] [8] [9, 6, 5, 4, 3]
[7, 2, 1] [8, 3] [9, 6, 5, 4]
[7, 2] [8, 3] [9, 6, 5, 4, 1]
[7] [8, 3, 2] [9, 6, 5, 4, 1]
[7] [8, 3, 2, 1] [9, 6, 5, 4]
[7, 4] [8, 3, 2, 1] [9, 6, 5]
[7, 4, 1] [8, 3, 2] [9, 6, 5]
[7, 4, 1] [8, 3] [9, 6, 5, 2]
[7, 4] [8, 3] [9, 6, 5, 2, 1]
[7, 4, 3] [8] [9, 6, 5, 2, 1]
[7, 4, 3] [8, 1] [9, 6, 5, 2]
[7, 4, 3, 2] [8, 1] [9, 6, 5]
[7, 4, 3, 2, 1] [8] [9, 6, 5]
[7, 4, 3, 2, 1] [8, 5] [9, 6]
[7, 4, 3, 2] [8, 5] [9, 6, 1]
[7, 4, 3] [8, 5, 2] [9, 6, 1]
[7, 4, 3] [8, 5, 2, 1] [9, 6]
[7, 4] [8, 5, 2, 1] [9, 6, 3]
[7, 4, 1] [8, 5, 2] [9, 6, 3]
[7, 4, 1] [8, 5] [9, 6, 3, 2]
[7, 4] [8, 5] [9, 6, 3, 2, 1]
[7] [8, 5, 4] [9, 6, 3, 2, 1]
[7] [8, 5, 4, 1] [9, 6, 3, 2]
[7, 2] [8, 5, 4, 1] [9, 6, 3]
[7, 2, 1] [8, 5, 4] [9, 6, 3]
[7, 2, 1] [8, 5, 4, 3] [9, 6]
[7, 2] [8, 5, 4, 3] [9, 6, 1]
[7] [8, 5, 4, 3, 2] [9, 6, 1]
[7] [8, 5, 4, 3, 2, 1] [9, 6]
[7, 6] [8, 5, 4, 3, 2, 1] [9]
[7, 6, 1] [8, 5, 4, 3, 2] [9]
[7, 6, 1] [8, 5, 4, 3] [9, 2]
[7, 6] [8, 5, 4, 3] [9, 2, 1]
[7, 6, 3] [8, 5, 4] [9, 2, 1]
[7, 6, 3] [8, 5, 4, 1] [9, 2]
[7, 6, 3, 2] [8, 5, 4, 1] [9]
[7, 6, 3, 2, 1] [8, 5, 4] [9]
[7, 6, 3, 2, 1] [8, 5] [9, 4]
[7, 6, 3, 2] [8, 5] [9, 4, 1]
[7, 6, 3] [8, 5, 2] [9, 4, 1]
[7, 6, 3] [8, 5, 2, 1] [9, 4]
[7, 6] [8, 5, 2, 1] [9, 4, 3]
[7, 6, 1] [8, 5, 2] [9, 4, 3]
[7, 6, 1] [8, 5] [9, 4, 3, 2]
[7, 6] [8, 5] [9, 4, 3, 2, 1]
[7, 6, 5] [8] [9, 4, 3, 2, 1]
[7, 6, 5] [8, 1] [9, 4, 3, 2]
[7, 6, 5, 2] [8, 1] [9, 4, 3]
[7, 6, 5, 2, 1] [8] [9, 4, 3]
[7, 6, 5, 2, 1] [8, 3] [9, 4]
[7, 6, 5, 2] [8, 3] [9, 4, 1]
[7, 6, 5] [8, 3, 2] [9, 4, 1]
[7, 6, 5] [8, 3, 2, 1] [9, 4]
[7, 6, 5, 4] [8, 3, 2, 1] [9]
[7, 6, 5, 4, 1] [8, 3, 2] [9]
[7, 6, 5, 4, 1] [8, 3] [9, 2]
[7, 6, 5, 4] [8, 3] [9, 2, 1]
[7, 6, 5, 4, 3] [8] [9, 2, 1]
[7, 6, 5, 4, 3] [8, 1] [9, 2]
[7, 6, 5, 4, 3, 2] [8, 1] [9]
[7, 6, 5, 4, 3, 2, 1] [8] [9]
[7, 6, 5, 4, 3, 2, 1] [] [9, 8]
[7, 6, 5, 4, 3, 2] [] [9, 8, 1]
[7, 6, 5, 4, 3] [2] [9, 8, 1]
[7, 6, 5, 4, 3] [2, 1] [9, 8]
[7, 6, 5, 4] [2, 1] [9, 8, 3]
[7, 6, 5, 4, 1] [2] [9, 8, 3]
[7, 6, 5, 4, 1] [] [9, 8, 3, 2]
[7, 6, 5, 4] [] [9, 8, 3, 2, 1]
[7, 6, 5] [4] [9, 8, 3, 2, 1]
[7, 6, 5] [4, 1] [9, 8, 3, 2]
[7, 6, 5, 2] [4, 1] [9, 8, 3]
[7, 6, 5, 2, 1] [4] [9, 8, 3]
[7, 6, 5, 2, 1] [4, 3] [9, 8]
[7, 6, 5, 2] [4, 3] [9, 8, 1]
[7, 6, 5] [4, 3, 2] [9, 8, 1]
[7, 6, 5] [4, 3, 2, 1] [9, 8]
[7, 6] [4, 3, 2, 1] [9, 8, 5]
[7, 6, 1] [4, 3, 2] [9, 8, 5]
[7, 6, 1] [4, 3] [9, 8, 5, 2]
[7, 6] [4, 3] [9, 8, 5, 2, 1]
[7, 6, 3] [4] [9, 8, 5, 2, 1]
[7, 6, 3] [4, 1] [9, 8, 5, 2]
[7, 6, 3, 2] [4, 1] [9, 8, 5]
[7, 6, 3, 2, 1] [4] [9, 8, 5]
[7, 6, 3, 2, 1] [] [9, 8, 5, 4]
[7, 6, 3, 2] [] [9, 8, 5, 4, 1]
[7, 6, 3] [2] [9, 8, 5, 4, 1]
[7, 6, 3] [2, 1] [9, 8, 5, 4]
[7, 6] [2, 1] [9, 8, 5, 4, 3]
[7, 6, 1] [2] [9, 8, 5, 4, 3]
[7, 6, 1] [] [9, 8, 5, 4, 3, 2]
[7, 6] [] [9, 8, 5, 4, 3, 2, 1]
[7] [6] [9, 8, 5, 4, 3, 2, 1]
[7] [6, 1] [9, 8, 5, 4, 3, 2]
[7, 2] [6, 1] [9, 8, 5, 4, 3]
[7, 2, 1] [6] [9, 8, 5, 4, 3]
[7, 2, 1] [6, 3] [9, 8, 5, 4]
[7, 2] [6, 3] [9, 8, 5, 4, 1]
[7] [6, 3, 2] [9, 8, 5, 4, 1]
[7] [6, 3, 2, 1] [9, 8, 5, 4]
[7, 4] [6, 3, 2, 1] [9, 8, 5]
[7, 4, 1] [6, 3, 2] [9, 8, 5]
[7, 4, 1] [6, 3] [9, 8, 5, 2]
[7, 4] [6, 3] [9, 8, 5, 2, 1]
[7, 4, 3] [6] [9, 8, 5, 2, 1]
[7, 4, 3] [6, 1] [9, 8, 5, 2]
[7, 4, 3, 2] [6, 1] [9, 8, 5]
[7, 4, 3, 2, 1] [6] [9, 8, 5]
[7, 4, 3, 2, 1] [6, 5] [9, 8]
[7, 4, 3, 2] [6, 5] [9, 8, 1]
[7, 4, 3] [6, 5, 2] [9, 8, 1]
[7, 4, 3] [6, 5, 2, 1] [9, 8]
[7, 4] [6, 5, 2, 1] [9, 8, 3]
[7, 4, 1] [6, 5, 2] [9, 8, 3]
[7, 4, 1] [6, 5] [9, 8, 3, 2]
[7, 4] [6, 5] [9, 8, 3, 2, 1]
[7] [6, 5, 4] [9, 8, 3, 2, 1]
[7] [6, 5, 4, 1] [9, 8, 3, 2]
[7, 2] [6, 5, 4, 1] [9, 8, 3]
[7, 2, 1] [6, 5, 4] [9, 8, 3]
[7, 2, 1] [6, 5, 4, 3] [9, 8]
[7, 2] [6, 5, 4, 3] [9, 8, 1]
[7] [6, 5, 4, 3, 2] [9, 8, 1]
[7] [6, 5, 4, 3, 2, 1] [9, 8]
[] [6, 5, 4, 3, 2, 1] [9, 8, 7]
[1] [6, 5, 4, 3, 2] [9, 8, 7]
[1] [6, 5, 4, 3] [9, 8, 7, 2]
[] [6, 5, 4, 3] [9, 8, 7, 2, 1]
[3] [6, 5, 4] [9, 8, 7, 2, 1]
[3] [6, 5, 4, 1] [9, 8, 7, 2]
[3, 2] [6, 5, 4, 1] [9, 8, 7]
[3, 2, 1] [6, 5, 4] [9, 8, 7]
[3, 2, 1] [6, 5] [9, 8, 7, 4]
[3, 2] [6, 5] [9, 8, 7, 4, 1]
[3] [6, 5, 2] [9, 8, 7, 4, 1]
[3] [6, 5, 2, 1] [9, 8, 7, 4]
[] [6, 5, 2, 1] [9, 8, 7, 4, 3]
[1] [6, 5, 2] [9, 8, 7, 4, 3]
[1] [6, 5] [9, 8, 7, 4, 3, 2]
[] [6, 5] [9, 8, 7, 4, 3, 2, 1]
[5] [6] [9, 8, 7, 4, 3, 2, 1]
[5] [6, 1] [9, 8, 7, 4, 3, 2]
[5, 2] [6, 1] [9, 8, 7, 4, 3]
[5, 2, 1] [6] [9, 8, 7, 4, 3]
[5, 2, 1] [6, 3] [9, 8, 7, 4]
[5, 2] [6, 3] [9, 8, 7, 4, 1]
[5] [6, 3, 2] [9, 8, 7, 4, 1]
[5] [6, 3, 2, 1] [9, 8, 7, 4]
[5, 4] [6, 3, 2, 1] [9, 8, 7]
[5, 4, 1] [6, 3, 2] [9, 8, 7]
[5, 4, 1] [6, 3] [9, 8, 7, 2]
[5, 4] [6, 3] [9, 8, 7, 2, 1]
[5, 4, 3] [6] [9, 8, 7, 2, 1]
[5, 4, 3] [6, 1] [9, 8, 7, 2]
[5, 4, 3, 2] [6, 1] [9, 8, 7]
[5, 4, 3, 2, 1] [6] [9, 8, 7]
[5, 4, 3, 2, 1] [] [9, 8, 7, 6]
[5, 4, 3, 2] [] [9, 8, 7, 6, 1]
[5, 4, 3] [2] [9, 8, 7, 6, 1]
[5, 4, 3] [2, 1] [9, 8, 7, 6]
[5, 4] [2, 1] [9, 8, 7, 6, 3]
[5, 4, 1] [2] [9, 8, 7, 6, 3]
[5, 4, 1] [] [9, 8, 7, 6, 3, 2]
[5, 4] [] [9, 8, 7, 6, 3, 2, 1]
[5] [4] [9, 8, 7, 6, 3, 2, 1]
[5] [4, 1] [9, 8, 7, 6, 3, 2]
[5, 2] [4, 1] [9, 8, 7, 6, 3]
[5, 2, 1] [4] [9, 8, 7, 6, 3]
[5, 2, 1] [4, 3] [9, 8, 7, 6]
[5, 2] [4, 3] [9, 8, 7, 6, 1]
[5] [4, 3, 2] [9, 8, 7, 6, 1]
[5] [4, 3, 2, 1] [9, 8, 7, 6]
[] [4, 3, 2, 1] [9, 8, 7, 6, 5]
[1] [4, 3, 2] [9, 8, 7, 6, 5]
[1] [4, 3] [9, 8, 7, 6, 5, 2]
[] [4, 3] [9, 8, 7, 6, 5, 2, 1]
[3] [4] [9, 8, 7, 6, 5, 2, 1]
[3] [4, 1] [9, 8, 7, 6, 5, 2]
[3, 2] [4, 1] [9, 8, 7, 6, 5]
[3, 2, 1] [4] [9, 8, 7, 6, 5]
[3, 2, 1] [] [9, 8, 7, 6, 5, 4]
[3, 2] [] [9, 8, 7, 6, 5, 4, 1]
[3] [2] [9, 8, 7, 6, 5, 4, 1]
[3] [2, 1] [9, 8, 7, 6, 5, 4]
[] [2, 1] [9, 8, 7, 6, 5, 4, 3]
[1] [2] [9, 8, 7, 6, 5, 4, 3]
[1] [] [9, 8, 7, 6, 5, 4, 3, 2]
[] [] [9, 8, 7, 6, 5, 4, 3, 2, 1]"""
assert hanoi_solver(9) == expected
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def hanoi_solver(n):
    # Initialize the three rods
    A = list(range(n, 0, -1))
    B = []
    C = []
    
    # List to store all moves as strings
    moves = [f"{A} {B} {C}"]
    
    def move(num_disks, source, auxiliary, target):
        if num_disks <= 0:
            return
        
        # Move n-1 disks from source to auxiliary, so they are out of the way
        move(num_disks - 1, source, target, auxiliary)
        
        # Move the nth disk from source to target
        target.append(source.pop())
        
        # Record the move
        moves.append(f"{A} {B} {C}")
        
        # Move the n-1 disks that we left on auxiliary onto target
        move(num_disks - 1, auxiliary, source, target)
    
    # Start the recursive process
    move(n, A, B, C)
    
    return '\n'.join(moves)
```
