---
id: 6606b961ebcf04460f8af76e
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Jaribu kwamba njia ya `find_empty_cell` inafanya kazi ipasavyo kwa kuitwa kwenye `gameboard` na kuchapisha matokeo.

Kumbuka kwamba, ingawa `find_empty_cell` imefafanuliwa na kigezo kimoja, huwezi kumpa thamani kwa kupitisha hoja kwenye wito wa kitendakazi, kwa sababu `self` hupitishwa moja kwa moja kama kitu unachoitia wito njia hiyo.

# --hints--

Unapaswa kuchapisha `gameboard.find_empty_cell()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
