---
id: 6606c3fd5634684d48a7887b
title: Hatua 32
challengeType: 20
dashedName: step-32
---

# --description--

Usemi wa kizalishi uliouandika katika hatua iliyopita unazalisha orodha ya thamani za kweli au siyo kweli zinazowakilisha kama sharti `self.board[row][col] != num` ni `True` au `False` kwa kila kipengele katika safu iliyotajwa katika mistari yote.

Pitia usemi huo wa kizalishi kwa kitendakazi cha `all()` ili ukague kama vipengele vyote katika safu ni tofauti na `num` na urudishe matokeo.

# --hints--

Unapaswa kupitisha usemi wa kizalishi kama hoja katika wito wa `all()` na kurudisha matokeo.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("return all(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        (self.board[row][col] != num for row in range(9))
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
