---
id: 66068fb0bfddba2b7977eb60
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utafundishwa kuhusu madarasa ya vitu na vitu kwa kujenga suluhisho la fumbo la sudoku.

Katika Python, darasa la vitu ni ramani ya kuunda vitu. Vitu vinavyoundwa kutoka kwa darasa ni matukio ya darasa hilo. Unaweza kuunda darasa kwa kutumia sintaksia hii:

```py
class ClassName:
    pass
```

Ambapo `class` ni neno muhimu linalohitajika kufafanua darasa na `ClassName` ni jina la darasa, linaloandikwa kwa mkataba wa *PascalCase*.

Anza kwa kuunda darasa la `Board`.

# --hints--

Unapaswa kuunda darasa lenye jina `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
