---
id: 6606a4641ec48b3a9fe8c2fc
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Kitendakazi iliyo ndani ya `enumerate` huchukua iterable kama hoja yake na hurudisha kitu cha enumerate ambacho unaweza kurudia juu yake.
Hutoa hesabu (ambayo kwa chaguo-msingi huanza kwa sifuri) na thamani kutoka kwenye iterable.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

Mzunguko kutoka kwenye mfano hapo juu ungechapisha tuples `0, a`, `1, b`, na `2, c`.

Ndani ya njia ya `find_empty_cell`, badilisha `pass` kwa mzunguko wa `for` unaotumia kitendakazi cha `enumerate()` kurudia kila safu kwenye bodi ya sudoku. Tumia `row` kama kielezo cha safu ya sasa na `contents` kwa vipengele vya safu ya sasa.


# --hints--

Unapaswa kubadilisha `pass` kwa mzunguko wa `for` unaorudia juu ya `enumerate(self.board)`. Tumia `row` na `contents` kama vigezo vya mzunguko.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
