---
id: 6606b224a69a293f98f8db8f
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Njia ya `.index()` husababisha hitilafu ya `ValueError` wakati thamani haipatikani. Ili kuzuia programu kusitisha utekelezaji, utaweka mstari huu wa msimbo ndani ya kizuizi cha `try`. Kifungu cha `try` hutumika kufunika msimbo ambao unaweza kusababisha hitilafu. Kifungu cha `except`, kwa upande mwingine, kinatoa msimbo mbadala wa kutekeleza ikiwa hitilafu itatokea:

```py
try:
    <code>
except:
    <code>
```

Weka mgawo wa `col` ndani ya kizuizi cha `try`. Kisha, tengeneza kifungu cha `except` na jaza mwili wake na `pass`.

# --hints--

Unapaswa kuweka mgawo wa `col` ndani ya kizuizi cha `try` na kutengeneza kifungu cha `except`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
