---
id: 66069e5759b800364707988e
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Kurudi kwenye njia ya `__init__`, inahitaji kigezo kingine kinachowakilisha fumbo la kutatua.

Ongeza kigezo cha pili kinachoitwa `board` kwenye njia ya `__init__` na rekebisha uanzishaji wa `gameboard` kwa kumpitisha orodha ya `puzzle` kama unavyopitisha hoja kwenye wito wa kitendakazi.

# --hints--

Njia yako ya `__init__` inapaswa kuwa na vigezo viwili kwa mpangilio: `self`, na `board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

Kigezo chako cha `gameboard` kinapaswa kuwa na thamani ya `Board(puzzle)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
