---
id: 6606933d6813a8308c962dd1
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Ili kuwa njia ya mfano, njia inahitaji kigezo maalum, kinachoitwa `self` kwa mkataba. Kigezo hiki ni rejea kwa mfano wa darasa la vitu na lazima daima kiwe kigezo cha kwanza.

Ongeza kigezo cha `self` kwenye njia yako ya `spam`.

# --hints--

Njia yako ya `spam` inapaswa kuwa na kigezo cha `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
