---
id: 6606927d010be4300a4e5330
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Uundaji wa mfano huunda kitu tupu. Lakini madarasa ya vitu yanaweza kuwa na njia, ambazo ni kama vitendakazi vya ndani kwa kila mfano. Ndani ya darasa la vitu, njia hutangazwa kama ifuatavyo:

```py
class ClassName:
    def method_name():
        pass
```

Ndani ya darasa la vitu la `Board`, badilisha `pass` na njia tupu `spam`.

# --hints--

Unapaswa kutangaza njia yenye jina `spam` ndani ya darasa la vitu la `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
