---
id: 6606cc754a8834509cd0afb6
title: Hatua 36
challengeType: 20
dashedName: step-36
---

# --description--

Jambo lingine la kukagua ni kama nambari inaweza kuingizwa katika mraba wa 3x3.

Ndani ya darasa la `Board`, tengeneza njia iitwayo `valid_in_square` yenye vigezo vinne: `self`, `row`, `col`, na `num`. Ambapo `row`, `col`, na `num` wakilisha kielezo cha safu, kielezo cha safu wima, na nambari inayopaswa kukaguliwa, kwa mtiririko huo.

# --hints--

Unapaswa kutengeneza njia mpya iitwayo `valid_in_square` ndani ya darasa la `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

Njia yako ya `valid_in_square` inapaswa kuwa na vigezo vinne: `self`, `row`, `col`, na `num`, kwa mpangilio huu.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
