---
id: 66069167b3307b2f4067b22b
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Mfano mpya wa darasa la vitu huundwa kwa kutumia kitendo cha kitendakazi, ambacho kinahusisha kuongeza jozi ya mabano ya kawaida kwenye jina la darasa.

Nje ya ufafanuzi wa darasa, unda mfano wa darasa la `Board` na uweke thamani yake kwenye kigezo kinachoitwa `gameboard`.

# --hints--

Unapaswa kutangaza kigezo `gameboard` nje ya darasa la `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Kigezo chako cha `gameboard` kinapaswa kuwa na thamani ya `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
