---
id: 660699aabc59c532f2d556e5
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Ili kuita njia ya mfano, unahitaji kutumia uandishi wa nukta:

```py
instance_name.method_name()
```

Ambapo `instance_name` ni mfano au kitu, na `method_name` ni njia unayotaka kuita.

Ita njia ya `spam` ya kitu cha `gameboard`.

# --hints--

Unapaswa kuita njia ya `spam` ya kitu cha `gameboard` kwa `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
