---
id: 66069b992c1c5e3451f3deb0
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Uundaji huunda kitu tupu. Njia ya `__init__` ni njia maalum inayokuruhusu kuunda kitu katika hali iliyobinafsishwa. Wakati darasa la vitu linapotekeleza njia ya `__init__`, `__init__` huitwa moja kwa moja wakati wa uundaji.

Ndani ya darasa lako la `Board`, futa njia ya `spam` na uibadilishe na njia ya `__init__` inayojumuisha kigezo cha `self`.

# --hints--

Haupaswi kuwa na njia ya `spam` katika darasa lako la `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Unapaswa kufafanua njia ya `__init__` katika darasa lako la `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Njia yako ya `__init__` inapaswa kuwa na kigezo cha `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
