---
id: 660699119472f332798860ad
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa, badilisha `pass` kwa wito wa `print` na upitie mfuatano wa herufi `'Spam!'`.

# --hints--

Unapaswa futa `pass` na kuchapisha mfuatano wa herufi `'Spam!'` ndani ya njia ya `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
