---
id: 66069f86f58f85371d47123e
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sifa ya kipengele ni kigezo kinachohusishwa na kitu, ambacho hutumika kuhifadhi data kama vigezo vya kawaida.

Ndani ya njia ya `__init__`, weka thamani ya kigezo `board` (ambacho hupitishwa wakati wa kuunda mfano wa darasa la `Board`) kwa sifa ya kipengele cha mfano `board` kwa kutumia `self.board`.

`self.board` inarejelea sifa ya kipengele ya `board` ya mfano wa darasa. Ni kigezo kinachomilikiwa na kitu kilichoundwa kutoka darasa la `Board`.

# --hints--

Unapaswa futa `pass` na weka thamani ya kigezo `board` kwa `self.board` ndani ya njia ya `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
