---
id: 6606bc4e5535c0484990ccd5
title: Hatua 24
challengeType: 20
dashedName: step-24
---

# --description--

Ifuatayo, utafanya kazi kwenye njia inayokagua kama nambari fulani inaweza kuingizwa katika safu maalum ya ubao wa sudoku.

Ndani ya darasa la `Board`, tengeneza njia iitwayo `valid_in_row` na uipe vigezo vitatu: `self`, `row`, na `num`. Ambapo `self` inawakilisha mfano wa darasa, na `row` na `num` ni kielezo cha safu na nambari inayokaguliwa, mtawalia.


# --hints--

Unapaswa kutengeneza njia mpya iitwayo `valid_in_row` ndani ya darasa la `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Njia yako ya `valid_in_row` inapaswa kuwa na vigezo vitatu: `self`, `row`, na `num`, kwa mpangilio huu.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
