---
id: 660a737f0f72b51de361051c
title: Hatua 54
challengeType: 20
dashedName: step-54
---

# --description--

Kipengele cha `:=` kinakupa uwezo wa kuweka thamani kwa vigezo katikati ya usemi wa kihesabu. Sintaksia ni: `name := val`, ambapo `name` ni jina la kigezo na `val` ni thamani ya kigezo.

Muundo huu unajulikana rasmi kama *usemi wa mgawo wa thamani*, wakati kipengele cha `:=` kinajulikana kwa kawaida kama kipengele cha *walrus*.

Kwa kuwa utahitaji wito wa `self.find_empty_cell()` zaidi ya mara moja, weka thamani yake katika kigezo `next_empty` kwa kutumia kipengele cha walrus. Kisha, zungusha mgawo wa thamani kati ya jozi ya mabano ya kawaida.

Kwa njia hii, utaunganisha mgawo wa thamani na ukaguzi wa masharti katika mstari mmoja, na kufanya msimbo kuwa mfupi zaidi.

# --hints--

Unapaswa kubadilisha hali ya `if` kuwa `(next_empty := self.find_empty_cell()) is None`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("solver").find_ifs()[0].find_conditions()[0].is_equivalent("(next_empty := self.find_empty_cell()) is None")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True

    def is_valid(self, empty, num):
        row, col = empty
        valid_in_row = self.valid_in_row(row, num)
        valid_in_col = self.valid_in_col(col, num)
        valid_in_square = self.valid_in_square(row, col, num)
        return all([valid_in_row, valid_in_col, valid_in_square])
--fcc-editable-region--
    def solver(self):
        if self.find_empty_cell() is None:
            return True
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
