---
id: 697fc395f3d85f14b05cd9ca
title: 배열의 스토리 조각에서 일관된 내러티브 복원하기
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

이 실습에서는 손상된 스토리 조각 배열에서 일관된 내러티브를 복원할 것입니다.

기본 배열 알고리즘을 처음부터 구현하며 반복문 사용을 연습할 것입니다.

스토리 조각 객체 배열을 다룰 것입니다. 각 조각 객체는 다음 속성을 가집니다:

| 속성 | 설명 | 예시 값 |
|-|-|-|
| `id` | 스토리 내 조각의 위치를 나타내는 양의 정수 | `3` |
| `text` | 실제 스토리 내용 | `"and I use Arch btw.\""` |

이 실습에서는 `shuffledFragments`라는 미리 채워진 배열이 제공됩니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리**

1. 미리 채워진 `shuffledFragments` 배열을 변경하지 않아야 합니다.
2. 프래그먼트 배열을 받아 undefined 요소가 모두 제거된 새 배열을 반환하는 `compactFragments` 함수를 만들어야 합니다. 함수가 undefined 요소를 제거하면 콘솔에 메시지를 출력해야 합니다. 메시지는 접두사 `[COMPACTED]`로 시작해야 합니다.
3. `compactedShuffledFragments` 변수를 선언하고 `shuffledFragments` 배열로 `compactFragments`를 호출한 결과를 할당해야 합니다.
4. `undefined` 요소가 없는 프래그먼트 배열을 받아 `id` 속성을 기준으로 오름차순 정렬된 새 배열을 반환하는 `sortFragments` 함수를 만들어야 합니다. 같은 `id`를 가진 프래그먼트는 원래 순서를 유지해야 합니다. JavaScript의 내장 `sort` 메서드를 사용하지 않아야 합니다.
5. `sortedFragments` 변수를 선언하고 `compactedShuffledFragments` 배열로 `sortFragments`를 호출한 결과를 할당해야 합니다.
6. 정렬된 프래그먼트 배열을 받아 중복을 제거하고 첫 번째 항목만 남긴 새 배열을 반환하는 `dedupeFragments` 함수를 만들어야 합니다. 중복은 같은 `id`를 공유하는 두 개 이상의 프래그먼트로 정의해야 합니다. 중복이 제거된 `id`마다 함수는 콘솔에 메시지를 출력해야 합니다. 메시지는 접두사 `[DEDUPED]`로 시작해야 합니다.
7. `dedupedFragments` 변수를 선언하고 `sortedFragments` 배열로 `dedupeFragments`를 호출한 결과를 할당해야 합니다.
8. 정렬된 프래그먼트 배열을 받아 누락된 프래그먼트를 자리 표시자 객체로 채운 새 배열을 반환하는 `fillMissingFragments` 함수를 만들어야 합니다. 누락된 프래그먼트는 최저 및 최고 `id` 사이 수열의 빈 구간으로 정의해야 합니다. 자리 표시자 객체는 `{ id: missingId, text: "[...]" }` 형식이어야 합니다. 추가된 자리 표시자마다 함수는 콘솔에 메시지를 출력해야 합니다. 메시지는 접두사 `[FILLED]`로 시작해야 합니다.
9. `filledFragments` 변수를 선언하고 `dedupedFragments` 배열로 `fillMissingFragments`를 호출한 결과를 할당해야 합니다.
10. 정렬된 프래그먼트 배열을 받아 모든 프래그먼트 텍스트를 줄바꿈으로 구분해 담은 하나의 문자열을 반환하는 `assembleStory` 함수를 만들어야 합니다.
11. `filledFragments`와 함께 `assembleStory`를 사용하여 완성된 이야기를 콘솔에 표시해야 합니다.
12. `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments`, `assembleStory` 함수는 호출에 사용된 배열을 변경하지 않아야 합니다.

**예시**

다음은 스토리 조각을 포함하는 배열의 예시입니다:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

`exampleArray`에서 스토리를 복원한 후 결과는 다음과 같습니다:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

미리 채워진 `shuffledFragments` 배열을 변경하지 마세요.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

`compactFragments` 함수를 만들어야 합니다.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])`는 `[COMPACTED]`로 시작하는 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments` 함수가 `undefined` 요소를 포함한 배열로 호출되면, `[COMPACTED]`로 시작하는 콘솔 메시지를 기록해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])`는 `[COMPACTED]`로 시작하는 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments` 함수가 빈 슬롯을 포함한 배열로 호출되면 `[COMPACTED]`로 시작하는 콘솔 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments` 함수가 `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]` 배열로 호출되면 `[COMPACTED]`로 시작하는 메시지를 출력하지 않아야 합니다.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])`는 `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`를 반환해야 합니다.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments`는 전달된 배열을 변경하지 않아야 합니다.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments`는 `undefined` 요소가 없는 배열이어야 합니다.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

`sortFragments` 함수를 만들어야 합니다.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])`는 `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`를 반환해야 합니다.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments`는 전달된 배열을 변경하지 않아야 합니다.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments`의 각 요소 `id`는 다음 요소의 `id`보다 작거나 같아야 합니다.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

`dedupeFragments` 함수를 만들어야 합니다.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])`는 `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`를 반환해야 합니다.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])`는 중복된 각 `id`에 대해 `[DEDUPED]` 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])`는 마지막 `id`가 중복될 때 `[DEDUPED]` 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

`dedupeFragments`가 중복을 포함한 배열로 호출되면 중복된 각 `id`에 대해 `[DEDUPED]`로 시작하는 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments`는 전달된 배열을 변경하지 않아야 합니다.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments`에는 중복된 `id` 값이 없어야 합니다.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

`fillMissingFragments` 함수를 만들어야 합니다.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])`는 `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`를 반환해야 합니다.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])`는 `[FILLED]`로 시작하는 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments`가 자리 표시자를 추가할 때 `[FILLED]`로 시작하는 메시지를 출력해야 합니다.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments`는 전달된 배열을 변경하지 않아야 합니다.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments`에는 `id` 수열에 빈 구간이 없어야 합니다.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

`assembleStory` 함수를 만들어야 합니다.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])`는 문자열 `"Hello\nWorld"`를 반환해야 합니다.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory`는 전달된 배열을 변경하지 않아야 합니다.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

`assembleStory(filledFragments)`의 결과를 `console.log`로 출력해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
