---
id: 615f2abbe7d18d49a1e0e1c8
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

기본 HTML 보일러플레이트를 제공했습니다.

`h1` 요소 안에 `body` 요소를 만들고 텍스트로 `Nutrition Facts`를 넣으세요.

# --hints--

새로운 `h1` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소는 `body` 요소 안에 있어야 합니다.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

`h1` 요소는 텍스트로 `Nutrition Facts`를 가져야 합니다.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
