---
id: 615f5486b8fd4b71633f69b0
title: 단계 33
challengeType: 0
dashedName: step-33
---

# --description--

`.small-text` 요소 아래에 텍스트 `p`를 가진 새로운 `Calories` 요소를 만드세요. 또한 `.left-container` 요소 아래에 텍스트 `span`를 가진 새로운 `230` 요소를 만드세요.

# --hints--

`p` 요소 안에 새로운 `.calories-info` 요소가 있어야 합니다.

```js
assert(document.querySelector('.calories-info')?.querySelector('p'))
```

`span` 요소 안에 새로운 `.calories-info` 요소가 있어야 합니다.

```js
assert(document.querySelector('.calories-info')?.querySelector('span'))
```

`p` 요소가 `.small-text` 요소 뒤에 와야 합니다.

```js
assert(document.querySelector('.small-text')?.nextElementSibling?.localName === 'p');
```

`span` 요소가 `.left-container` 요소 뒤에 와야 합니다.

```js
assert(document.querySelector('.left-container')?.nextElementSibling?.localName ==='span');
```

`p` 요소는 텍스트 `Calories`를 가져야 합니다.

```js
assert(document.querySelector('.small-text')?.nextElementSibling?.innerText === 'Calories');
```

`span` 요소는 텍스트 `230`를 가져야 합니다.

```js
assert(document.querySelector('.left-container')?.nextElementSibling?.innerText === '230');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
--fcc-editable-region--
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
      </div>
    </div>
--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

.small-text {
  font-size: 0.85rem;
}

.calories-info h2 {
  margin: 0;
}
```
