---
id: 615f3cafd794015aa9547a6d
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

`h1`, `h2` 같은 태그 사용이 HTML의 의미 구조를 결정한다는 점을 기억하세요. 하지만 이 요소들의 CSS를 조정해 시각적 흐름과 계층을 제어할 수 있습니다.

`h1` 규칙을 만들고 `font-weight` 속성을 `800`로 설정하세요. 이렇게 하면 `h1` 텍스트가 더 굵게 표시됩니다.

# --hints--

`h1` 선택자를 만들어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` 선택자에 `font-weight` 속성을 `800`로 설정하세요.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
