---
id: 615f3ed16592445e57941aec
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

라인은 공간이 제한될 때 특히 중요한 내용을 구분하고 그룹화하는 데 도움이 됩니다.

`div` 요소 아래에 `h1` 요소를 만들고, `class` 속성을 `divider` 값으로 설정하세요.

# --hints--

새로운 `div`를 만들어야 합니다.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

새로운 `div`는 `class` 속성을 `divider` 값으로 설정해야 합니다.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

`.divider` 요소는 `.label` 요소 안에 있어야 합니다.

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

`.divider` 요소는 `h1` 요소 다음에 와야 합니다.

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
