---
id: 615f3b091162165948e1cb33
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

브라우저 개발자 도구로 `.label` 요소를 검사하면 실제 너비가 270픽셀이 아니라 288픽셀인 것을 알 수 있습니다. 이는 기본적으로 브라우저가 요소 크기를 결정할 때 테두리와 패딩을 포함하기 때문입니다.

이를 해결하려면 `*` 선택자를 만들고 `box-sizing` 속성에 `border-box` 값을 지정하여 박스 모델을 재설정하세요.

# --hints--

`*` 선택자를 만들어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` 선택자에 `box-sizing` 속성을 `border-box` 값으로 설정해야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
