---
id: 615f34ecc1091b4fd5a8a484
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`head` 요소 안에 `link` 요소를 추가하고 `rel` 속성을 `stylesheet`으로, `href` 속성을 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`로 설정하세요.

이렇게 하면 `Open Sans` 폰트 패밀리를 가져오며, 폰트 두께 값은 `400`, `700`, `800`입니다.

또한 `link` 요소를 추가해 `styles.css` 파일을 연결하세요.

# --hints--

코드에는 `link` 요소가 두 개 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

`link` 요소는 빈 요소여야 하며, 닫는 태그가 없어야 합니다 `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

두 `link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

두 `link` 요소 모두 `rel` 속성이 `stylesheet`로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

`link` 요소 중 하나는 `href` 속성이 `styles.css`로 설정되어야 합니다.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

`link` 요소 중 하나는 `href` 속성이 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`로 설정되어야 합니다.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
