---
id: 615f357957e370510f21ea16
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`body` 선택자를 만들고 `font-family`을 `Open Sans`로 설정하며 대체 수단으로 `sans-serif`을 지정하세요.

이름에 공백이 있는 폰트는 CSS에서 따옴표로 감싸야 한다는 점을 기억하세요.

# --hints--

`body` 선택자가 있어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 선택자의 `font-family` 속성은 `"Open Sans", sans-serif`로 설정되어야 합니다.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
