---
id: 68551b677a5f6898af364b85
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、Tailwind CSS の flexbox ユーティリティクラスを使ってレスポンシブなコールトゥアクション（CTA）コンポーネントを作成しながら、Tailwind CSS についてさらに学びます。

基本的な HTML のボイラープレートは用意されているので、まずは HTML 要素を作成してください。

`div` 要素を作成することから始めましょう。その `div` の中に、さらに別々の `div` 要素を2つ作成してください。

# --hints--

`div` 要素を作成してください。

```js
assert.exists(document.querySelector("div"));
```

`div` 要素の中に、さらに `div` 要素を2つ作成してください。

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
