---
id: 68552464ef378e1051a55c28
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

2番目の`div`の中に、2つのアンカー要素を作成してください。1つ目はテキストが`Learn more`で、2つ目はテキストが`Start listening`であるべきです。

2つのアンカー要素の`href`属性は、どちらも`#`に設定してください。

# --hints--

2番目の`a`要素の中に、2つの`div`要素を作成する必要があります。

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

1つ目の`a`要素は、テキストが`Learn more`で、2番目の`href`要素の中に`#`属性が`div`であるべきです。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

2つ目の`a`要素は、テキストが`Start listening`で、2番目の`href`要素の中に`#`属性が`div`であるべきです。

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
