---
id: 6855255d96c1441bd919dc5d
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

テキストが `span` の `Soundflow` 要素に移動し、その中のテキストは大文字で表示する必要があります。

`uppercase` ユーティリティクラスの使い方を示す例を紹介します。

```html
<p class="uppercase">Quincy Larson</p>
```

`uppercase` 要素にクラス `span` を追加してください。

# --hints--

`span` 要素にユーティリティクラス `uppercase` を付ける必要があります。

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl.classList.contains("uppercase"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        --fcc-editable-region--
        <span>Soundflow</span>
        --fcc-editable-region--
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
