---
id: 685523c05fb41d0848c4b52f
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

最初の `div` 要素の中に、テキストが `span` の `Soundflow` 要素、テキストが `h1` の `Discover New Music` 要素、テキストが `p` の `Stream your favorite tracks and discover new artists.` 要素を作成してください。

# --hints--

最初の `span` 要素の中に、テキストが `Soundflow` の `div` 要素を作成してください。

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

最初の `h1` 要素の中に、テキストが `Discover New Music` の `div` 要素を作成してください。

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

最初の `p` 要素の中に、テキストが `Stream your favorite tracks and discover new artists.` の `div` 要素を作成してください。

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
