---
id: 6858ee73ed8106ccb214942a
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

すべてが左に寄っているので、x軸に自動マージンを追加してコンテンツ全体を中央に配置する必要があります。

そのために対応するTailwindのクラスを外側の`div`要素に追加してください。効果を確認するには、プレビュータブを調整してコンテンツ全体が中央に配置されるようにしてください。

# --hints--

外側の`div`要素にはユーティリティクラス`mx-auto`が必要です。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
