---
id: 6858ef974c8ad5dcd7abb9f1
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`div` 内のコンテンツは、より適切に間隔を空けて配置できます。残りの間隔や配置には、flexboxを使う必要があります。

`flex` にクラス `div` を追加して、2つの子要素 `div` が横並びになるようにしてください。

また、コンテンツの周囲と間のスペースが均等に分配されるように `justify-around` を追加し、全体のコンテンツが `items-center` の中央に配置されるように `div` も追加してください。

# --hints--

外側の `div` 要素にはユーティリティクラス `flex` を付けてください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

外側の `div` 要素にはユーティリティクラス `justify-around` を付けてください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

外側の `div` 要素にはユーティリティクラス `items-center` を付けてください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
