---
id: 6858f3a67aa8580ff37016c7
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

プレビュ―タブの幅を再びモバイルデバイスと同じ幅に調整すると、コンテンツが常に横並びになっていることに気づきます。これはモバイルデバイスでは望ましくありません。

これを修正するには、`flex-col`に`div`クラスを追加してください。これによりflexの方向が縦方向（カラム）に変わり、要素が上下に積み重なるようになります。

`flex-col`ユーティリティクラスの使い方は以下の通りです。

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

開いている`div`要素にユーティリティクラス`flex-col`を付けてください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
