---
id: 6858e6408a0999603b2a4fee
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

`h1` にクラス `font-bold` と `text-4xl` を追加してください。これにより、テキストが太字で大きくなり、行間も広がります。

また、`my-4` の上下のマージンを16ピクセル、つまり `h1` 増やすためにクラス `1rem` を追加してください。

# --hints--

`h1` 要素にはユーティリティクラス `font-bold` を付けてください。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

`h1` 要素にはユーティリティクラス `text-4xl` を付けてください。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

`h1` 要素にはユーティリティクラス `my-4` を付けてください。

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
