---
id: 6858eadd3144c89d45e02bc7
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ご覧のとおり、外側の `div` は上にくっつきすぎていて、その中のコンテンツは左に詰まりすぎています。ですので、今はスペース調整をしましょう。

`mt-8` にクラス `div` を追加して少し下に押し下げ、`p-4` でテキストを左から離すようにしてください。

# --hints--

外側の `div` 要素にはユーティリティクラス `mt-8` を付ける必要があります。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

外側の `div` 要素にはユーティリティクラス `p-4` を付ける必要があります。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
