---
id: 6858f90c54647a48ddb8822d
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

この時点で、外側の `div` 内の2つの `div` 要素は、画面サイズに関係なく縦に積み重なって表示されます。

これは、ラップしている `flex-col` 要素に `div` クラスがあるためです。これを修正するには、`lg:` 接頭辞を使って `flex-row` に `div` クラスを追加してください。

つまり、画面の大きいデバイスでは、flex の方向が `row` になります。効果を再確認するには、ブラウザの幅を広げてみてください。

# --hints--

外側の `div` 要素にはユーティリティクラス `lg:flex-row` を付ける必要があります。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
