---
id: 6858e7e23bf9e475e14cb931
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

HTML構造が完成したので、Tailwindを使った適切なスタイリングに進みましょう。

コンテンツ全体を包む`div`要素から始めます。`class`に`bg-indigo-600`を設定して背景色を指定してください。

# --hints--

外側の`div`には`class`属性を`bg-indigo-600`に設定してください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
