---
id: 6858ebf2552046ab3de9f17c
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

この時点で、`div` のコンテンツの右側にスペースが広すぎることに気づくでしょう。幅を指定するとこれを最小限にできます。

`md:w-1/2` にクラス `div` を追加して、ミディアム以上の画面サイズ（768ピクセル以上）で幅を50%に縮小してください。

効果を確認するには、プレビュータブを展開してください。

# --hints--

外側の `div` 要素にユーティリティクラス `md:w-1/2` を付けてください。

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
