---
id: 68597114730d7c6b8838d204
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`h1` の下の段落は問題ないので、これ以上スタイリングする必要はありません。つまり、`div` の2番目の要素内にあるアンカー要素に移ってください。

最初のアンカー要素にはクラス `bg-white`、`hover:bg-gray-100`、`text-indigo-600` を付けてください。これにより、背景が白くなり、テキストカラーがインディゴになり、ホバー時に少し明るい背景になります。

2番目のアンカー要素にはクラス `bg-pink-500`、`hover:bg-pink-600`、`text-white` を付けてください。

# --hints--

最初のアンカー要素にはユーティリティクラス `bg-white` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

最初のアンカー要素にはユーティリティクラス `hover:bg-gray-100` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

最初のアンカー要素にはユーティリティクラス `text-indigo-600` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

2番目のアンカー要素にはユーティリティクラス `bg-pink-500` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

2番目のアンカー要素にはユーティリティクラス `hover:bg-pink-600` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

2番目のアンカー要素にはユーティリティクラス `text-white` を付けてください。

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
