---
id: 60fad1cafcde010995e15306
title: Passo 38
challengeType: 0
dashedName: step-38
---

# --description--

Com envios de formulário, é útil e boa prática fornecer a cada elemento enviado um atributo `name`. Este atributo é usado para identificar o elemento no envio do formulário.

Excetuando as duas entradas `radio` (que você já nomeou), dê a cada elemento que possa ser enviado um atributo `name` exclusivo e de sua escolha.

# --hints--

Você deve fornecer o `input` esperando um primeiro nome um atributo `name`. *PS Eu teria escolhido `first-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Você deve fornecer o `input` esperando um sobrenome um atributo `name`. *PS Eu teria escolhido `last-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Você deve dar ao `email` um atributo `name`. *PS Eu teria escolhido `email`*

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Você deve dar ao `password` um atributo `name`. *PS Eu teria escolhido `password`*

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Você deve dar ao `checkbox` um atributo `name`. *PS Eu teria escolhido `terms`*

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Você deve dar ao `file` um atributo `name`. *PS Eu teria escolhido `file`*

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Você deve dar ao `number` um atributo `name`. *PS Eu teria escolhido `age`*

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Você deve dar ao elemento `select` um atributo `name`. *PS Eu teria escolhido `referrer`*

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Você deve dar ao elemento `textarea` um atributo `name`. *PS Eu teria escolhido `bio`*

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Você deve não dar atributos `name` para nenhum elemento `option`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Você deve não dar atributos `name` para nenhum elemento `label`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Você deve não dar atributos `name` para nenhum elemento `fieldset`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
      <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
      <label for="referrer">How did you hear about us?
        <select id="referrer">
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
      <label for="bio">Provide a bio:
        <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
      </label>
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
