---
id: 66c057041df6394ca796bf33
title: Construye un Creador de Oraciones
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

En este laboratorio, crearás dos historias diferentes utilizando una plantilla de oraciones. Usarás variables para almacenar diferentes partes de la historia y luego mostrarás las historias en la consola.

**Objetivo:** Cumplir con los siguientes requisitos del usuario y lograr que pasen todas las pruebas para completar el laboratorio.

**Requisitos del usuario:**

1. Debes declarar las siguientes variables usando `let`:

    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. Debes asignar a las variables anteriores algunas cadenas de texto de tu elección.

1. Debes declarar una variable `firstStory`.

1. Debes usar la siguiente plantilla de historia para crear la primera y asignarla a la variable `firstStory`: `"Había una vez, un(a) [adjetivo] [sustantivo] que le encantaba comer [sustantivo2]. El [sustantivo] vivía en un [lugar] y tenía fosas nasales [adjetivo2] que echaban fuego cuando estaba [verbo]."`;

1. Debes mostrar tu primera historia en la consola usando el mensaje `"Primera historia: [firstStory]"`.

1. Debes asignar nuevos valores a tus variables `adjetivo`, `sustantivo`, `verbo`, `lugar`, `adjetivo2` y `sustantivo2`.

1. Debes declarar una variable `secondStory`.

1. Crea otra historia utilizando la misma plantilla y asígnala a la variable `secondStory`.

1. Debes mostrar tu segunda historia en la consola usando el mensaje `"Segunda historia: [secondStory]"`.


# --hints--

Debes declarar una variable `adjetivo`.

```js
assert.isNotNull(adjective);
```

Debes declarar una variable `sustantivo`.

```js
assert.isNotNull(noun);
```

Debes declarar una variable `verbo`.

```js
assert.isNotNull(verb);
```

Debes declarar una variable `lugar`.

```js
assert.isNotNull(place);
```

Debes declarar una variable `adjetivo2`.

```js
assert.isNotNull(adjective2);
```

Debes declarar una variable `sustantivo2`.

```js
assert.isNotNull(noun2);
```

Debes asignar un valor de cadena de texto a la variable `adjetivo`.

```js
assert.isString(adjective);
```

Debes asignar un valor de cadena de texto a la variable `sustantivo`.

```js
assert.isString(noun);
```

Debes asignar un valor de cadena de texto a la variable `verbo`.

```js
assert.isString(verb);
```

Debes asignar un valor de cadena de texto a la variable `lugar`.

```js
assert.isString(place);
```

Debes asignar un valor de cadena de texto a la variable `adjetivo2`.

```js
assert.isString(adjective2);
```

Debes asignar un valor de cadena de texto a la variable `sustantivo2`.

```js
assert.isString(noun2);
```

Debes declarar una variable `firstStory`.

```js
assert.isNotNull(firstStory);
```

Debes usar el formato de historia correcto para la primera historia: `"Había una vez, un(a) [adjetivo] [sustantivo] que le encantaba comer [sustantivo2]. El [sustantivo] vivía en un [lugar] y tenía fosas nasales [adjetivo2] que echaban fuego cuando estaba [verbo]."`. Presta atención a los espacios.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Debes registrar tu primera historia usando el mensaje `"Primera historia: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Debes declarar una variable `secondStory`.

```js
assert.isNotNull(secondStory);
```

Debes reasignar la variable `adjetivo` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Debes reasignar la variable `sustantivo` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Debes reasignar la variable `verbo` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Debes reasignar la variable `lugar` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Debes reasignar la variable `adjetivo2` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Debes reasignar la variable `sustantivo2` para la segunda historia.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Debes usar el formato de historia correcto para la segunda historia: `"Había una vez, un(a) [adjetivo] [sustantivo] que le encantaba comer [sustantivo2]. El [sustantivo] vivía en un [lugar] y tenía fosas nasales [adjetivo2] que echaban fuego cuando estaba [verbo]."`. Presta atención a los espacios.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Debes registrar tu segunda historia usando el formato `"Segunda historia: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

La `firstStory` no debe ser la misma que la `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```

