---
id: afcc8d540bea9ea2669306b6
title: 문자열 반복 함수 만들기
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

이 실습에서는 주어진 문자열을 특정 횟수만큼 반복하는 함수를 만드세요. 이 실습에서는 내장 `.repeat()` 메서드를 사용하지 마세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `repeatStringNumTimes`라는 이름의 함수를 만들고, 두 개의 매개변수를 받으세요: 문자열과 숫자.
2. 함수는 지정된 횟수만큼 문자열을 반복하여 반환해야 합니다.
3. 숫자가 0 이하일 경우, 함수는 빈 문자열을 반환해야 합니다.

# --hints--

`repeatStringNumTimes`라는 이름의 함수를 만드세요.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes`는 두 개의 매개변수를 받아야 합니다.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

`repeatStringNumTimes` 함수는 항상 문자열을 반환해야 합니다.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)`는 문자열 `***`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)`는 문자열 `abcabcabc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)`는 문자열 `abcabcabcabc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)`는 문자열 `abc`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)`는 문자열 `********`를 반환해야 합니다.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)`는 빈 문자열(`""`)을 반환해야 합니다.

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)`는 `""`를 반환해야 합니다.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

내장 `repeat()` 메서드는 사용하지 마세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
