---
id: a103376db3ba46b2d50db289
title: Implemente um Conversor de Spinal Case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case é um formato de string onde todas as palavras estão em letras minúsculas e separadas por hífens. `"this-is-spinal-tap"` é um exemplo de uma string em spinal case.

Neste laboratório, você criará uma função que converte uma string dada para spinal case.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `spinalCase`.
2. A função `spinalCase` deve receber um único argumento, uma string.
3. A função `spinalCase` deve retornar a string no formato spinal case. Por exemplo, se o argumento for `ProductLanding page`, a função deve retornar `product-landing-page`.

# --hints--

Você deve criar uma função chamada `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` deve receber um único argumento.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` deve retornar a string `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` deve retornar a string `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` deve retornar a string `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` deve retornar a string `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` deve retornar a string `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
