---
id: 64220e8cb589f61e625bf453
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

数字と `dollars` というテキストの間に、百や千の位を表す語が入る場合も検出できるようにしましょう。

`|` の記号を使用して、数字と `dollars` の間に `hundred`、`thousand`、`million`、または `billion` が入ることを許容しましょう。

# --hints--

正規表現 `dollarRegex` には `|` 記号を使用してください。

```js
assert.match(dollarRegex.source, /\|/);
```

正規表現 `dollarRegex` には `|` 記号が 3 つ必要です。

```js
assert.lengthOf(dollarRegex.source.match(/\|/g), 3);
```

正規表現 `dollarRegex` で `|` 記号を使用して、`hundred`、`thousand`、`million`、または `billion` にマッチさせてください。

```js
const placeValues = dollarRegex.source.replace("[0-9]+ ", "").replace(" dollars", "").split('|');
assert.include(placeValues, 'hundred');
assert.include(placeValues, 'thousand');
assert.include(placeValues, 'million');
assert.include(placeValues, 'billion');
```

正規表現 `dollarRegex` が `1 hundred dollars` にマッチするようにしてください。

```js
assert.match('1 hundred dollars', dollarRegex);
```

正規表現 `dollarRegex` が `1 thousand dollars` にマッチするようにしてください。

```js
assert.match('1 thousand dollars', dollarRegex);
```

正規表現 `dollarRegex` が `1 million dollars` にマッチするようにしてください。

```js
assert.match('1 million dollars', dollarRegex);
```

正規表現 `dollarRegex` が `1 billion dollars` にマッチするようにしてください。

```js
assert.match('1 billion dollars', dollarRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /[0-9]+ dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
