---
id: a6b0bb188d873cb2c8729495
title: Implementar um Conversor de Entidades HTML
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Este laboratório trata da conversão de caracteres especiais em uma string para suas respectivas entidades HTML.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `convertHTML` que aceita uma string como argumento.
1. A função `convertHTML` deve retornar uma nova string convertendo os caracteres especiais na string de argumento para suas entidades HTML correspondentes.

   - `&` deve ser convertido para `&amp;`.
   - `<` deve ser convertido para `&lt;`.
   - `>` deve ser convertido para `&gt;`.
   - `&quot;` deve ser convertido para `&quot;`.
   - '` should be convertido para `&apos;`.

# --hints--

Você deve ter uma função `convertHTML`.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` deve retornar a string `Dolce & Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` deve retornar a string `Hamburgers < Pizza < Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sessenta > doze")` deve retornar a string `Sessenta > doze`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` deve retornar a string `Stuff in "quotation marks"`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("A Lista de Schindler")` deve retornar a string `Schindler's List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` deve retornar a string `<>`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` deve retornar a string `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
