---
id: 587d7b85367417b2b2512b3a
title: "Виявлення аргументів, переданих у неправильному порядку при виклику функції"
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

Продовжуючи обговорення виклику функцій, наступною помилкою, на яку слід звернути увагу, є передача аргументів функції у неправильному порядку. Якщо аргументи мають різні типи, наприклад, функція очікує масив і ціле число, це, ймовірно, спричинить помилку виконання. Якщо аргументи одного типу (наприклад, усі цілі числа), тоді логіка коду не матиме сенсу. Переконайтеся, що всі необхідні аргументи передані у правильному порядку, щоб уникнути цих проблем.

# --instructions--

Функція `raiseToPower` підносить основу до степеня. На жаль, вона викликається неправильно — виправте код так, щоб значення `power` було очікуваним 8.

# --hints--

Ваш код має виправити змінну `power` так, щоб вона дорівнювала 2 у степені 3, а не 3 у степені 2.

```js
assert(power == 8);
```

Ваш код має використовувати правильний порядок аргументів для виклику функції `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
