---
id: 587d7b86367417b2b2512b3c
title: Будьте обережні при повторній ініціалізації змінних у циклі
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

Іноді потрібно зберігати інформацію, збільшувати фішки або повторно встановлювати змінні всередині циклу. Потенційна проблема виникає, коли змінні або мають бути повторно ініціалізовані, але цього не відбувається, або навпаки. Це особливо небезпечно, якщо ви випадково скидаєте змінну, яка використовується для кінцевої умови, що призводить до безкінечного циклу.

Виведення значень змінних на кожній ітерації циклу за допомогою `console.log()` може допомогти виявити помилки, пов’язані зі скиданням або відсутністю скидання змінної.

# --instructions--

Наступна функція має створити двовимірний масив з `m` рядками та `n` стовпцями, заповненими нулями. На жаль, вона не дає очікуваного результату, бо змінна `row` не повторно ініціалізується (не встановлюється назад у порожній масив) у зовнішньому циклі. Виправте код, щоб він повертав правильний масив 3x2, заповнений нулями, який виглядає як `[[0, 0], [0, 0], [0, 0]]`.

# --hints--

Ваш код має встановити змінну `matrix` у масив, що містить 3 рядки по 2 стовпці нулів у кожному.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

Змінна `matrix` має містити 3 рядки.

```js
assert(matrix.length == 3);
```

Змінна `matrix` має містити 2 стовпці в кожному рядку.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` має повертати масив, що містить 4 рядки по 3 стовпці нулів у кожному.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
