---
id: 587d7b84367417b2b2512b36
title: "Виявлення незакритих круглих дужок, дужок, фігурних дужок і лапок"
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

Ще одна синтаксична помилка, про яку варто пам’ятати, — це те, що всі відкриваючі круглі дужки, дужки, фігурні дужки та лапки мають мати пару для закриття. Забути про одну з них часто трапляється, коли ви редагуєте існуючий код і вставляєте елементи з одним із типів пар. Також будьте уважні, коли вкладаєте блоки коду один в одного, наприклад, додаючи функцію зворотного виклику як аргумент до методу.

Один зі способів уникнути цієї помилки — одразу після введення відкриваючого символу додати відповідний закриваючий, потім повернути курсор між ними і продовжити писати код. На щастя, більшість сучасних редакторів коду автоматично генерують другу половину пари.

# --instructions--

Виправте дві помилки з парами у коді.

# --hints--

Ваш код має виправити відсутню частину масиву.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Ваш код має виправити відсутню частину методу `.reduce()`. Вихідні дані в консолі мають показати, що `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
