---
id: 587d7b86367417b2b2512b3b
title: Виправлення помилок з індексацією на один елемент
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>Помилки на один елемент</dfn> (іноді їх називають OBOE) виникають, коли ви намагаєтеся звернутися до конкретного індексу рядка або масиву (щоб вирізати або отримати сегмент), або коли перебираєте їх індекси в циклі. Індексація в JavaScript починається з нуля, а не з одиниці, тому останній індекс завжди на одиницю менший за довжину елемента. Якщо ви спробуєте звернутися до індексу, рівного довжині, програма може викинути помилку "index out of range" або вивести `undefined`.

Коли ви використовуєте методи рядків або масивів, які приймають діапазони індексів як аргументи, корисно прочитати документацію і зрозуміти, чи вони включають цей індекс (елемент за заданим індексом входить у результат) чи ні. Ось кілька прикладів помилок на один елемент:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

Перший приклад тут виконує цикл на один раз більше, ніж потрібно, а другий — на один раз менше (пропускає перший індекс, 0). Третій приклад правильний.

# --instructions--

Виправте дві помилки з індексацією у наведеній функції, щоб усі числа від 1 до 5 виводилися в консоль.

# --hints--

Ваш код має встановити початкову умову циклу так, щоб він починався з першого індексу.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

Ваш код має виправити початкову умову циклу так, щоб індекс починався з 0.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

Ваш код має встановити кінцеву умову циклу так, щоб він зупинявся на останньому індексі.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

Ваш код має виправити кінцеву умову циклу так, щоб він зупинявся за 1 до довжини.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```
