---
id: 587d7b84367417b2b2512b37
title: Виявлення змішаного використання одинарних і подвійних лапок
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript дозволяє використовувати як одинарні (`'`), так і подвійні (`"`) лапки для оголошення рядка. Вибір між ними зазвичай залежить від особистих уподобань, з деякими винятками.

Наявність двох варіантів дуже зручна, коли рядок містить скорочення або інший текст у лапках. Просто будьте уважні, щоб не закрити рядок занадто рано, інакше виникне синтаксична помилка.

Ось кілька прикладів змішування лапок:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

Перші два приклади правильні, а третій — неправильний.

Звісно, можна використовувати лише один стиль лапок. Ви можете екранувати лапки всередині рядка, використовуючи символ екранування зворотної скісної риски (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

Виправте рядок так, щоб для значення `href` використовувалися різні лапки, або екрануйте існуючі. Залиште подвійні лапки навколо всього рядка.

# --hints--

Ваш код має виправити лапки навколо значення `href` `#Home`, змінивши їх або екранувавши.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

Ваш код має зберегти подвійні лапки навколо всього рядка.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```
