---
id: 587d7b85367417b2b2512b39
title: Виявлення відсутніх відкриваючих і закриваючих круглих дужок після виклику функції
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

Коли функція або метод не приймає жодних аргументів, ви можете забути додати (порожні) відкриваючі та закриваючі круглі дужки під час її виклику. Часто результат виклику функції зберігається у змінній для подальшого використання у вашому коді. Цю помилку можна виявити, вивівши значення змінних (або їхні типи) у консоль і побачивши, що одна зі змінних містить референс на функцію, замість очікуваного значення, яке функція повертає.

Змінні у наведеному прикладі різні:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Тут `varOne` — це функція `myFunction`, а `varTwo` — рядок `You rock!`.

# --instructions--

Виправте код так, щоб змінна `result` містила значення, яке повертає виклик функції `getNine`.

# --hints--

Ваш код має виправити змінну `result`, щоб вона містила число, яке повертає функція `getNine`.

```js
assert(result == 9);
```

Ваш код має викликати функцію `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
