---
id: 587d7b83367417b2b2512b33
title: Використання консолі JavaScript для перевірки значення змінної
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

І Chrome, і Firefox мають чудові консолі JavaScript, також відомі як інструменти розробника, для налагодження вашого JavaScript.

Ви можете знайти інструменти розробника в меню Chrome або веб-консоль у меню Firefox. Якщо ви користуєтеся іншим браузером або мобільним телефоном, ми наполегливо рекомендуємо перейти на десктопну версію Firefox або Chrome.

Метод `console.log()`, який "виводить" у консоль результат того, що знаходиться в його круглих дужках, ймовірно, буде найкориснішим інструментом для налагодження. Розміщення його у стратегічних місцях у вашому коді може показати проміжні значення змінних. Добра практика — мати уявлення про те, яким має бути вихідний результат, перш ніж дивитися, яким він є насправді. Наявність контрольних точок для перевірки стану ваших обчислень протягом усього коду допоможе звузити місце, де виникла помилка.

Ось приклад, як вивести рядок `Hello world!` у консоль:

```js
console.log('Hello world!');
```

# --instructions--

Використайте метод `console.log()`, щоб вивести значення змінної `a` у вказаному місці коду.

# --hints--

Ваш код має використовувати `console.log()`, щоб перевірити значення змінної `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
