---
id: 587d7b84367417b2b2512b35
title: Виявлення неправильно написаних імен змінних і функцій
challengeType: 1
forumTopicId: 301186
dashedName: catch-misspelled-variable-and-function-names
---

# --description--

Методи `console.log()` і `typeof` — це два основні способи перевірити проміжні значення та типи вихідних даних програми. Тепер настав час розглянути загальні форми, які приймають помилки. Одна з проблем на рівні синтаксису, з якою можуть зрозумітися швидкі друкарі, — це звичайна орфографічна помилка.

Переставлені, пропущені або неправильно написані великі літери в імені змінної чи функції призведуть до того, що браузер шукатиме об’єкт, якого не існує, — і повідомить про це у вигляді помилки референсу. Імена змінних і функцій у JavaScript чутливі до регістру.

# --instructions--

Виправте дві орфографічні помилки в коді, щоб обчислення `netWorkingCapital` працювало.

# --hints--

Перевірте написання двох змінних, які використовуються в обчисленні netWorkingCapital, у консолі має відобразитися "Net working capital is: 2".

```js
assert(netWorkingCapital === 2);
```

У коді не повинно бути неправильно написаних змінних.

```js
assert(!__helpers.removeJSComments(code).match(/recievables/g));
```

Змінна `receivables` має бути оголошена і використана правильно в коді.

```js
assert(__helpers.removeJSComments(code).match(/receivables/g).length == 2);
```

У коді не повинно бути неправильно написаних змінних.

```js
assert(!__helpers.removeJSComments(code).match(/payable;/g));
```

Змінна `payables` має бути оголошена і використана правильно в коді.

```js
assert(__helpers.removeJSComments(code).match(/payables/g).length == 2);
```

# --seed--

## --seed-contents--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = recievables - payable;
console.log(`Net working capital is: ${netWorkingCapital}`);
```

# --solutions--

```js
let receivables = 10;
let payables = 8;
let netWorkingCapital = receivables - payables;
console.log(`Net working capital is: ${netWorkingCapital}`);
```
