---
id: 68e448ec4ac4205968c0a805
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Sarebbe bello se la parola `"love"` fosse ripetuta tre volte nella frase `"I love learning!"`.

Come ricordi da una lezione precedente, puoi ripetere una stringa un numero specifico di volte usando il metodo `repeat()`. Ecco un esempio:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Poiché le stringhe sono immutabili, questo metodo non modificherà la stringa originale. Restituirà una nuova stringa con il contenuto ripetuto.

Crea una variabile chiamata `repeatedLove` e assegnale `"love ".repeat(3)`. Poi stampa nella console la variabile `repeatedLove` per vedere il risultato.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile chiamata `repeatedLove`.

```js
assert.exists(repeatedLove);
```

La tua variabile `repeatedLove` dovrebbe essere una stringa.

```js
assert.isString(repeatedLove);
```

Dovresti assegnare `"love ".repeat(3)` alla variabile `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Dovresti stampare nella console la variabile `repeatedLove`.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
