---
id: 68e42de6396f5c49d90d1951
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

In una lezione precedente, hai imparato il metodo `replace()` che si usa per trovare una sottostringa in una stringa e sostituirla con un altro valore.

Ecco un esempio:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Ricorda che le stringhe sono immutabili, il che significa che la `text` originale non viene modificata in quell'esempio. Viene invece creata una nuova stringa.

Crea una nuova variabile chiamata `replacedString` e assegnale `originalString.replace("cats", "dogs")`.

# --hints--

Dovresti avere una variabile chiamata `replacedString`.

```js
assert.exists(replacedString);
```

La tua variabile `replacedString` dovrebbe essere una stringa.

```js
assert.isString(replacedString);
```

Dovresti assegnare `originalString.replace("cats", "dogs")` alla variabile `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
