---
id: 68e43e884ef09d4d2ab05631
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora è il momento di vedere il risultato dell'uso del metodo `replace()`.

Inizia aggiungendo un `console.log()` con il valore `"After using the replace() method:"`.

Poi, sotto quel `console.log()`, aggiungi un altro `console.log()` con la variabile `replacedString`.

Dai un'occhiata alla console e dovresti vedere che la nuova frase dice `"I love dogs."` invece dell'originale `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti registrare `"After using the replace() method:"` nella console.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Dovresti registrare la variabile `replacedString` nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
