---
id: 68e44abea678a65b877864dc
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Sarebbe bello vedere quella parola ripetuta in una frase. Crea una nuova variabile chiamata `newSentence` e assegnale il risultato della concatenazione della parola `"I "` con la variabile `repeatedLove` seguita dalla concatenazione della stringa `" learning."`. Poi stampa la variabile `newSentence` nella console. Puoi scegliere di usare l'operatore `+` oppure i template literal per la concatenazione delle stringhe.

Ora dovresti vedere il testo `"I love love love  learning."`.

*NOTA*: C'è uno spazio extra dopo l'ultimo `"love"` che è corretto. Nell'ultimo passaggio, rimuoverai quello spazio extra.

# --hints--

Dovresti avere una variabile chiamata `newSentence`.

```js
assert.exists(newSentence);
```

La tua variabile `newSentence` dovrebbe essere una stringa.

```js
assert.isString(newSentence);
```

Dovresti concatenare la parola `"I "` con la variabile `repeatedLove` seguita dalla concatenazione della stringa `" learning."`. Questo risultato dovrebbe essere assegnato alla variabile `newSentence`. Fai attenzione agli spazi.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Non dovresti assegnare la stringa letterale `"I love love love  learning."` alla variabile `newSentence`. Rileggi le istruzioni per l'approccio corretto.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
