---
id: 68e4417b7317835208145323
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Successivamente, crea una nuova variabile chiamata `dogsOnlySentence` e assegnale `exampleSentence.replace("cats", "dogs")`.

Poi, sotto quella variabile, stampa nella console `"Replacing all occurrences of cats with dogs:"`. Sotto `console.log()`, aggiungi un altro `console.log()` per la variabile `dogsOnlySentence`.

Ora dovresti vedere il testo `"I love dogs and cats are so much fun!"` stampato nella console. Ma questo non è il risultato desiderato. Nel passaggio successivo imparerai come correggerlo in modo che il testo sia `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

La tua variabile `dogsOnlySentence` dovrebbe essere una stringa.

```js
assert.isString(dogsOnlySentence);
```

Dovresti assegnare `exampleSentence.replace("cats", "dogs")` alla tua variabile `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Dovresti stampare `"Replacing all occurrences of cats with dogs:"` nella console.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Dovresti stampare la variabile `dogsOnlySentence` nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
