---
id: 68e42b62d7093345c6975f22
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Sotto la tua variabile `originalString`, aggiungi una `console.log()` con la stringa `"Original string:"`. 

Poi, sotto quella `console.log()`, aggiungi un'altra `console.log()` con la variabile `originalString`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti stampare la stringa `"Original string:"` nella console.

```js
assert.equal(getLogs()[0], "Original string:");
```

Dovresti stampare la variabile `originalString` nella console. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
