---
id: 68e44519c769bf545bb3f22a
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Per ottenere l'output desiderato di `"I love dogs and dogs are so much fun!"`, dovrai usare il metodo `replaceAll()` invece del metodo `replace()`. Questo metodo restituisce una nuova stringa per tutte le corrispondenze della sottostringa.

Aggiorna la riga `exampleSentence.replace("cats", "dogs")` per usare il metodo `replaceAll()` invece. Ora dovresti vedere il testo corretto nella console.

# --hints--

Dovresti aggiornare la riga `exampleSentence.replace("cats", "dogs")` del codice per usare il metodo `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
