---
id: 68e41084bb275539f72de9fe
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso dei metodi `replace()`, `replaceAll()` e `repeat()` costruendo un'app per trasformare stringhe.

Per iniziare, crea una variabile chiamata `originalString` e assegnale il valore di `"I love cats."`.

# --hints--

Dovresti avere una variabile chiamata `originalString`.

```js
assert.exists(originalString);
```

La tua variabile `originalString` dovrebbe essere una stringa.

```js
assert.isString(originalString);
```

Dovresti assegnare la stringa `"I love cats."` alla variabile `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
