---
id: 68e43f94365a964f26c88f78
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora è il momento di lavorare con un nuovo metodo per le stringhe.

Inizia creando una nuova variabile chiamata `exampleSentence` e assegnale la stringa `"I love cats and cats are so much fun!"`.

Poi, sotto quella variabile, stampa nella console `"Original sentence:"`. Poi, sotto quella, `console.log()`, aggiungi un altro `console.log()` con la variabile `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile `exampleSentence`.

```js
assert.exists(exampleSentence);
```

La tua variabile `exampleSentence` dovrebbe essere una stringa.

```js
assert.isString(exampleSentence);
```

Dovresti assegnare la stringa `"I love cats and cats are so much fun!"` alla variabile `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Dovresti stampare `"Original sentence:"` nella console.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Dovresti stampare la variabile `exampleSentence` nella console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
