---
id: 68e44c305f1f935d949d76bb
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Come accennato nel passaggio precedente, c’è uno spazio extra dopo l’ultimo `"love"` nella frase `"I love love love  learning."`. Per rimuovere quello spazio extra, puoi usare il metodo `trimEnd()`.

Alla fine del metodo `"love ".repeat(3)`, concatena il metodo `trimEnd()`. Puoi concatenare i metodi in questo modo:

```js
.firstMethod().secondMethod()
```

Ora, quando controlli la console, dovresti vedere che lo spazio extra è stato rimosso.

E con quest’ultima modifica, il tuo workshop è completo!

# --hints--

Dovresti concatenare il metodo `trimEnd()` al metodo `"love ".repeat(3)`.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
