---
id: 68e4465fba0bd256868b7deb
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Ora è il momento di lavorare con un nuovo metodo per le stringhe.

Inizia creando una variabile chiamata `learningSentence` e assegnale la stringa `"I love learning!"`.

Poi, sotto quella variabile, stampa nella console `"Original learning sentence:"`. Sotto `console.log()`, aggiungi un'altra `console.log()` per la variabile `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile `learningSentence`.

```js
assert.exists(learningSentence);
```

La tua variabile `learningSentence` dovrebbe essere una stringa.

```js
assert.isString(learningSentence);
```

Dovresti assegnare la stringa `"I love learning!"` alla variabile `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Dovresti stampare `"Original learning sentence:"` nella console.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Dovresti stampare la variabile `learningSentence` nella console.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
