---
id: 655cd899f8de09431eabb40c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Сортування злиттям — це алгоритм сортування, який використовує принцип розділяй і володарюй для сортування колекцій даних. Тобто він «ділить» колекцію на менші підчастини, а потім «перемагає» ці підчастини, сортувавши їх окремо, після чого об’єднує відсортовані підчастини.

Почніть зі створення функції з назвою `merge_sort`, яка приймає параметр `array`. Ця функція виконуватиме завдання сортування списку чисел. Використайте ключове слово `pass` у тілі функції.

# --hints--

Ви повинні оголосити функцію з назвою `merge_sort` з одним параметром: `array`. Не забудьте використати ключове слово `pass` у тілі функції.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
