---
id: 6564aee9c077774ea49c3faf
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Алгоритм сортування злиттям виконує три дії:

- Розділити несортовану послідовність елементів на підчастини
- Відсортувати елементи в підчастинах
- Об’єднати відсортовані підчастини

Вищезазначене відбувається рекурсивно, доки підчастини не будуть об’єднані у повну відсортовану послідовність. Почнемо з розділення послідовності.

Спершу замініть ключове слово `pass` на змінну `middle_point`. Використайте оператор цілочисельного ділення (`//`) для ділення довжини списку `array` на 2 і призначте результат новій змінній `middle_point`. Не забудьте відступити код.

# --hints--

У тілі функції не повинно бути `pass`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

Потрібно оголосити змінну `middle_point` і призначити їй `len(array) // 2`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
