---
id: 656627b47bd2d2a4afbc945d
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У попередньому кроці ви отримали середину. Ви можете використати її, щоб розділити `array` на дві частини та призначити кожну частину новим змінним.

Використайте синтаксис slice, щоб витягти ліву половину `array` і призначити її змінній з ім’ям `left_part`.

# --hints--

У вашій функції `left_part` має бути змінна з ім’ям `merge_sort`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Ви маєте використати синтаксис slice, щоб витягти ліву половину масиву і призначити її змінній `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
