---
id: 656680b0fc79f2c38a34d90e
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер настав час відсортувати та об’єднати списки (`left_part` і `right_part`) у початковий список.

Ви можете зробити це, порівнюючи елементи обох списків і додаючи менший елемент до головного списку. Ви будете робити це порівняння для всіх індексів у `left_part` і `right_part`.

Створіть три змінні: `left_array_index`, `right_array_index` і `sorted_index` та встановіть їхні значення рівними `0`. Ці змінні допоможуть вам відстежувати кожен індекс під час процесу сортування.

# --hints--

У вашій функції `left_array_index` має бути змінна з ім’ям `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

Значення `left_array_index` має бути `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

У вашій функції `right_array_index` має бути змінна з ім’ям `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

Значення `right_array_index` має бути `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

У вашій функції `sorted_index` має бути змінна з ім’ям `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

Значення `sorted_index` має бути `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
