---
id: 656680b0fc79f2c38a34d90e
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Now it's time to sort and merge the lists (`left_part` and `right_part`) into the original list.

Для цього можна порівняти елементи обох списків та додати найменший елемент до головного списку. Ви будете виконувати це порівняння для всіх індексів у `left_part` та `right_part`.

Створіть три змінні: `left_array_index`, `right_array_index` та `sorted_index`, і встановіть їхні значення на `0`. Ці змінні допоможуть вам відстежувати кожен індекс під час процесу сортування.

# --hints--

Функція `merge_sort` повинна містити змінну під назвою `left_array_index`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

The value of `left_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

Функція `merge_sort` повинна містити змінну під назвою `right_array_index`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

The value of `right_array_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

Функція `merge_sort` повинна містити змінну під назвою `sorted_index`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

The value of `sorted_index` should be `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--

--fcc-editable-region--
```
