---
id: 656639771fed09b5c0e8fe71
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер, коли ви розділили список `array` на два окремі списки, продовжуйте ділити кожен список, доки кожен елемент не опиниться у власному списку. Список з одним числом завжди відсортований.

Для цього рекурсивно викликайте `merge_sort` у своїй функції та передавайте `left_part` як аргумент виклику.

# --hints--

Ви повинні викликати `merge_sort` з аргументом `left_part` внизу тіла вашої функції.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
