---
id: 656702f8b4cbd8cbf0a433c6
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Далі потрібно створити цикл, який триватиме, поки залишаються елементи в обох `left_part` і `right_part`.

Для цього створіть `while` цикл із двома умовами: одна перевіряє, чи `left_array_index` менше за довжину `left_part`, а інша перевіряє, чи `right_array_index` менше за довжину `right_part`.

# --hints--

Ваш `while` цикл має перевіряти, чи `left_array_index` менше за довжину `left_part` і чи `right_array_index` менше за довжину `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Використовуйте ключове слово `pass` у тілі циклу.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
