---
id: 6569b743630ee592a65a7e2f
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Перед тестуванням функції `merge_sort()` потрібно створити базовий випадок, який зупиняє виконання функції, коли довжина `array` менша або дорівнює `1`.

Цей базовий випадок зупинить рекурсивний виклик. Без нього операція сортування злиттям продовжувала б виконуватися навіть тоді, коли список уже відсортований або не містить елементів.

Одразу після оголошення функції створіть інструкцію `if` з умовою: `len(array) <= 1`. У тілі інструкції `if` додайте інструкцію `return`, щоб зупинити виконання функції.

# --hints--

Ви повинні створити інструкцію `if`, щоб перевірити, чи `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

У тілі інструкції `return` має бути інструкція `if`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
