---
id: 6569b20f829b7e69d43c232a
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Останнє, що потрібно зробити для циклу `while`, — перейти до наступного індексу в відсортованому масиві.

Додайте `1` до значення `sorted_index`.

# --hints--

Ви повинні збільшити поточне значення `sorted_index` на одиницю всередині циклу `while`. Зверніть увагу, що збільшення має відбуватися після призначення елемента лівої частини відсортованому масиву

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
