---
id: 6569b4e0bd29d17d4b53b16c
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Тепер ви маєте відтворити ту саму логіку циклу `while` для `right_part`.

Створіть цикл `while`, який запускається, коли `right_array_index` менше за `len(right_part)`. Усередині циклу `while` призначте `right_part[right_array_index]` `array[sorted_index]`. Наприкінці збільшіть `right_array_index` і `sorted_index` на `1`.

# --hints--

Ваш цикл `while` має перевіряти, чи `right_array_index` менше за `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

У тілі циклу слід призначити `right_part[right_array_index]` `array[sorted_index]`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

У циклі `right_array_index` слід збільшувати `1` на `while`. Зверніть увагу, що збільшення має відбуватися після призначення елемента правої частини до відсортованого масиву.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

У циклі `sorted_index` слід збільшувати поточне значення `1` на `while`. Зверніть увагу, що збільшення має відбуватися після призначення елемента правої частини до відсортованого масиву.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
