---
id: 6577254891048c8f2c19e961
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Внизу тіла вашої функції знову викличте функцію `merge_sort()`. Цього разу передайте `right_part` як аргумент виклику функції.

# --hints--

Ви маєте викликати `merge_sort()` з аргументом `right_part` внизу тіла вашої функції.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
