---
id: 657f59751d5a5c9b069ae0f3
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Як ви дізналися в попередньому уроці, ви можете використати змінну `__name__`, щоб визначити, чи виконується скрипт Python як основна програма, чи він імпортується як модуль (код, написаний в іншому файлі Python).

Якщо значення `__name__` встановлено в `'__main__'`, це означає, що поточний скрипт є основною програмою, а не модулем.

У цьому проєкті ви використаєте поточний скрипт як основну програму.

Створіть інструкцію `if`, яка перевіряє, чи значення `__name__` дорівнює `'__main__'`.
Усередині інструкції `if` створіть список з назвою `numbers` і призначте йому такі значення: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

Ваша інструкція `if` має перевіряти, чи має `__name__` значення `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

У вас має бути список `numbers` із такими значеннями: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
