---
id: 657f3a661730891aa64f3568
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Використайте синтаксис slice, щоб витягти праву половину `array` і призначити її змінній з ім’ям `right_part`.

# --hints--

У вас має бути змінна з ім’ям `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Використайте синтаксис slice, щоб витягти праву половину масиву і призначити її змінній `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
