---
id: 56533eb9ac21ba0edf2244b3
title: Erstellen Sie einen Celsius-zu-Fahrenheit-Konverter
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

In diesem Labor schreiben Sie eine Funktion, die die Temperatur von Celsius in Fahrenheit umwandelt. Die Formel zur Umrechnung von Celsius in Fahrenheit lautet:

```js
fahrenheit = celsius * (9/5) + 32
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `convertCtoF` erstellen.
2. `convertCtoF` sollte ein einzelnes numerisches Argument annehmen, das die Temperatur in Celsius ist.
3. `convertCtoF` sollte eine Zahl zurückgeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `convertCtoF` erstellen.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` sollte einen einzelnen Parameter annehmen.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` sollte eine Zahl zurückgeben.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` sollte einen Wert von `-22` zurückgeben.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` sollte einen Wert von `14` zurückgeben.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` sollte einen Wert von `32` zurückgeben.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` sollte einen Wert von `68` zurückgeben.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` sollte einen Wert von `86` zurückgeben.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
