---
id: 6735a7370e0ae93a4577c689
title: 設計一個雜誌版面
challengeType: 25
dashedName: design-a-magazine-layout
demoType: onClick
---

# --description--

在本實驗室中，你將使用 CSS Grid 設計雜誌佈局，包括網格行和網格列等概念。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應該有一個類爲 `magazine-cover` 的 `main`元素。

2. 你的 `.magazine-cover` 應包含一個 `header` 元素，其類爲 `title`。

3. 你的頁面中應該有四個 `section` 元素。 前三個部分用於雜誌文章，第四個 `section` 元素用於封面圖片。

4. 四個 `section` 元素的 `class` 屬性分別爲 `feature-article`、`small-article1`、`small-article2` 和 `cover-image`。

5. 你應該將 `.magazine-cover` 的 `display` 屬性設置爲 `grid`。

6. `.magazine-cover` 應包含一個 `grid-template-areas` 屬性，用於定義網格佈局：

   - 標題橫跨頂行。
   - 在第二行，專題文章橫跨兩欄，然後是封面圖片。
   - 第三行是兩篇小文章和封面圖片。

7. `.magazine-cover` 應具有一個 `grid-template-columns` 屬性，可將空間分成三個相等的部分。

8. `.magazine-cover` 應具有一個 `grid-template-rows` 屬性。 對於第一行，應將其設置爲 `auto`，並平分其餘兩行所佔的空間。

9. 你應在網格項之間添加間隙，並將 `.magazine-cover` 的寬度設置爲佔據大部分視口寬度。

**注意：** 確保在 HTML 中鏈接樣式表以應用 CSS。

# --hints--

你應該有一個 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 元素應該有一個 `class`，值爲 `magazine-cover`。

```js
assert.exists(document.querySelector('main.magazine-cover'));
```

你的 `main` 元素中應該有一個 `header` 元素。

```js
assert.exists(document.querySelector("main > header"));
```

你的 `header` 應該有一個 `class`，值爲 `title`，幷包含一些文本。

```js
const header = document.querySelector("header.title");
assert.exists(header);
assert.isAbove(header.innerText.length, 0);
```

你應該爲每篇文章和封面圖片設置一個 `section` 元素, 類分別爲 `feature-article`、`small-article1`、`small-article2` 和 `cover-image`。

```js
const featureArticle = document.querySelector('section.feature-article');
const smallArticle1 = document.querySelector('section.small-article1');
const smallArticle2 = document.querySelector('section.small-article2');
const coverImage = document.querySelector('section.cover-image');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(coverImage);
```

`.magazine-cover` 的 `display` 屬性應該設置爲 `grid`。

```js
assert.equal(getComputedStyle(document.querySelector('.magazine-cover')).display, 'grid');
```

你的 `.magazine-cover` 應該有一個 `grid-template-areas` 屬性， 設置 `title` 佔據第一行的全部三列。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /^("|')title title title\1/);
```

網格模板的第二行應包含一個跨兩列的 `feature article` 和 一個 `cover image`。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /.("|')feature-article feature-article cover-image\1./);
```

網格模板的第三行應包含 `small-article1`、`small-article2` 和 `cover-image`。

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /("|')small-article1 small-article2 cover-image\1$/);
```

你的 `.magazine-cover` 選擇器應使用 `grid-template-columns` 屬性將空間劃分爲三個相等的部分。 請記住，你可以使用 `repeat` 函數或手動定義每個分數單位。

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

你的 `.magazine-cover` 類應將 `grid-template-rows` 屬性設置爲 `auto 1fr 1fr` 。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateRows, 'auto 1fr 1fr');
```

你應該在網格項之間添加 `10px` 的間隙。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gap, '10px');
```

`.title` 類應有一個同名的網格區域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` 類應有一個同名的網格區域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` 類應有一個同名的網格區域。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Magazine Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Adventure Magazine</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <main class="magazine-cover">
      <header class="title">
          <h1>Adventure Awaits</h1>
      </header>
      <section class="feature-article">
          <h2>Top 10 Exotic Destinations for 2024</h2>
          <p>Discover the most breathtaking places to visit this year, from hidden beaches to mountain retreats. Our guide takes you through the best spots for your next adventure.</p>
      </section>
      <section class="small-article1">
          <h3>Gear Up: Must-Have Gadgets</h3>
          <p>Check out the latest tech and gear to make your travels more exciting and comfortable.</p>
      </section>
      <section class="small-article2">
          <h3>Meet the Explorers</h3>
          <p>Get to know the modern adventurers who are pushing the boundaries of exploration.</p>
      </section>
      <section class="cover-image">
          <img src="https://cdn.freecodecamp.org/curriculum/labs/magazine-cover.png" alt="Adventure Cover Image">
      </section>
  </main>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.magazine-cover {
  display: grid;
  grid-template-areas: 
      "title title title"
      "feature-article feature-article cover-image"
      "small-article1 small-article2 cover-image";
  grid-template-columns: 1fr 1fr 1fr;
  grid-template-rows: auto 1fr 1fr;
  gap: 10px;
  width: 80%;
  background-color: #fff;
}

.title {
  grid-area: title;
  background-color: #1B1B32;
  color: white;
  padding: 10px;
  text-align: center;
}

.feature-article {
  grid-area: feature-article;
  background-color: #f9f9f9;
  padding: 10px;
}

.small-article1, .small-article2 {
  background-color: #e9e9e9;
  padding: 10px;
}

.cover-image {
  grid-area: cover-image;
  display: flex;
  justify-content: center;
  align-items: center;
  background-color: #ddd;
  padding: 10px;
}

.cover-image img {
  max-width: 100%;
  height: auto;
} 
```
