---
id: 67cd49e05693111a41771f29
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man Zustandsvariablen mit dem `useState`-Hook und der Destrukturierung von Arrays erstellt:

```js
const [user, setUser] = useState("Jessica");
```

Die Konvention zur Benennung von Zustandsvariablen ist die Verwendung des Formats `example` und `setExample`.

`useState` gibt ein Array mit zwei Werten zurück. Der erste Wert ist der aktuelle Zustand. Der zweite Wert ist die Set-Funktion, mit der der Zustand aktualisiert wird.

Erstellen Sie eine neue Zustandsvariable und eine Set-Funktion namens `isVisible` und `setIsVisible`. Denken Sie daran, die Syntax der Destrukturierung von Arrays wie im Beispiel gezeigt zu verwenden.

Weisen Sie dann dem `useState`-Hook einen Anfangswert von `false` zu.

# --hints--

Sie sollten die Syntax der Destrukturierung von Arrays für die Variablen `isVisible` und `setIsVisible` verwenden. Schauen Sie sich das Beispiel noch einmal an.

```js
assert.match(code, /(const|let|var)\s+\[\s*(isVisible)\s*,\s*(setIsVisible)\s*\]\s*/);
```

Sie sollten den `useState`-Hook verwenden.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.isAtLeast(abuseState.calls.length, 1);
```

Ihr `useState`-Hook sollte einen Anfangswert von `false` haben.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.equal(abuseState.calls[0]?.[0], false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  --fcc-editable-region--

  --fcc-editable-region--
  
  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      <p id="message">I love freeCodeCamp!</p>
    </div>
  );
};
```
