---
id: 67cd45eb92c40b176d987d10
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie innerhalb des `toggle-container`-Elements ein `button`-Element mit einem `id` namens `toggle-button` hinzu. Der Text des Buttons soll `Message` lauten.

# --hints--

Sie sollten ein `button`-Element haben.

```js
assert.exists(document.querySelector("button"));
```

Ihr `button`-Element sollte ein `id` namens `toggle-button` enthalten.

```js
const buttonEl = document.querySelector("button");
assert.exists(document.getElementById("toggle-button"))
```

Der Text Ihres Buttons sollte `Message` sein.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl.textContent, "Message");
```

Ihr `button`-Element sollte sich innerhalb des `toggle-container`-Elements befinden.

```js
const buttonEl = document.querySelector("#toggle-container > button");
assert.exists(buttonEl);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  return (
    <div id="toggle-container">
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  );
};
```
