---
id: 67cd525b0e575b1e5abf4c03
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Wie bereits erwähnt, ist die set-Funktion dafür verantwortlich, den Zustand zu aktualisieren. Ein gängiger Ansatz beim Umschalten zwischen Boolean-Zustandswerten besteht darin, den Zustand auf diese Weise zu aktualisieren:

```js
const [isUserLoggedIn, setIsUserLoggedIn] = useState(false);

// updating state
setIsUserLoggedIn(!isUserLoggedIn);
```

Im obigen Beispiel wird die set-Funktion mit dem Namen `setIsUserLoggedIn` verwendet, um den Anmeldestatus des Benutzers zu aktualisieren. Es ist üblich, den logischen NICHT-Operator (`!`) zu verwenden, um zwischen `true` und `false` umzuschalten. Da der Anfangswert `false` ist, ändert sich durch Anwenden des `!`-Operators der Wert auf `true`.

Verwenden Sie innerhalb der `handleToggleVisibility`-Funktion `setIsVisible`, um den Sichtbarkeits-State mit dem logischen NOT-Operator (`!`) zu aktualisieren.

# --hints--

Sie sollten `setIsVisible` verwenden, um die Sichtbarkeit des Texts zu aktualisieren. Schauen Sie sich das Beispiel an, wenn Sie Hilfe benötigen.

```js
assert.match(code, /setIsVisible\s*\(\s*!isVisible\s*\)/);  
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);
  
  const handleToggleVisibility = () => {
  --fcc-editable-region--
    
  --fcc-editable-region--
  };
  
  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
