---
id: 694bf6eb530e19a9c48c59f7
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Unaweza pia kutumia nambari hasi kuchukua sehemu ya mfuatano wa herufi kutoka mwisho. Kwa mfano, `-1` inawakilisha herufi ya mwisho, `-2` inawakilisha herufi ya pili kutoka mwisho, na kadhalika.

Ili kupata herufi tatu za mwisho za mfuatano wa herufi, unaweza kutumia `string[-3:]`. Angalia jinsi kigezo cha kuacha kinavyobakia kutojumuishwa baada ya nukta ya wima. Hii ina maana kwamba kuchukua sehemu ya mfuatano kunapaswa kuendelea hadi mwisho wa mfuatano.

Tengeneza kigezo kinachoitwa `last_three`. Tumia kielezo hasi kuchukua herufi tatu za mwisho kutoka `employee_code` (zinazo wakilisha nambari ya kitambulisho). Mwishowe, chapisha `last_three` kwenye terminali.

Kwa hayo, warsha ya kadi ya mfanyakazi imekamilika.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Unapaswa kuchukua herufi tatu za mwisho za `employee_code` kwa kutumia kielezo hasi na kuziweka kwenye kigezo chako cha `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Unapaswa kuchapisha `last_three` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
