---
id: 694bf6eb530e19a9c48c59f6
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Unaweza kata sehemu yoyote ya mfuatano wa herufi, si mwanzo tu. Na hii ni msaada katika hali nyingi.

Tengeneza kigezo `year_code` na weka thamani yake kama sehemu ya `employee_code` kuanzia kielezo `4` hadi `8`. Hii itatoa `2026`.

Kisha tengeneza kigezo `initials` na weka thamani yake kama sehemu ya `employee_code` kuanzia kielezo `9` hadi `11`. Hii itatoa `JD`.

Mwishowe, chapisha vigezo vyote viwili kwenye terminali.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Unapaswa kata `employee_code` kuanzia kielezo `4` hadi `8` na weka matokeo kwenye kigezo chako cha `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Unapaswa kuwa na kigezo kinachoitwa `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Unapaswa kata `employee_code` kuanzia kielezo `9` hadi `11` na weka matokeo kwenye kigezo chako cha `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Unapaswa kuchapisha `year_code` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Unapaswa kuchapisha `initials` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
