---
id: 694bf6ea530e19a9c48c59f0
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ikiwa unachanganya mfuatano wa herufi miwili kama `'John' + 'Doe'`, matokeo ni `'JohnDoe'` bila nafasi. Ili kurekebisha hili, unahitaji kuchanganya mfuatano wa herufi unao na nafasi (`' '`) kati yao.

Sahihisha mabadiliko yako ya `full_name` ili ichanganye `first_name`, nafasi, na `last_name`.

# --hints--

Unapaswa kuongeza nafasi (`' '`) kati ya `first_name` na `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
