---
id: 694bf6eb530e19a9c48c59f4
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Kuunganisha mfuatano wa herufi mingi kwa kutumia `+` na kubadilisha nambari kwa kutumia `str()` kunaweza kuwa ngumu na kusababisha mkanganyiko wa kusoma.

Python 3.6 ilizindua **f-strings** kutatua hili. Kwa kuongeza herufi `f` kabla ya nukuu ya kufungua, unaweza kuweka vigezo na usemi wa kihesabu ndani ya sehemu za uingizaji zinazowakilishwa na mabano ya mviringo `{}`. Kwa mfano:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Tengeneza kigezo `employee_card` na ukaweke thamani ya f-string inayonyesha `Employee:` ikifuatiwa na nafasi na thamani ya kigezo `full_name`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Kigezo chako cha `employee_card` kinapaswa kuwa na thamani `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
