---
id: 694bf3af05d10fd89dab05cc
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Katika Python, unaweza kuunganisha mfuatano wa herufi kwa kutumia opereta `+`. Hii huitwa concatenation. Hapa kuna mfano:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Tengeneza kigezo `full_name` kwa kuunganisha `first_name` na `last_name`. Kisha chapisha `full_name`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi ulioundwa kwa kuunganisha `first_name` na `last_name` katika kigezo chako cha `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Unapaswa kuchapisha kigezo cha `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
