---
id: 694bf6eb530e19a9c48c59f1
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa, anwani yako inaonekana haijakamilika. Pia unataka kuongeza nambari ya ghorofa ambapo mfanyakazi anaishi, kwa hivyo unapaswa kubadilisha mabadiliko.

Unapotaka kuongeza maudhui mwishoni mwa mfuatano wa herufi uliopo, unaweza kutumia kielezi cha **mgawo wa thamani ulioboreshwa** ,`+=`. Hii ni fupi zaidi kuliko kuandika `var = var + 'new text'`. Kwa mfano:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Kumbuka kuwa mfuatano wa herufi hauwezi kubadilika, kwa hivyo operesheni hii haibadilishi mfuatano wa herufi wa awali. Badala yake huunda mfuatano mpya wa herufi na kuuweka tena kwenye mabadiliko.

Tumia kielezi cha `+=` kuongeza mfuatano wa herufi `, Apartment 4B` kwenye mabadiliko yako ya `address`.

# --hints--

Unapaswa kutumia kielezi cha `+=` kuongeza mfuatano wa herufi `, Apartment 4B` kwenye mabadiliko ya `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
