---
id: 694bf6eb530e19a9c48c59f2
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Sasa kamilisha sentensi kwa kusahisha mgawo wa thamani wa `employee_info` ili pia uunganishe mfuatano wa herufi ` years old` mwishoni, ukitumia kiendeshaji cha `+`. Kumbuka kujumuisha nafasi mwanzoni mwa mfuatano wako wa herufi.

Mwishowe, chapisha `employee_info`.

# --hints--

Unapaswa kuunganisha `' years old'` mwishoni mwa kigezo `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

Unapaswa kuchapisha `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
