---
id: 694bf6eb530e19a9c48c59f3
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Sasa utatumia kitendakazi cha `str()` mara moja zaidi. Kama ilivyo kwa umri, lazima ubadilishe thamani yoyote ya nambari kuwa mfuatano wa herufi kabla ya kuunganisha na maandishi mengine.

Tengeneza kigezo kinachoitwa `experience_years` na weka thamani ya nambari kamili `5`.

Kisha, tengeneza kigezo `experience_info`. Weka thamani ya mfuatano wa herufi ulioundwa kwa kuunganisha `'Experience: '`, kigezo cha `experience_years` (ambacho kimebadilishwa kuwa mfuatano wa herufi), na `' years'`, kwa kutumia kiendeshaji cha `+`. Chapisha matokeo kwenye terminali.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Unapaswa kuweka thamani ya nambari kamili `5` kwenye kigezo chako cha `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

Unapaswa kuwa na kigezo kinachoitwa `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi unaopatikana kwa kuunganisha `'Experience: '`, `str(experience_years)`, na `' years'` kwenye kigezo chako cha `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Unapaswa kuchapisha `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
