---
id: 695c96f82af4fb78e59fa9ed
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Mfuatano wa herufi ni mfululizo wa herufi unaotumika kuhifadhi data ya maandishi. Kama unavyoweza kukumbuka kutoka kwa mafundisho ya awali, unaweza kuunda mfuatano wa herufi kwa kuweka maandishi ndani ya alama za nukuu moja (`'`) au alama za nukuu maradufu (`"`). Kwa mfano:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Unda kigezo `first_name` kinachohifadhi mfuatano wa herufi `John` na kigezo `last_name` kinachohifadhi mfuatano wa herufi `Doe`. Kisha chapisha `first_name` na `last_name`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `John` kwenye kigezo chako cha `first_name`. Kumbuka kuweka maandishi kati ya alama za nukuu moja au alama za nukuu maradufu.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

Unapaswa kuwa na kigezo kinachoitwa `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `Doe` kwenye kigezo chako cha `last_name`. Kumbuka kuweka maandishi kati ya alama za nukuu moja au alama za nukuu maradufu.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Unapaswa kuchapisha kigezo chako cha `first_name` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Unapaswa kuchapisha kigezo chako cha `last_name` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
