---
id: 695c9c7c61d446f1451e0ae7
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Ifuatayo, tengeneza mabadiliko `address` kuhifadhi anwani ya mfanyakazi. Weka thamani yake kuwa mfuatano wa herufi `123 Main Street`, na hatimaye chapisha `address`.

# --hints--

Unapaswa kuwa na mabadiliko yenye jina `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `123 Main Street` kwa mabadiliko yako ya `address`. Kumbuka kuweka maandishi kati ya nukuu moja au alama za nukuu maradufu.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Unapaswa kuchapisha mabadiliko ya `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
