---
id: 695c9f21a8b1d53cf1b97841
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa tengeneza kigezo kinachoitwa `employee_age` na weka thamani yake kuwa nambari kamili `28`. 


# --hints--

Unapaswa kuwa na kigezo kinachoitwa `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Unapaswa weka thamani ya nambari kamili `28` kwa kigezo chako cha `employee_age`. Usizungushe maandishi na nukuu.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
