---
id: 695df6480de23758af5a29cd
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Sasa, unataka kuunda mfuatano wa herufi unaoonyesha umri wa mfanyakazi.

Anza kwa kuunda kigezo `employee_info` na ukaweke thamani ya matokeo ya kuunganisha:

- kigezo `full_name`.
- mfuatano wa herufi unaojumuisha herufi `is` uliotanguliwa na kufuatwa na nafasi.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Unapaswa kuweka thamani ya `full_name + ' is '` kwa kigezo chako cha `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
