---
id: 695ea1e55c0e619e597fc5c9
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Kama unavyoona Python ilitoa `TypeError: can only concatenate str (not "int") to str`. Hii hutokea kwa sababu Python haikuruhusu kuunganisha maandishi (mfuatano wa herufi) na nambari (nambari kamili) moja kwa moja.

Ili kurekebisha hili, lazima ubadilishe nambari kuwa mfuatano wa herufi kwanza kwa kutumia kitendakazi cha `str()`, ambacho hurudisha toleo la mfuatano wa herufi la kitu:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Sahihisha mgawo wa `employee_info` kwa kubadilisha `employee_age` kuwa mfuatano wa herufi kwa kutumia `str(employee_age)`.

# --hints--

Unapaswa kubadilisha mgawo wa thamani wa kigezo chako cha `employee_info` kwa kubadilisha `employee_age` kuwa mfuatano wa herufi kwa kutumia `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
