---
id: 695fd3dde03273875e5d6709
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Sahihisha mgawo wa thamani wa `employee_info` ili pia uunganishishe `employee_age` mwishoni, ukitumia kiendeshaji cha `+`.

Mara utakapofanya hivyo, utaona `TypeError` kwenye terminal. Katika hatua inayofuata, utafanya kazi ya kuirekebisha.

# --hints--

Unapaswa kusahihisha mgawo wa thamani wa `employee_info` ili kuunganisha `employee_age` mwishoni.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
