---
id: 69727ea5c9606bedfe36d3d2
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Sasa ni wakati wa kuongeza maelezo ya mwisho kwenye kadi.

Tengeneza kiwakilishi kinachoitwa `position` chenye thamani ya mfuatano wa herufi `Data Analyst` na kiwakilishi kinachoitwa `salary` chenye thamani ya nambari kamili `75000`.

Kisha, sasisha f-string yako ya `employee_card` ili ijumuishe nafasi na mshahara. Inapaswa kufuata muundo huu halisi: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Badilisha viwakilishi na vigezo vinavyolingana.

Mwishowe, chapisha `employee_card` kuona matokeo. 

# --hints--

Unapaswa kuwa na kiwakilishi kinachoitwa `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi `Data Analyst` kwenye kiwakilishi chako cha `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

Unapaswa kuwa na kiwakilishi kinachoitwa `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Unapaswa kuweka thamani ya nambari kamili `75000` kwenye kiwakilishi chako cha `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Unapaswa kusasisha `employee_card` kutumia f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Unapaswa kuchapisha `employee_card` kwenye terminali.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
