---
id: 6980d3ecfe800d8a54130f30
title: Hatua 14
challengeType: 20
dashedName: step-14
---

# --description--

Kwa sasa, `employee_card` inaonyesha tu jina la mfanyakazi. Sasa utaongeza taarifa zaidi ndani yake.

Sahihisha mgawo wa thamani wa `employee_card` ili ujumuishe umri wa mfanyakazi. Mfuatano wa herufi wa mwisho unapaswa kuonekana hivi: `Employee: [name] | Age: [age]` ambapo `[name]` imebadilishwa na jina la mfanyakazi, na `[age]` imebadilishwa na umri wa mfanyakazi.

# --hints--

Unapaswa kusahihisha `employee_card` ili itumie f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
