---
id: a0b5010f579e69b815e7c5d6
title: Criar uma Função Inteligente de Substituição de Palavras
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

Neste laboratório, você criará uma função que realiza uma busca e substituição na sentença usando os argumentos fornecidos e então retorna a nova sentença.

**Nota:** Preserve a capitalização do primeiro caractere na palavra original ao substituí-la. Por exemplo, se você quiser substituir a palavra `Book` com a palavra `dog`, deve ser substituído com `Dog`

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `myReplace`.
1. A função `myReplace` deve receber três argumentos: uma string, uma palavra a ser substituída e a palavra para substituí-la.
1. A função `myReplace` deve retornar a nova string com a substituição.
1. Você deve preservar a capitalização do primeiro caractere da palavra original quando estiver substituindo-a.

# --hints--

Você deve criar uma função chamada `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` deve receber três argumentos.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` deve retornar a string `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` deve retornar a string `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` deve retornar a string `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` deve retornar a string `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` deve retornar a string `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` deve retornar a string `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
