---
id: 587d8259367417b2b2512c86
title: реалізувати сортування вставками
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм сортування вставками. Цей метод працює, поступово формуючи відсортований масив на початку списку. Він починає відсортований масив із першого елемента. Потім перевіряє наступний елемент і переставляє його назад у відсортовану частину, доки він не опиниться на правильній позиції. Процес повторюється для всього списку, і нові елементи переставляються назад у відсортовану частину, доки не буде досягнуто кінця. Цей алгоритм має квадратичну часову складність у середньому та гіршому випадках.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію `insertionSort`.
1. Функція `insertionSort` повинна приймати масив цілих чисел і повертати масив із тими ж числами, відсортованими від найменшого до найбільшого.

# --hints--

`insertionSort` має бути функцією.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` має повертати масив, який не змінений, окрім порядку елементів.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` має повертати відсортований масив (від найменшого до найбільшого).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` має повертати `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
