---
id: 6814d8e1516e86b171929de4
title: "Desafio 1: Equilíbrio de Vogais"
challengeType: 28
dashedName: challenge-1
---

# --description--

Dada uma string, determine se o número de vogais na primeira metade da string é igual ao número de vogais na segunda metade.

- A string pode conter quaisquer caracteres.
- As letras `a`, `e`, `i`, `o` e `u`, tanto maiúsculas quanto minúsculas, são consideradas vogais.
- Se houver um número ímpar de caracteres na string, ignore o caractere central.

# --hints--

`isBalanced("racecar")` deve retornar `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` deve retornar `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` deve retornar `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` deve retornar `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` deve retornar `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` deve retornar `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` deve retornar `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
