---
id: 681cb1afdab50c87ddb2e517
title: "Desafio 6: Verificador de Anagramas"
challengeType: 28
dashedName: challenge-6
---

# --description--

Dadas duas strings, determine se elas são anagramas uma da outra (contêm os mesmos caracteres em qualquer ordem).

- Ignore maiúsculas, minúsculas e espaços em branco.

# --hints--

`areAnagrams("listen", "silent")` deve retornar `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` deve retornar `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` deve retornar `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` deve retornar `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` deve retornar `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` deve retornar `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
