---
id: 681cb1b0dab50c87ddb2e51b
title: "Desafio 10: 3 Strikes"
challengeType: 28
dashedName: challenge-10
---

# --description--

Dado um inteiro entre 1 e 10.000, retorne a contagem de quantos números de 1 até esse inteiro cujo quadrado contém pelo menos um dígito 3.

# --hints--

`squaresWithThree(1)` deve retornar `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` deve retornar `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` deve retornar `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` deve retornar `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` deve retornar `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
