---
id: 681cb1afdab50c87ddb2e516
title: "Desafio 5: Texto Jbelmud"
challengeType: 28
dashedName: challenge-5
---

# --description--

Dada uma string, retorne uma versão embaralhada dessa string onde cada palavra é transformada usando as seguintes restrições:

- As primeiras e últimas letras das palavras permanecem no lugar
- Todas as letras entre a primeira e a última letra são ordenadas alfabeticamente.
- as strings de entrada não conterão pontuação e serão totalmente em minúsculas

# --hints--

`jbelmu("hello world")` deve retornar `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` deve retornar `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` deve retornar `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` deve retornar `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
