---
id: 681cb1afdab50c87ddb2e515
title: "Desafio 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

Dada uma string, remova todos os espaços da string, insira dois espaços entre cada caractere, converta todas as letras alfabéticas para maiúsculas e retorne o resultado.

- Caracteres não alfabéticos devem permanecer inalterados (exceto espaços).

# --hints--

`spaceJam("freeCodeCamp")` deve retornar `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` deve retornar `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Olá Mundo?!")` deve retornar `"O  L  Á  M  U  N  D  O  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` deve retornar `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` deve retornar `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
