---
id: 681cb1b0dab50c87ddb2e519
title: "Desafio 8: Factorializer"
challengeType: 28
dashedName: challenge-8
---

# --description--

Dado um inteiro de zero a 20, retorne o fatorial desse número. O fatorial de um número é o produto de todos os números entre 1 e o número dado.

- O fatorial de zero é 1.

# --hints--

`factorial(0)` deve retornar `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` deve retornar `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` deve retornar `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
