---
id: 681cb1b0dab50c87ddb2e518
title: "Desafio 7: Soma Alvo"
challengeType: 28
dashedName: challenge-7
---

# --description--

Dado um array de números e um inteiro target, encontre dois números únicos no array que somem o valor target. Retorne um array com os índices desses dois números ou `"Target not found"` se nenhum par de números somar o target.

- O array retornado deve ter os índices em ordem crescente.

# --hints--

`findTarget([2, 7, 11, 15], 9)` deve retornar `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` deve retornar `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` deve retornar `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` deve retornar `"Alvo não encontrado"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
