---
id: 681cb1a2dab50c87ddb2e514
title: "Desafio 3: Sequência de Fibonacci"
challengeType: 28
dashedName: challenge-3
---

# --description--

A sequência de Fibonacci é uma série de números onde cada número é a soma dos dois anteriores. Ao começar com `0` e `1`, os primeiros 10 números da sequência são `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21` e `34`.

Dado um array contendo os dois primeiros números de uma sequência de Fibonacci e um inteiro representando o comprimento da sequência, retorne um array contendo a sequência do comprimento dado.

- Sua função deve lidar com sequências de qualquer comprimento maior ou igual a zero.
- Se o comprimento for zero, retorne um array vazio.
- Observe que os números iniciais fazem parte da sequência.

# --hints--

`fibonacciSequence([0, 1], 20)` deve retornar `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` deve retornar `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` deve retornar `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` deve retornar `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` deve retornar `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
