---
id: 681cb1b0dab50c87ddb2e51a
title: "Desafio 9: Soma dos Quadrados"
challengeType: 28
dashedName: challenge-9
---

# --description--

Dado um número inteiro positivo até 1.000, retorne a soma de todos os inteiros ao quadrado de 1 até o número.

# --hints--

`sumOfSquares(5)` deve retornar `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` deve retornar `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` deve retornar `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` deve retornar `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` deve retornar `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
