---
id: 68216ef80f957572e7c340c5
title: "Desafio 12: Decodificador de Mensagens"
challengeType: 28
dashedName: challenge-12
---

# --description--

Dada uma string de mensagem secreta e um inteiro representando o número de letras que foram usadas para deslocar a mensagem para codificá-la, retorne a string decodificada.

- Um número positivo significa que a mensagem foi deslocada para frente no alfabeto.
- Um número negativo significa que a mensagem foi deslocada para trás no alfabeto.
- Maiúsculas e minúsculas importam, caracteres decodificados devem manter a capitalização de seus equivalentes codificados.
- Caracteres não alfabéticos não devem ser decodificados.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` deve retornar `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` deve retornar `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` deve retornar `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` deve retornar `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
