---
id: 6821ebce237de8297eaee790
title: "Desafio 14: Batalha de Personagens"
challengeType: 28
dashedName: challenge-14
---

# --description--

Dadas duas strings representando seu exército e um exército oposto, cada caractere do seu exército batalha contra o caractere na mesma posição do exército oposto usando as seguintes regras:

- Os caracteres `a-z` têm uma força de `1-26`, respectivamente.
- Os caracteres `A-Z` têm uma força de `27-52`, respectivamente.
- Dígitos `0-9` têm uma força igual ao seu valor nominal.
- Todos os outros caracteres têm valor zero.
- Cada personagem pode lutar em apenas uma batalha.

Para cada batalha, o personagem mais forte vence. O exército com mais vitórias vence a guerra. Retorne os seguintes valores:

- "`Opponent retreated` se seu exército tiver mais personagens do que o exército oposto."
- "Recuamos" se o exército adversário tiver mais personagens do que o seu.
- "Ganharam" se seu exército venceu mais batalhas.
- "We lost" se o exército adversário venceu mais batalhas.
- "Foi um empate" se ambos os exércitos venceram o mesmo número de batalhas.

# --hints--

`battle("Hello", "World")` deve retornar `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` deve retornar `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` deve retornar `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` deve retornar `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` deve retornar `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` deve retornar `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` deve retornar `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
