---
id: 6821ebea237de8297eaee795
title: "Desafio 19: Candlelight"
challengeType: 28
dashedName: challenge-19
---

# --description--

Dado um inteiro representando o número de velas com que você começa, e um inteiro representando quantas velas queimadas são necessárias para criar uma nova, retorne o número de velas que você terá usado após criar e queimar o máximo que puder.

Por exemplo, se forem dadas 7 velas e forem necessárias 2 velas queimadas para fazer uma nova:

1. Queime 7 velas para obter 7 sobras,
2. Recicle 6 sobras em 3 novas velas (1 sobra permanece),
3. Queime 3 velas para obter 3 sobras a mais (4 no total),
4. Recicle 4 sobras em 2 novas velas,
5. Queime 2 velas para obter 2 sobras,
6. Recicle 2 sobras em 1 nova vela,
7. Queime 1 vela.

Você terá queimado 13 velas no total no exemplo.

# --hints--

`burnCandles(7, 2)` deve retornar `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` deve retornar `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` deve retornar `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` deve retornar `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` deve retornar `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
