---
id: 6821ebf8237de8297eaee798
title: "Desafio 22: Sequência Tribonacci"
challengeType: 28
dashedName: challenge-22
---

# --description--

A sequência Tribonacci é uma série de números onde cada número é a soma dos três anteriores. Quando iniciada com `0`, `0` e `1`, os primeiros 10 números da sequência são `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24` e `44`.

Dada uma matriz contendo os três primeiros números de uma sequência Tribonacci e um inteiro representando o comprimento da sequência, retorne uma matriz contendo a sequência do comprimento dado.

- Sua função deve lidar com sequências de qualquer comprimento maior ou igual a zero.
- Se o comprimento for zero, retorne um array vazio.
- Observe que os números iniciais fazem parte da sequência.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` deve retornar `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` deve retornar `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` deve retornar `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` deve retornar `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` deve retornar `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` deve retornar `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
