---
id: 6821ebda237de8297eaee792
title: "Desafio 16: Parênteses Invertidos"
challengeType: 28
dashedName: challenge-16
---

# --description--

Dada uma string que contém parênteses devidamente aninhados, retorne a versão decodificada da string usando as seguintes regras:

- Todos os caracteres dentro de cada parêntese devem ser invertidos.
- Os parênteses devem ser removidos do resultado final.
- Se parênteses estiverem aninhados, o par mais interno deve ser invertido primeiro e então seu resultado deve ser incluído na inversão do par externo.
- Assuma que todos os parênteses estão equilibrados e corretamente aninhados.

# --hints--

`decode("(f(b(dc)e)a)")` deve retornar `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` deve retornar `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` deve retornar `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
