---
id: 6821ebe4237de8297eaee794
title: "Desafio 18: Segundo Melhor"
challengeType: 28
dashedName: challenge-18
---

# --description--

Dado um array de inteiros representando o preço de diferentes laptops e um inteiro representando seu orçamento, retorne:

1. O segundo laptop mais caro se estiver dentro do seu orçamento, ou
2. O laptop mais caro que está dentro do seu orçamento, ou
3. `0` se nenhum laptop estiver dentro do seu orçamento.

- Preços duplicados devem ser ignorados.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` deve retornar `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` deve retornar `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` deve retornar `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` deve retornar `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` deve retornar `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
