---
id: 68216eb60f957572e7c340c4
title: "Desafio 11: Ritmo por Milha"
challengeType: 28
dashedName: challenge-11
---

# --description--

Dado um número de milhas percorridas e um tempo em `"MM:SS"` (minutos:segundos) que levou para correr essas milhas, retorne uma string com o tempo médio que levou para correr cada milha no formato `"MM:SS"`.

- Adicione zeros à esquerda quando necessário.

# --hints--

`milePace(3, "24:00")` deve retornar `"08:00"`.

```js
assert.equal(milePace(3, "24:00"), "08:00");
```

`milePace(1, "06:45")` deve retornar `"06:45"`.

```js
assert.equal(milePace(1, "06:45"), "06:45");
```

`milePace(2, "07:00")` deve retornar `"03:30"`.

```js
assert.equal(milePace(2, "07:00"), "03:30");
```

`milePace(26.2, "120:35")` deve retornar `"04:36"`.

```js
assert.equal(milePace(26.2, "120:35"), "04:36");
```

# --seed--

## --seed-contents--

```js
function milePace(miles, duration) {

  return miles;
}
```

# --solutions--

```js
function milePace(miles, duration) {
  const [minutes, seconds] = duration.split(':').map(Number);
  const totalSeconds = minutes * 60 + seconds;
  const avgSecondsPerMile = totalSeconds / miles;

  const avgMinutes = Math.floor(avgSecondsPerMile / 60);
  const avgSeconds = Math.round(avgSecondsPerMile % 60);

  const paddedMinutes = avgMinutes.toString().padStart(2, '0');
  const paddedSeconds = avgSeconds.toString().padStart(2, '0');

  return `${paddedMinutes}:${paddedSeconds}`;
}
```
