---
id: 6821ec02237de8297eaee79a
title: "Desafio 24: Pangrama"
challengeType: 28
dashedName: challenge-24
---

# --description--

Dada uma palavra ou frase e uma string de letras minúsculas, determine se a palavra ou frase usa todas as letras do conjunto fornecido pelo menos uma vez e nenhuma outra letra.

- Ignore caracteres não alfabéticos na palavra ou frase.
- Ignore a diferenciação entre maiúsculas e minúsculas na palavra ou frase.

# --hints--

`isPangram("hello", "helo")` deve retornar `true`

```js
assert.isTrue(isPangram("hello", "helo"));
```

`isPangram("hello", "hel")` deve retornar `false`

```js
assert.isFalse(isPangram("hello", "hel"));
```

`isPangram("hello", "helow")` deve retornar `false`

```js
assert.isFalse(isPangram("hello", "helow"));
```

`isPangram("hello world", "helowrd")` deve retornar `true`

```js
assert.isTrue(isPangram("hello world", "helowrd"));
```

`isPangram("Hello World!", "helowrd")` deve retornar `true`

```js
assert.isTrue(isPangram("Hello World!", "helowrd"));
```

`isPangram("Hello World!", "heliowrd")` deve retornar `false`

```js
assert.isFalse(isPangram("Hello World!", "heliowrd"));
```

`isPangram("freeCodeCamp", "frcdmp")` deve retornar `false`

```js
assert.isFalse(isPangram("freeCodeCamp", "frcdmp"));
```

`isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` deve retornar `true`

```js
assert.isTrue(isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"));
```

# --seed--

## --seed-contents--

```js
function isPangram(sentence, letters) {

  return sentence;
}
```

# --solutions--

```js
function isPangram(sentence, letters) {
  const usedLetters = [];
  for (let i = 0; i < sentence.length; i++) {
    const letter = sentence[i].toLowerCase();
    if (!usedLetters.includes(letter) && /[a-z]/.test(letter)) {
      usedLetters.push(letter);
    }
  }

  const sortedLetters = letters.split('').sort().join('');
  const sortedUsedLetters = usedLetters.sort().join('');

  return sortedLetters === sortedUsedLetters;
}
```
