---
id: 6821ebd4237de8297eaee791
title: "Desafio 15: camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

Dada uma string, retorne sua versão em camel case usando as seguintes regras:

- As palavras no argumento da string são separadas por um ou mais caracteres do seguinte conjunto: espaço (` `), hífen (`-`) ou sublinhado (`_`). Considere qualquer sequência desses como uma quebra de palavra.
- A primeira palavra deve estar toda em minúsculas.
- Cada palavra subsequente deve começar com uma letra maiúscula, com o restante em minúsculas.
- Todos os espaços e separadores devem ser removidos.

# --hints--

`toCamelCase("hello world")` deve retornar `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` deve retornar `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("agente secreto-X")` deve retornar `"agenteSecretoX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` deve retornar `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` deve retornar `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
