---
id: 6821ebc9237de8297eaee78f
title: "Desafio 13: Primo Não Natural"
challengeType: 28
dashedName: challenge-13
---

# --description--

Dado um inteiro, determine se esse número é um número primo ou um número primo negativo.

- Um número primo é um inteiro positivo maior que 1 que é divisível apenas por 1 e ele mesmo.
- Um número primo negativo é a versão negativa de um número primo positivo.
- `1` e `0` não são considerados números primos.

# --hints--

`isUnnaturalPrime(1)` deve retornar `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` deve retornar `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` deve retornar `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` deve retornar `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` deve retornar `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` deve retornar `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` deve retornar `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` deve retornar `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` deve retornar `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
