---
id: 6821ebee237de8297eaee796
title: "Desafio 20: Duplicatas em Array"
challengeType: 28
dashedName: challenge-20
---

# --description--

Dado um array de inteiros, retorne um array de inteiros que aparecem mais de uma vez no array inicial, ordenados em ordem crescente. Se nenhum valor aparecer mais de uma vez, retorne um array vazio.

- Inclua apenas uma instância de cada valor no array retornado.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` deve retornar `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` deve retornar `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` deve retornar `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
