---
id: 6821ebfd237de8297eaee799
title: "Desafio 23: RGB para Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

Dada uma string de cor CSS `rgb(r, g, b)`, retorne seu equivalente hexadecimal.

Aqui estão alguns exemplos de saídas para uma dada entrada:

| Input   | Output   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Deixe quaisquer letras em minúsculas.
- Retorna um `#` seguido de seis caracteres. Não use valores abreviados.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` deve retornar `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` deve retornar `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` deve retornar `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` deve retornar `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
