---
id: 6821ebdf237de8297eaee793
title: "Desafio 17: Ordem das Operações"
challengeType: 28
dashedName: challenge-17
---

# --description--

Dado um array de inteiros e um array de operadores string, aplique as operações aos números sequencialmente da esquerda para a direita. Repita as operações conforme necessário até que todos os números sejam usados. Retorne o resultado final.

Por exemplo, dado `[1, 2, 3, 4, 5]` e `['+', '*']`, retorne o resultado de avaliar `1 + 2 * 3 + 4 * 5` da esquerda para a direita ignorando a ordem padrão das operações.

- Operadores válidos são `+`, `-`, `*`, `/` e `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` deve retornar `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` deve retornar `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` deve retornar `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` deve retornar `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` deve retornar `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
