---
id: 68adce01c0e1144d0a90295c
title: "Desafio 28: Parser de Numerais Romanos"
challengeType: 28
dashedName: challenge-28
---

# --description--

Dada uma string que representa um numeral romano, retorne seu valor inteiro.

Os numerais romanos consistem nos seguintes símbolos e valores:

| Symbol | Value |
|--------|-------|
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

- Os numerais são lidos da esquerda para a direita. Se um numeral menor aparecer antes de um maior, o valor é subtraído. Caso contrário, os valores são somados.

# --hints--

`parseRomanNumeral("III")` deve retornar `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` deve retornar `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` deve retornar `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` deve retornar `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` deve retornar `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` deve retornar `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` deve retornar `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
