---
id: 68adce01c0e1144d0a90295e
title: "Desafio 29: Construtor de Acrônimos"
challengeType: 28
dashedName: challenge-29
---

# --description--

Dada uma string contendo uma ou mais palavras, retorne um acrônimo das palavras usando as seguintes restrições:

- A sigla deve consistir na primeira letra de cada palavra em maiúscula, a menos que indicado de outra forma.
- A sigla deve ignorar a primeira letra dessas palavras a menos que sejam a primeira palavra da string dada: `a`, `for`, `an`, `and`, `by` e `of`.
- As letras do acrônimo devem ser retornadas na ordem em que são fornecidas.
- O acrônimo não deve conter espaços.

# --hints--

`buildAcronym("Otimização para Motores de Busca")` deve retornar `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Perguntas Frequentes")` deve retornar `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("Administração Nacional de Aeronáutica e Espaço")` deve retornar `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` deve retornar `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("Para sua informação")` deve retornar `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("A propósito")` deve retornar `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("Uma manada imparável de pinguins cambaleantes ultrapassa as montanhas geladas e canta feliz")` deve retornar `"UMIDPCUMGECF"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
