---
id: 68adce01c0e1144d0a902956
title: "Desafio 25: Repetidor de Vogais"
challengeType: 28
dashedName: challenge-25
---

# --description--

Dada uma string, retorne uma nova versão da string onde cada vogal é duplicada uma vez a mais do que a vogal anterior que você encontrou. Por exemplo, a primeira vogal na frase deve permanecer inalterada. A segunda vogal deve aparecer duas vezes seguidas. A terceira vogal deve aparecer três vezes seguidas e assim por diante.

- As letras `a`, `e`, `i`, `o` e `u`, tanto maiúsculas quanto minúsculas, são consideradas vogais.
- A vogal original deve manter seu caso.
- Vogais repetidas devem estar em minúsculas.
- Todos os caracteres não vogais devem manter sua caixa original.

# --hints--

`repeatVowels("hello world")` deve retornar `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` deve retornar `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` deve retornar `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("Eu gosto de comer sorvete na Islândia")` deve retornar `"Eu goostooo deee coomeeer soooorveeeete na Isllââââândia"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
