---
id: 68adce01c0e1144d0a90295a
title: "Desafio 27: Matrix Rotate"
challengeType: 28
dashedName: challenge-27
---

# --description--

Dada uma matriz (um array de arrays), rotacione a matriz 90 graus no sentido horário e retorne-a. Por exemplo, dada `[[1, 2], [3, 4]]`, que se parece com isto:

| 1 | 2 |
|---|---|
| 3 | 4 |

Você deve retornar `[[3, 1], [4, 2]]`, que se parece com isto:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

`rotate([[1]])` deve retornar `[[1]]`.

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` deve retornar `[[3, 1], [4, 2]]`.

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` deve retornar `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` deve retornar `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
