---
id: 68adce01c0e1144d0a902958
title: "Desafio 26: Validador IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

Dada uma string, determine se ela é um endereço IPv4 válido. Um endereço IPv4 válido consiste em quatro números inteiros separados por pontos (`.`). Cada número deve satisfazer as seguintes condições:

- Está entre 0 e 255 inclusive.
- Não possui zeros à esquerda (por exemplo, 0 é permitido, 01 não é).
- Apenas caracteres numéricos são permitidos.

# --hints--

`isValidIPv4("192.168.1.1")` deve retornar `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` deve retornar `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` deve retornar `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` deve retornar `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` deve retornar `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` deve retornar `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` deve retornar `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
