---
id: 68af0687ef34c76c28ffa54d
title: "Desafio 33: Tempo de Tela"
challengeType: 28
dashedName: challenge-33
---

# --description--

Dado um array de entrada com sete inteiros, representando o tempo de uma semana, onde cada inteiro é a quantidade de horas gastas no seu telefone naquele dia, determine se é tempo de tela demais com base nestas restrições:

- Se qualquer dia tiver 10 horas ou mais, é demais.
- Se a média de quaisquer três dias consecutivos for maior ou igual a 8 horas, é demais.
- Se a média dos sete dias for maior ou igual a 6 horas, é tempo demais.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` deve retornar `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` deve retornar `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` deve retornar `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` deve retornar `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` deve retornar `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` deve retornar `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` deve retornar `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
