---
id: 68af0687ef34c76c28ffa54f
title: "Desafio 34: Números Faltando"
challengeType: 28
dashedName: challenge-34
---

# --description--

Dado um array de inteiros de 1 até `n`, inclusive, retorne um array com todos os inteiros faltantes entre 1 e `n` (onde `n` é o maior número no array fornecido).

- O array fornecido pode estar desordenado e pode conter duplicatas.
- O array retornado deve estar em ordem crescente.
- Se nenhum inteiro estiver faltando, retorne um array vazio.

# --hints--

`findMissingNumbers([1, 3, 5])` deve retornar `[2, 4]`.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` deve retornar `[]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` deve retornar `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` deve retornar `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` deve retornar `[2, 6, 7, 8]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` deve retornar `[11]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```
