---
id: 68af0687ef34c76c28ffa54b
title: "Desafio 32: Inverter Sentença"
challengeType: 28
dashedName: challenge-32
---

# --description--

Dada uma string de palavras, retorne uma nova string com as palavras em ordem inversa. Por exemplo, a primeira palavra deve estar no final da string retornada e a última palavra deve estar no início da string retornada.

- Na string fornecida, as palavras podem ser separadas por um ou mais espaços.
- A string retornada deve ter apenas um espaço entre as palavras.

# --hints--

`reverseSentence("world hello")` deve retornar `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` deve retornar `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` deve retornar `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` deve retornar `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
