---
id: 68af0687ef34c76c28ffa547
title: "Desafio 30: Caracteres Únicos"
challengeType: 28
dashedName: challenge-30
---

# --description--

Dada uma string, determine se todos os caracteres na string são únicos.

- Letras maiúsculas e minúsculas devem ser consideradas caracteres diferentes.

# --hints--

`allUnique("abc")` deve retornar `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` deve retornar `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` deve retornar `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` deve retornar `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` deve retornar `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` deve retornar `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` deve retornar `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
