---
id: 68b06e589bf2273243814771
title: "Desafio 36: Ajustador de Termostato"
challengeType: 28
dashedName: challenge-36
---

# --description--

Dada a temperatura atual de uma sala e uma temperatura alvo, retorne uma string indicando como ajustar a temperatura da sala com base nessas restrições:

- Retorne `"heat"` se a temperatura atual estiver abaixo do alvo.
- Retorne `"cool"` se a temperatura atual estiver acima do alvo.
- Retorne `"hold"` se a temperatura atual for igual à temperatura alvo.

# --hints--

`adjustThermostat(68, 72)` deve retornar `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` deve retornar `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` deve retornar `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` deve retornar `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` deve retornar `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` deve retornar `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
