---
id: 68b06e589bf227324381476f
title: "Desafio 35: Frequência de Palavras"
challengeType: 28
dashedName: challenge-35
---

# --description--

Dado um parágrafo, retorne um array com as três palavras que ocorrem com mais frequência.

- Palavras no parágrafo serão separadas por espaços.
- Ignore maiúsculas e minúsculas no parágrafo dado. Por exemplo, trate `Hello` e `hello` como a mesma palavra.
- Ignore a pontuação no parágrafo fornecido. A pontuação consiste em vírgulas (`,`), pontos (`.`) e pontos de exclamação (`!`).
- O array retornado deve conter todas as palavras em letras minúsculas.
- O array retornado deve estar em ordem decrescente com a palavra que ocorre com mais frequência primeiro.

# --hints--

`getWords("Programar em Python é divertido porque programar em Python permite programar em Python facilmente enquanto programa")` deve retornar `["programar", "python", "em"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` deve retornar `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` deve retornar `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
