---
id: 68b06e589bf2273243814775
title: "Desafio 38: Gerador de Slug"
challengeType: 28
dashedName: challenge-38
---

# --description--

Dada uma string, retorne uma versão amigável para URL da string usando as seguintes restrições:

- Todas as letras devem estar em minúsculas.
- Todos os caracteres que não são letras, números ou espaços devem ser removidos.
- Todos os espaços devem ser substituídos pelo código de espaço codificado para URL `%20`.
- Espaços consecutivos devem ser substituídos por um único `%20`.
- A string retornada não deve ter `%20` no início ou no final.

# --hints--

`generateSlug("helloWorld")` deve retornar `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` deve retornar `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` deve retornar `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` deve retornar `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` deve retornar `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
