---
id: 68b06e589bf2273243814777
title: "Desafio 39: Encha o Tanque"
challengeType: 28
dashedName: challenge-39
---

# --description--

Dado o tamanho de um tanque de combustível, o nível atual de combustível e o preço por galão, retorne o custo para encher o tanque completamente.

- `tankSize` é a capacidade total do tanque em galões.
- `fuelLevel` é a quantidade atual de combustível no tanque em galões.
- `pricePerGallon` é o custo de um galão de combustível.
- O valor retornado deve ser arredondado para duas casas decimais no formato: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` deve retornar `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` deve retornar `"$17,50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` deve retornar `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` deve retornar `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` deve retornar `"$21,89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
