---
id: 68b1f72371a5ac895ac70a0a
title: "Desafio 44: Espelho de String"
challengeType: 28
dashedName: challenge-44
---

# --description--

Dadas duas strings, determine se a segunda string é um espelho da primeira.

- Uma string é considerada um espelho se contiver as mesmas letras na ordem inversa.
- Trate letras maiúsculas e minúsculas como distintas.
- Ignore todos os caracteres não alfabéticos.

# --hints--

`isMirror("helloworld", "helloworld")` deve retornar `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` deve retornar `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` deve retornar `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` deve retornar `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Espelho", "rorrim")` deve retornar `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` deve retornar `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` deve retornar `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
