---
id: 68b1f72371a5ac895ac70a08
title: "Desafio 43: Dígitos vs Letras"
challengeType: 28
dashedName: challenge-43
---

# --description--

Dada uma string, retorne `"digits"` se a string tiver mais dígitos do que letras, `"letters"` se tiver mais letras do que dígitos e `"tie"` se tiver a mesma quantidade de dígitos e letras.

- Dígitos consistem em `0-9`.
- Letras consistem em `a-z` em maiúsculas ou minúsculas.
- Ignore quaisquer outros caracteres.

# --hints--

`digitsOrLetters("abc123")` deve retornar `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` deve retornar `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` deve retornar `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` deve retornar `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` deve retornar `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` deve retornar `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
