---
id: 68b1f72371a5ac895ac70a04
title: "Desafio 41: Armazenamento de Arquivos"
challengeType: 28
dashedName: challenge-41
---

# --description--

Dado um tamanho de arquivo, uma unidade para o tamanho do arquivo e a capacidade do disco rígido em gigabytes (GB), retorne o número de arquivos que o disco rígido pode armazenar usando as seguintes restrições:

- A unidade para o tamanho do arquivo pode ser bytes (`"B"`), kilobytes (`"KB"`) ou megabytes (`"MB"`).
- Retorna o número de arquivos inteiros que o drive pode armazenar.
- Use as seguintes conversões:

| Unidade | Equivalente |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

Por exemplo, dado `500`, `"KB"` e `1` como argumentos, determine quantos arquivos de 500 KB podem caber em um disco rígido de 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` deve retornar `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` deve retornar `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` deve retornar `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` deve retornar `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` deve retornar `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` deve retornar `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
