---
id: 68b1f72371a5ac895ac70a02
title: "Desafio 40: Armazenamento de Fotos"
challengeType: 28
dashedName: challenge-40
---

# --description--

Dado um tamanho de foto em megabytes (MB) e a capacidade do disco rígido em gigabytes (GB), retorne o número de fotos que o disco rígido pode armazenar usando as seguintes restrições:

- 1 gigabyte equivale a 1000 megabytes.
- Retorna o número de fotos inteiras que o drive pode armazenar.

# --hints--

`numberOfPhotos(1, 1)` deve retornar `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` deve retornar `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` deve retornar `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` deve retornar `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` deve retornar `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
