---
id: 68b7687dded630607aceccaf
title: "Desafio 47: Pegando Excesso de Velocidade"
challengeType: 28
dashedName: challenge-47
---

# --description--

Dado um array de números representando a velocidade com que os veículos foram observados viajando e um número representando o limite de velocidade, retorne um array com dois itens, o número de veículos que estavam acima da velocidade permitida seguido pela média do quanto esses veículos ultrapassaram o limite de velocidade.

- Se não houver veículos acelerando, retorne `[0, 0]`.

# --hints--

`speeding([50, 60, 55], 60)` deve retornar `[0, 0]`.

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` deve retornar `[2, 4]`.

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` deve retornar `[4, 8.5]`.

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` deve retornar `[2, 3.5]`.

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` deve retornar `[1, 57]`.

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
