---
id: 68b7687dded630607aceccb1
title: "Desafio 48: Detector de Spam"
challengeType: 28
dashedName: challenge-48
---

# --description--

Dado um número de telefone no formato `"+A (BBB) CCC-DDDD"`, onde cada letra representa um dígito conforme a seguir:

- `A` representa o código do país e pode ser qualquer número de dígitos.
- `BBB` representa o código de área e sempre terá três dígitos.
- `CCC` e `DDDD` representam o número local e terão sempre três e quatro dígitos, respectivamente.

Determine se é um número de spam com base nos seguintes critérios:

- O código do país tem mais de 2 dígitos ou não começa com zero (`0`).
- O código de área é maior que 900 ou menor que 200.
- A soma dos três primeiros dígitos do número local aparece dentro dos últimos quatro dígitos do número local.
- O número tem o mesmo dígito quatro ou mais vezes seguidas (ignorando os caracteres de formatação).

# --hints--

`isSpam("+0 (200) 234-0182")` deve retornar `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` deve retornar `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` deve retornar `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` deve retornar `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` deve retornar `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` deve retornar `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` deve retornar `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` deve retornar `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
