---
id: 68b7687dded630607aceccb3
title: "Desafio 49: Parser de Cabeçalho CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

Dada a primeira linha de um arquivo de valores separados por vírgula (CSV), retorne um array contendo os cabeçalhos.

- A primeira linha de um arquivo CSV contém cabeçalhos separados por vírgulas.
- Remova quaisquer espaços em branco no início ou no final de cada título.

# --hints--

`getHeadings("nome,idade,cidade")` deve retornar `["nome", "idade", "cidade"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` deve retornar `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` deve retornar `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
