---
id: 68b7cadffed0e75a517da673
title: "Desafio 52: Binário para Decimal"
challengeType: 28
dashedName: challenge-52
---

# --description--

Dada uma string que representa um número binário, retorne seu equivalente decimal como um número.

Um número binário usa apenas os dígitos `0` e `1` para representar qualquer número. Para converter binário em decimal, multiplique cada dígito por uma potência de `2` e some-os. Comece multiplicando o dígito mais à direita por `2^0`, o próximo dígito à esquerda por `2^1` e assim por diante. Depois que todos os dígitos forem multiplicados por uma potência de `2`, some o resultado.

Por exemplo, o número binário `101` equivale a `5` em decimal porque:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` deve retornar `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` deve retornar `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` deve retornar `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` deve retornar `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
