---
id: 68b7cadffed0e75a517da677
title: "Desafio 54: Força da Senha P@ssw0rd Str3ngth!"
challengeType: 28
dashedName: challenge-54
---

# --description--

Dada uma string de senha, retorne `"weak"`, `"medium"` ou `"strong"` com base na força da senha.

Uma senha é avaliada de acordo com as seguintes regras:

- Tem pelo menos 8 caracteres.
- Ele contém letras maiúsculas e minúsculas.
- Ele contém pelo menos um número.
- Ele contém pelo menos um caractere especial deste conjunto: `!`, `@`, `#`, `$`, `%`, `^`, `&` ou `*`.

Retorne `"weak"` se a senha atender a menos de duas das regras.  
Retorne `"medium"` se a senha atender 2 ou 3 das regras.  
Retorne `"strong"` se a senha atender todas as 4 regras.

# --hints--

`checkStrength("123456")` deve retornar `"fraco"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` deve retornar `"fraco"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` deve retornar `"fraco"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("SENHA")` deve retornar `"fraca"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` deve retornar `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` deve retornar `"médio"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` deve retornar `"médio"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` deve retornar `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` deve retornar `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
