---
id: 68b7687dded630607aceccab
title: "Desafio 45: Quadrado Perfeito"
challengeType: 28
dashedName: challenge-45
---

# --description--

Dado um inteiro, determine se ele é um quadrado perfeito.

- Um número é um quadrado perfeito se você pode multiplicar um inteiro por ele mesmo para obter o número. Por exemplo, 9 é um quadrado perfeito porque você pode multiplicar 3 por ele mesmo para obtê-lo.

# --hints--

`isPerfectSquare(9)` deve retornar `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` deve retornar `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` deve retornar `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` deve retornar `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` deve retornar `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` deve retornar `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` deve retornar `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` deve retornar `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
