---
id: 68c1a929005bf54d342aa8d4
title: "Desafio 57: Semana do Espaço Dia 3: Phone Home"
challengeType: 28
dashedName: challenge-57
---

# --description--

No terceiro dia da Semana do Espaço, você recebe um array de números que representam distâncias (em quilômetros) entre você, satélites e seu planeta natal em uma rota de comunicação. Determine quanto tempo levará para uma mensagem enviada através da rota alcançar seu planeta de destino usando as seguintes restrições:

- O primeiro valor no array é a distância da sua localização até o primeiro satélite.
- Cada valor subsequente, exceto o último, é a distância até o próximo satélite.
- O último valor no array é a distância do satélite anterior ao seu planeta natal.
- A mensagem viaja a 300.000 km/s.
- Cada satélite pelo qual a mensagem **passa** adiciona um atraso de transmissão de 0,5 segundo.
- Retorna um número arredondado para 4 casas decimais, com zeros à direita removidos.

# --hints--

`sendMessage([300000, 300000])` deve retornar `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` deve retornar `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` deve retornar `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` deve retornar `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` deve retornar `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` deve retornar `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
