---
id: 68c1a929005bf54d342aa8d2
title: "Desafio 55: Semana do Espaço Dia 1: Classificação Estelar"
challengeType: 28
dashedName: challenge-55
---

# --description--

4 de outubro marca o início da Semana Mundial do Espaço. Os próximos sete dias trarão desafios de programação com tema de astronomia.

Para o desafio de hoje, você recebe a temperatura da superfície de uma estrela em Kelvin (K) e precisa determinar sua classificação estelar com base nas seguintes faixas:

- "O": 30.000 K ou mais
- "B": 10.000 K - 29.999 K
- "A": 7.500 K - 9.999 K
- "F": 6.000 K - 7.499 K
- "G": 5.200 K - 5.999 K
- "K": 3.700 K - 5.199 K
- "M": 0 K - 3,699 K

- Retorna a classificação da `star` fornecida.

# --hints--

`classification(5778)` deve retornar `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` deve retornar `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` deve retornar `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` deve retornar `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` deve retornar `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` deve retornar `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` deve retornar `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` deve retornar `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
