---
id: 68c1a929005bf54d342aa8d6
title: "Desafio 59: Semana do Espaço Dia 5: Zona Goldilocks"
challengeType: 28
dashedName: challenge-59
---

# --description--

No quinto dia da Semana do Espaço, você calculará a "zona Goldilocks" de uma estrela - a região ao redor de uma estrela onde as condições são "perfeitas" para a existência de água líquida.

Dada a massa de uma estrela, retorne um array com as distâncias inicial e final da sua Zona Habitável em Unidades Astronômicas.

Para calcular a Goldilocks Zone:

1. Encontre a luminosidade da estrela elevando sua massa à potência de 3,5.
2. O início da zona é 0,95 vezes a raiz quadrada da sua luminosidade.
3. O final da zona é 1,37 vezes a raiz quadrada da sua luminosidade.

- Retorne as distâncias arredondadas para duas casas decimais.

Por exemplo, dado `1` como uma massa, retorne `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` deve retornar `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` deve retornar `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` deve retornar `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` deve retornar `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` deve retornar `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
