---
id: 68c1a929005bf54d342aa8d3
title: "Desafio 56: Semana do Espaço Dia 2: Busca por Exoplanetas"
challengeType: 28
dashedName: challenge-56
---

# --description--

No segundo dia da Semana do Espaço, você recebe uma string onde cada caractere representa a leitura de luminosidade de uma estrela. Determine se as leituras detectaram um exoplaneta usando o método de trânsito. O método de trânsito é quando um planeta passa na frente de uma estrela, reduzindo sua luminosidade observada.

- As leituras de luminosidade compreendem apenas os caracteres `0-9` e `A-Z` onde cada leitura corresponde aos seguintes valores numéricos:
- Os caracteres `0-9` correspondem aos níveis de luminosidade `0-9`.
- Os caracteres `A-Z` correspondem aos níveis de luminosidade `10-35`.

Uma estrela é considerada como tendo um exoplaneta se qualquer leitura única for menor ou igual a 80% da média de todas as leituras. Por exemplo, se a luminosidade média de uma estrela for 10, ela seria considerada como tendo um exoplaneta se qualquer leitura única for 8 ou menos.

# --hints--

`hasExoplanet("665544554")` deve retornar `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` deve retornar `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` deve retornar `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` deve retornar `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` deve retornar `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` deve retornar `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
