---
id: 68c1a929005bf54d342aa8d5
title: "Desafio 58: Semana do Espaço Dia 4: Local de Pouso"
challengeType: 28
dashedName: challenge-58
---

# --description--

No quarto dia da Semana do Espaço, você recebe uma matriz de números (um array de arrays), representando possíveis locais de pouso para seu rover. Encontre o local de pouso mais seguro com base nas seguintes regras:

- Cada posição na matriz conterá um número de `0-9`, inclusive.
- Qualquer `0` representa um possível ponto de aterrissagem.
- Qualquer número diferente de `0` é perigoso demais para pousar. Quanto maior o número, mais perigoso.
- O local mais seguro é definido como a célula `0` cujas células ao redor (até 4 vizinhos, ignorando diagonais) têm o menor perigo total.
- Ignore vizinhos fora dos limites (cantos e bordas simplesmente têm menos vizinhos).
- Retorna os índices do local de pouso mais seguro. Sempre haverá apenas um local mais seguro.

Por exemplo, dado:

```js
[
  [1, 0],
  [2, 0]
]
```

Retorna `[0, 1]`, os índices para o `0` no primeiro array.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` deve retornar `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` deve retornar `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` deve retornar `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` deve retornar `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
