---
id: 68c497f3aaefc9fd9f1b0e24
title: "Desafio 60: Semana do Espaço Dia 6: Fase da Lua"
challengeType: 28
dashedName: challenge-60
---

# --description--

Para o sexto dia da Semana do Espaço, você receberá uma data no formato `"YYYY-MM-DD"` e precisará determinar a fase da lua para esse dia usando as seguintes regras:

Use um ciclo lunar simplificado de 28 dias, dividido em quatro fases iguais:

- "New": dias 1 a 7
- "Waning": dias 8 - 14
- "Full": dias 15 - 21
- "Waning": dias 22 - 28

Após o dia 28, o ciclo se repete com o dia 1, uma lua nova.

- Use `"2000-01-06"` como uma lua nova de referência (dia 1 do ciclo) para determinar a fase do dia dado.
- Você não receberá nenhuma data anterior à data de referência.
- Retorne a fase correta como uma string.

# --hints--

`moonPhase("2000-01-12")` deve retornar `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` deve retornar `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` deve retornar `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` deve retornar `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` deve retornar `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
