---
id: 68c497f3aaefc9fd9f1b0e26
title: "Desafio 62: Hexadecimal para Decimal"
challengeType: 28
dashedName: challenge-62
---

# --description--

Dada uma string que representa um número hexadecimal (base 16), retorne seu valor decimal (base 10) como um inteiro.

Hexadecimal é um sistema numérico que usa 16 dígitos:

- `0-9` representam valores de `0` a `9`.
- `A-F` representam os valores de `10` a `15`.

Aqui está uma tabela de conversão parcial:

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- A string conterá apenas os caracteres `0–9` e `A–F`.

# --hints--

`hexToDecimal("A")` deve retornar `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` deve retornar `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` deve retornar `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` deve retornar `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` deve retornar `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
