---
id: 68cae5b538ff798bbd4da005
title: "Desafio 67: Validador de Email"
challengeType: 28
dashedName: challenge-67
---

# --description--

Dada uma string, determine se ela é um endereço de email válido usando as seguintes restrições:

- Deve conter exatamente um símbolo `@`.
- A parte local (antes do `@`):
  - Só pode conter letras (`a-z`, `A-Z`), dígitos (`0-9`), pontos (`.`), underscores (`_`) ou hífens (`-`).
  - Não pode começar ou terminar com um ponto.
- A parte do domínio (após o `@`):
  - Deve conter pelo menos um ponto.
  - Deve terminar com um ponto seguido de pelo menos duas letras.
- Nem a parte local nem a parte do domínio podem ter dois pontos seguidos.

# --hints--

`validate("a@b.cd")` deve retornar `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` deve retornar `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` deve retornar `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` deve retornar `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` deve retornar `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` deve retornar `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` deve retornar `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` deve retornar `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("git@commit@push.io")` deve retornar `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
