---
id: 68cae5b538ff798bbd4da002
title: "Desafio 64: 24 para 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

Dada uma string que representa um horário no formato 24 horas `"HHMM"`, retorne o horário no seu equivalente formato 12 horas `"H:MM AM"` ou `"H:MM PM"`.

- A entrada fornecida será sempre uma string de quatro dígitos no formato de hora 24 horas, de `"0000"` a `"2359"`.

# --hints--

`to12("1124")` deve retornar `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` deve retornar `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` deve retornar `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` deve retornar `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` deve retornar `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
