---
id: 68cae5b538ff798bbd4da007
title: "Desafio 69: Meias Desaparecidas"
challengeType: 28
dashedName: challenge-69
---

# --description--

Dado um inteiro representando o número de pares de meias com os quais você começou e outro inteiro representando quantos ciclos de lavagem você realizou, retorne o número de pares completos de meias que você tem atualmente usando as seguintes restrições:

- A cada 2 ciclos de lavagem, você perde uma única meia.
- A cada 3 ciclos de lavagem, você encontra uma meia faltando.
- A cada 5 ciclos de lavagem, uma única meia se desgasta e deve ser descartada.
- A cada 10 ciclos de lavagem, você compra um par de meias.
- Você nunca pode ter menos que zero meias no total.
- Regras podem se sobrepor. Por exemplo, no ciclo de lavagem 10, você perderá uma meia, jogará fora uma meia e comprará um novo par de meias.
- Retorna o número de pares completos de meias.

# --hints--

`sockPairs(2, 5)` deve retornar `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` deve retornar `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` deve retornar `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` deve retornar `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` deve retornar `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
