---
id: 68cae5b538ff798bbd4da001
title: "Desafio 63: Batalha de Palavras"
challengeType: 28
dashedName: challenge-63
---

# --description--

Dadas duas sentenças representando seu time e um time adversário, onde cada palavra do seu time batalha contra a palavra correspondente do time adversário, determine qual time vence usando as seguintes regras:

- As sentenças fornecidas sempre conterão o mesmo número de palavras.
- As palavras são separadas por um único espaço e conterão apenas letras.
- O valor de cada palavra é a soma de suas letras.
- As letras `a` a `z` correspondem aos valores de `1` a `26`. Por exemplo, `a` é `1` e `z` é `26`.
- Uma letra maiúscula dobra o valor da letra. Por exemplo, `A` é `2` e `Z` é `52`.
- Palavras batalham em ordem: a primeira palavra do seu time batalha com a primeira palavra do time adversário e assim por diante.
- Uma palavra vence se seu valor for maior que o valor da palavra oposta.
- A equipe com mais palavras vencedoras é a vencedora.

Retorne `"We win"` se seu time for o vencedor, `"We lose"` se seu time perder e `"Draw"` se ambos os times tiverem o mesmo número de vitórias.

# --hints--

`battle("hello world", "hello word")` deve retornar `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` deve retornar `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` deve retornar `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` deve retornar `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` deve retornar `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` deve retornar `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` deve retornar `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
