---
id: 68cae5b538ff798bbd4da003
title: "Desafio 65: Contagem de Strings"
challengeType: 28
dashedName: challenge-65
---

# --description--

Dadas duas strings, determine quantas vezes a segunda string aparece na primeira.

- A string do padrão pode se sobrepor na primeira string. Por exemplo, `"aaa"` contém `"aa"` duas vezes. Os dois primeiros `a` e os dois segundos.

# --hints--

`count('abcdefg', 'def')` deve retornar `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` deve retornar `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` deve retornar `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` deve retornar `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` deve retornar `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
