---
id: 68cae5b538ff798bbd4da00a
title: "Desafio 72: Ajustador de Termostato 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Dada a temperatura atual de uma sala em Fahrenheit e uma temperatura alvo em Celsius, retorne uma string indicando como ajustar a temperatura da sala com base nessas restrições:

- Retorne `"Heat: X degrees Fahrenheit"` se a temperatura atual estiver abaixo do alvo. Com `X` sendo o número de graus em Fahrenheit para aquecer o ambiente até o alvo, arredondado para 1 casa decimal.
- Retorne `"Cool: X degrees Fahrenheit"` se a temperatura atual estiver acima do alvo. Com `X` sendo o número de graus em Fahrenheit para resfriar o ambiente até atingir o alvo, arredondado para 1 casa decimal.
- Retorne `"Hold"` se a temperatura atual for igual à temperatura alvo.

Para converter Celsius em Fahrenheit, multiplique a temperatura em Celsius por 1.8 e some 32 ao resultado (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` deve retornar `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` deve retornar `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` deve retornar `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` deve retornar `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` deve retornar `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
