---
id: 68cae5b538ff798bbd4da009
title: "Desafio 71: Calculadora de Gorjeta"
challengeType: 28
dashedName: challenge-71
---

# --description--

Dado o preço da sua refeição e uma porcentagem de gorjeta personalizada, retorne um array com três valores de gorjeta; 15%, 20% e o valor personalizado.

- Os preços serão fornecidos no formato: `"$N.NN"`.
- Porcentagens personalizadas de dicas serão fornecidas neste formato: `"25%"`.
- Retorne os valores no mesmo formato `"$N.NN"`, arredondados para duas casas decimais.

Por exemplo, dado um preço de refeição de `"$10.00"` e um valor de gorjeta personalizada de `"25%"`, retorne `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` deve retornar `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` deve retornar `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` deve retornar `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
