---
id: 68cae5b538ff798bbd4da008
title: "Desafio 70: Extrator de Atributos HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

Dada uma string de um elemento HTML válido, retorne os atributos do elemento usando os seguintes critérios:

- Você receberá apenas um elemento.
- Os atributos estarão no formato: `attribute="value"`.
- Retorna um array de strings com cada propriedade e valor do atributo, separados por vírgula, neste formato: `["attribute1, value1", "attribute2, value2"]`.
- Retorne os atributos na ordem em que são fornecidos.
- Se nenhum atributo for encontrado, retorne um array vazio.

# --hints--

`extractAttributes('<span class="red"></span>')` deve retornar `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` deve retornar `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` deve retornar `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` deve retornar `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` deve retornar `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
