---
id: 68cae5b538ff798bbd4da004
title: "Desafio 66: Removedor de Tags HTML"
challengeType: 28
dashedName: challenge-66
---

# --description--

Dada uma string de código HTML, remova as tags e retorne o conteúdo de texto simples.

- A string de entrada conterá apenas HTML válido.
- As tags HTML podem ser aninhadas.
- Remova as tags e quaisquer atributos.

Por exemplo, `'<a href="#">Click here</a>'` deve retornar `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` deve retornar `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` deve retornar `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` deve retornar uma string vazia (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` deve retornar `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
