---
id: 68d2ba1468508398389487d0
title: "Desafio 75: Tesouro Escondido"
challengeType: 28
dashedName: challenge-75
---

# --description--

Dada uma matriz 2D que representa um mapa do fundo do oceano que inclui um tesouro escondido, e uma matriz com as coordenadas (`[row, column]`) para o próximo mergulho da sua busca pelo tesouro, retorne `"Empty"`, `"Found"` ou `"Recovered"` usando as seguintes regras:

- A matriz 2D fornecida conterá exatamente um tesouro não recuperado, que ocupará várias células.
- Cada célula na matriz 2D conterá um dos seguintes valores:
  - `"-"`: Sem tesouro.
  - "O": Uma parte do tesouro que não foi encontrada.
  - "X": Uma parte do tesouro que já foi encontrada.

- Se o local do mergulho não tiver tesouro, retorne `"Empty"`.
- Se o local do mergulho encontrar tesouro, mas pelo menos outra parte do tesouro permanecer não encontrada, retorne `"Found"`.
- Se o local do mergulho encontrar a última parte do tesouro não encontrada, retorne `"Recovered"`.

Por exemplo, dado:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

E `[2, 1]` para as coordenadas do local do mergulho, retorne `"Recovered"` porque o mergulho encontrou a última parte do tesouro que não havia sido encontrada.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` deve retornar `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` deve retornar `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` deve retornar `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` deve retornar `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` deve retornar `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` deve retornar `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
