---
id: 68d2ba1468508398389487cf
title: "Desafio 74: Músicas Favoritas"
challengeType: 28
dashedName: challenge-74
---

# --description--

Lembra dos iPods? O primeiro modelo foi lançado há 24 anos hoje, em 23 de outubro de 2001.

Dado um array de objetos de música que representam sua playlist do iPod, retorne um array com os títulos das duas músicas mais tocadas, com a música mais tocada primeiro.

- Cada objeto terá uma propriedade `"title"` (string) e uma propriedade `"plays"` (inteiro).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` deve retornar `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` deve retornar `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` deve retornar `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
