---
id: 68d2ba1468508398389487ce
title: "Desafio 73: Fale com Sabedoria, Você Deve"
challengeType: 28
dashedName: challenge-73
---

# --description--

Dada uma frase, retorne uma versão dela que soe como um conselho de um sábio mestre usando as seguintes regras:

- Palavras são separadas por um único espaço.
- Encontre a primeira ocorrência de uma das seguintes palavras na frase: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Mova todas as palavras antes e incluindo essa palavra para o final da frase e:
  - Preserve a ordem das palavras quando você as mover.
  - Deixe todos em letras minúsculas.
  - E adicione uma vírgula e espaço antes deles.
- Coloque em maiúscula a primeira letra da nova primeira palavra da frase.
- Todas as sentenças fornecidas terminarão com um único sinal de pontuação. Mantenha a pontuação original da sentença e coloque-a no final da nova sentença.
- Retorne a nova sentença, certifique-se de que haja um único espaço entre cada palavra e nenhum espaço no início e no final da sentença.

Por exemplo, dado `"You must speak wisely."` retorne `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` deve retornar `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` deve retornar `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` deve retornar `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` deve retornar `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` deve retornar `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
