---
id: 68d2ba1468508398389487d0
title: "Desafio 75: Ataque de Míssil"
challengeType: 28
dashedName: challenge-75
---

# --description--

Dada uma matriz 2D representando a localização do seu navio no oceano e uma matriz com as coordenadas (`[row, column]`) de um ataque de míssil, retorne `"Hit"`, `"Miss"` ou `"Sunk"` usando as seguintes regras:

- O array 2D fornecido conterá exatamente um navio não afundado.
- Cada célula na matriz 2D conterá um dos seguintes valores:
  - `"-"`: Água (sem navio).
  - "O": Uma parte da sua nave que não foi atingida.
  - "X": Uma parte do seu navio que já foi atingida.

- Se o golpe atingir a água, retorne `"Miss"`.
- Se o strike atingir sua nave, mas pelo menos uma outra parte da nave permanecer não atingida, retorne `"Hit"`.
- Se o ataque atingir a última parte não atingida do seu navio, retorne `"Sunk"`.

Por exemplo, dado:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

E `[2, 1]` para as coordenadas do ataque de míssil, retorne `"Sunk"` porque atingiu a última parte não atingida do seu navio.

# --hints--

`strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` deve retornar `"Sunk"`.

```js
assert.equal(strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Sunk");
```

`strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` deve retornar `"Miss"`.

```js
assert.equal(strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Miss");
```

`strike([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` deve retornar `"Hit"`.

```js
assert.equal(strike([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Hit");
```

`strike([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` deve retornar `"Hit"`.

```js
assert.equal(strike([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Hit");
```

`strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` deve retornar `"Sunk"`.

```js
assert.equal(strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Sunk");
```

`strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` deve retornar `"Miss"`.

```js
assert.equal(strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Miss");
```

# --seed--

## --seed-contents--

```js
function strike(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function strike(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Miss";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Hit";
        }
      }
    }
    return "Sunk";
  }

  return "Hit";
}
```
