---
id: 68d30845cc08266018fc46bd
title: "Desafio 77: Formatador de Duração"
challengeType: 28
dashedName: challenge-77
---

# --description--

Dado um número inteiro de segundos, retorne uma string que represente a mesma duração no formato `"H:MM:SS"`, onde `"H"` é o número de horas, `"MM"` é o número de minutos e `"SS"` é o número de segundos. Retorne o tempo usando as seguintes regras:

- Segundos: Deve sempre ter dois dígitos.
- Minutos: Deve omitir zeros à esquerda quando não forem necessários. Use `"0"` se a duração for menor que um minuto.
- Horas: Deve ser incluído apenas se for maior que zero.

# --hints--

`format(500)` deve retornar `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` deve retornar `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` deve retornar `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` deve retornar `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` deve retornar `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
