---
id: 68d30845cc08266018fc46bc
title: "Desafio 76: DNA Complementar"
challengeType: 28
dashedName: challenge-76
---

# --description--

Dada uma string que representa uma sequência de DNA, retorne sua fita complementar usando as seguintes regras:

- O DNA consiste nas letras `"A"`, `"C"`, `"G"` e `"T"`.
- As letras `"A"` e `"T"` se complementam.
- As letras `"C"` e `"G"` se complementam.

Por exemplo, dado `"ACGT"`, retorne `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` deve retornar `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` deve retornar `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` deve retornar `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` deve retornar `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
