---
id: 68d30fc57588d97fd3027b30
title: "Desafio 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

Em 28 de outubro de 1994, o Netscape Navigator foi lançado, ajudando milhões a explorar a web inicial.

Dada uma matriz de comandos do navegador que você executou no Netscape Navigator, retorne a página atual em que você está após executar todos os comandos usando as seguintes regras:

- Você sempre começa na página `"Home"`, que não será incluída no array de comandos.
- Os comandos válidos são:
  - "Visit Page": Onde "Page" é o nome da página que você está visitando. Por exemplo, "Visit About" leva você para a página "About". Quando você visita uma nova página, certifique-se de descartar qualquer histórico de avanço que você tenha.
  - `"Back"`: Leva você para a página anterior no seu histórico ou permanece na página atual se não houver uma.
  - "Forward": Leva você para frente no histórico até a página de onde você veio ou permanece na página atual se não houver uma.

Por exemplo, dado `["Visit About Us", "Back", "Forward"]`, retorne `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` deve retornar `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` deve retornar `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` deve retornar `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` deve retornar `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` deve retornar `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` deve retornar `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` deve retornar `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
