---
id: 68d30845cc08266018fc46be
title: "Desafio 78: Sequência Inteira"
challengeType: 28
dashedName: challenge-78
---

# --description--

Dado um número inteiro positivo, retorne uma string com todos os números inteiros de `1` até, e incluindo, o número fornecido, em ordem numérica.

Por exemplo, dado `5`, retorne `"12345"`.

# --hints--

`sequence(5)` deve retornar `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` deve retornar `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` deve retornar `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` deve retornar `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
