---
id: 68e39ed6106dac2f0a98fd64
title: "Desafio 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

Dada uma string que representa um nome de variável, converta-a para "spooky case" usando as seguintes restrições:

- Substitua todos os underscores (`_`) e hífens (`-`) por um til (~).
- Coloque a primeira letra da string em maiúscula e todas as outras letras depois disso em minúscula. Ignore o caractere til ao contar.

Por exemplo, dado `hello_world`, retorne `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` deve retornar `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` deve retornar `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` deve retornar `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` deve retornar `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` deve retornar `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
