---
id: 68e39ed6106dac2f0a98fd63
title: "Desafio 81: Nth Prime"
challengeType: 28
dashedName: challenge-81
---

# --description--

Um número primo é um inteiro positivo maior que 1 que é divisível apenas por 1 e por ele mesmo. Os cinco primeiros números primos são `2`, `3`, `5`, `7` e `11`.

Dado um inteiro positivo `n`, retorne o `n`th número primo. Por exemplo, dado `5` retorne o 5th número primo: `11`.

# --hints--

`nthPrime(5)` deve retornar `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` deve retornar `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` deve retornar `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` deve retornar `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` deve retornar `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
