---
id: 68e39ed6106dac2f0a98fd65
title: "Desafio 83: Validação de Assinatura"
challengeType: 28
dashedName: challenge-83
---

# --description--

Dada uma string de mensagem, uma string de chave secreta e um número de assinatura, determine se a assinatura é válida usando este método de codificação:

- Letras na mensagem e na chave secreta têm esses valores:
  - `a` a `z` têm valores de `1` a `26` respectivamente.
  - `A` a `Z` têm valores de `27` a `52` respectivamente.
- Todos os outros caracteres não têm valor.
- Calcule a assinatura somando a mensagem e somando a chave secreta.

Por exemplo, dado a mensagem `"foo"` e a chave secreta `"bar"`, a assinatura seria `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Verifique se a assinatura calculada corresponde à assinatura fornecida.

# --hints--

`verify("foo", "bar", 57)` deve retornar `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` deve retornar `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` deve retornar `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` deve retornar `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` deve retornar `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` deve retornar `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
