---
id: 68ee9e3066cfd4eb2328e8a4
title: "Desafio 85: Contador de Palavras"
challengeType: 28
dashedName: challenge-85
---

# --description--

Dada uma string de sentença, retorne o número de palavras que estão na sentença.

- Palavras são qualquer sequência de caracteres sem espaço e são separadas por um único espaço.

# --hints--

`countWords("Hello world")` deve retornar `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` deve retornar `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` deve retornar `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` deve retornar `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` deve retornar `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
