---
id: 68ee9e3066cfd4eb2328e8a6
title: "Desafio 87: Construtor de Matrizes"
challengeType: 28
dashedName: challenge-87
---

# --description--

Dado dois inteiros (um número de linhas e um número de colunas), retorne uma matriz (um array de arrays) preenchida com zeros (`0`) do tamanho especificado.

Por exemplo, dado `2` e `3`, retorne:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` deve retornar `[[0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` deve retornar `[[0, 0], [0, 0], [0, 0]]`.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` deve retornar `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` deve retornar `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
