---
id: 68ee9e3066cfd4eb2328e8a7
title: "Desafio 88: Localizador de Dia da Semana"
challengeType: 28
dashedName: challenge-88
---

# --description--

Dada uma string de data no formato `YYYY-MM-DD`, retorne o dia da semana.

Os dias válidos para retorno são:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Certifique-se de ignorar os fusos horários.

# --hints--

`getWeekday("2025-11-06")` deve retornar `Thursday`.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` deve retornar `Friday`.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` deve retornar `Saturday`.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` deve retornar `Wednesday`.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` deve retornar `Monday`.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```
