---
id: 68f6587287ad1f4ad39b0c7f
title: "Desafio 93: Vogais e Consoantes"
challengeType: 28
dashedName: challenge-93
---

# --description--

Dada uma string, retorne um array com o número de vogais e o número de consoantes na string.

- As vogais consistem em `a`, `e`, `i`, `o`, `u` em qualquer caixa.
- Consoantes consistem em todas as outras letras em qualquer caixa.
- Ignore quaisquer caracteres que não sejam letras.

Por exemplo, dado `"Hello World"`, retorne `[3, 7]`.

# --hints--

`count("Hello World")` deve retornar `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` deve retornar `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` deve retornar `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` deve retornar `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` deve retornar `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` deve retornar `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
