---
id: 68f6587287ad1f4ad39b0c84
title: "Desafio 98: Contagem de Retângulos"
challengeType: 28
dashedName: challenge-98
---

# --description--

Dado dois inteiros positivos representando a largura e a altura de um retângulo, determine quantos retângulos podem caber no retângulo dado.

- Conte apenas retângulos com largura e altura inteiras.

Por exemplo, dado `1` e `3`, retorne `6`. Três retângulos 1x1, dois retângulos 1x2 e um retângulo 1x3.

# --hints--

`countRectangles(1, 3)` deve retornar `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` deve retornar `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` deve retornar `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` deve retornar `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` deve retornar `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
