---
id: 68f6587287ad1f4ad39b0c81
title: "Desafio 95: Array Shift"
challengeType: 28
dashedName: challenge-95
---

# --description--

Dado um array e um inteiro representando quantas posições deslocar o array, retorne o array deslocado.

- Um número inteiro positivo desloca o array para a esquerda.
- Um número inteiro negativo desloca o array para a direita.
- O deslocamento envolve o array.

Por exemplo, dado `[1, 2, 3]` e `1`, desloque o array 1 para a esquerda, retornando `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` deve retornar `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` deve retornar `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` deve retornar `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` deve retornar `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` deve retornar `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
