---
id: 68f6587287ad1f4ad39b0c80
title: "Desafio 94: Gerador de Assinatura de Email"
challengeType: 28
dashedName: challenge-94
---

# --description--

Dadas strings para o nome, cargo e empresa de uma pessoa, retorne uma assinatura de email como uma única string usando as seguintes regras:

- O nome deve aparecer primeiro, precedido por um prefixo que depende da primeira letra do nome. Para nomes que começam com (case-insensitive):
  - `A-I`: Use `>>` como prefixo.
  - `J-R`: Use `--` como prefixo.
  - `S-Z`: Use `::` como prefixo.
- Uma vírgula e um espaço (`, `) devem seguir o nome.
- O título e a empresa devem seguir a vírgula e o espaço, separados por `" at "` (com espaços ao redor).

Por exemplo, dado `"Quinn Waverly"`, `"Founder and CEO"` e `"TechCo"` retorne `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` deve retornar `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` deve retornar `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` deve retornar `"::Tina Vaughn,Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` deve retornar `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` deve retornar `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
