---
id: 68f6587287ad1f4ad39b0c7e
title: "Desafio 92: Extension Extractor"
challengeType: 28
dashedName: challenge-92
---

# --description--

Dada uma string que representa um nome de arquivo, retorne a extensão do arquivo.

- A extensão é a parte do nome do arquivo que vem após o último ponto (`.`).
- Se o nome do arquivo não contiver um ponto ou terminar com um ponto, retorne `"none"`.
- A extensão deve ser retornada como está, preservando maiúsculas e minúsculas.

# --hints--

`getExtension("document.txt")` deve retornar `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` deve retornar `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` deve retornar `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` deve retornar `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` deve retornar `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` deve retornar `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
