---
id: 68f6587287ad1f4ad39b0c82
title: "Desafio 96: É fim de semana?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Dada uma data no formato `"YYYY-MM-DD"`, retorne o número de dias restantes até o final de semana.

- O fim de semana começa no sábado.
- Se a data fornecida for sábado ou domingo, retorne `"It's the weekend!"`.
- Caso contrário, retorne `"X days until the weekend."`, onde `X` é o número de dias até sábado.
- Se `X` for `1`, use `"day"` (singular) em vez de `"days"` (plural).
- Certifique-se de que o cálculo ignore seu fuso horário local.

# --hints--

`daysUntilWeekend("2025-11-14")` deve retornar `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` deve retornar `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` deve retornar `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` deve retornar `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` deve retornar `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` deve retornar `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
