---
id: 68f6587287ad1f4ad39b0c7d
title: "Desafio 91: Caça-palavras"
challengeType: 28
dashedName: challenge-91
---

# --description--

Dada uma matriz (um array de arrays) de letras únicas e uma palavra para encontrar, retorne os índices inicial e final da palavra na matriz.

- A matriz fornecida será preenchida com todas as letras minúsculas (`a-z`).
- A palavra a ser encontrada estará sempre na matriz exatamente uma vez.
- A palavra a ser encontrada estará sempre em linha reta em uma destas direções:
  - da esquerda para a direita
  - da direita para a esquerda
  - de cima para baixo
  - de baixo para cima

Por exemplo, dada a matriz:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

E a palavra `"cat"`, retorne:

```md
[[0, 1], [2, 1]]
```

Onde `[0, 1]` são os índices para o `"c"` (início da palavra) e `[2, 1]` são os índices para o `"t"` (fim da palavra).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` deve retornar `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` deve retornar `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` deve retornar `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` deve retornar `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
