---
id: 68ffb91507a5b645769328c9
title: "Desafio 106: Validador de Mensagens"
challengeType: 28
dashedName: challenge-106
---

# --description--

Dada uma string de mensagem e uma string de validação, determine se a mensagem é válida.

- Uma mensagem é válida se cada palavra na mensagem começar com a letra correspondente na string de validação, na ordem.
- Letras não diferenciam maiúsculas de minúsculas.
- As palavras na mensagem são separadas por espaços simples.

# --hints--

`isValidMessage("hello world", "hw")` deve retornar `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` deve retornar `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` deve retornar `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` deve retornar `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` deve retornar `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
