---
id: 68ffb91507a5b645769328c6
title: "Desafio 103: LCM"
challengeType: 28
dashedName: challenge-103
---

# --description--

Dado dois inteiros, retorne o mínimo múltiplo comum (MMC) dos dois números.

O MMC de dois números é o menor inteiro positivo que é múltiplo de ambos os números. Por exemplo, dado `4` e `6`, retorne `12` porque:

- Múltiplos de `4` são `4`, `8`, `12` e assim por diante.
- Múltiplos de `6` são `6`, `12`, `18` e assim por diante.
- `12` é o menor número que é múltiplo de ambos.

# --hints--

`lcm(4, 6)` deve retornar `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` deve retornar `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` deve retornar `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` deve retornar `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` deve retornar `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
