---
id: 68ffb91507a5b645769328cc
title: "Desafio 109: Qual é a Minha Idade Mesmo?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Dada a data de aniversário de alguém no formato `YYYY-MM-DD`, retorne a idade da pessoa em 27 de novembro de 2025.

- Assuma que todos os aniversários são datas válidas antes de 27 de novembro de 2025.
- Retorna a idade como um inteiro.
- Certifique-se de considerar se a pessoa já fez aniversário em 2025.

# --hints--

`calculateAge("2000-11-20")` deve retornar `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` deve retornar `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` deve retornar `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` deve retornar `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` deve retornar `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
