---
id: 68ffb91507a5b645769328c5
title: "Desafio 102: Palavra Mais Longa"
challengeType: 28
dashedName: challenge-102
---

# --description--

Dada uma string de sentença, retorne a palavra mais longa na sentença.

- Palavras são separadas por um único espaço.
- Apenas letras (`a-z`, case-insensitive) contam para o comprimento da palavra.
- Se houver várias palavras com o mesmo comprimento, retorne a primeira que aparecer.
- Retorne a palavra como ela aparece na string fornecida, com a pontuação removida.

# --hints--

`longestWord("The quick red fox")` deve retornar `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` deve retornar `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` deve retornar `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` deve retornar `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` deve retornar `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` deve retornar `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
