---
id: 69162d64f96574d9bb629eff
title: "Desafio 114: Camel para Snake"
challengeType: 28
dashedName: challenge-114
---

# --description--

Dada uma string em camel case, retorne a versão em snake case da string usando as seguintes regras:

- A string de entrada conterá apenas letras (`A-Z` e `a-z`) e sempre começará com uma letra minúscula.
- Cada letra maiúscula na string camel case inicia uma nova palavra.
- Converta todas as letras para minúsculas.
- Separe as palavras com um underscore (`_`).

# --hints--

`toSnake("helloWorld")` deve retornar `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` deve retornar `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` deve retornar `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
