---
id: 69162d64f96574d9bb629efe
title: "Desafio 113: Milhas para Quilômetros"
challengeType: 28
dashedName: challenge-113
---

# --description--

Dada uma distância em milhas como um número, retorne a distância equivalente em quilômetros.

- A entrada será sempre um número não negativo.
- 1 milha equivale a 1.60934 quilômetros.
- Arredonde o resultado para duas casas decimais.
- Remova zeros finais desnecessários do resultado arredondado.

# --hints--

`convertToKm(1)` deve retornar `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` deve retornar `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` deve retornar `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` deve retornar `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` deve retornar `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
