---
id: 69162d64f96574d9bb629f01
title: "Desafio 116: Contagem de Permutação"
challengeType: 28
dashedName: challenge-116
---

# --description--

Dada uma string, retorne o número de permutações distintas que podem ser formadas a partir de seus caracteres.

- Uma permutação é qualquer reordenação dos caracteres na string.
- Não conte permutações duplicadas.
- Se a string contiver caracteres repetidos, arranjos repetidos devem ser contados apenas uma vez.
- A string conterá apenas letras (`A-Z`, `a-z`).

Por exemplo, dado `"abb"`, retorne `3` porque há três maneiras únicas de organizar as letras: `"abb"`, `"bab"` e `"bba"`.

# --hints--

`countPermutations("abb")` deve retornar `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` deve retornar `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` deve retornar `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` deve retornar `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
