---
id: 69162d64f96574d9bb629efb
title: "Desafio 110: Adivinhador de Palavras"
challengeType: 28
dashedName: challenge-110
---

# --description--

Dadas duas strings do mesmo tamanho, uma palavra secreta e uma tentativa, compare a tentativa com a palavra secreta usando as seguintes regras:

- A palavra secreta e a tentativa consistirão apenas em letras maiúsculas (`"A"` a `"Z"`);
- For each letter in the guess, replace it with a number according to how it matches the secret word:
  - "2" se a letra estiver na palavra secreta e na posição correta.
  - "1" se a letra estiver na palavra secreta mas na posição errada.
  - "0" se a letra não estiver na palavra secreta.
- Cada letra na palavra secreta pode ser usada no máximo uma vez.
- Correspondências exatas (`"2"`) são atribuídas primeiro, depois correspondências parciais (`"1"`) são atribuídas da esquerda para a direita para as letras restantes.
- Se uma letra ocorrer várias vezes na tentativa, ela pode corresponder apenas ao número de vezes que aparece na palavra secreta.

Por exemplo, dado uma palavra secreta `"APPLE"` e uma tentativa `"POPPA"`, retorne `"10201"`:

O primeiro `"P"` não está na localização correta (`"1"`), o `"O"` não está na palavra secreta (`"0"`), o segundo `"P"` está na localização correta (`"2"`), o terceiro `"P"` é um zero (`"0"`) porque os dois `"P"` na palavra secreta já foram usados e o `"A"` não está na localização correta (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` deve retornar `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` deve retornar `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` deve retornar `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` deve retornar `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` deve retornar `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` deve retornar `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
