---
id: 69162d64f96574d9bb629efd
title: "Desafio 112: Detector de IA"
challengeType: 28
dashedName: challenge-112
---

# --description--

O desafio de hoje é inspirado pelo lançamento do ChatGPT em 30 de novembro de 2022.

Dada uma string de uma ou mais sentenças, determine se ela foi provavelmente gerada por IA usando as seguintes regras:

- Ele contém dois ou mais traços (`-`).
- Ele contém dois ou mais conjuntos de parênteses (`()`). O texto pode estar dentro dos parênteses.
- Ele contém três ou mais palavras com 7 ou mais letras.

- As palavras são separadas por um único espaço e consistem apenas de letras (`A-Z`). Não inclua pontuação ou outros caracteres que não sejam letras como parte de uma palavra.

Se a frase dada atender a alguma das regras acima, retorne `"AI"`, caso contrário, retorne `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` deve retornar `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` deve retornar `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` deve retornar `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` deve retornar `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` deve retornar `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
