---
id: 69162d64f96574d9bb629f02
title: "Desafio 117: Diferença Simétrica"
challengeType: 28
dashedName: challenge-117
---

# --description--

Dadas duas arrays, retorne uma nova array contendo a diferença simétrica entre elas.

- A diferença simétrica entre dois conjuntos é o conjunto de valores que aparecem em um dos conjuntos, mas não em ambos.
- Retorne os valores na ordem em que aparecem pela primeira vez nos arrays de entrada.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` deve retornar `[1, 2, 4, 5]`.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` deve retornar `["a", "c"]`.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` deve retornar `[1, "b"]`.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` deve retornar `[2, 4, 6, 8]`.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
