---
id: 69162d64f96574d9bb629f04
title: "Desafio 119: Compressão de String"
challengeType: 28
dashedName: challenge-119
---

# --description--

Dada uma string `sentence`, retorne uma versão comprimida da sentença onde palavras duplicadas consecutivas são substituídas pela palavra seguida do número de vezes que ela se repete entre parênteses.

- Apenas duplicatas consecutivas são comprimidas.
- Palavras são separadas por espaços simples.

Por exemplo, dado `"yes yes yes please"`, retorne `"yes(3) please"`.

# --hints--

`compressString("yes yes yes please")` deve retornar `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` deve retornar `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` deve retornar `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` deve retornar `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
