---
id: 69162d64f96574d9bb629efc
title: "Desafio 111: Trajetória da Bola"
challengeType: 28
dashedName: challenge-111
---

# --description--

O desafio de hoje é inspirado no videogame Pong, que foi lançado em 29 de novembro de 1972.

Dada uma matriz (array de arrays) que inclui a localização da bola (`2`) e a localização anterior da bola (`1`), retorne os índices da matriz para a próxima localização da bola.

- A bola sempre se move em linha reta.
- A direção do movimento é determinada por como a bola se moveu de `1` para `2`.
- As bordas da matriz são consideradas paredes. Se a bola bater em um:
  - parede superior ou inferior, ele rebate invertendo sua direção vertical.
  - parede esquerda ou direita, ele rebate invertendo sua direção horizontal.
  - canto, ele rebate invertendo ambas as direções.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` deve retornar `[2, 3]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` deve retornar `[3, 0]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` deve retornar `[1, 2]`.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` deve retornar `[1, 1]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` deve retornar `[2, 2]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
