---
id: 69162d64f96574d9bb629f00
title: "Desafio 115: Conversor de Item de Lista Ordenada em Markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

Dada uma string que representa um item de lista ordenada em Markdown, retorne a string HTML equivalente.

Um item válido de lista ordenada em Markdown deve:

- Comece com zero ou mais espaços, seguido por
- Um número (1 ou maior) e um ponto (`.`), seguido de
- Pelo menos um espaço e então
- O texto do item da lista.

Se a string não tiver o formato exato acima, retorne `"Invalid format"`. Caso contrário, envolva o texto do item da lista em tags `li` e retorne a string.

Por exemplo, dado `"1. My item"`, retorne `"<li>My item</li>"`

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`convertListItem("1. My item")` deve retornar `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` deve retornar `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` deve retornar `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` deve retornar `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` deve retornar `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` deve retornar `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
