---
id: 691b559495c5cb5a37b9b484
title: "Desafio 124: Atualização de Inventário"
challengeType: 28
dashedName: challenge-124
---

# --description--

Dada uma matriz 2D que representa o inventário da sua loja e outra matriz 2D que representa uma remessa que você recebeu, retorne seu inventário atualizado.

- Cada elemento nos arrays terá o formato: `[quantity, "item"]`, onde `quantity` é um inteiro e `"item"` é uma string.
- Atualize os itens no inventário adicionando a quantidade de quaisquer itens correspondentes da remessa.
- Se um item recebido não existir no inventário atual, adicione-o como uma nova entrada ao final do inventário.
- Retorne o inventário na ordem em que foi fornecido com os novos itens no final na ordem em que aparecem no envio.

Por exemplo, dado um inventário de `[[2, "apples"], [5, "bananas"]]` e um envio de `[[1, "apples"], [3, "bananas"]]`, retorne `[[3, "apples"], [8, "bananas"]]`.

# --hints--

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` deve retornar `[[3, "apples"], [8, "bananas"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]]);
```

`updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` deve retornar `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
assert.deepEqual(updateInventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]]);
```

`updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` deve retornar `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
assert.deepEqual(updateInventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]]);
```

`updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` deve retornar `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
assert.deepEqual(updateInventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]);
```

# --seed--

## --seed-contents--

```js
function updateInventory(inventory, shipment) {
  return inventory;
}
```

# --solutions--

```js
function updateInventory(inventory, shipment) {
  const inventoryMap = new Map();
  inventory.forEach(([qty, item], index) => {
    inventoryMap.set(item, index);
  });

  shipment.forEach(([qty, item]) => {
    if (inventoryMap.has(item)) {
      const index = inventoryMap.get(item);
      inventory[index][0] += qty;
    } else {
      inventory.push([qty, item]);
      inventoryMap.set(item, inventory.length - 1);
    }
  });

  return inventory;
}
```
