---
id: 691b559495c5cb5a37b9b486
title: "Desafio 126: Coloque em Maiúscula"
challengeType: 28
dashedName: challenge-126
---

# --description--

Dada uma string title, retorne uma nova string formatada em title case usando as seguintes regras:

- Coloque a primeira letra de cada palavra em maiúscula.
- Deixe todas as outras letras de cada palavra em minúsculas.
- As palavras estão sempre separadas por um único espaço.

# --hints--

`titleCase("hello world")` deve retornar `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` deve retornar `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` deve retornar `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` deve retornar `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
