---
id: 691b559495c5cb5a37b9b480
title: "Desafio 120: Libras para Quilogramas"
challengeType: 28
dashedName: challenge-120
---

# --description--

Dado um peso em libras como um número, retorne a string `"(lbs) pounds equals (kgs) kilograms."`.

- Substitua `"(lbs)"` pelo número de entrada.
- Substitua `"(kgs)"` pela entrada convertida para quilogramas, arredondada para duas casas decimais e sempre incluindo duas casas decimais no valor.
- 1 libra equivale a 0,453592 quilogramas.
- Se a entrada for `1`, use `"pound"` em vez de `"pounds"`.
- Se o valor convertido for `1`, use `"kilogram"` em vez de `"kilograms"`.

# --hints--

`convertToKgs(1)` deve retornar `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` deve retornar `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` deve retornar `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` deve retornar `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` deve retornar `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
