---
id: 691b559495c5cb5a37b9b487
title: "Desafio 127: Verificação de Velocidade"
challengeType: 28
dashedName: challenge-127
---

# --description--

Dada a velocidade que você está viajando em milhas por hora (MPH) e um limite de velocidade em quilômetros por hora (KPH), determine se você está acima da velocidade permitida e se receberá um aviso ou uma multa.

- 1 milha equivale a 1.60934 quilômetros.
- Se você estiver viajando a uma velocidade menor ou igual ao limite, retorne `"Not Speeding"`.
- Se você estiver viajando a 5 KPH ou menos acima do limite de velocidade, retorne `"Warning"`.
- Se você estiver viajando mais de 5 KPH acima do limite de velocidade, retorne `"Ticket"`.

# --hints--

`speedCheck(30, 70)` deve retornar `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` deve retornar `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` deve retornar `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` deve retornar `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` deve retornar `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` deve retornar `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
