---
id: 691b559495c5cb5a37b9b488
title: "Desafio 128: Contagem de Consoantes"
challengeType: 28
dashedName: challenge-128
---

# --description--

Dada uma string e um número alvo, determine se a string contém exatamente o número alvo de consoantes.

- Consoantes são todos os caracteres alfabéticos exceto `"a"`, `"e"`, `"i"`, `"o"` e `"u"` em qualquer caixa.
- Ignore dígitos, pontuação, espaços e outros caracteres que não sejam letras ao contar.

# --hints--

`hasConsonantCount("helloworld", 7)` deve retornar `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` deve retornar `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` deve retornar `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` deve retornar `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` deve retornar `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
