---
id: 691b559495c5cb5a37b9b485
title: "Desafio 125: Game of Life"
challengeType: 28
dashedName: challenge-125
---

# --description--

Dada uma matriz (array de arrays) que representa o estado atual no Jogo da Vida de Conway, retorne o próximo estado da matriz usando estas regras:

- Cada célula é `1` (viva) ou `0` (morta).
- Os vizinhos de uma célula são até oito células ao redor (verticalmente, horizontalmente e diagonalmente).
- Células nas bordas têm menos de oito vizinhos.

Regras para atualizar cada célula:

- Qualquer célula viva com menos de dois vizinhos vivos morre (subpopulação).
- Qualquer célula viva com dois ou três vizinhos vivos continua viva.
- Qualquer célula viva com mais de três vizinhos vivos morre (superpopulação).
- Qualquer célula morta com exatamente três vizinhos vivos se torna viva (reprodução).

Por exemplo, dado:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

retornar:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Cada célula é atualizada de acordo com o número de vizinhos vivos. Por exemplo, `[0][0]` permanece morta (2 vizinhos vivos), `[0][1]` permanece viva (2 vizinhos vivos), `[0][2]` morre (3 vizinhos vivos) e assim por diante.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` deve retornar `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` deve retornar `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` deve retornar `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` deve retornar `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
