---
id: 691b559495c5cb5a37b9b482
title: "Desafio 122: Parser de Negrito em Markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

Dada uma string que pode incluir algum texto em negrito no Markdown, retorne a string HTML equivalente.

- Texto em negrito no Markdown é qualquer texto que começa e termina com dois asteriscos (`**`) ou dois underlines (`__`).
- Não pode haver espaços entre o texto e os asteriscos ou sublinhados, mas pode haver espaços no próprio texto.
- Converta todas as ocorrências em **negrito** para tags HTML `b` e retorne a string.

Por exemplo, dado `"**This is bold**"`, retorne `"<b>This is bold</b>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseBold("**This is bold**")` deve retornar `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` deve retornar `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` deve retornar `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` deve retornar `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` deve retornar `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
