---
id: 691b559495c5cb5a37b9b489
title: "Desafio 129: Parser de Bloco de Citação Markdown"
challengeType: 28
dashedName: challenge-129
---

# --description--

Dada uma string que inclui um blockquote em Markdown, retorne a string HTML equivalente.

Um blockquote em Markdown é qualquer linha que:

- Começa com zero ou mais espaços
- >
- Então, um ou mais espaços
- E finalmente, o texto do blockquote.

Retorne o texto do bloco de citação cercado pelas tags HTML de abertura e fechamento `blockquote`.

Por exemplo, dado `"> This is a quote"`, retorne `<blockquote>This is a quote</blockquote>`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseBlockquote("> This is a quote")` deve retornar `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` deve retornar `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` deve retornar `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
