---
id: 6925e2068081f40f549ced1c
title: "Desafio 138: Soma dos Divisores"
challengeType: 28
dashedName: challenge-138
---

# --description--

Dado um número inteiro positivo, retorne a soma de todos os seus divisores.

- Um divisor é qualquer número inteiro que divide o número de forma exata (o resto é `0`).
- Conte cada divisor apenas uma vez.

Por exemplo, dado `6`, retorne `12` porque os divisores de `6` são `1`, `2`, `3` e `6` e a soma deles é `12`.

# --hints--

`sumDivisors(6)` deve retornar `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` deve retornar `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` deve retornar `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` deve retornar `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` deve retornar `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` deve retornar `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
