---
id: 6925e2068081f40f549ced1a
title: "Desafio 136: Parser de Imagem Markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Dada uma string de uma imagem em Markdown, retorne a string HTML equivalente.

Uma imagem Markdown tem o seguinte formato: `"![alt text](image_url)"`. Onde:

- `alt text` é a descrição da imagem (o valor do atributo `alt`).
- `image_url` é a URL de origem da imagem (o valor do atributo `src`).

Retorne uma string da tag HTML `img` com o `src` definido para a URL da imagem e o `alt` definido para o texto alternativo.

Por exemplo, dado `"![Cute cat](cat.png)"` retorne `'<img src="cat.png" alt="Cute cat">'`;

- Certifique-se de que a tag, a ordem dos atributos, o espaçamento e o uso de aspas sejam os mesmos que acima.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens — verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseImage("![Cute cat](cat.png)")` deve retornar `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` deve retornar `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` deve retornar `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
