---
id: 6925e2068081f40f549ced1b
title: "Desafio 137: Gerador de Snowflake"
challengeType: 28
dashedName: challenge-137
---

# --description--

Dada uma string multilinha que usa caracteres de nova linha (`\n`) para representar uma quebra de linha, retorne uma nova string onde cada linha é espelhada horizontalmente e anexada ao final da linha original.

- Espelhe uma linha invertendo todos os seus caracteres, incluindo espaços.

Por exemplo, dado `"* \n *\n* "`, que registra no console como:

```sh
* 
 *
* 
```

Retorna `"*  *\n ** \n*  *"`, que é exibido no console como:

```sh
*  *
 ** 
*  *
```

Preste muita atenção aos espaços em branco nas strings fornecidas e retornadas. Certifique-se de não remover nenhum deles.

# --hints--

`generateSnowflake("* \n *\n* ")` deve retornar `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` deve retornar `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` deve retornar `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` deve retornar `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` deve retornar `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
