---
id: 6925e2068081f40f549ced1d
title: "Desafio 139: Pedra, Papel e Tesoura"
challengeType: 28
dashedName: challenge-139
---

# --description--

Dadas duas strings, a primeira representando o Jogador 1 e a segunda representando o Jogador 2, determine o vencedor de uma partida de Pedra, Papel e Tesoura.

- As strings de entrada serão sempre `"Rock"`, `"Paper"` ou `"Scissors"`.
- "Pedra" vence "Tesoura".
- `"Paper"` vence `"Rock"`.
- `"Tesoura"` vence `"Papel"`.

Retornar:

- "Jogador 1 vence" se o Jogador 1 vencer.
- "Jogador 2 vence" se o Jogador 2 vencer.
- "Tie" se ambos os jogadores escolherem a mesma opção.

# --hints--

`rockPaperScissors("Rock", "Rock")` deve retornar `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` deve retornar `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` deve retornar `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` deve retornar `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` deve retornar `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` deve retornar `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
