---
id: 69272dcf1c24b44fd79137c5
title: "Desafio 142: Some a String"
challengeType: 28
dashedName: challenge-142
---

# --description--

Dada uma string contendo dígitos e outros caracteres, retorne a soma de todos os números na string.

- Trate dígitos consecutivos como um único número. Por exemplo, `"13"` conta como 13, não 1 + 3.
- Ignore quaisquer caracteres que não sejam dígitos.

# --hints--

`stringSum("3apples2bananas")` deve retornar `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` deve retornar `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` deve retornar `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` deve retornar `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` deve retornar `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return arg;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
