---
id: 69272dcf1c24b44fd79137c4
title: "Desafio 141: Combustível de Decolagem"
challengeType: 28
dashedName: challenge-141
---

# --description--

Dado o número de galões de combustível atualmente no seu avião e o número necessário de litros de combustível para alcançar seu destino, determine quantos galões adicionais de combustível você deve adicionar.

- 1 galão equivale a 3.78541 litros.
- Se o avião já tiver combustível suficiente, retorne `0`.
- Você só pode adicionar galões inteiros.
- Não inclua decimais no número retornado.

# --hints--

`fuelToAdd(0, 1)` deve retornar `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` deve retornar `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` deve retornar `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` deve retornar `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` deve retornar `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
