---
id: 69272dcf1c24b44fd79137c6
title: "Desafio 143: Parser de Itálico em Markdown"
challengeType: 28
dashedName: challenge-143
---

# --description--

Dada uma string que pode incluir algum texto em itálico no Markdown, retorne a string HTML equivalente.

- Texto em itálico no Markdown é qualquer texto que começa e termina com um asterisco simples (`*`) ou um sublinhado simples (`_`).
- Não pode haver espaços entre o texto e o asterisco ou sublinhado, mas pode haver espaços no próprio texto.
- Converta todas as ocorrências em itálico para tags HTML `i` e retorne a string.

Por exemplo, dado `"*This is italic*"`, retorne `"<i>This is italic</i>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseItalics("*This is italic*")` deve retornar `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` deve retornar `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` deve retornar `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` deve retornar `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` deve retornar `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
