---
id: 69272dcf1c24b44fd79137c3
title: "Desafio 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Dada uma string que representa um nome de variável, retorne o nome da variável convertido para SCREAMING_SNAKE_CASE.

Os nomes de variáveis fornecidos serão escritos em um dos seguintes formatos:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

Nos formatos acima, as palavras são separadas por um underscore (`_`), um hífen (`-`) ou uma nova palavra começa com uma letra maiúscula.

Para converter para SCREAMING_SNAKE_CASE:

- Deixe todas as letras em maiúsculas
- Separe as palavras com um underscore (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` deve retornar `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` deve retornar `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` deve retornar `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` deve retornar `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` deve retornar `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();;
}
```
