---
id: 69306364df283fcaff2e1ad6
title: "Desafio 145: N-ésimo Número de Fibonacci"
challengeType: 28
dashedName: challenge-145
---

# --description--

Dado um inteiro `n`, retorne o `n`-ésimo número na sequência de fibonacci.

A sequência de Fibonacci é uma série de números onde cada número é a soma dos dois anteriores. Os 10 primeiros números da sequência são `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21` e `34`.

# --hints--

`nthFibonacci(4)` deve retornar `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` deve retornar `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` deve retornar `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` deve retornar `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` deve retornar `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
