---
id: 69306364df283fcaff2e1ada
title: "Desafio 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Dada uma string, retorne uma nova string onde todas as vogais são convertidas para maiúsculas e todos os outros caracteres alfabéticos são convertidos para minúsculas.

- As vogais são `"a"`, `"e"`, `"i"`, `"o"` e `"u"` em qualquer caso.
- Caracteres não alfabéticos devem permanecer inalterados.

# --hints--

`vowelCase("vowelcase")` deve retornar `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` deve retornar `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` deve retornar `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` deve retornar `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` deve retornar `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
