---
id: 69306364df283fcaff2e1ad8
title: "Desafio 147: Calculadora de Ano Bissexto"
challengeType: 28
dashedName: challenge-147
---

# --description--

Dado um ano inteiro, determine se ele é um ano bissexto.

Um ano é bissexto se satisfizer as seguintes regras:

- O ano é divisível por 4 e
- O ano não é divisível por 100, a menos que
- O ano é divisível por 400 de forma exata.

# --hints--

`isLeapYear(2024)` deve retornar `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` deve retornar `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` deve retornar `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` deve retornar `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` deve retornar `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` deve retornar `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` deve retornar `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
