---
id: 69306364df283fcaff2e1ad5
title: "Desafio 144: Sequência de Resoluções"
challengeType: 28
dashedName: challenge-144
---

# --description--

Dado um array de arrays, onde cada sub-array representa um dia das suas atividades de resolução e contém três inteiros: o número de passos dados naquele dia, os minutos de tempo de tela naquele dia e o número de páginas lidas naquele dia; determine se você está cumprindo suas resoluções.

- O primeiro subarray é o dia 1 e o segundo o dia 2 e assim por diante.

Um dia é considerado bem-sucedido se todas as três metas a seguir forem alcançadas:

- Você caminhou pelo menos 10.000 passos.
- Você não teve mais do que 120 minutos de tempo de tela.
- Você leu pelo menos cinco páginas.

Se todos os dias fornecidos forem bem-sucedidos, retorne `"Resolution on track: N day streak."` Onde `N` é o número de dias bem-sucedidos.

Se um ou mais dias não forem um sucesso, retorne `"Resolution failed on day X: N day streak."`. Onde `X` é o número do dia do primeiro dia sem sucesso e `N` é o número de dias bem-sucedidos antes do primeiro dia sem sucesso.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` deve retornar `"Resolução em dia: sequência de 3 dias."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` deve retornar `"Falha na resolução no dia 2: sequência de 1 dia."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` deve retornar `"Falha na resolução no dia 3: sequência de 2 dias."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` deve retornar `"Falha na resolução no dia 4: sequência de 3 dias."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` deve retornar `"Resolução em dia: sequência de 10 dias."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
