---
id: 69306364df283fcaff2e1ad7
title: "Desafio 146: Assento Canhoto na Mesa"
challengeType: 28
dashedName: challenge-146
---

# --description--

Dada uma matriz 4x2 (array de arrays) representando a disposição dos assentos para uma refeição, determine em quantos assentos uma pessoa canhota pode se sentar.

- Uma pessoa canhota não pode sentar onde uma pessoa destra estaria no assento imediatamente à esquerda dela.

Na matriz dada:

- Um `"R"` é um assento ocupado por uma pessoa destra.
- Um `"L"` é um assento ocupado por uma pessoa canhota.
- Um `"U"` é um assento desocupado.
- Apenas assentos desocupados estão disponíveis para sentar.
- Os assentos na fileira superior estão voltados para "baixo" e os assentos na fileira inferior estão voltados para "cima" (como uma mesa), então esquerda e direita são relativos à orientação do assento.
- Os assentos de canto têm apenas um assento ao lado deles.

Por exemplo, na matriz dada:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

O assento no canto superior esquerdo não pode ser ocupado porque há uma pessoa destra à esquerda. Os outros dois assentos disponíveis podem ser ocupados porque não há uma pessoa destra à esquerda.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` deve retornar `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` deve retornar `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` deve retornar `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` deve retornar `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` deve retornar `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
