---
id: 69306364df283fcaff2e1ad9
title: "Desafio 148: Pressão dos Pneus"
challengeType: 28
dashedName: challenge-148
---

# --description--

Dado um array com quatro números representando as pressões dos pneus em psi dos quatro pneus do seu veículo e outro array com dois números representando a pressão mínima e máxima para seus pneus em bar, retorne um array com quatro strings descrevendo o status de cada pneu.

- 1 bar equivale a 14.5038 psi.

Retorne um array com os seguintes valores para cada pneu:

- "Baixo" se a pressão do pneu estiver abaixo do mínimo permitido.
- "Good" se estiver entre o mínimo e o máximo permitido.
- "Alto" se estiver acima do máximo permitido.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` deve retornar `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` deve retornar `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` deve retornar `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` deve retornar `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` deve retornar `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
