---
id: 69373793f5a867f769cde138
title: "Desafio 153: Jogo da Velha"
challengeType: 28
dashedName: challenge-153
---

# --description--

Dada uma matriz 3×3 (um array de arrays) representando um jogo de Tic-Tac-Toe concluído, determine o vencedor.

- Cada elemento na matriz dada é ou um `"X"` ou um `"O"`.

Um jogador vence se tiver três de seus personagens em uma linha - horizontalmente, verticalmente ou diagonalmente.

Retornar:

- "`X wins` se o jogador X tiver três em linha."
- "`O wins` se o jogador O tiver três em linha."
- "Empate" se nenhum jogador tiver três em linha.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` deve retornar `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` deve retornar `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` deve retornar `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]])` deve retornar `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` deve retornar `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` deve retornar `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
