---
id: 69373793f5a867f769cde136
title: "Desafio 151: Array Ordenado?"
challengeType: 28
dashedName: challenge-151
---

# --description--

Dado um array de números, determine se os números estão ordenados em ordem crescente, decrescente ou nenhuma das duas.

Se o array fornecido for:

- Em ordem crescente (do menor para o maior), retorne `"Ascending"`.
- Em ordem decrescente (do maior para o menor), retorne `"Descending"`.
- Não ordenado em ordem crescente ou decrescente, retorna `"Not sorted"`.

# --hints--

`isSorted([1, 2, 3, 4, 5])` deve retornar `"Ascending"`.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` deve retornar `"Descending"`.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` deve retornar `"Not sorted"`.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` deve retornar `"Descending"`.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` deve retornar `"Ascending"`.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` deve retornar `"Not sorted"`.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
