---
id: 69373793f5a867f769cde137
title: "Desafio 152: Primo Circular"
challengeType: 28
dashedName: challenge-152
---

# --description--

Dado um inteiro, determine se ele é um primo circular.

Um primo circular é um número inteiro onde todas as rotações de seus dígitos são primos.

Por exemplo, `197` é um primo circular porque todas as rotações de seus dígitos: `197`, `971` e `719`, são números primos.

# --hints--

`isCircularPrime(197)` deve retornar `true`.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` deve retornar `false`.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` deve retornar `true`.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` deve retornar `false`.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` deve retornar `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
