---
id: 69373793f5a867f769cde13a
title: "Desafio 155: Plantar a Colheita"
challengeType: 28
dashedName: challenge-155
---

# --description--

Dado um inteiro representando o tamanho do seu campo agrícola, e `"acres"` ou `"hectares"` representando a unidade para o tamanho do seu campo agrícola, e um tipo de cultura, determine quantas plantas desse tipo você pode acomodar no seu campo.

- 1 acre equivale a 4046,86 metros quadrados.
- 1 hectare equivale a 10.000 metros quadrados.

Aqui está uma lista de culturas que serão fornecidas como entrada e quanto espaço uma única planta ocupa:

|Crop|Space per plant|
|-|-|
|`"corn"`|1 square meter|
|`"wheat"`|0.1 square meters|
|`"soybeans"`|0.5 square meters|
|`"tomatoes"`|0.25 square meters|
|`"lettuce"`|0.2 square meters|

Retorna o número de plantas que cabem no campo, arredondado para baixo até a planta inteira mais próxima.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` deve retornar `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` deve retornar `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` deve retornar `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` deve retornar `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` deve retornar `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
