---
id: 69373793f5a867f769cde135
title: "Desafio 150: Parser de Lista Não Ordenada em Markdown"
challengeType: 28
dashedName: challenge-150
---

# --description--

Dada a string de uma lista não ordenada válida em Markdown, retorne a string HTML equivalente.

Uma lista não ordenada consiste em um ou mais itens de lista. Um item de lista válido aparece em sua própria linha e:

- - Começa com um traço (`"-"`), seguido por
- Pelo menos um espaço e então
- O texto do item da lista.

A lista é fornecida como uma única string com linhas novas separadas pelo caractere de nova linha (`"\n"`). Não inclua os caracteres de nova linha no texto do item.

Envolva cada item da lista em tags HTML `li` e a lista inteira de itens em tags `ul`.

Por exemplo, dado `"- Item A\n- Item B"`, retorne `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` deve retornar `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` deve retornar `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` deve retornar `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` deve retornar `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
