---
id: 69373793f5a867f769cde139
title: "Desafio 154: Par para o Buraco"
challengeType: 28
dashedName: challenge-154
---

# --description--

Dado dois inteiros, o par para um buraco de golfe e o número de tacadas que um golfista deu nesse buraco, retorne a pontuação do golfista usando termos de golfe.

Retornar:

- "Hole in one!" se levou uma tacada.
- "Eagle" se levou dois golpes a menos que o par.
- "Birdie" se levou uma tacada a menos que o par.
- "Par" se levou o mesmo número de tacadas que o par.
- “Bogey” se levou uma tacada a mais que o par.
- "Double bogey" se levou duas tacadas a mais que o par.

# --hints--

`golfScore(3, 3)` deve retornar `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` deve retornar `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` deve retornar `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` deve retornar `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` deve retornar `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` deve retornar `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
