---
id: 6939b873185d8e00d453563f
title: "Desafio 160: Movimentos do Cavaleiro"
challengeType: 28
dashedName: challenge-160
---

# --description--

Dada a posição de um cavalo em um tabuleiro de xadrez, retorne o número de casas válidas para as quais o cavalo pode se mover.

Um tabuleiro de xadrez padrão é 8x8, com colunas rotuladas de `A` a `H` (da esquerda para a direita) e linhas rotuladas de `1` a `8` (de baixo para cima). Ele se parece com isto:

**A8****B8****C8****D8****E8****F8****G8****H8**
|-|-|-|-|-|-|-|-|
**A7****B7****C7****D7****E7****F7****G7****H7**
**A6****B6****C6****D6****E6****F6****G6****H6**
**A5****B5****C5****D5****E5****F5****G5****H5**
**A4****B4****C4****D4****E4****F4****G4****H4**
**A3****B3****C3****D3****E3****F3****G3****H3**
**A2****B2****C2****D2****E2****F2****G2****H2**
**A1****B1****C1****D1****E1****F1****G1****H1**

Um cavaleiro se move em forma de "L": dois quadrados em uma direção (horizontal ou vertical) e um quadrado na direção perpendicular.

Isso significa que um cavaleiro pode se mover para até oito posições possíveis, mas menos quando está perto das bordas do tabuleiro. Por exemplo, se um cavaleiro estivesse em `A1`, ele poderia se mover apenas para `B3` ou `C2`.

# --hints--

`knightMoves("A1")` deve retornar `2`.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` deve retornar `8`.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` deve retornar `6`.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` deve retornar `3`.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` deve retornar `4`.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
