---
id: 6939b873185d8e00d453563c
title: "Desafio 157: Parser de Link Markdown"
challengeType: 28
dashedName: challenge-157
---

# --description--

Dada a string de um link em Markdown, retorne a string HTML equivalente.

Uma imagem Markdown tem o seguinte formato: `"[link_text](link_url)"`. Retorne a string da tag HTML `a` com o `href` definido para o `link_url` e o `link_text` como o conteúdo da tag.

Por exemplo, dado `"[freeCodeCamp](https://freecodecamp.org/)"` retorne `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens — verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` deve retornar `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` deve retornar `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribua para o nosso repositório em https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` deve retornar `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribua para o nosso repositório em https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
