---
id: 6939b873185d8e00d453563e
title: "Desafio 159: Hipotenusa Inteira"
challengeType: 28
dashedName: challenge-159
---

# --description--

Dado dois inteiros positivos representando os comprimentos dos dois catetos (os dois lados curtos) de um triângulo retângulo, determine se a hipotenusa é um inteiro.

O comprimento da hipotenusa é calculado somando os quadrados dos comprimentos dos dois catetos e depois extraindo a raiz quadrada desse total (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` deve retornar `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` deve retornar `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` deve retornar `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` deve retornar `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` deve retornar `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` deve retornar `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
