---
id: 6939b873185d8e00d453563b
title: "Desafio 156: Ímpar ou Par?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Dado um número inteiro positivo, retorne `"Odd"` se for um número ímpar e `"Even"` se for par.

# --hints--

`oddOrEven(1)` deve retornar `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` deve retornar `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` deve retornar `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` deve retornar `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` deve retornar `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
