---
id: 6939b873185d8e00d453563d
title: "Desafio 158: Troca de Array"
challengeType: 28
dashedName: challenge-158
---

# --description--

Dado um array com dois valores, retorne um array com os valores trocados.

Por exemplo, dado `["A", "B"]` retorne `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` deve retornar `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` deve retornar `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` deve retornar `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` deve retornar `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
