---
id: 694596b0585c11170ac7c7fa
title: "Desafio 162: Consumo de Energia"
challengeType: 28
dashedName: challenge-162
---

# --description--

Dado o número de Calorias queimadas durante um treino e o número de watt-horas usados pelos seus dispositivos eletrônicos durante esse treino, determine qual deles usou mais energia.

Para compará-los, converta ambos os valores para joules usando as seguintes conversões:

- 1 Caloria equivale a 4184 joules.
- 1 watt-hora equivale a 3600 joules.

Retornar:

- "Workout" se o treino usou mais energia.
- "Dispositivos" se o dispositivo usou mais energia.
- "Igual" se ambos usaram a mesma quantidade de energia.

# --hints--

`compareEnergy(250, 50)` deve retornar `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` deve retornar `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` deve retornar `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` deve retornar `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` deve retornar `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` deve retornar `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
