---
id: 694596b0585c11170ac7c7fc
title: "Desafio 164: Parser de Código Inline em Markdown"
challengeType: 28
dashedName: challenge-164
---

# --description--

Dada uma string de Markdown que inclui um ou mais blocos de código inline, retorne a string HTML equivalente.

Blocos de código inline em Markdown usam uma única crase (```) no início e no fim do texto do bloco de código.

Retorne a string fornecida com todos os blocos de código convertidos para tags HTML `code`.

Por exemplo, dada a string `"Use `let` to declare the variable."`, retorne `"Use <code>let</code> to declare the variable."`.

Nota: O console pode não exibir tags HTML em strings ao registrar mensagens. Verifique o console do navegador para ver os logs com as tags incluídas.

# --hints--

`parseInlineCode("Use `let` to declare the variable.")` deve retornar `"Use <code>let</code> para declarar a variável."`.

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

`parseInlineCode("Use `let` or `const` to declare a variável.")` deve retornar `"Use <code>let</code> ou <code>const</code> para declarar uma variável."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

`parseInlineCode("Run `npm install` then `npm start`.")` deve retornar `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
