---
id: 694596b0585c11170ac7c7fb
title: "Desafio 163: Caso das Consoantes"
challengeType: 28
dashedName: challenge-163
---

# --description--

Dada uma string que representa o nome de uma variável, converta-a para consonant case usando as seguintes regras:

- Todas as consoantes devem ser convertidas para maiúsculas.
- Todas as vogais (`a`, `e`, `i`, `o`, `u` em qualquer caso) devem ser convertidas para minúsculas.
- Todos os hífens (`-`) devem ser convertidos em underscores (`_`).

# --hints--

`toConsonantCase("helloworld")` deve retornar `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` deve retornar `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` deve retornar `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` deve retornar `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
