---
id: 696655d24b614176d4c9b78b
title: "Desafio 168: Imagem Escalada"
challengeType: 28
dashedName: challenge-168
---

# --description--

Dada uma string que representa a largura e altura de uma imagem e um número para escalar a imagem, retorne a largura e altura escaladas.

- A string de entrada está no formato `"WxH"`. Por exemplo, `"800x600"`.
- A escala é um número para multiplicar a largura e a altura.

Retorna as dimensões escaladas no mesmo formato `"WxH"`.

# --hints--

`scaleImage("800x600", 2)` deve retornar `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` deve retornar `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` deve retornar `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` deve retornar `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
