---
id: 696655d24b614176d4c9b78c
title: "Desafio 169: FizzBuzz Mini"
challengeType: 28
dashedName: challenge-169
---

# --description--

Dado um inteiro, retorne uma string com base nas seguintes regras:

- Se o número for divisível por 3, retorne `"Fizz"`.
- Se o número for divisível por 5, retorne `"Buzz"`.
- Se o número for divisível por 3 e 5, retorne `"FizzBuzz"`.
- Caso contrário, retorne o número fornecido como uma string.

# --hints--

`fizzBuzzMini(3)` deve retornar `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` deve retornar `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` deve retornar `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` deve retornar `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` deve retornar `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
