---
id: 696655d24b614176d4c9b789
title: "Desafio 166: Validador Hexadecimal"
challengeType: 28
dashedName: challenge-166
---

# --description--

Dada uma string, determine se ela é uma cor hex válida em CSS. Uma cor hex válida em CSS deve:

- Comece com um `#`, e
- deve ser seguido por 3 ou 6 caracteres hexadecimais.

Caracteres hexadecimais são números `0` a `9` e letras `a` a `f` (case-insensitive).

# --hints--

`isValidHex("#123")` deve retornar `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` deve retornar `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` deve retornar `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` deve retornar `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` deve retornar `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` deve retornar `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` deve retornar `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` deve retornar `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
