---
id: 696655d24b614176d4c9b78a
title: "Desafio 167: Bingo! Letra"
challengeType: 28
dashedName: challenge-167
---

# --description--

Dado um número, retorne a letra do bingo associada a ele (em maiúscula). Os números do bingo são agrupados da seguinte forma:

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)` deve retornar `"O"`.

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)` deve retornar `"G"`.

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)` deve retornar `"I"`.

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)` deve retornar `"N"`.

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)` deve retornar `"B"`.

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
