---
id: 696655d24b614176d4c9b78d
title: "Desafio 170: Dia Ímpar ou Par"
challengeType: 28
dashedName: challenge-170
---

# --description--

Dado um timestamp (número de milissegundos desde a época Unix), retorne:

- `"odd"` se o dia do mês para esse timestamp for ímpar.
- "even" se o dia do mês para esse timestamp for par.

Por exemplo, dado `1769472000000`, um timestamp para 27 de janeiro de 2026, retorne `"odd"` porque o número do dia (`27`) é um número ímpar.

Nota: O timestamp está em milissegundos e você deve usar a data no fuso horário UTC, não no seu horário local.

# --hints--

`oddOrEvenDay(1769472000000)` deve retornar `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` deve retornar `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` deve retornar `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` deve retornar `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` deve retornar `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
