---
id: 69738771fb5a7b8b24cca2a3
title: "Desafio 177: Espelho de String"
challengeType: 28
dashedName: challenge-177
---

# --description--

Dada uma string, retorne uma nova string que consiste na string dada com uma cópia invertida dela mesma anexada ao final.

# --hints--

`mirror("freeCodeCamp")` deve retornar `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` deve retornar `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` deve retornar `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` deve retornar `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
