---
id: 69738771fb5a7b8b24cca29f
title: "Desafio 173: Movimentos Válidos do Peão"
challengeType: 28
dashedName: challenge-173
---

# --description--

Dada a posição de um dos seus peões em um tabuleiro de xadrez, retorne um array de todas as casas válidas para as quais ele pode se mover em ordem crescente.

Um tabuleiro de xadrez padrão é 8x8, com colunas rotuladas de `A` a `H` (da esquerda para a direita) e linhas rotuladas de `1` a `8` (de baixo para cima). Ele se parece com isto:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Para este desafio:

- Você é o jogador na parte inferior do tabuleiro.
- Peões só podem se mover uma casa "para cima".
- A menos que o peão esteja na linha inicial (linha 2), então ele pode mover uma ou duas casas para cima.

Por exemplo, dado `"D4"`, retorne `["D5"]`, a única casa para a qual seu peão pode se mover. Dado `"B2"`, retorne `["B3", "B4"]`, porque está na linha inicial e precisa estar em ordem crescente.

# --hints--

`findPawnMoves("D4")` deve retornar `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` deve retornar `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` deve retornar `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` deve retornar `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` deve retornar `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
