---
id: 69738771fb5a7b8b24cca2a5
title: "Desafio 179: Troco de Bolso"
challengeType: 28
dashedName: challenge-179
---

# --description--

Dado um array de inteiros representando as moedas no seu bolso, com cada inteiro sendo o valor de uma moeda em centavos, retorne o valor total no formato `"$D.CC"`.

- 100 centavos equivalem a 1 dólar.
- No valor de retorno, inclua um zero à esquerda para valores menores que um dólar e sempre exatamente dois dígitos para os centavos.

# --hints--

`countChange([25, 10, 5, 1])` deve retornar `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` deve retornar `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` deve retornar `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` deve retornar `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` deve retornar `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` deve retornar `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
