---
id: 69738771fb5a7b8b24cca2a2
title: "Desafio 176: Dia da Marmota"
challengeType: 28
dashedName: challenge-176
---

# --description--

Hoje é o Dia da Marmota, no qual uma marmota prevê o tempo com base em se ela vê ou não sua sombra.

Dado um valor que representa a aparência da marmota, retorne a previsão correta:

- Se o valor fornecido for o booleano `true` (a marmota viu sua sombra), retorne `"Looks like we'll have six more weeks of winter."`.
- Se o valor for o booleano `false` (a marmota não viu sua sombra), retorne `"It's going to be an early spring."`.
- Se o valor for qualquer outra coisa (a marmota não apareceu), retorne `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` deve retornar `"Parece que teremos mais seis semanas de inverno."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` deve retornar `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` deve retornar `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` deve retornar `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` deve retornar `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
