---
id: 69738771fb5a7b8b24cca29d
title: "Desafio 171: Achatar o Array"
challengeType: 28
dashedName: challenge-171
---

# --description--

Dado um array que contém arrays aninhados, retorne um novo array com todos os valores achatados em um único array unidimensional. Mantenha a ordem original dos itens nos arrays.

# --hints--

`flatten([1, [2, 3], 4])` deve retornar `[1, 2, 3, 4]`.

```js
assert.deepEqual(flatten([1, [2, 3], 4]), [1, 2, 3, 4]);
```

`flatten([5, [4, [3, 2]], 1])` deve retornar `[5, 4, 3, 2, 1]`.

```js
assert.deepEqual(flatten([5, [4, [3, 2]], 1]), [5, 4, 3, 2, 1]);
```

`flatten(["A", [[[["B"]]]], "C"])` deve retornar `["A", "B", "C"]`.

```js
assert.deepEqual(flatten(["A", [[[["B"]]]], "C"]), ["A", "B", "C"]);
```

`flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]])` deve retornar `["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]`.

```js
assert.deepEqual(flatten([["L", "M", "N"], ["O", ["P", "Q", ["R", ["S", ["T", "U"]]]]], "V", ["W", ["X", ["Y", ["Z"]]]]]), ["L", "M", "N", "O", "P", "Q", "R", "S", "T", "U", "V", "W", "X", "Y", "Z"]);
```

`flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]])` deve retornar `["red","blue","green","yellow","purple","orange","pink","brown"]`.

```js
assert.deepEqual(flatten([["red", ["blue", ["green", ["yellow", ["purple"]]]]], "orange", ["pink", ["brown"]]]), ["red","blue","green","yellow","purple","orange","pink","brown"]);
```

# --seed--

## --seed-contents--

```js
function flatten(arr) {

  return arr;
}
```

# --solutions--

```js
function flatten(arr) {
  let result = [];

  for (const item of arr) {
    if (Array.isArray(item)) {
      result = result.concat(flatten(item));
    } else {
      result.push(item);
    }
  }

  return result;
}
```
