---
id: 69738771fb5a7b8b24cca2a4
title: "Desafio 178: Truncar o Texto"
challengeType: 28
dashedName: challenge-178
---

# --description--

Dada uma string, retorne-a como está se tiver 20 caracteres ou menos. Se for maior que 20 caracteres, trunque-a para os primeiros 17 caracteres e adicione `"..."` ao final (para que tenha 20 caracteres no total) e retorne o resultado.

# --hints--

`truncateText("Hello, world!")` deve retornar `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` deve retornar `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` deve retornar `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` deve retornar `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
