---
id: 697a49e6ff50d756c9b69362
title: "Desafio 185: Jogos de Inverno 2026 Dia 6: Patinação Artística"
challengeType: 28
dashedName: challenge-185
---

# --description--

Dado um array de pontuações dos juízes e penalidades opcionais, calcule a pontuação final para uma rotina de patinação artística.

O primeiro argumento é um array de 10 pontuações de juízes, cada uma um número de 0 a 10. Remova as pontuações mais alta e mais baixa dos juízes e some as 8 pontuações restantes para obter a pontuação base.

Quaisquer argumentos adicionais passados para a função são penalidades. Subtraia todas as penalidades da pontuação base para obter a pontuação final.

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` deve retornar `64`.

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` deve retornar `80`.

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` deve retornar `67`.

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` deve retornar `67.5`.

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` deve retornar `59`.

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```
