---
id: 697a49e6ff50d756c9b69364
title: "Desafio 187: Jogos de Inverno 2026 Dia 8: Luge"
challengeType: 28
dashedName: challenge-187
---

# --description--

Dado um array de cinco números, cada um representando o tempo (em segundos) que um luger levou para completar um segmento de uma pista, determine qual segmento teve a velocidade mais rápida e qual foi essa velocidade.

A trilha está dividida nos seguintes segmentos:

- Segmento 1: 320 metros
- Segmento 2: 280 metros
- Segmento 3: 350 metros
- Segmento 4: 300 metros
- Segmento 5: 250 metros

O primeiro valor no array fornecido corresponde ao tempo para o segmento 1, o segundo valor para o segmento 2 e assim por diante.

Para calcular a velocidade (em metros por segundo) de um segmento, divida a distância pelo tempo.

Retorne `"A velocidade máxima do luger foi de X m/s no segmento Y."`. Onde `X` é a velocidade máxima, arredondada para duas casas decimais, e `Y` é o número do segmento onde a velocidade máxima ocorreu.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])` deve retornar `"A maior velocidade do luger foi 35.07 m/s no segmento 5."`

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])` deve retornar `"A velocidade mais rápida do luger foi 37.75 m/s no segmento 2."`

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])` deve retornar `"A maior velocidade do luger foi 38.49 m/s no segmento 3."`

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])` deve retornar `"A maior velocidade do luger foi 37.69 m/s no segmento 1."`

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])` deve retornar `"A maior velocidade do luger foi 39.24 m/s no segmento 4."`

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
