---
id: 697a49e6ff50d756c9b6935f
title: "Desafio 182: Jogos de Inverno 2026 Dia 3: Biatlo"
challengeType: 28
dashedName: challenge-182
---

# --description--

Dado um array de inteiros, onde cada valor representa o número de alvos atingidos em uma única rodada de um biatlo, retorne a distância total de penalidade que o atleta deve esquiar.

- Cada rodada consiste em 5 alvos.
- Cada alvo perdido resulta em um loop de penalidade de 150 metros.

# --hints--

`calculatePenaltyDistance([4, 4])` deve retornar `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` deve retornar `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` deve retornar `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` deve retornar `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` deve retornar `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
