---
id: 697a49e9860d24853adef67c
title: "Desafio 191: Jogos de Inverno 2026 Dia 12: Bobsled"
challengeType: 28
dashedName: challenge-191
---

# --description--

Dado um array representando os pesos dos atletas em uma equipe de bobsled e um número representando o peso do bobsled, determine se a equipe está elegível para competir.

- O tamanho do array determina o tamanho da equipe: equipes de 1, 2 ou 4 pessoas.
- Todos os valores de peso fornecidos estão em quilogramas (kg).

O trenó (trenó por si só) deve ter um peso mínimo de:

- 162 kg para uma equipe de 1 pessoa
- 170 kg para uma equipe de 2 pessoas
- 210 kg para uma equipe de 4 pessoas

O peso total do trenó (atletas mais trenó) não deve exceder:

- 247 kg para uma equipe de 1 pessoa
- 390 kg para uma equipe de 2 pessoas
- 630 kg para uma equipe de 4 pessoas

Retorne "Eligible" se a equipe atender a todos os requisitos, ou `"Not Eligible"` se a equipe não cumprir um ou mais dos requisitos.

# --hints--

`checkEligibility([78], 165)` deve retornar `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` deve retornar `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` deve retornar `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` deve retornar `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
