---
id: 697a49e6ff50d756c9b6935e
title: "Desafio 181: Jogos de Inverno 2026 Dia 2: Snowboard"
challengeType: 28
dashedName: challenge-181
---

# --description--

Dada a posição inicial de um snowboarder e uma rotação em graus, determine a posição de aterrissagem dele.

- A posição de um snowboarder é `"Regular"` ou `"Goofy"`.
- As rotações de truque são múltiplos de 90 graus. Positivo indica rotação no sentido horário e negativo indica rotação no sentido anti-horário.
- A posição de aterrissagem inverte a cada 180 graus de rotação.

Por exemplo, dado `"Regular"` e `90`, retorne `"Regular"`. Dado `"Regular"` e `180` graus, retorne `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` deve retornar `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` deve retornar `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` deve retornar `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` deve retornar `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` deve retornar `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` deve retornar `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
