---
id: 697a49e6ff50d756c9b69366
title: "Desafio 189: Jogos de Inverno 2026 Dia 10: Esqui Alpino"
challengeType: 28
dashedName: challenge-189
---

# --description--

Dada a queda vertical, distância horizontal e tipo de uma pista de esqui, determine a classificação de dificuldade da pista.

Para determinar a classificação:

- Calcule a inclinação da colina dividindo a queda pela distância.
- Então, calcule a inclinação ajustada com base no tipo de colina:
  - "Downhill" multiplique a inclinação por 1.2
  - "Slalom": multiplique a inclinação por 0.9
  - "Giant Slalom": multiplique a inclinação por 1.0

Retornar:

- "Green" se a inclinação ajustada for menor ou igual a 0.1
- "Blue" se a inclinação ajustada for maior que 0,1 e menor ou igual a 0,25
- "Preto" se a inclinação ajustada for maior que 0,25

# --hints--

`getHillRating(95, 900, "Slalom")` deve retornar `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` deve retornar `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` deve retornar `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` deve retornar `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` deve retornar `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` deve retornar `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
