---
id: 697a49e6ff50d756c9b69365
title: "Desafio 188: Jogos de Inverno 2026 Dia 9: Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

Dada uma string que representa as curvas em uma pista de esqueleto, determine a dificuldade da pista.

- The given string will only consist of the letters:
  - "L" para uma curva à esquerda
  - `"R"` para uma curva à direita
  - "S" para um segmento reto

- Cada mudança de direção adiciona 15 pontos (um `"L"` seguido de um `"R"` ou vice-versa).
- Todas as outras curvas adicionam 5 pontos cada (todos os outros caracteres `"L"` ou `"R"`).
- Segmentos retos adicionam 0 pontos.

A dificuldade da pista é baseada na pontuação total. Retorne:

- "Fácil" se o total for 0 - 100
- "Médio" se o total for 101-200
- "Hard" se o total for superior a 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` deve retornar `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` deve retornar `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` deve retornar `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` deve retornar `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` deve retornar `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` deve retornar `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
