---
id: 697a49e6ff50d756c9b6935d
title: "Desafio 180: Jogos de inverno dia 1: dia de abertura"
challengeType: 28
dashedName: challenge-180
---

# --description--

Os próximos 17 dias vão trazer desafios de código inspirados nos Jogos de inverno.

Para o primeiro, você recebe um código de país de duas letras e precisa retornar o emoji da bandeira desse país.

Use esta lista:

| País | Código | Flag |
| - | -: | :--: |
| Albânia                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armênia                    |            `"AM"` | `"🇦🇲"` |
| Austrália                  |            `"AU"` | `"🇦🇺"` |
| Áustria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijão                 |            `"AZ"` | `"🇦🇿"` |
| Bélgica                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolívia                    |            `"BO"` | `"🇧🇴"` |
| Bósnia e Herzegovina       |            `"BA"` | `"🇧🇦"` |
| Brasil                     |            `"BR"` | `"🇧🇷"` |
| Bulgária                   |            `"BG"` | `"🇧🇬"` |
| Canadá                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colômbia                   |            `"CO"` | `"🇨🇴"` |
| Croácia                    |            `"HR"` | `"🇭🇷"` |
| Chipre                     |            `"CY"` | `"🇨🇾"` |
| República Tcheca           |            `"CZ"` | `"🇨🇿"` |
| Dinamarca                  |            `"DK"` | `"🇩🇰"` |
| Equador                    |            `"EC"` | `"🇪🇨"` |
| Eritreia                   |            `"ER"` | `"🇪🇷"` |
| Estônia                    |            `"EE"` | `"🇪🇪"` |
| Finlândia                  |            `"FI"` | `"🇫🇮"` |
| França                     |            `"FR"` | `"🇫🇷"` |
| Geórgia                    |            `"GE"` | `"🇬🇪"` |
| Alemanha                   |            `"DE"` | `"🇩🇪"` |
| Grã-Bretanha               |            `"GB"` | `"🇬🇧"` |
| Grécia                     |            `"GR"` | `"🇬🇷"` |
| Guiné-Bissau               |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungria                    |            `"HU"` | `"🇭🇺"` |
| Islândia                   |            `"IS"` | `"🇮🇸"` |
| Índia                      |            `"IN"` | `"🇮🇳"` |
| Irã                        |            `"IR"` | `"🇮🇷"` |
| Irlanda                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Itália                     |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japão                      |            `"JP"` | `"🇯🇵"` |
| Cazaquistão                |            `"KZ"` | `"🇰🇿"` |
| Quênia                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Quirguistão                |            `"KG"` | `"🇰🇬"` |
| Letônia                    |            `"LV"` | `"🇱🇻"` |
| Líbano                     |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lituânia                   |            `"LT"` | `"🇱🇹"` |
| Luxemburgo                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malásia                    |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| México                     |            `"MX"` | `"🇲🇽"` |
| Moldávia                   |            `"MD"` | `"🇲🇩"` |
| Mônaco                     |            `"MC"` | `"🇲🇨"` |
| Mongólia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Marrocos                   |            `"MA"` | `"🇲🇦"` |
| Países Baixos              |            `"NL"` | `"🇳🇱"` |
| Nova Zelândia              |            `"NZ"` | `"🇳🇿"` |
| Nigéria                    |            `"NG"` | `"🇳🇬"` |
| Macedônia do Norte         |            `"MK"` | `"🇲🇰"` |
| Noruega                    |            `"NO"` | `"🇳🇴"` |
| Paquistão                  |            `"PK"` | `"🇵🇰"` |
| Filipinas                  |            `"PH"` | `"🇵🇭"` |
| Polônia                    |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Porto Rico                 |            `"PR"` | `"🇵🇷"` |
| Romênia                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Arábia Saudita             |            `"SA"` | `"🇸🇦"` |
| Sérvia                     |            `"RS"` | `"🇷🇸"` |
| Singapura                  |            `"SG"` | `"🇸🇬"` |
| Eslováquia                 |            `"SK"` | `"🇸🇰"` |
| Eslovênia                  |            `"SI"` | `"🇸🇮"` |
| África do Sul              |            `"ZA"` | `"🇿🇦"` |
| Coreia do Sul              |            `"KR"` | `"🇰🇷"` |
| Espanha                    |            `"ES"` | `"🇪🇸"` |
| Suécia                     |            `"SE"` | `"🇸🇪"` |
| Suíça                      |            `"CH"` | `"🇨🇭"` |
| Tailândia                  |            `"TH"` | `"🇹🇭"` |
| Trinidad e Tobago          |            `"TT"` | `"🇹🇹"` |
| Turquia                    |            `"TR"` | `"🇹🇷"` |
| Ucrânia                    |            `"UA"` | `"🇺🇦"` |
| Emirados Árabes Unidos     |            `"AE"` | `"🇦🇪"` |
| Estados Unidos             |            `"US"` | `"🇺🇸"` |
| Uruguai                    |            `"UY"` | `"🇺🇾"` |
| Uzbequistão                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` deve retornar `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` deve retornar `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` deve retornar `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` deve retornar `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` deve retornar `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` deve retornar `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` deve retornar `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` deve retornar `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` deve retornar `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` deve retornar `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` deve retornar `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` deve retornar `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` deve retornar `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` deve retornar `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` deve retornar `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` deve retornar `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` deve retornar `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` deve retornar `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` deve retornar `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` deve retornar `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` deve retornar `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` deve retornar `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` deve retornar `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` deve retornar `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` deve retornar `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` deve retornar `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` deve retornar `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` deve retornar `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` deve retornar `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` deve retornar `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` deve retornar `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` deve retornar `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` deve retornar `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` deve retornar `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` deve retornar `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` deve retornar `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` deve retornar `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` deve retornar `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` deve retornar `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` deve retornar `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` deve retornar `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` deve retornar `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` deve retornar `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` deve retornar `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` deve retornar `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` deve retornar `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` deve retornar `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` deve retornar `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` deve retornar `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` deve retornar `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` deve retornar `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` deve retornar `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` deve retornar `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` deve retornar `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` deve retornar `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` deve retornar `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` deve retornar `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` deve retornar `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` deve retornar `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` deve retornar `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` deve retornar `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` deve retornar `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` deve retornar `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` deve retornar `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` deve retornar `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` deve retornar `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` deve retornar `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` deve retornar `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` deve retornar `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` deve retornar `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` deve retornar `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` deve retornar `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` deve retornar `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` deve retornar `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` deve retornar `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` deve retornar `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` deve retornar `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` deve retornar `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` deve retornar `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` deve retornar `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` deve retornar `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` deve retornar `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` deve retornar `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` deve retornar `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` deve retornar `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` deve retornar `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` deve retornar `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` deve retornar `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` deve retornar `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` deve retornar `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` deve retornar `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
