---
id: 697a49e9860d24853adef67b
title: "Desafio 190: Jogos de Inverno 2026 Dia 11: Hóquei no Gelo"
challengeType: 28
dashedName: challenge-190
---

# --description--

Dado um array de 6 times de hóquei no gelo e seus registros após os jogos do round robin, determine os confrontos para a rodada semifinal.

- Each array item will have a team and their record in the format `"TEAM: W-OTW-OTL-L"`. Where:
  - "W" é o número de vitórias no tempo regulamentar, valendo 3 pontos cada
  - "OTW" é o número de vitórias em prorrogação, valendo 2 pontos cada
  - "OTL" é o número de derrotas na prorrogação, valendo 1 ponto cada uma
  - "L" é o número de derrotas, valendo 0 pontos cada
  
Por exemplo, `"FIN: 2-2-1-0"` teria 11 pontos após somar seu registro.
  
Encontre o número total de pontos para cada time e retorne `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. Por exemplo, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` deve retornar `"Os jogos da semi-final serão FIN vs SWE e CAN vs USA."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` deve retornar `"Os jogos da semi-final serão USA vs CZE e CAN vs FIN."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` deve retornar `"Os jogos da semi-final serão CAN vs ITA e USA vs CZE."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` deve retornar `"Os jogos da semi-final serão JPN vs ITA e AUT vs KOR."`.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
