---
id: 697a49e6ff50d756c9b69360
title: "Desafio 183: Jogos de Inverno 2026 Dia 4: Salto de Esqui"
challengeType: 28
dashedName: challenge-183
---

# --description--

Dadas as distâncias, os pontos de estilo, um valor de compensação do vento e um valor de bônus do ponto K, calcule sua pontuação para o salto de esqui e determine se você ganhou uma medalha ou não.

- Sua pontuação é calculada somando os quatro valores acima.

As pontuações totais atuais dos outros saltadores são:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Se sua pontuação for a melhor, retorne `"Gold"`
- Se for o segundo melhor, retorne `"Silver"`
- Se for o terceiro melhor, retorne `"Bronze"`
- Caso contrário, retorne `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` deve retornar `"Gold"`.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` deve retornar `"Bronze"`.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` deve retornar `"Silver"`.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` deve retornar `"No Medal"`.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` deve retornar `"Gold"`.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` deve retornar `"No Medal"`.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
