---
id: 698a1a73ade5ac0e19180fa7
title: "Desafio 204: Some as Letras"
challengeType: 28
dashedName: challenge-204
---

# --description--

Dada uma `string`, retorne a soma de suas letras.

- Letras são A-Z em maiúsculas ou minúsculas
- Os valores das letras são: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Letras maiúsculas e minúsculas têm o mesmo valor.
- Ignore todos os caracteres que não sejam letras.

# --hints--

`sumLetters("Hello")` deve retornar `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` deve retornar `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` deve retornar `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` deve retornar `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` deve retornar `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
