---
id: 698a1a73ade5ac0e19180fa1
title: "Desafio 198: Contagem de Dias Úteis"
challengeType: 28
dashedName: challenge-198
---

# --description--

Dada uma data de início e uma data de término, retorne o número de dias úteis entre as duas.

- As datas fornecidas estão no formato `"YYYY-MM-DD"`.
- Dias úteis são dias de trabalho (segunda a sexta-feira).
- Finais de semana não são dias úteis (sábado e domingo).
- Inclua tanto as datas de início e término ao contar.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` deve retornar `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` deve retornar `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` deve retornar `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` deve retornar `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` deve retornar `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
