---
id: 698a1a73ade5ac0e19180fa0
title: "Desafio 197: Compatibilidade de Tipo Sanguíneo"
challengeType: 28
dashedName: challenge-197
---

# --description--

Dado um tipo sanguíneo do doador e um tipo sanguíneo do receptor, determine se o doador pode doar sangue para o receptor.

Cada tipo sanguíneo consiste em:

- Uma letra: `"A"`, `"B"`, `"AB"` ou `"O"`
- E um fator Rh: `"+"` ou `"-"`

Os tipos sanguíneos serão uma das letras válidas seguidas por um fator Rh. Por exemplo, `"AB+"` e `"O-"` são tipos sanguíneos válidos.

Regras de Letras:

- `"O"` pode doar para outro tipo de letra.
- `"A"` pode doar para `"A"` e `"AB"`.
- `"B"` pode doar para `"B"` e `"AB"`.
- `"AB"` pode doar apenas para `"AB"`.

Regras Rh:

- Negativo (`"-"`) pode doar para ambos `"-"` e `"+"`.
- Positivo (`"+"`) pode doar apenas para `"+"`.

Tanto a regra da letra quanto a regra Rh devem ser aprovadas para que um doador possa doar para o receptor.

# --hints--

`canDonate("B+", "B+")` deve retornar `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` deve retornar `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` deve retornar `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` deve retornar `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` deve retornar `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` deve retornar `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` deve retornar `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` deve retornar `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` deve retornar `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` deve retornar `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
