---
id: 698a1a73ade5ac0e19180fa3
title: "Desafio 200: Contagem de Letras e Números"
challengeType: 28
dashedName: challenge-200
---

# --description--

Dada uma `string`, retorne uma mensagem com a contagem de quantas letras e números ela contém.

- As letras são `A-Z` e `a-z`.
- Números são `0-9`.
- Ignore todos os outros caracteres.

Retorne `"The string has X letters and Y numbers."`, onde `"X"` é a contagem de letras e `"Y"` é a contagem de números. Se qualquer contagem for 1, use a forma singular para esse item. Exemplo: `"1 letter"` em vez de `"1 letters"` e `"1 number"` em vez de `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` deve retornar `"The string has 10 letters and 3 numbers."`.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` deve retornar `"The string has 5 letters and 2 numbers."`.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` deve retornar `"The string has 1 letter and 1 number."`.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` deve retornar `"The string has 0 letters and 5 numbers."`.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` deve retornar `"The string has 8 letters and 0 numbers."`.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
