---
id: 698a1a863194f1f4e63f6460
title: "Desafio 209: Tamanho do elemento"
challengeType: 28
dashedName: challenge-209
---

# --description--

Dado um tamanho de janela, a largura de um elemento em unidades de largura do viewport `"vw"` e a altura de um elemento em unidades de altura do viewport `"vh"`, determine o tamanho do elemento em pixels.

- O tamanho da janela fornecido e o tamanho do elemento retornado são strings no formato `"width x height"`, `"1200 x 800"` por exemplo.

- As unidades `"vw"` são a porcentagem da largura da janela. `"50vw"`, por exemplo, é 50% da largura da janela.
- As unidades `"vh"` são a porcentagem da altura da janela. `"50vh"`, por exemplo, é 50% da altura da janela.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` deve retornar `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` deve retornar `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` deve retornar `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` deve retornar `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` deve retornar `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` deve retornar `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
