---
id: 698a1a73ade5ac0e19180fa8
title: "Desafio 205: Contagem de Cubos Perfeitos"
challengeType: 28
dashedName: challenge-205
---

# --description--

Dado dois inteiros, determine quantos cubos perfeitos existem no intervalo entre e incluindo os dois números.

- O número menor não é garantido para ser o primeiro argumento.
- Um número é um cubo perfeito se existir um inteiro (`n`) onde `n * n * n = number`. Por exemplo, 27 é um cubo perfeito porque `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` deve retornar `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` deve retornar `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` deve retornar `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` deve retornar `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` deve retornar `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
