---
id: 698a1a73ade5ac0e19180fa9
title: "Desafio 206: Valores das Cartas de Baralho"
challengeType: 28
dashedName: challenge-206
---

# --description--

Dado um array de cartas de baralho, retorne um novo array com o valor numérico de cada carta.

Valores das Cartas:

- Um Ás (`"A"`) tem um valor de 1.
- Cartas numeradas (`"2"` - `"10"`) têm seu valor facial: 2 - 10, respectivamente.
- Cartas de figura: Valete (`"J"`), Dama (`"Q"`) e Rei (`"K"`) valem 10 cada.

Trajes:

- Cada carta tem um naipe: Espadas (`"S"`), Paus (`"C"`), Ouros (`"D"`), ou Copas (`"H"`).

Formato de Cartão:

- Cada carta é representada como uma string: `"valueSuit"`. Por exemplo: `"AS"` é o Ás de Espadas, `"10H"` é o Dez de Copas e `"QC"` é a Dama de Paus.

# --hints--

`cardValues(["3H", "4D", "5S"])` deve retornar `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` deve retornar `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` deve retornar `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` deve retornar `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` deve retornar `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
