---
id: 698a1a863194f1f4e63f645e
title: "Desafio 207: Menor Gap"
challengeType: 28
dashedName: challenge-207
---

# --description--

Dada uma string, retorne a substring entre os dois caracteres idênticos que têm o menor número de caracteres entre eles (menor gap).

- Sempre haverá pelo menos um par de caracteres correspondentes.
- A substring retornada deve excluir os caracteres correspondentes.
- Se dois ou mais gaps tiverem o mesmo comprimento, retorne os caracteres do primeiro.

Por exemplo, dado `"ABCDAC"`, retorne `"DA"`.

- Apenas `"A"` e `"C"` se repetem na `string`.
- O número de caracteres entre os dois `"A"` é 3 e entre os caracteres `"C"` é 2.
- Então retorne a string entre os dois caracteres `"C"`.

# --hints--

`smallestGap("ABCDAC")` deve retornar `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` deve retornar `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` deve retornar `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` deve retornar `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` deve retornar `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
