---
id: 698a1a73ade5ac0e19180fa6
title: "Desafio 203: Flattened"
challengeType: 28
dashedName: challenge-203
---

# --description--

Dado um array, determine se ele é plano.

- Um array é plano se nenhum de seus elementos for um array.

# --hints--

`isFlat([1, 2, 3, 4])` deve retornar `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` deve retornar `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` deve retornar `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` deve retornar `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` deve retornar `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
