---
id: 698a1a863194f1f4e63f645f
title: "Desafio 208: Travessia da Trilha"
challengeType: 28
dashedName: challenge-208
---

# --description--

Dado um array de strings que representam seu map de trilha, retorne uma string dos movimentos necessários para chegar ao seu objetivo.

As strings fornecidas conterão os valores:

- `"C"`: Sua localização atual
- `"G"`: Seu objetivo
- `"T"`: Partes percorríveis da trilha
- `"-"`: Partes intransitáveis do map

Retorne uma string com os movimentos necessários para seguir a trilha do seu local até seu objetivo onde:

- `"R"` é um movimento para a direita
- `"D"` é um movimento para baixo
- `"L"` é um movimento para a esquerda
- `"U"` é um movimento para cima

- Sempre haverá uma única trilha contínua, sem ramificações, do seu local atual até seu objetivo.
- Cada local da trilha terá no máximo dois locais atravessáveis tocando-o.

Por exemplo, dado:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Retorne `"RDDRDD"`.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` deve retornar `"RDDRDD"`.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` deve retornar `"RRUUURR"`.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` deve retornar `"DLDDRRRRD"`.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` deve retornar `"RRRDDL"`.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` deve retornar `"ULLLUUURRRRRRDDDR"`.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```
