---
id: 6994cff2290543b3aec9f510
title: "Desafio 211: Soma de Array"
challengeType: 28
dashedName: challenge-211
---

# --description--

Dado um array de números, retorne a soma de todos os números.

# --hints--

`sumArray([1, 2, 3, 4, 5])` deve retornar `15`.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` deve retornar `42`.

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` deve retornar `7`.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` deve retornar `9395`.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` deve retornar `0`.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
