---
id: 6994cff2290543b3aec9f50f
title: "Desafio 210: Validador HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Dada uma string, determine se é um valor válido de cor CSS `hsl()`.

- Um valor HSL válido deve estar no formato `"hsl(h, s%, l%)"`, onde:
  - `h` (matiz) deve ser um número entre 0 e 360 (inclusive).
  - `s` (saturation) deve ser uma porcentagem entre 0% e 100%.
  - `l` (lightness) deve ser uma porcentagem entre 0% e 100%.

- Espaços são permitidos apenas:
  - Após o parêntese de abertura
  - Antes e/ou depois das vírgulas
  - Antes e/ou depois do parêntese de fechamento

- Opcionalmente, o valor pode terminar com um ponto e vírgula (`";"`).

Por exemplo, `"hsl(240, 50%, 50%)"` é um valor HSL válido.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` deve retornar `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` deve retornar `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` deve retornar `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` deve retornar `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` deve retornar `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` deve retornar `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` deve retornar `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` deve retornar `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` deve retornar `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
