---
id: 699c8e045ee7cb94ed2322dc
title: "Desafio 221: Matriz Invertida"
challengeType: 28
dashedName: challenge-221
---

# --description--

Dada uma matriz (um array de arrays) preenchida com dois valores distintos, retorne uma nova matriz onde todas as ocorrências de um valor são trocadas pelo outro.

Por exemplo, dado:

```js
[
  ["a", "b"],
  ["a", "a"]
]
```

Retornar:

```js
[
  ["b", "a"],
  ["b", "b"]
]
```

# --hints--

`invertMatrix([["a", "b"], ["a", "a"]])` deve retornar `[["b", "a"], ["b", "b"]]`.

```js
assert.deepEqual(invertMatrix([["a", "b"], ["a", "a"]]), [["b", "a"], ["b", "b"]]);
```

`invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]])` deve retornar `[[0, 1, 0], [0, 0, 0], [1, 0, 1]]`.

```js
assert.deepEqual(invertMatrix([[1, 0, 1], [1, 1, 1], [0, 1, 0]]), [[0, 1, 0], [0, 0, 0], [1, 0, 1]]);
```

`invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]])` deve retornar `[["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]`.

```js
assert.deepEqual(invertMatrix([["apple", "banana", "banana", "apple"], ["banana", "apple", "apple", "banana"], ["banana", "banana", "banana", "apple"]]), [["banana", "apple", "apple", "banana"], ["apple", "banana", "banana", "apple"], ["apple", "apple", "apple", "banana"]]);
```

`invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]])` deve retornar `[[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]`.

```js
assert.deepEqual(invertMatrix([[6, 7, 7, 7, 6], [7, 6, 7, 6, 7], [7, 7, 6, 7, 7], [7, 6, 7, 6, 7], [6, 7, 7, 7, 6]]), [[7, 6, 6, 6, 7], [6, 7, 6, 7, 6], [6, 6, 7, 6, 6], [6, 7, 6, 7, 6], [7, 6, 6, 6, 7]]);
```

`invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]])` deve retornar `[[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]`.

```js
assert.deepEqual(invertMatrix([[1.2, 2.1, 2.1, 2.1], [2.1, 1.2, 2.1, 1.2], [1.2, 1.2, 2.1, 2.1]]), [[2.1, 1.2, 1.2, 1.2], [1.2, 2.1, 1.2, 2.1], [2.1, 2.1, 1.2, 1.2]]);
```

# --seed--

## --seed-contents--

```js
function invertMatrix(matrix) {

  return matrix;
}
```

# --solutions--

```js
function invertMatrix(matrix) {
  if (matrix.length === 0) return [];

  const values = [...new Set(matrix.flat())];
  const [val1, val2] = values;

  return matrix.map(row =>
    row.map(cell => (cell === val1 ? val2 : val1))
  );
}
```
