---
id: 699c8e045ee7cb94ed2322d5
title: "Desafio 214: Validador de Corrente de Dominó"
challengeType: 28
dashedName: challenge-214
---

# --description--

Dado um array 2D que representa uma sequência de dominós, determine se ele forma uma cadeia válida.

- Cada elemento no array representa um dominó e será um array de dois números de 1 a 6 (inclusive).
- Para a cadeia ser válida, o segundo número de cada dominó deve corresponder ao primeiro número do próximo dominó.
- O primeiro número do primeiro dominó e o último número do último dominó não precisam corresponder a nada.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` deve retornar `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` deve retornar `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` deve retornar `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` deve retornar `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` deve retornar `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
