---
id: 699c8e045ee7cb94ed2322db
title: "Desafio 220: Maior Número"
challengeType: 28
dashedName: challenge-220
---

# --description--

Dada uma string de números separados por várias pontuações, retorne o maior número.

- A string fornecida conterá apenas números e separadores.
- Os separadores podem ser vírgulas (`","`), pontos de exclamação (`"!"`), pontos de interrogação (`"?"`), dois-pontos (`":"`) ou ponto e vírgula (`";"`).

# --hints--

`largestNumber("1,2")` deve retornar `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` deve retornar `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` deve retornar `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` deve retornar `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
