---
id: 699c8e045ee7cb94ed2322d6
title: "Desafio 215: Calculadora de Taxa de Estacionamento"
challengeType: 28
dashedName: challenge-215
---

# --description--

Dadas duas `strings` representando o horário em que você estacionou seu carro e o horário em que o retirou, calcule a taxa de estacionamento.

- As strings fornecidas estarão no formato `"HH:MM"` usando um relógio de 24 horas. Então `"14:00"` é 14h por exemplo.
- A primeira `string` será a hora que você estacionou seu carro e a segunda será a hora que você o pegou.
- Se o horário de retirada for anterior ao horário de entrada, isso significa que o estacionamento ultrapassou a meia-noite e entrou no dia seguinte.

Regras de taxa:

- Cada hora estacionada custa $3.
- Horas parciais são arredondadas para a próxima hora cheia.
- Se o estacionamento se estender durante a noite (após a meia-noite), uma taxa adicional de $10 por pernoite será aplicada.
- Há uma taxa mínima de $5 (usada apenas se o total for menor que $5).

Retorne o custo total no formato `"$cost"`, `"$5"` por exemplo.

# --hints--

`calculateParkingFee("09:00", "11:00")` deve retornar `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` deve retornar `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` deve retornar `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` deve retornar `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` deve retornar `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` deve retornar `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
