---
id: 699c8e045ee7cb94ed2322da
title: "Desafio 219: Marcos do Aniversário"
challengeType: 28
dashedName: challenge-219
---

# --description--

Dado um inteiro representando o número de anos que um casal está casado, retorne seu marco de aniversário mais recente de acordo com este gráfico:

| Years Married | Milestone    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Se eles não alcançaram o primeiro marco, retorne `"Newlyweds"`.

# --hints--

`getMilestone(0)` deve retornar `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` deve retornar `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` deve retornar `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` deve retornar `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` deve retornar `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` deve retornar `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` deve retornar `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` deve retornar `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` deve retornar `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
