---
id: 699c8e045ee7cb94ed2322d8
title: "Desafio 217: Peças de Xadrez Capturadas"
challengeType: 28
dashedName: challenge-217
---

# --description--

Dado um array de strings que representam as peças de xadrez que você ainda tem no tabuleiro, calcule o valor das peças que seu oponente capturou.

No xadrez, você começa com 16 peças:

| Piece  | Abbreviation   | Quantity | Value |
| ------ | -------------- | -------- | ----- |
| Pawn   | `"P"`          | 8        | 1     |
| Rook   | `"R"`          | 2        | 5     |
| Knight | `"N"`          | 2        | 3     |
| Bishop | `"B"`          | 2        | 3     |
| Queen  | `"Q"`          | 1        | 9     |
| King   | `"K"`          | 1        | 0     |

- O array fornecido conterá apenas as abreviações acima.
- Qualquer uma das 16 peças que não estão incluídas no `array` fornecido foram capturadas.
- Retorne o valor total de todas as peças capturadas, a menos que...
- Se o Rei foi capturado, retorne `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` deve retornar `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` deve retornar `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` deve retornar `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` deve retornar `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` deve retornar `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` deve retornar `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` deve retornar `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
