---
id: 699c8e045ee7cb94ed2322d9
title: "Desafio 218: Divisível Uniformemente"
challengeType: 28
dashedName: challenge-218
---

# --description--

Dado dois inteiros, determine se você pode dividir o primeiro pelo segundo de forma exata.

# --hints--

`isEvenlyDivisible(4, 2)` deve retornar `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` deve retornar `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` deve retornar `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` deve retornar `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` deve retornar `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` deve retornar `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
