---
id: 699c8e045ee7cb94ed2322dd
title: "Desafio 222: Sombras do Equinócio"
challengeType: 28
dashedName: challenge-222
---

# --description--

Hoje é o equinócio, quando o sol está diretamente acima do equador e perfeitamente no zênite ao meio-dia. Dado um horário, determine a sombra projetada por um poste vertical de 4 pés.

- O horário será uma string no formato 24 horas `"HH:MM"` (por exemplo, `"15:00"` é 15h).
- Você receberá apenas um horário em incrementos de 30 minutos.

Regras:

- O sol nasce às 6h diretamente `"east"`, e se põe às 18h diretamente `"west"`.
- Uma sombra sempre aponta na direção oposta ao sol.
- O comprimento da sombra (em pés) é o número de horas distante do meio-dia, elevado ao cubo.
- Não há sombra antes do nascer do sol (antes das 6h), depois do pôr do sol (6h ou mais tarde) ou ao meio-dia.

Retornar:

- Se uma sombra existir, retorne `"(length)ft (direction)"`. Por exemplo, `"8ft west"`.
- Caso contrário, retorne `"No shadow"`.

Por exemplo, dado `"10:00"`, retorne `"8ft west"` porque 10h é 2 horas antes do meio-dia, então 2<sup>3</sup> = 8 pés, e a sombra aponta para o oeste porque o sol está no leste às 10h.

# --hints--

`getShadow("10:00")` deve retornar `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` deve retornar `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` deve retornar `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` deve retornar `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` deve retornar `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` deve retornar `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` deve retornar `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` deve retornar `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` deve retornar `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
