---
id: 69a890af247de743333bd4d1
title: "Desafio 228: Noite de Filme"
challengeType: 28
dashedName: challenge-228
---

# --description--

Dada uma string para o dia da semana, outra string para o horário da sessão e um número inteiro de ingressos, retorne o custo total dos ingressos de cinema para essa sessão.

O dia dado será um dos seguintes:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

O horário será dado no formato `"H:MMam"` ou `"H:MMpm"`. Por exemplo, `"10:00am"` ou `"10:00pm"`.

Retorne o custo total no formato `"$D.CC"` usando estas regras:

- Fim de semana (sexta-feira - domingo): $12.00 por ingresso.
- Dia da semana (segunda a quinta): $10.00 por ingresso.
- Matinê (antes das 17:00): subtraia $2,00 por ingresso (exceto às terças-feiras).
- Terças-feiras: todos os ingressos custam $5,00 cada.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` deve retornar `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` deve retornar `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` deve retornar `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` deve retornar `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` deve retornar `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` deve retornar `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` deve retornar `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
