---
id: 69a890af247de743333bd4cd
title: "Desafio 224: Detector de Torra de Café"
challengeType: 28
dashedName: challenge-224
---

# --description--

Dada uma `string` representando os grãos usados para fazer uma xícara de café, determine a torra da xícara.

- A string fornecida conterá os seguintes caracteres, cada um representando um tipo de feijão:

  - Um apóstrofo (`'`) é um grão de torra clara que vale 1 ponto cada.
  - Um traço (`-`) é um grão de torra médio que vale 2 pontos cada.
  - Um ponto (`.`) é um grão de torra escura que vale 3 pontos cada.

- O nível de torra é determinado pela média de todos os grãos.

Retornar:

- `"Light"` se a média for menor que 1.75.
- `"Medium"` se a média estiver entre 1.75 e 2.5.
- `"Dark"` se a média for maior que 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` deve retornar `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` deve retornar `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` deve retornar `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` deve retornar `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` deve retornar `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` deve retornar `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` deve retornar `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
