---
id: 69a890af247de743333bd4d0
title: "Desafio 227: Tempo de Resfriamento"
challengeType: 28
dashedName: challenge-227
---

# --description--

Dadas duas marcas de tempo, a primeira representando quando um usuário terminou um exame e a segunda representando o tempo atual, determine se o usuário pode fazer um exame novamente.

- Ambos os timestamps terão o formato: `"YYYY-MM-DDTHH:MM:SS"`, por exemplo `"2026-03-25T14:00:00"`. Note que o horário está no formato de 24 horas.
- Um usuário deve esperar pelo menos 48 horas antes de refazer um exame.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` deve retornar `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` deve retornar `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` deve retornar `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` deve retornar `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
