---
id: 69b1028d6e265413d0198a2f
title: "Desafio 236: histórico do navegador"
challengeType: 28
dashedName: challenge-236
---

# --description--

Dado um array de comandos do navegador, retorne um array com dois valores: o histórico como um array de URLs e o índice da página atual.

Os comandos válidos são:

- `"URL"` - onde URL é um endereço web (`"freecodecamp.org"` por exemplo). Navega para a URL dada, adiciona-a ao histórico na próxima posição e descarta qualquer histórico à frente.
- `"Back"` - move para a página anterior no histórico, ou permanece na página atual se não houver uma anterior.
- `"Forward"` - move para a próxima página no histórico, ou permanece na página atual se não houver uma próxima.

Por exemplo, dado `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`, retorne `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` deve retornar `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` deve retornar `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` deve retornar `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` deve retornar `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` deve retornar `[["example.com", "example.com/about"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` deve retornar `[["example.com", "example.com/about"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
