---
id: 69b1028d6e265413d0198a2c
title: "Desafio 233: alarme para acordar"
challengeType: 28
dashedName: challenge-233
---

# --description--

Dada uma string que representa o horário em que você programou seu alarme e uma string que representa o horário em que você realmente acordou, determine se você acordou cedo, na hora ou atrasado.

- Ambos os horários serão fornecidos no formato `"HH:MM"` de 24 horas.

Retorne:

- `"early"` se você acordou antes do horário do alarme.
- `"on time"` se você acordou no horário do alarme ou dentro da janela de soneca de 10 minutos após o horário do alarme.
- `"late"` se você acordou mais de 10 minutos depois do horário do alarme.

Ambos os horários são no mesmo dia.

# --hints--

`alarmCheck("07:00", "06:45")` deve retornar `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` deve retornar `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` deve retornar `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` deve retornar `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` deve retornar `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` deve retornar `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` deve retornar `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
