---
id: 69b1028d6e265413d0198a2a
title: "Desafio 231: validador de ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Dada uma string, determine se ela é um ISBN-10 válido.

Um ISBN-10 consiste em hífens (`"-"`) e 10 outros caracteres. Após remover os hífens (`"-"`):

- Os primeiros 9 caracteres devem ser dígitos, e
- O caractere final pode ser um dígito ou a letra `"X"`, que representa o número 10.

Para validá-lo:

- Multiplique cada dígito (ou valor) pela sua posição (multiplique o primeiro dígito por 1, o segundo por 2 e assim por diante).
- Some todos os resultados.
- Se o total for divisível por 11, é válido.

# --hints--

`isValidIsbn10("0-306-40615-2")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` deve retornar `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` deve retornar `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
