---
id: 69b1028d6e265413d0198a29
title: "Desafio 230: linha do triângulo de Pascal"
challengeType: 28
dashedName: challenge-230
---

# --description--

Dado um inteiro `n`, retorne a `n`ª linha do triângulo de Pascal como um array.

No triângulo de Pascal, cada linha começa e termina com 1, e cada valor interno é a soma dos dois valores diretamente acima dele.

Aqui estão as primeiras 5 linhas do triângulo:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` deve retornar `[1, 4, 6, 4, 1]`.

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` deve retornar `[1, 2, 1]`.

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` deve retornar `[1]`.

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` deve retornar `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`.

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` deve retornar `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`.

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
