---
id: 69b58ce40693f140c84c855b
title: "Desafio 242: próximo número do bingo"
challengeType: 28
dashedName: challenge-242
---

# --description--

Dado um número de bingo, retorne o próximo número de bingo na sequência.

Um número de bingo é uma única letra seguida por um número dentro do seu intervalo conforme este gráfico:

| Letra | Intervalo de números |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Por exemplo, dado `"B10"`, retorne `"B11"`, o próximo número de bingo. Se for dado o último número de bingo, retorne `"B1"`.

# --hints--

`getNextBingoNumber("B10")` deve retornar `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` deve retornar `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` deve retornar `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` deve retornar `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` deve retornar `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
