---
id: 69b58ce40693f140c84c8559
title: "Desafio 240: caracteres de palíndromo"
challengeType: 28
dashedName: challenge-240
---

# --description--

Dada uma string, determine se ela é um palíndromo e retorne o caractere do meio (se o comprimento for ímpar) ou os dois caracteres do meio (se for par).

- Um palíndromo é uma string que é igual da frente para trás e de trás para frente.
- Se não for um palíndromo, retorne `"none"`.

# --hints--

`palindromeLocator("racecar")` deve retornar `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` deve retornar `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` deve retornar `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` deve retornar `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` deve retornar `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
