---
id: 69b559d2903b9e4afe9075f9
title: "Desafio 239: que dia é hoje?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Dado um timestamp Unix em milissegundos, retorne o dia da semana.

Os dias válidos para retorno são:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Certifique-se de ignorar os fusos horários.

# --hints--

`getDayOfWeek(1775492249000)` deve retornar `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` deve retornar `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` deve retornar `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` deve retornar `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` deve retornar `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
