---
id: 69b5b2be76ec8135a7fbe973
title: "Desafio 244: ataque da torre e do bispo"
challengeType: 28
dashedName: challenge-244
---

# --description--

Dadas duas strings, uma para a localização de uma torre em um tabuleiro de xadrez e outra para a localização de um bispo, determine se uma peça pode atacar a outra.

Um tabuleiro de xadrez padrão tem 8x8 casas, com colunas rotuladas de `A` até `H` (da esquerda para a direita) e linhas rotuladas de `1` até `8` (de baixo para cima). Ele se parece com isto:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Torres podem se mover quantas casas quiserem na direção horizontal ou vertical.
- Bispos podem se mover quantas casas quiserem em qualquer direção diagonal.
- Uma peça pode atacar outra se puder se mover para a localização dessa peça.

Retorne:

- `"rook"` se a torre puder atacar o bispo.
- `"bishop"` se o bispo puder atacar a torre.
- `"neither"` se nenhuma das peças puder atacar a outra.

# --hints--

`rookBishopAttack("A1", "A5")` deve retornar `"rook"`.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` deve retornar `"bishop"`.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` deve retornar `"rook"`.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` deve retornar `"bishop"`.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` deve retornar `"neither"`.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` deve retornar `"neither"`.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
