---
id: 69b559d2903b9e4afe9075f8
title: "Desafio 238: escape de rotação de dígito"
challengeType: 28
dashedName: challenge-238
---

# --description--

Dado um número inteiro positivo, determine se ele, ou qualquer uma de suas rotações, é divisível de forma exata pela quantidade de seus dígitos.

Uma rotação significa mover o primeiro dígito para o final. Por exemplo, após 1 rotação, 123 se torna 231.

- Verifique a rotação `0` (o número dado) primeiro.
- Os números fornecidos não conterão zeros.
- Retorne o primeiro número de rotação encontrado, ou `"none"` se não houver nenhum.

# --hints--

`getRotation(123)` deve retornar `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` deve retornar `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` deve retornar `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` deve retornar `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
