---
id: 69bc6cb30c1d112a2e110a04
title: "Desafio 247: última letra"
challengeType: 28
dashedName: challenge-247
---

# --description--

Dada uma string, retorne a letra da string que aparece por último no alfabeto.

- Se duas ou mais letras empatarem por serem as últimas no alfabeto, retorne a primeira delas.
- Ignore todos os caracteres que não sejam letras.

# --hints--

`getLastLetter("world")` deve retornar `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` deve retornar `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` deve retornar `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` deve retornar `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` deve retornar `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
