---
id: 69bc6cb30c1d112a2e110a06
title: "Desafio 249: matemática com string"
challengeType: 28
dashedName: challenge-249
---

# --description--

Dada uma string com números e outros caracteres, realize operações matemáticas nos números com base na contagem de caracteres não numéricos entre eles.

- Se a quantidade de caracteres separando dois números for par, use adição.
- Se for ímpar, use subtração.
- Dígitos consecutivos formam um único número.
- As operações são aplicadas da esquerda para a direita.
- Ignore caracteres iniciais e finais que não sejam dígitos.

Por exemplo, dada `"3ab10c8"`, retorne `5`. Some 3 e 10 para obter 13 porque há um número par de caracteres entre eles. Depois subtraia 8 de 13 porque há um número ímpar de caracteres entre o resultado e 8.

# --hints--

`doMath("3ab10c8")` deve retornar `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` deve retornar `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` deve retornar `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` deve retornar `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` deve retornar `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
