---
id: 69bc6cb30c1d112a2e110a03
title: "Desafio 246: iniciais do nome"
challengeType: 28
dashedName: challenge-246
---

# --description--

Dado um nome completo como uma string, retorne suas iniciais.

- Os nomes para inicializar são separados por um espaço.
- As iniciais devem ser maiúsculas.
- As iniciais devem ser separadas por pontos.

Por exemplo, `"Tommy Millwood"` retorna `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` deve retornar `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` deve retornar `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` deve retornar `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` deve retornar `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` deve retornar `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
