---
id: 69c5f3d787b1725d5f00c8ba
title: "Desafio 256: direção do tempo mais próxima"
challengeType: 28
dashedName: challenge-256
---

# --description--

Dadas duas horas, determine se é mais rápido ir da primeira para a segunda avançando ou retrocedendo.

- As horas são dadas no formato 24 horas (`"HH:MM"`)
- O relógio é circular (23:59 vai para 00:00 ao avançar, e 00:00 vai para 23:59 ao retroceder)

Retorne:

- `"forward"` se avançar for mais rápido
- `"backward"` se retroceder for mais rápido
- `"equal"` se ambas as direções levarem o mesmo tempo

# --hints--

`getDirection("10:00", "12:00")` deve retornar `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` deve retornar `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` deve retornar `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` deve retornar `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` deve retornar `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` deve retornar `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
