---
id: 69c5f3d787b1725d5f00c8bd
title: "Desafio 259: explosão fizzbuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--

Dado um inteiro, retorne o número de passos necessários para transformar a palavra `"fizzbuzz"` em uma string com pelo menos o número dado de `"z"` usando as seguintes regras:

- Comece com a string `"fizzbuzz"`.
- A cada passo, aplique as regras padrão do FizzBuzz usando a posição da letra na string (a primeira `"f"` está na posição 1).
  - Se a posição da letra for divisível por 3, substitua a letra por `"fizz"`
  - Se for divisível por 5, substitua a letra por `"buzz"`
  - Se for divisível por 3 e 5, substitua a letra por `"fizzbuzz"`

Então, após 1 passo, `"fizzbuzz"` se transforma em `"fifizzzbuzzfizzzz"`, que tem 9 `"z"`'s.

# --hints--

`explodeFizzbuzz(9)` deve retornar `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` deve retornar `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` deve retornar `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` deve retornar `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` deve retornar `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` deve retornar `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` deve retornar `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` deve retornar `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
