---
id: 69c5f3d787b1725d5f00c8b8
title: "Desafio 254: palavras com número ímpar de letras"
challengeType: 28
dashedName: challenge-254
---

# --description--

Dada uma string de palavras, retorne apenas as palavras com um número ímpar de letras.

- As palavras na string dada estarão separadas por um único espaço.
- Retorne as palavras separadas por um único espaço.

# --hints--

`getOddWords("This is a super good test")` deve retornar `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` deve retornar `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` deve retornar `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` deve retornar `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
