---
id: 69c5f3d787b1725d5f00c8b9
title: "Desafio 255: equipe de limpeza do dia da terra"
challengeType: 28
dashedName: challenge-255
---

# --description--

Hoje é o dia da terra. Dado um array de itens que você limpou, retorne sua pontuação total de limpeza com base nas regras abaixo.

Os itens dados serão um dos seguintes:

| Item | Valor base |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

Um item raro é representado como `["rare", value]`. Por exemplo, `["rare", 80]`. Itens raros não recebem bônus de sequência.

- Bônus de sequência: se o mesmo item aparecer consecutivamente, ele recebe pontos de bônus crescentes.
  - Primeira ocorrência consecutiva: valor base
  - Segunda: valor base + 1
  - Terceira: valor base + 2
  - etc.
- Multiplicador do quinto item: todo quinto item coletado recebe um multiplicador.
  - Quinto item: *2
  - Décimo item: *3
  - etc.

- Aplique o multiplicador após calcular quaisquer bônus.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` deve retornar `44`.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` deve retornar `58`.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` deve retornar `79`.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` deve retornar `358`.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` deve retornar `383`.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
