---
id: 69c6ff713a52713463aa7929
title: "Desafio 263: palavras cruzadas binárias"
challengeType: 28
dashedName: challenge-263
---

# --description--

Dado um caractere, determine se sua representação binária de 8 bits pode ser encontrada na grade a seguir, horizontalmente ou verticalmente em qualquer direção:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Por exemplo, `"A"` tem a representação binária `01000001`, que aparece na primeira linha da esquerda para a direita.

# --hints--

`isInCrossword("I")` deve retornar `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` deve retornar `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` deve retornar `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` deve retornar `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` deve retornar `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` deve retornar `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` deve retornar `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` deve retornar `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
