---
id: 69c6ff713a52713463aa7928
title: "Desafio 262: analisador de consulta de URL"
challengeType: 28
dashedName: challenge-262
---

# --description--

Dada uma URL que contém uma string de consulta, analise a string de consulta em um objeto (ou dicionário) de pares chave-valor.

- A string de consulta começa após o `"?"`,
- cada parâmetro é separado por `"&"`,
- cada par chave/valor é separado por `"="`

Por exemplo, dado `"https://example.com/search?name=Alice&age=30"`, retorne:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Todos os valores devem ser retornados como strings.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` deve retornar `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` deve retornar `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` deve retornar `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` deve retornar `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
