---
id: 69c6ff713a52713463aa7927
title: "Desafio 261: palavras numéricas"
challengeType: 28
dashedName: challenge-261
---

# --description--

Dado um inteiro de 0 a 99, retorne sua representação em palavras em inglês.

- `0` retorna `"zero"`.
- Os números de 1 a 19 têm nomes únicos (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Múltiplos de 10 de 20 a 90 têm seus próprios nomes (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Números de 21 a 99 que não são múltiplos de 10 são escritos como duas palavras unidas por um hífen. Por exemplo `"forty-two"` e `"fifty-three"`.

# --hints--

`getNumberWords(0)` deve retornar `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` deve retornar `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` deve retornar `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` deve retornar `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` deve retornar `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` deve retornar `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` deve retornar `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` deve retornar `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` deve retornar `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` deve retornar `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
