---
id: 69cfca90e8a0a6d4d6871c51
title: "Desafio 267: conversor de parsec"
challengeType: 28
dashedName: challenge-267
---

# --description--

Em uma galáxia distante, parsecs são usados para medir tanto tempo quanto distância. Dado um número inteiro de parsecs, retorne seu equivalente em tempo ou distância.

- Se o número inteiro dado for ímpar, ele representa tempo. Se for par, representa distância.

Use estas taxas de conversão:

| Parsecs | Tempo/Distância |
| - | - |
| 1 | 2 horas |
| 2 | 6 anos-luz |

Retorne o valor convertido como um inteiro.

# --hints--

`convertParsecs(1)` deve retornar `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` deve retornar `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` deve retornar `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` deve retornar `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` deve retornar `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` deve retornar `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
