---
id: 69cfca90e8a0a6d4d6871c50
title: "Desafio 266: bom dia"
challengeType: 28
dashedName: challenge-266
---

# --description--

Dada uma string de horário no formato `"HH:MM"` (relógio de 24 horas), retorne:

- `"Good morning"` para horários de `05:00` a `11:59`
- `"Good afternoon"` para horários de `12:00` a `17:59`
- `"Good evening"` para horários de `18:00` a `21:59`
- `"Good night"` para horários de `22:00` a `04:59`

# --hints--

`getGreeting("06:30")` deve retornar `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` deve retornar `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` deve retornar `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` deve retornar `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` deve retornar `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
