---
id: 69cfca90e8a0a6d4d6871c54
title: "Desafio 270: maior substring comum"
challengeType: 28
dashedName: challenge-270
---

# --description--

Dada uma string, retorne a maior substring que aparece mais de uma vez.

- As substrings podem se sobrepor.

# --hints--

`getLongestSubstring("abracadabra")` deve retornar `"abra"`.

```js
assert.equal(getLongestSubstring("abracadabra"), "abra");
```

`getLongestSubstring("hello world hello")` deve retornar `"hello"`.

```js
assert.equal(getLongestSubstring("hello world hello"), "hello");
```

`getLongestSubstring("mississippi")` deve retornar `"issi"`.

```js
assert.equal(getLongestSubstring("mississippi"), "issi");
```

`getLongestSubstring("ha ha ha ha ha ha ha")` deve retornar `"ha ha ha ha ha ha"`.

```js
assert.equal(getLongestSubstring("ha ha ha ha ha ha ha"), "ha ha ha ha ha ha");
```

`getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over")` deve retornar `"the quick brown fox jumped over"`.

```js
assert.equal(getLongestSubstring("the quick brown fox jumped over the lazy dog that the quick brown fox jumped over"), "the quick brown fox jumped over");
```

# --seed--

## --seed-contents--

```js
function getLongestSubstring(str) {

  return str;
}
```

# --solutions--

```js
function getLongestSubstring(str) {
  let longest = '';

  for (let len = str.length - 1; len >= 1; len--) {
    for (let i = 0; i <= str.length - len; i++) {
      const sub = str.slice(i, i + len);
      if (str.indexOf(sub) !== str.lastIndexOf(sub)) {
        return sub;
      }
    }
  }

  return longest;
}
```
