---
id: 69cfca90e8a0a6d4d6871c53
title: "Desafio 269: refeições sem alérgenos"
challengeType: 28
dashedName: challenge-269
---

# --description--

Dado um array de refeições e um array de alérgenos a evitar, retorne os nomes de todas as refeições que não contenham nenhum dos alérgenos indicados.

- Cada refeição está no formato `[meal, allergens]`, onde `meal` é o nome da refeição e `allergens` é um array dos alérgenos que a refeição contém. Por exemplo, `["pasta", ["wheat", "milk"]]`.
- Os alérgenos a evitar serão um array de strings.

Retorne os nomes das refeições seguras na mesma ordem dada. Se nenhuma refeição for segura, retorne um array vazio.

# --hints--

`getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` deve retornar `["salad"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"]);
```

`getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` deve retornar `["fried rice", "chicken parmesan"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` deve retornar `["oatmeal", "granola", "toast"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"]);
```

`getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` deve retornar `["granola", "yogurt", "eggs"]`.

```js
assert.deepEqual(getAllergenFriendlyMeals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"]);
```

# --seed--

## --seed-contents--

```js
function getAllergenFriendlyMeals(meals, allergens) {

  return meals;
}
```

# --solutions--

```js
function getAllergenFriendlyMeals(meals, allergens) {
  return meals
    .filter(([name, mealAllergens]) =>
      !allergens.some(a => mealAllergens.includes(a)))
    .map(([name]) => name);
}
```
