---
id: 69cfca90e8a0a6d4d6871c4f
title: "Desafio 265: colchetes mais profundos"
challengeType: 28
dashedName: challenge-265
---

# --description--

Dada uma string contendo colchetes balanceados, retorne o conteúdo dos colchetes mais profundamente aninhados.

- Os colchetes podem ser de qualquer um dos três tipos: `()`, `[]` e `{}`.
- A entrada sempre terá um único grupo mais profundo.

Por exemplo, dada `"(hello (world))"`, retorne `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` deve retornar `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` deve retornar `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` deve retornar `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` deve retornar `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` deve retornar `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
