---
id: 69cfca90e8a0a6d4d6871c52
title: "Desafio 268: número narcisista"
challengeType: 28
dashedName: challenge-268
---

# --description--

Dado um inteiro positivo, determine se ele é um número narcisista.

- Um número é narcisista se a soma de cada um de seus dígitos elevados à potência do total de dígitos for igual ao próprio número.

Por exemplo, 153 tem 3 dígitos, e 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, então ele é narcisista.

# --hints--

`isNarcissistic(153)` deve retornar `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` deve retornar `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` deve retornar `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` deve retornar `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` deve retornar `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` deve retornar `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` deve retornar `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` deve retornar `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
