---
id: 69d03e549613fbdbe21d11e0
title: "Desafio 271: lembrete de medicação"
challengeType: 28
dashedName: challenge-271
---

# --description--

Dado um array de medicamentos e uma string representando o horário atual, retorne o próximo medicamento que você precisa tomar e quanto tempo falta para tomá-lo.

- Cada medicamento está no formato `[name, lastTaken]`, onde `name` é o nome do medicamento e `lastTaken` é o horário em que foi tomado pela última vez.
- Todos os horários fornecidos estarão no formato `"HH:MM"` (24 horas).

Use o seguinte cronograma de medicação:

| Nome | Cronograma |
| - | - |
| Deployxitrin | 08:00, 16:00 |
| Debuggamanizole | 07:00, 13:00, 21:00 |
| Mergeflictamine | a cada 4 horas |

Retorne uma string no formato `"{name} in Hh Mm"`. Por exemplo, `"Debuggamanizole in 2h 0m"` ou `"Deployxitrin in 1h 5m"`.

# --hints--

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` deve retornar `"Debuggamanizole in 2h 0m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` deve retornar `"Deployxitrin in 1h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` deve retornar `"Mergeflictamine in 0h 45m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` deve retornar `"Debuggamanizole in 0h 1m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` deve retornar `"Debuggamanizole in 0h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` deve retornar `"Mergeflictamine in 3h 30m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

# --seed--

## --seed-contents--

```js
function medicationReminder(medications, currentTime) {

  return medications;
}
```

# --solutions--

```js
function medicationReminder(medications, currentTime) {
  const schedule = {
    Deployxitrin: { type: 'fixed', times: ['08:00', '16:00'] },
    Debuggamanizole: { type: 'fixed', times: ['07:00', '13:00', '21:00'] },
    Mergeflictamine: { type: 'interval', hours: 4 }
  };

  const toMinutes = t => {
    const [h, m] = t.split(':').map(Number);
    return h * 60 + m;
  };

  const current = toMinutes(currentTime);
  let earliest = Infinity;
  let nextMed = '';

  for (const [name, lastTaken] of medications) {
    const med = schedule[name];
    let nextDue;

    if (med.type === 'fixed') {
      const futureTimes = med.times
        .map(toMinutes)
        .filter(t => t > current);
      if (futureTimes.length === 0) continue;
      nextDue = Math.min(...futureTimes);
    } else {
      nextDue = toMinutes(lastTaken) + med.hours * 60;
    }

    const diff = nextDue - current;
    if (diff > 0 && diff < earliest) {
      earliest = diff;
      nextMed = name;
    }
  }

  const h = Math.floor(earliest / 60);
  const m = earliest % 60;
  return `${nextMed} in ${h}h ${m}m`;
}
```
