---
id: 69e2383af7832c8032603b92
title: "Desafio 276: elemento fora do lugar"
challengeType: 28
dashedName: challenge-276
---

# --description--

Dado um array de inteiros que está ordenado em ordem crescente, exceto por um elemento fora do lugar, retorne o índice desse elemento.

- Se mais de um elemento puder ser considerado fora do lugar, retorne o índice do primeiro.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` deve retornar `1`.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` deve retornar `3`.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` deve retornar `0`.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` deve retornar `2`.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` deve retornar `6`.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
