---
id: 69e2383af7832c8032603b93
title: "Desafio 277: imagem espelhada"
challengeType: 28
dashedName: challenge-277
---

# --description--

Dadas duas strings, determine se a segunda string é uma imagem espelhada da primeira.

Uma imagem espelhada é formada invertendo a string e substituindo cada caractere pelo seu equivalente espelhado.

- Caracteres simétricos parecem eles mesmos no espelho:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, e o espaço (` `).

- Pares espelhados trocam entre si no espelho:

| Caractere | Troca com |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Se qualquer uma das strings incluir um caractere que não esteja nas listas acima, ela não terá uma imagem espelhada que possa ser criada a partir dos caracteres.

Por exemplo, a imagem espelhada de `"[HOW]"` é `"[WOH]"`.

# --hints--

`isMirrorImage("[HOW]", "[WOH]")` deve retornar `true`.

```js
assert.isTrue(isMirrorImage("[HOW]", "[WOH]"));
```

`isMirrorImage("MOM", "MOM")` deve retornar `true`.

```js
assert.isTrue(isMirrorImage("MOM", "MOM"));
```

`isMirrorImage("vow", "wov")` deve retornar `true`.

```js
assert.isTrue(isMirrorImage("vow", "wov"));
```

`isMirrorImage("TIM", "TIM")` deve retornar `false`.

```js
assert.isFalse(isMirrorImage("TIM", "TIM"));
```

`isMirrorImage("{WOW}", "}WOW{")` deve retornar `false`.

```js
assert.isFalse(isMirrorImage("{WOW}", "}WOW{"));
```

`isMirrorImage("XXVII", "IIV%X")` deve retornar `false`.

```js
assert.isFalse(isMirrorImage("XXVII", "IIV%X"));
```

`isMirrorImage("><(((*>", "<*)))><")` deve retornar `true`.

```js
assert.isTrue(isMirrorImage("><(((*>", "<*)))><"));
```

`isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` deve retornar `true`.

```js
assert.isTrue(isMirrorImage("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"));
```

# --seed--

## --seed-contents--

```js
function isMirrorImage(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirrorImage(str1, str2) {
  const symmetric = new Set('WTYUIOHA XVMwoxv08=+:|-_*^!.');
  const pairs = { '[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<', 'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '(' };
  
  const mirrored = str1.split('').reverse().map(c => {
    if (symmetric.has(c)) return c;
    if (pairs[c]) return pairs[c];
    return null;
  });
  
  if (mirrored.includes(null)) return false;
  return mirrored.join('') === str2;
}
```
