---
id: 69e2383af7832c8032603b8f
title: "Desafio 273: validador de ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

Dada uma string, determine se ela é um número ISBN-13 válido.

Um ISBN-13 válido:

- Contém apenas dígitos e hífens
- Tem exatamente 13 dígitos após remover os hífens
- Passa na seguinte verificação:
  1. Multiplique cada dígito por 1 ou 3, alternando (multiplique o primeiro dígito por 1, o segundo por 3, o terceiro por 1 e assim por diante).
  2. A soma dos resultados deve ser divisível por 10.

# --hints--

`isValidIsbn13("9780306406157")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` deve retornar `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` deve retornar `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` deve retornar `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
