---
id: 69e2383af7832c8032603b91
title: "Desafio 275: frequência de caracteres"
challengeType: 28
dashedName: challenge-275
---

# --description--

Dada uma string, retorne um objeto (JavaScript) ou dicionário (Python) que mapeie cada caractere para o número de vezes que ele aparece.

# --hints--

`getFrequency("test")` deve retornar `{t: 2, e: 1, s: 1}`.

```js
assert.deepEqual(getFrequency("test"), {t: 2, e: 1, s: 1});
```

`getFrequency("mississippi")` deve retornar `{m: 1, i: 4, s: 4, p: 2}`.

```js
assert.deepEqual(getFrequency("mississippi"), {m: 1, i: 4, s: 4, p: 2});
```

`getFrequency("hello world")` deve retornar `{h: 1, e: 1, l: 3, o: 2, " ": 1, w: 1, r: 1, d: 1}`.

```js
assert.deepEqual(getFrequency("hello world"), {h: 1, e: 1, l: 3, o: 2, " ": 1, w: 1, r: 1, d: 1});
```

`getFrequency("She sells seashells by the seashore.")` deve retornar `{S: 1, h: 4, e: 7, " ": 5, s: 7, l: 4, a: 2, b: 1, y: 1, t: 1, o: 1, r: 1, ".": 1}`.

```js
assert.deepEqual(getFrequency("She sells seashells by the seashore."), {S: 1, h: 4, e: 7, " ": 5, s: 7, l: 4, a: 2, b: 1, y: 1, t: 1, o: 1, r: 1, ".": 1});
```

`getFrequency("The quick brown fox jumps over the lazy dog.")` deve retornar `{T: 1, h: 2, e: 3, " ": 8, q: 1, u: 2, i: 1, c: 1, k: 1, b: 1, r: 2, o: 4, w: 1, n: 1, f: 1, x: 1, j: 1, m: 1, p: 1, s: 1, v: 1, t: 1, l: 1, a: 1, z: 1, y: 1, d: 1, g: 1, ".": 1}`.

```js
assert.deepEqual(getFrequency("The quick brown fox jumps over the lazy dog."), {T: 1, h: 2, e: 3, " ": 8, q: 1, u: 2, i: 1, c: 1, k: 1, b: 1, r: 2, o: 4, w: 1, n: 1, f: 1, x: 1, j: 1, m: 1, p: 1, s: 1, v: 1, t: 1, l: 1, a: 1, z: 1, y: 1, d: 1, g: 1, ".": 1});
```

# --seed--

## --seed-contents--

```js
function getFrequency(str) {

  return str;
}
```

# --solutions--

```js
function getFrequency(str) {
  return str.split('').reduce((acc, char) => {
    acc[char] = (acc[char] || 0) + 1;
    return acc;
  }, {});
}
```
