---
id: 69f35a5bb823ed620fcb7cbb
title: "Desafio 282: dívida de sono"
challengeType: 28
dashedName: challenge-282
---

# --description--

Dado um array com as horas dormidas a cada noite até hoje, e um número alvo de horas por noite, retorne quantas horas de sono você precisa dormir hoje à noite para eliminar sua dívida de sono.

- Inclua as horas de hoje à noite no total de tempo necessário para compensar.
- Se você já dormiu o suficiente para cobrir o alvo de hoje à noite ou mais, retorne `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` deve retornar `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` deve retornar `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` deve retornar `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` deve retornar `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` deve retornar `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
