---
id: 69f35a5bb823ed620fcb7cbd
title: "Desafio 284: I antes de E"
challengeType: 28
dashedName: challenge-284
---

# --description--

Dada uma palavra ou frase, retorne uma versão corrigida onde cada palavra siga a regra "I antes de E, exceto depois de C".

- Se uma palavra contiver `"ei"` não precedida por `"c"`, substitua-a por `"ie"`.
- Se uma palavra contiver `"ie"` precedida por `"c"`, substitua-a por `"ei"`.
- Todas as outras palavras permanecem inalteradas.

# --hints--

`iBeforeE("beleive")` deve retornar `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` deve retornar `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` deve retornar `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` deve retornar `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` deve retornar `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
