---
id: 69f8c998d78ad3171a0713bf
title: "Desafio 292: proporção de aspecto mais ampla"
challengeType: 28
dashedName: challenge-292
---

# --description--

Dadas duas strings para diferentes dimensões de imagem, retorne a proporção de aspecto da imagem com a maior razão largura-altura.

- As strings fornecidas estarão no formato `"WxH"`, por exemplo, `"1920x1080"`.
- A proporção de aspecto é a razão entre largura e altura, reduzida aos menores números inteiros. Por exemplo, `"1920x1080"` reduz para `"16:9"`.
- Retorne uma string no formato `"W:H"`, por exemplo, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` deve retornar `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` deve retornar `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` deve retornar `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` deve retornar `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` deve retornar `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` deve retornar `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
