---
id: 69f8c998d78ad3171a0713b9
title: "Desafio 286: issues abertas"
challengeType: 28
dashedName: challenge-286
---

# --description--

Dado um array de números de issues e outro array de números de pull requests (PR), retorne um array das issues que permanecem abertas depois que todos os PRs foram mesclados.

- Um PR fecha uma issue se seus dígitos forem rotações um do outro. Por exemplo, a issue 123 seria fechada pelo PR 231 ou 312.
- Um PR não fecha uma issue com o mesmo número exato. Por exemplo, o PR 123 não fecha a issue 123. Então uma issue com todos os dígitos iguais não pode ser fechada.
- Qualquer número pode ter zeros à esquerda removidos. Por exemplo, o PR 201 fecharia a issue 12 (012, uma rotação de 201). Similarmente, a issue 201 seria fechada pelo PR 12.

Retorne as issues abertas restantes na ordem em que foram dadas.

# --hints--

`getOpenIssues([123, 234], [231])` deve retornar `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` deve retornar `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` deve retornar `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` deve retornar `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` deve retornar `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
