---
id: 69f8c998d78ad3171a0713be
title: "Desafio 291: contagem fizzbuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--

Dado um número inicial e final, conte o número de aparições de fizz e buzz no intervalo (inclusivo).

- Números divisíveis por 3 contam como fizz.
- Números divisíveis por 5 contam como buzz.
- Números divisíveis por 3 e 5 contam como fizz e buzz.

Retorne um objeto ou dicionário com as contagens no formato: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` deve retornar `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` deve retornar `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` deve retornar `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` deve retornar `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` deve retornar `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
