---
id: 69f8c998d78ad3171a0713bb
title: "Desafio 288: número secreto"
challengeType: 28
dashedName: challenge-288
---

# --description--

Dado um número secreto e um palpite, determine se o palpite está correto.

Retorne:

- `"higher"` se o número secreto for maior que o palpite.
- `"lower"` se o número secreto for menor que o palpite.
- `"you got it!"` se o palpite estiver correto.

# --hints--

`guessNumber(50, 30)` deve retornar `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` deve retornar `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` deve retornar `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` deve retornar `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` deve retornar `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` deve retornar `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
