---
id: 69f8c998d78ad3171a0713bd
title: "Desafio 290: festa de pizza"
challengeType: 28
dashedName: challenge-290
---

# --description--

Dado um array com as horas trabalhadas hoje por pessoa, retorne o número de pizzas para pedir em uma festa de pizza.

- Divida as horas trabalhadas de cada pessoa por 3 para obter a quantidade de fatias.
- Você não pode comer uma fatia parcial, então arredonde a quantidade de fatias de cada pessoa para cima até o número inteiro mais próximo.
- Cada pessoa recebe no mínimo duas fatias.
- Cada pizza tem 8 fatias. Arredonde o número total de pizzas para cima até a pizza inteira mais próxima.

# --hints--

`getPizzasToOrder([8, 8, 8])` deve retornar `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` deve retornar `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` deve retornar `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` deve retornar `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` deve retornar `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` deve retornar `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
