---
id: 69f90c1329a94b37e2a2086d
title: "Desafio 294: combinações de parênteses"
challengeType: 28
dashedName: challenge-294
---

# --description--

Dado um número inteiro, `n`, retorne o número de combinações válidas de `n` pares de parênteses.

- Uma combinação válida é uma string onde todo parêntese de abertura tem um parêntese de fechamento correspondente, e nenhum parêntese de fechamento aparece antes do seu parêntese de abertura correspondente.

Por exemplo, dado `2`, existem 2 combinações válidas:

```json
(())
()()
```

# --hints--

`getCombinations(2)` deve retornar `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` deve retornar `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` deve retornar `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` deve retornar `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` deve retornar `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
