---
id: 6a0dcd03ee4e68698080ef6c
title: "Desafio 308: validador de cartão de crédito"
challengeType: 28
dashedName: challenge-308
---

# --description--

Dada uma string de dígitos para um número de cartão de crédito, determine se é um número válido usando o seguinte método:

- Começando pelo penúltimo dígito, dobre a cada outro dígito indo para a esquerda.
- Se dobrar um dígito resultar em um número maior que 9, subtraia 9.
- Some todos os dígitos (dobrados e não dobrados).
- Se o total for divisível por 10, o número é válido.

# --hints--

`isValidCard("4532015112830366")` deve retornar `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` deve retornar `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` deve retornar `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` deve retornar `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` deve retornar `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` deve retornar `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` deve retornar `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
