---
id: 6a0dcd03ee4e68698080ef6b
title: "Desafio 307: regulamentos de zoneamento"
challengeType: 28
dashedName: challenge-307
---

# --description--

Dada uma grade 2D (array de arrays) que representa o layout dos prédios de uma cidade, retorne as coordenadas de todos os prédios que estão violando as regras de zoneamento.

Cada célula na grade contém um dos rótulos da tabela abaixo. Um prédio está em violação se qualquer um dos seus (até) 4 vizinhos, horizontal ou vertical, for de um tipo ao qual ele não pode estar adjacente.

| Rótulo | Tipo | Não pode estar adjacente a |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | agrícola | `"C"` |
| `"R"` | residencial | `"i"`, `"C"` |
| `"I"` | institucional | `"i"` |
| `"C"` | comercial | `"R"`, `"A"` |
| `""` (string vazia) | não desenvolvido | sem restrições |

Retorne as coordenadas de todas as células em violação como um array de pares `[row, col]`, em qualquer ordem. Se não houver violações, retorne um array vazio.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` deve retornar `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` deve retornar `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` deve retornar `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` deve retornar `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` deve retornar `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
