---
id: 6a0dcd03ee4e68698080ef68
title: "Desafio 304: arranjos de itinerário"
challengeType: 28
dashedName: challenge-304
---

# --description--

Dado um array com pelo menos duas paradas opcionais para um passeio de um dia, retorne o número de arranjos válidos de itinerário.

O itinerário sempre inclui `"breakfast"`, `"lunch"` e `"dinner"`, que não serão passados como argumentos. As paradas opcionais podem ser colocadas em qualquer lugar do itinerário, sujeitas às seguintes regras:

- `"breakfast"` é sempre o primeiro, com pelo menos uma parada antes de `"lunch"`.
- `"lunch"` deve aparecer antes de `"dinner"`, com pelo menos uma parada entre eles.
- No máximo, uma parada opcional pode aparecer depois de `"dinner"`.

Retorne o número de arranjos válidos.

# --hints--

`getItineraryCount(["library", "park"])` deve retornar `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` deve retornar `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` deve retornar `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` deve retornar `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` deve retornar `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
