---
id: 6a0dcc730cb92a616f86f0c5
title: "Desafio 301: última carga"
challengeType: 28
dashedName: challenge-301
---

# --description--

Dado o número de colheres de detergente para roupa que você tem restantes e um array com quantas colheres você usou em cada um dos dias anteriores, retorne o número de dias completos de detergente que ainda restam.

Calcule sua média diária de uso a partir do histórico de uso e assuma essa quantidade de uso para cada dia daqui para frente.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` deve retornar `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` deve retornar `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` deve retornar `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` deve retornar `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` deve retornar `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
