---
id: 6a0dcd03ee4e68698080ef66
title: "Desafio 302: jet lag"
challengeType: 28
dashedName: challenge-302
---

# --description--

Dadas uma cidade de partida, uma cidade de chegada, uma duração de voo em horas e uma direção de viagem, retorne o número de horas de jet lag que o viajante está experimentando.

As cidades fornecidas serão da seguinte lista que inclui seu offset UTC:

| Cidade | Offset |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Para calcular as horas de jet lag:

1. Busque a diferença de fuso horário em horas entre as duas cidades.
2. Determine o multiplicador de direção. Se estiver viajando `"east"`, é 1.5, caso contrário, é 1.0.
3. Obtenha as horas de jet lag com a fórmula: diferença de fuso horário + (duração do voo * 0.1) * multiplicador de direção

Retorne as horas de jet lag arredondadas para uma casa decimal.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` deve retornar `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` deve retornar `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` deve retornar `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` deve retornar `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` deve retornar `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` deve retornar `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` deve retornar `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
