---
id: 6a0dcc730cb92a616f86f0c2
title: "Desafio 298: validador de esquema parte 4"
challengeType: 28
dashedName: challenge-298
---

# --description--

Dado um objeto (JavaScript) ou dicionário (Python), determine se ele corresponde ao seguinte esquema:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean
}
```

- O símbolo pipe (`|`) significa "ou". `role` deve ser um dos valores listados em `Roles`.
- O ponto de interrogação (`?`) após `supporter` significa que o campo é opcional, mas deve ser do tipo especificado se existir.
- Chaves extras são permitidas

# --hints--

`isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true })` deve retornar `true`.

```js
assert.isTrue(isValidSchema({ username: "vivian", posts: 1, verified: false, role: "user", supporter: true }));
```

`isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" })` deve retornar `true`.

```js
assert.isTrue(isValidSchema({ username: "rudolph", posts: 15, verified: true, role: "creator" }));
```

`isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 })` deve retornar `true`.

```js
assert.isTrue(isValidSchema({ username: "hernandez", posts: 35, verified: true, role: "moderator", supporter: false, followers: 55 }));
```

`isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" })` deve retornar `false`.

```js
assert.isFalse(isValidSchema({ username: "julia", posts: 50, verified: true, role: "admin", supporter: "true" }));
```

`isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true })` deve retornar `false`.

```js
assert.isFalse(isValidSchema({ username: "bernard", posts: 0, verified: true, role: "friend", supporter: true }));
```

`isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false })` deve retornar `false`.

```js
assert.isFalse(isValidSchema({ username: "felix", posts: 40, verified: "yes", role: "staff", supporter: false }));
```

`isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true })` deve retornar `false`.

```js
assert.isFalse(isValidSchema({ username: "jimmy", posts: true, verified: false, role: "creator", supporter: true }));
```

`isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false })` deve retornar `false`.

```js
assert.isFalse(isValidSchema({ username: true, posts: 30, verified: true, role: "moderator", supporter: false }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role) &&
    (obj.supporter === undefined || typeof obj.supporter === 'boolean')
  );
}
```
