---
id: 6a15cadf5f240d05a2649558
title: "Desafio 314: fatoração em números primos"
challengeType: 28
dashedName: challenge-314
---

# --description--

Dado um inteiro maior que 1, retorne sua fatoração em números primos como um array de números em ordem crescente.

Uma fatoração em números primos é o conjunto de números primos que se multiplicam para produzir o inteiro dado. Cada número tem exatamente um conjunto. Por exemplo, a fatoração em números primos de 20 é `[2, 2, 5]` porque 2 * 2 * 5 = 20.

Se o inteiro dado for ele mesmo primo, retorne-o em um array de um único elemento.

# --hints--

`primeFactorization(20)` deve retornar `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` deve retornar `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` deve retornar `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` deve retornar `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` deve retornar `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` deve retornar `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` deve retornar `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
