---
id: 6a151d32d772271dd2448c2c
title: "Desafio 310: do inglês britânico para o americano"
challengeType: 28
dashedName: challenge-310
---

# --description--

Dada uma frase, converta quaisquer grafias do inglês britânico para seus equivalentes no inglês americano usando a tabela de consulta abaixo e retorne a frase atualizada:

| British | American |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- As substituições devem ser case-insensitive. Por exemplo, `"Colour"` deve se tornar `"Color"`.
- A entrada pode conter palavras que derivam da grafia exata de uma raiz na tabela e que também precisam ser alteradas. Por exemplo, `"colouring"` deve se tornar `"coloring"`, e `"disorganised"` deve se tornar `"disorganized"`.

# --hints--

`britishToAmerican("I love the colour blue.")` deve retornar `"I love the color blue."`

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` deve retornar `"The fiber optic cable is new."`

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` deve retornar `"It's an honor to meet someone with such humor."`

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` deve retornar `"The unrecognized artist analyzed his color palette at the center."`

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` deve retornar `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
