---
id: 6a15cadf5f240d05a264955c
title: "Desafio 318: mutações no DNA"
challengeType: 28
dashedName: challenge-318
---

# --description--

Dadas duas fitas de DNA com o mesmo comprimento, retorne um array com os índices onde as fitas diferem (mutações).

- As fitas de DNA são strings compostas pelos caracteres `"A"`, `"T"`, `"C"` e `"G"`
- Retorne os índices em ordem crescente
- Se não houver mutações, retorne um array vazio

# --hints--

`detectMutations("ATCG", "ATGG")` deve retornar `[2]`.

```js
assert.deepEqual(detectMutations("ATCG", "ATGG"), [2]);
```

`detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` deve retornar `[4, 9, 11]`.

```js
assert.deepEqual(detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11]);
```

`detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` deve retornar `[]`.

```js
assert.deepEqual(detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), []);
```

`detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` deve retornar `[15, 16, 17, 18]`.

```js
assert.deepEqual(detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18]);
```

`detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` deve retornar `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`.

```js
assert.deepEqual(detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]);
```

# --seed--

## --seed-contents--

```js
function detectMutations(strand1, strand2) {

  return strand1;
}
```

# --solutions--

```js
function detectMutations(strand1, strand2) {
  return strand1
    .split("")
    .reduce((mutations, base, i) => {
      if (base !== strand2[i]) mutations.push(i);
      return mutations;
    }, []);
}
```
