---
id: 6a15cadf5f240d05a264955d
title: "Desafio 319: analisador de frontmatter"
challengeType: 28
dashedName: challenge-319
---

# --description--

Dada uma string que representa um bloco de frontmatter, analise-a e retorne um objeto (JavaScript) ou dicionário (Python) com as chaves e valores.

O frontmatter é delimitado por `---` e contém pares `key: value` dentro dele, um por linha. Por exemplo:

```md
---
title: My Post
draft: false
views: 100
---
```

Deve retornar:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Números, booleanos e strings devem ser retornados como seus respectivos tipos.
- A string fornecida terá quebras de linha separadas pelo caractere de nova linha (`"\n"`). O exemplo acima seria dado como: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` deve retornar `{ title: "My Post", draft: false, views: 100 }`.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` deve retornar `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` deve retornar `{ version: "1.0.0", url: "https://example.com", private: true }`.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` deve retornar `{ rating: 4.5, price: 9.99 }`.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
