---
id: 6a15cadf5f240d05a264955e
title: "Desafio 320: banco de sangue"
challengeType: 28
dashedName: challenge-320
---

# --description--

Dado um array do estoque de um banco de sangue e um array com os pedidos de tipos sanguíneos dos pacientes, retorne uma string no formato `"X of Y patients served"`. Onde `X` é o número máximo de pacientes que podem receber sangue do estoque do banco, e `Y` é o número total de pacientes.

Cada entrada em ambos os arrays é um dos seguintes tipos sanguíneos: `"AB"`, `"A"`, `"B"` ou `"O"`.

Regras de compatibilidade:

- `"AB"` pode receber de qualquer tipo sanguíneo.
- `"A"` pode receber de `"A"` e `"O"`.
- `"B"` pode receber de `"B"` e `"O"`.
- `"O"` só pode receber de `"O"`.

Entradas duplicadas nos arrays dados representam quantidade.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` deve retornar `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` deve retornar `"3 of 5 patients served"`.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` deve retornar `"4 of 5 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` deve retornar `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` deve retornar `"8 of 11 patients served"`.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` deve retornar `"13 of 14 patients served"`.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
