---
id: 6a151d32d772271dd2448c2d
title: "Desafio 311: conjurador"
challengeType: 28
dashedName: challenge-311
---

# --description--

Dada uma string de códigos de feitiços que você está lançando, calcule a pontuação total.

Cada caractere na string representa um feitiço:

| Código | Feitiço | Categoria | Pontuação base |
| - | - | - | - |
| `"f"` | Fogo | Destruição | 3 |
| `"l"` | Raio | Destruição | 3 |
| `"i"` | Gelo | Controle | 2 |
| `"w"` | Vento | Controle | 2 |
| `"h"` | Cura | Restauração | 1 |
| `"s"` | Escudo | Restauração | 1 |

Um multiplicador de combo é aplicado com base em quantos feitiços consecutivos foram lançados de categorias diferentes:

- O primeiro feitiço sempre pontua o valor base.
- Cada feitiço consecutivo de uma categoria diferente da anterior aumenta o multiplicador em 1.
- Lançar um feitiço da mesma categoria do anterior reseta o multiplicador para 1.
- A pontuação de cada feitiço é sua pontuação base multiplicada pelo multiplicador atual.

Retorne a pontuação total da sequência de feitiços.

# --hints--

`cast("fihwl")` deve retornar `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` deve retornar `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` deve retornar `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` deve retornar `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` deve retornar `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
