---
id: 6a15cadf5f240d05a264955b
title: "Desafio 317: calculadora de IMC"
challengeType: 28
dashedName: challenge-317
---

# --description--

Dado um peso em libras e uma altura em polegadas, retorne o IMC (Índice de Massa Corporal) arredondado para uma casa decimal.

Para obter o IMC: divida o peso pela altura ao quadrado, depois multiplique o resultado por 703.

# --hints--

`calculateBmi(180, 70)` deve retornar `25.8`.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` deve retornar `24.0`.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` deve retornar `19.5`.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` deve retornar `39.1`.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` deve retornar `22.8`.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
