---
id: 6a19b1062d1b153d8ac76d73
title: "Desafio 325: número da sorte"
challengeType: 28
dashedName: challenge-325
---

# --description--

Dada uma string com o primeiro e o último nome de uma pessoa, calcule o número da sorte dela usando as seguintes regras:

- O primeiro e o último nome são separados por um espaço
- Busque a contagem de vogais e consoantes para cada nome
- Multiplique as menores contagens de vogais e consoantes entre si e depois pelo comprimento do nome menor
- Faça o mesmo para as duas maiores contagens e o nome maior
- Subtraia o valor menor do maior para obter o número da sorte

Se o valor final for zero (`0`), retorne `13`.

# --hints--

`getLuckyNumber("John Doe")` deve retornar `21`.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` deve retornar `52`.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` deve retornar `18`.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` deve retornar `81`.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` deve retornar `32`.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` deve retornar `13`.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
