---
id: 6a19b1062d1b153d8ac76d72
title: "Desafio 324: contagem de caracteres duplicados"
challengeType: 28
dashedName: challenge-324
---

# --description--

Dadas duas strings, retorne a contagem de caracteres da segunda string que podem ser encontrados na primeira.

- Caracteres duplicados na segunda string são contados separadamente.

# --hints--

`duplicateCharacterCount("aloha", "hei")` deve retornar `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` deve retornar `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` deve retornar `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` deve retornar `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` deve retornar `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` deve retornar `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` deve retornar `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
