---
id: 6a1d9f98e819ed70a0e994dc
title: "Desafio 330: palavras em minúsculas"
challengeType: 28
dashedName: challenge-330
---

# --description--

Dada uma string, retorne apenas as palavras que estão inteiramente em minúsculas, na ordem original e com um espaço entre cada palavra.

# --hints--

`getLowercaseWords("hello GOOD world")` deve retornar `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` deve retornar `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` deve retornar `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` deve retornar `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` deve retornar `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
