---
id: 6a1d9f98e819ed70a0e994e0
title: "Desafio 334: Troco exato"
challengeType: 28
dashedName: challenge-334
---

# --description--

Dado um valor inteiro em centavos, retorne o número de maneiras distintas de fazer o troco exato usando pennies (1 centavo), nickels (5 centavos), dimes (10 centavos) e quarters (25 centavos).

# --hints--

`exactChange(3)` deve retornar `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` deve retornar `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` deve retornar `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` deve retornar `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` deve retornar `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` deve retornar `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
