---
id: 6a1d9f98e819ed70a0e994da
title: "Desafio 328: a rotina de Kaprekar"
challengeType: 28
dashedName: challenge-328
---

# --description--

Dado um número de 4 dígitos, retorne quantas vezes você precisa aplicar a rotina de Kaprekar até chegar em 6174.

A rotina de Kaprekar funciona da seguinte forma:

- Organize os dígitos em ordem decrescente para formar o maior número
- Organize os dígitos em ordem crescente para formar o menor número (complete com zeros à esquerda se necessário)
- Subtraia o menor do maior
- Repita com o novo número

# --hints--

`kaprekar(1234)` deve retornar `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` deve retornar `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` deve retornar `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` deve retornar `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` deve retornar `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
