---
id: 6a1d9f98e819ed70a0e994dd
title: "Desafio 331: múltiplo mais próximo"
challengeType: 28
dashedName: challenge-331
---

# --description--

Dado dois inteiros, arredonde o primeiro para o múltiplo mais próximo do segundo.

# --hints--

`roundToNearestMultiple(5, 3)` deve retornar `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` deve retornar `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` deve retornar `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` deve retornar `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` deve retornar `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
