---
id: 6a1d9f98e819ed70a0e994df
title: "Desafio 333: Triagem de issue 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

Dado um título de issue e um array dos rótulos atuais, retorne um array atualizado de rótulos baseado nas seguintes regras:

Se a issue não tiver nenhum rótulo, adicione:

- `"bug"` e `"needs triage"` se o título contiver `"error"` ou `"bug"`
- `"enhancement"` e `"discussing"` se o título contiver `"feature"` ou `"add"`

Caso contrário, se os rótulos dados contiverem:

- `"needs triage"` e o título contiver `"simple"` ou `"easy"`, remova `"needs triage"` e adicione `"good first issue"`
- `"discussing"` e o título contiver `"planned"` ou `"next"`, remova `"discussing"` e adicione `"on the roadmap"`
- Caso contrário, se `"needs triage"` ou `"discussing"` estiver presente, remova e adicione `"help wanted"`

Se o título contiver:

- `"security"`, adicione um rótulo `"critical"`

# --hints--

`triageIssue("app crashes with error", [])` deve retornar `["bug", "needs triage"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` deve retornar `["bug", "help wanted"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` deve retornar `["enhancement", "discussing"]`.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` deve retornar `["enhancement", "help wanted"]`.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` deve retornar `["bug", "needs triage", "critical"]`.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` deve retornar `["critical"]`.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` deve retornar `["bug", "good first issue"]`.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` deve retornar `["enhancement", "on the roadmap"]`.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` deve retornar `["enhancement", "help wanted", "critical"]`.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
