---
id: 6a1d9f98e819ed70a0e994de
title: "Desafio 332: triagem de issue"
challengeType: 28
dashedName: challenge-332
---

# --description--

Dado um número de milissegundos desde a última postagem em uma issue e a última mensagem postada na issue, determine o que você deve fazer com a issue de acordo com estas regras:

- Se a última mensagem foi há menos de 7 dias, retorne `"leave it"`
- Se a última mensagem foi há 7 ou mais dias e seu conteúdo contém `"bump"` (sem diferenciar maiúsculas de minúsculas), retorne `"close it"`
- Caso contrário, retorne `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` deve retornar `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` deve retornar `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` deve retornar `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` deve retornar `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` deve retornar `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` deve retornar `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
