---
id: 6a2037a68a0bc2aef0006004
title: "Desafio 342: probabilidades com dados"
challengeType: 28
dashedName: challenge-342
---

# --description--

Dado um número de dados de seis lados para rolar e uma soma alvo, retorne as probabilidades de tirar essa soma como uma string no formato `"1 in X"`.

- O número de dados estará entre 1 e 6.
- A soma alvo é sempre alcançável com o número dado de dados.
- Arredonde `"X"` para o número inteiro mais próximo.

# --hints--

`getOdds(1, 5)` deve retornar `"1 in 6"`.

```js
assert.equal(getOdds(1, 5), "1 in 6");
```

`getOdds(2, 4)` deve retornar `"1 in 12"`.

```js
assert.equal(getOdds(2, 4), "1 in 12");
```

`getOdds(3, 10)` deve retornar `"1 in 8"`.

```js
assert.equal(getOdds(3, 10), "1 in 8");
```

`getOdds(4, 7)` deve retornar `"1 in 65"`.

```js
assert.equal(getOdds(4, 7), "1 in 65");
```

`getOdds(5, 26)` deve retornar `"1 in 111"`.

```js
assert.equal(getOdds(5, 26), "1 in 111");
```

`getOdds(6, 35)` deve retornar `"1 in 7776"`.

```js
assert.equal(getOdds(6, 35), "1 in 7776");
```

# --seed--

## --seed-contents--

```js
function getOdds(dice, target) {

  return dice;
}
```

# --solutions--

```js
function getOdds(dice, target) {
  let validCombos = 0;
  const totalCombos = Math.pow(6, dice);

  function roll(diceLeft, sum) {
    if (diceLeft === 0) {
      if (sum === target) validCombos++;
      return;
    }
    for (let face = 1; face <= 6; face++) {
      roll(diceLeft - 1, sum + face);
    }
  }

  roll(dice, 0);
  return `1 in ${Math.round(totalCombos / validCombos)}`;
}
```
