---
id: 6a2037a68a0bc2aef0006002
title: "Desafio 340: conversor de pig latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

Dada uma string, converta-a para Pig Latin usando as seguintes regras:

- Se uma palavra começar com uma vogal (`"a"`, `"e"`, `"i"`, `"o"` ou `"u"`), adicione `"way"` ao final. Por exemplo, `"universe"` se converte em `"universeway"`.
- Se uma palavra começar com uma ou mais consoantes, mova-as para o final e adicione `"ay"`. Por exemplo, `"hello"` se converte em `"ellohay"`.
- Preserve o caso da primeira letra. Por exemplo, `"Hello"` se converte em `"Ellohay"`.

# --hints--

`pigLatin("universe")` deve retornar `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` deve retornar `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` deve retornar `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` deve retornar `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` deve retornar `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` deve retornar `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
