---
id: 6a2037a68a0bc2aef0005fff
title: "Desafio 337: contador de marcações"
challengeType: 28
dashedName: challenge-337
---

# --description--

Dada uma string de marcações, retorne o total representado.

- Cada barra vertical `"|"` representa uma contagem.
- A cada quinta marca é representada por uma barra inclinada para frente `"/"`, completando um grupo de cinco (`"||||/"`).
- Os grupos são separados por um espaço.

# --hints--

`getTallyCount("||||")` deve retornar `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` deve retornar `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` deve retornar `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` deve retornar `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` deve retornar `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
