---
id: 6a2037a68a0bc2aef0006006
title: "Desafio 344: razão áurea"
challengeType: 28
dashedName: challenge-344
---

# --description--

Dado dois números, determine se a razão entre eles se aproxima da razão áurea.

- Use uma razão áurea de `1.618`
- Permita uma tolerância de `0.01`

# --hints--

`isGoldenRatio(21, 34)` deve retornar `true`.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` deve retornar `false`.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` deve retornar `true`.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` deve retornar `false`.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` deve retornar `true`.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` deve retornar `false`.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
