---
id: 6a2037a68a0bc2aef0006000
title: "Desafio 338: calculadora de idade de pet"
challengeType: 28
dashedName: challenge-338
---

# --description--

Dado um tipo de pet e a idade em anos humanos, retorne a idade equivalente em anos de pet usando a seguinte tabela de conversão:

| Pet | Multiplicador |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` deve retornar `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` deve retornar `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` deve retornar `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` deve retornar `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` deve retornar `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` deve retornar `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` deve retornar `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
