---
id: 6a22d77ddf034bc4e35b1d58
title: "Desafio 351: número prônico"
challengeType: 28
dashedName: challenge-351
---

# --description--

Dado um número, determine se ele é um número prônico.

Um número prônico é o produto de dois inteiros consecutivos. Por exemplo, 6 é prônico porque 2 * 3 = 6.

# --hints--

`isPronic(6)` deve retornar `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` deve retornar `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` deve retornar `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` deve retornar `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` deve retornar `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` deve retornar `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
