---
id: 6a22d77ddf034bc4e35b1d5b
title: "Desafio 354: avaliação de contraste 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

Dado dois arrays representando valores RGB e um booleano indicando se o texto é grande, retorne a avaliação de contraste WCAG usando o seguinte método:

Primeiro, converta cada valor RGB para luminância relativa:

- Divida cada canal `[R, G, B]` por 255 para obter um valor entre `0` e `1`
- Aplique a fórmula de correção gama em cada canal:
  - Se o valor do canal for menor ou igual a `0.04045`: `channel / 12.92`
  - Caso contrário: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcule a luminância: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Depois, calcule a taxa de contraste somando 0.05 a cada valor de luminância, depois dividindo o valor mais claro pelo mais escuro. O mais claro será sempre o primeiro argumento.

Retorne a avaliação baseada na taxa de contraste usando a tabela a seguir:

| Avaliação | Texto normal | Texto grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | abaixo de 4.5 | abaixo de 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
