---
id: 6a22d77ddf034bc4e35b1d54
title: "Desafio 347: teoria dos jogos"
challengeType: 28
dashedName: challenge-347
---

# --description--

Dadas duas strings de mesmo comprimento que representam as estratégias de dois jogadores para um jogo, retorne as pontuações como um array `[player1, player2]`.

- As strings dadas conterão apenas uma de duas letras: `"C"` (cooperar) ou `"D"` (trair).
- Cada caractere representa uma rodada, pontuada da seguinte forma:
  - Se ambos os jogadores cooperarem, cada um marca 3.
  - Se ambos os jogadores traírem, cada um marca 1.
  - Se um jogador trair e o outro cooperar, o traidor marca 5 e o cooperador marca 0.

# --hints--

`playGame("CCCC", "CCCC")` deve retornar `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` deve retornar `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` deve retornar `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` deve retornar `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` deve retornar `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
