---
id: 6a22d77ddf034bc4e35b1d59
title: "Desafio 352: avaliação de taxa de contraste 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

Dada uma taxa de contraste e um booleano indicando se o texto é grande, retorne a avaliação WCAG usando a seguinte tabela:

| Avaliação | Texto normal | Texto grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | abaixo de 4.5 | abaixo de 3.0 |

# --hints--

`getContrastRating("7.5", false)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
