---
id: 6a22d77ddf034bc4e35b1d5c
title: "Desafio 355: código Morse"
challengeType: 28
dashedName: challenge-355
---

# --description--

Dada uma string em código Morse, retorne a mensagem decodificada usando a tabela a seguir:

| Código | Letra | Código | Letra |
|------|--------|------|--------|
| `.-` | `A` | `-.` | `N` |
| `-...` | `B` | `---` | `O` |
| `-.-.` | `C` | `.--.` | `P` |
| `-..` | `D` | `--.-` | `Q` |
| `.` | `E` | `.-.` | `R` |
| `..-.` | `F` | `...` | `S` |
| `--.` | `G` | `-` | `T` |
| `....` | `H` | `..-` | `U` |
| `..` | `I` | `...-` | `V` |
| `.---` | `J` | `.--` | `W` |
| `-.-` | `K` | `-..-` | `X` |
| `.-..` | `L` | `-.--` | `Y` |
| `--` | `M` | `--..` | `Z` |

- As letras são separadas por um único espaço
- As palavras são separadas por três espaços

# --hints--

`decodeMorse("--..")` deve retornar `"Z"`.

```js
assert.equal(decodeMorse("--.."), "Z");
```

`decodeMorse("... --- ...")` deve retornar `"SOS"`.

```js
assert.equal(decodeMorse("... --- ..."), "SOS");
```

`decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--.")` deve retornar `"FREECODECAMP"`.

```js
assert.equal(decodeMorse("..-. .-. . . -.-. --- -.. . -.-. .- -- .--."), "FREECODECAMP");
```

`decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -..")` deve retornar `"HELLO WORLD"`.

```js
assert.equal(decodeMorse(".... . .-.. .-.. ---   .-- --- .-. .-.. -.."), "HELLO WORLD");
```

`decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --.")` deve retornar `"THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG"`.

```js
assert.equal(decodeMorse("- .... .   --.- ..- .. -.-. -.-   -... .-. --- .-- -.   ..-. --- -..-   .--- ..- -- .--. . -..   --- ...- . .-.   - .... .   .-.. .- --.. -.--   -.. --- --."), "THE QUICK BROWN FOX JUMPED OVER THE LAZY DOG");
```

# --seed--

## --seed-contents--

```js
function decodeMorse(code) {

  return code;
}
```

# --solutions--

```js
function decodeMorse(code) {
  const table = {
    '.-': 'A', '-...': 'B', '-.-.': 'C', '-..': 'D', '.': 'E',
    '..-.': 'F', '--.': 'G', '....': 'H', '..': 'I', '.---': 'J',
    '-.-': 'K', '.-..': 'L', '--': 'M', '-.': 'N', '---': 'O',
    '.--.': 'P', '--.-': 'Q', '.-.': 'R', '...': 'S', '-': 'T',
    '..-': 'U', '...-': 'V', '.--': 'W', '-..-': 'X', '-.--': 'Y',
    '--..': 'Z'
  };

  return code.split('   ')
    .map(word => word.split(' ').map(c => table[c]).join(''))
    .join(' ');
}
```
