---
id: 6a22d77ddf034bc4e35b1d57
title: "Desafio 350: distância entre letras"
challengeType: 28
dashedName: challenge-350
---

# --description--

Dadas duas strings de mesmo comprimento, retorne a soma das menores distâncias entre cada par de caracteres.

- A entrada conterá apenas letras minúsculas
- O alfabeto é tratado como um círculo, então a distância entre `a` e `z` é 1.

# --hints--

`letterDistance("abc", "bcd")` deve retornar `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` deve retornar `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` deve retornar `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` deve retornar `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` deve retornar `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` deve retornar `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
