---
id: 6a22d77ddf034bc4e35b1d5a
title: "Desafio 353: avaliação de contraste 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

Dado dois valores de luminância relativa e um booleano indicando se o texto é grande, retorne a avaliação de contraste WCAG usando o seguinte método:

Calcule a taxa de contraste adicionando 0,05 a cada valor de luminância, depois dividindo o valor mais claro pelo mais escuro. O valor mais claro será sempre o primeiro argumento.

Retorne a avaliação com base na taxa de contraste usando a seguinte tabela:

| Avaliação | Texto normal | Texto grande |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | abaixo de 4,5 | abaixo de 3,0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` deve retornar `"AAA"`.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` deve retornar `"AA"`.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` deve retornar `"Fail"`.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
