---
id: 6a26df3e2988bcdded204897
title: "Desafio 361: hora falada"
challengeType: 28
dashedName: challenge-361
---

# --description--

Dadas as posições dos ponteiros das horas e dos minutos de um relógio analógico em graus (no sentido horário a partir do 12), retorne a hora em inglês falado.

Converta o ângulo do ponteiro dos minutos em minutos (360° = 60 minutos), depois use as seguintes regras:

| Minutos | Falado |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (excluindo 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (excluindo 45) | `"X minutes to Z"` (onde `X` é 60 - minutos) |

Onde `Y` é a hora atual e `Z` é a próxima hora, ambos derivados do ângulo do ponteiro das horas (360° = 12 horas).

Nota: os ângulos dos ponteiros podem não coincidir exatamente com um número, considere arredondá-los de alguma forma.

# --hints--

`getSpokenTime(90, 0)` deve retornar `"3 o'clock"`.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` deve retornar `"20 minutes past 5"`.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` deve retornar `"half past 8"`.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` deve retornar `"quarter past 2"`.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` deve retornar `"20 minutes to 7"`.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` deve retornar `"quarter to 11"`.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` deve retornar `"5 minutes to 4"`.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
