---
id: 6a26df3e2988bcdded204893
title: "Desafio 357: cadeia alimentar"
challengeType: 28
dashedName: challenge-357
---

# --description--

Dado um array de pares `[predator, prey]`, retorne a cadeia alimentar do predador de topo até o fundo.

- O predador de topo é o animal que nunca é presa de outro animal.
- Retorne a cadeia como um array de strings.

# --hints--

`getFoodChain([["cat", "mouse"]])` deve retornar `["cat", "mouse"]`.

```js
assert.deepEqual(getFoodChain([["cat", "mouse"]]), ["cat", "mouse"]);
```

`getFoodChain([["wolf", "deer"], ["deer", "grass"]])` deve retornar `["wolf", "deer", "grass"]`.

```js
assert.deepEqual(getFoodChain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"]);
```

`getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` deve retornar `["hawk", "snake", "frog", "fly"]`.

```js
assert.deepEqual(getFoodChain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"]);
```

`getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` deve retornar `["eagle", "fox", "rabbit", "grass"]`.

```js
assert.deepEqual(getFoodChain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"]);
```

`getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` deve retornar `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]`.

```js
assert.deepEqual(getFoodChain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"]);
```

# --seed--

## --seed-contents--

```js
function getFoodChain(pairs) {

  return pairs;
}
```

# --solutions--

```js
function getFoodChain(pairs) {
  const prey = new Set(pairs.map(([, p]) => p));
  const map = Object.fromEntries(pairs);
  const apex = pairs.map(([p]) => p).find(p => !prey.has(p));
  const chain = [apex];
  while (map[chain.at(-1)]) chain.push(map[chain.at(-1)]);
  return chain;
}
```
