---
id: 6a26df95efa55a2524399745
title: "Desafio 364: Entre dois baldes"
challengeType: 28
dashedName: challenge-364
---

# --description--

Dado dois baldes de tinta, cada um com uma cor RGB e um nível de preenchimento, retorne a cor RGB misturada como um array de três inteiros.

- Cada balde é um objeto (JavaScript) ou dicionário (Python) com uma propriedade `color` (um array de três inteiros `[r, g, b]`) e uma propriedade `fullness` (0–100).
- A cor misturada é uma média ponderada de cada canal nas duas cores com base no nível de preenchimento, com cada canal arredondado para o inteiro mais próximo.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` deve retornar `[125, 125, 125]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` deve retornar `[200, 200, 200]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` deve retornar `[100, 150, 200]`.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` deve retornar `[114, 184, 170]`.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` deve retornar `[68, 163, 123]`.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
