---
id: 6a26df95efa55a2524399743
title: "Desafio 362: validador de nonogram"
challengeType: 28
dashedName: challenge-362
---

# --description--

Dado um array de números de pistas e um array de células, determine se as células satisfazem a pista do nonogram.

- A pista é um array de números que representam os comprimentos de células preenchidas consecutivas, em ordem. Por exemplo, uma pista de `[3, 2]` significa que deve haver 3 células preenchidas consecutivas seguidas por 2 células preenchidas consecutivas, separadas por pelo menos uma célula vazia.
- A linha é um array de 1s (preenchidos) e 0s (vazios).

# --hints--

`isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1])` deve retornar `true`.

```js
assert.isTrue(isValidNonogram([3, 2], [1, 1, 1, 0, 1, 1]));
```

`isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1])` deve retornar `false`.

```js
assert.isFalse(isValidNonogram([3, 2], [0, 1, 1, 1, 1, 1]));
```

`isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` deve retornar `false`.

```js
assert.isFalse(isValidNonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]));
```

`isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` deve retornar `true`.

```js
assert.isTrue(isValidNonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` deve retornar `true`.

```js
assert.isTrue(isValidNonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]));
```

`isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` deve retornar `false`.

```js
assert.isFalse(isValidNonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]));
```

# --seed--

## --seed-contents--

```js
function isValidNonogram(clue, cells) {

  return clue;
}
```

# --solutions--

```js
function isValidNonogram(clue, cells) {
  const runs = [];
  let count = 0;

  for (let i = 0; i <= cells.length; i++) {
    if (cells[i] === 1) {
      count++;
    } else if (count > 0) {
      runs.push(count);
      count = 0;
    }
  }

  if (runs.length !== clue.length) return false;
  return runs.every((run, i) => run === clue[i]);
}
```
