---
id: 6a26df3e2988bcdded204896
title: "Desafio 360: duração falada"
challengeType: 28
dashedName: challenge-360
---

# --description--

Dado um número de segundos, retorne a duração em inglês falado.

- Divida a duração em horas, minutos e segundos.
- Ignore quaisquer valores zero.
- Use singular ou plural conforme apropriado (`"1 hour"`, `"2 hours"`).
- Se presentes, una as duas últimas unidades com `"and"` e a segunda e terceira últimas unidades com uma vírgula (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`getSpokenDuration(3723)` deve retornar `"1 hour, 2 minutes and 3 seconds"`.

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` deve retornar `"2 hours, 1 minute and 35 seconds"`.

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` deve retornar `"2 hours, 22 minutes and 1 second"`.

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` deve retornar `"7 minutes and 15 seconds"`.

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` deve retornar `"4 hours and 55 seconds"`.

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` deve retornar `"20 hours"`.

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` deve retornar `"1 second"`.

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
