---
id: 680a82806c4c6bc3567f2708
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Declare uma nova variável chamada `text` e atribua a ela a string `hello world`. Esta será a mensagem a ser criptografada.

# --hints--

Você deve declarar uma variável chamada `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

Você deve atribuir a string `hello world` a `text`. Lembre-se de colocar a string entre aspas simples ou duplas e prestar atenção à diferenciação entre maiúsculas e minúsculas.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
