---
id: 680a5e7f3d47715196b12ca0
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora que você tem todo o seu alfabeto deslocado, remova a chamada `print(shifted_alphabet)`.

# --hints--

Você deve remover `print(shifted_alphabet)` do seu código.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
print(shifted_alphabet)
--fcc-editable-region--
```
