---
id: 680a836235461fce1ff53b3f
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

O método `translate()` recebe como argumento a tabela de tradução gerada por `str.maketrans()`.É chamado em uma string e retorna uma cópia da string original onde os caracteres foram substituídos com base na tabela de tradução:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Chame o método `translate()` em `text` passando a `translation_table` como argumento e atribua o resultado a uma variável chamada `encrypted_text`.

# --hints--

Você deve ter uma variável chamada `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Você deve atribuir `text.translate(translation_table)` a `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'

--fcc-editable-region--
```
