---
id: 680a7e200400089967b8b1f4
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

O método `str.maketrans()` recebe duas strings de comprimento igual e retorna uma tabela de tradução que mapeia cada caractere da primeira string com o caractere correspondente da segunda string.Cada caractere na tabela de tradução é armazenado como um ordinal Unicode, um número que identifica exclusivamente o caractere.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Você aprenderá mais sobre o tipo de estrutura retornada por `str.maketrans()` mais adiante no currículo.Por enquanto, crie uma tabela de tradução que mapeie os caracteres em `alphabet` para os caracteres em `shifted_alphabet` e atribua-a a uma variável chamada `translation_table`.

# --hints--

Você deve declarar uma variável chamada `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

Você deve chamar `str.maketrans()` passando `alphabet` e `shifted_alphabet` como argumentos e atribuir o resultado a `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]

--fcc-editable-region--
```
