---
id: 680a549f60f5a3fedfc0edd9
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Como você pode ver na saída, o alfabeto deslocado começa na letra `f` porque `shift` tem o valor `5`. Mas agora as cinco primeiras letras do alfabeto – `a`, `b`, `c`, `d` e `e` – estão faltando no alfabeto deslocado, então você precisará adicioná-las no final do alfabeto deslocado.

O operador `+` é usado para combinar duas ou mais strings juntas em um processo chamado concatenação assim:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Modifique a atribuição existente da variável `shifted_alphabet`: use a sintaxe de slicing para extrair a primeira parte faltante de `alphabet` e concatená-la a `alphabet[shift:]`.

Como lembrete, `sentence[start:stop]` retorna os caracteres de `sentence` da posição `start` (incluída) até `stop` (excluída).

# --hints--

Você não deve modificar diretamente a variável `alphabet`.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

Você deve atribuir a concatenação de `alphabet[shift:]` e a primeira parte faltante de `alphabet` à variável `shifted_alphabet` na mesma linha em que ela foi declarada. Use `shift` para especificar onde parar o fatiamento.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
