---
id: 680a431692a5f863f726c438
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Nesta oficina, você vai construir uma *cifra de César*. Uma cifra de César criptografa o texto substituindo cada letra do texto original por uma letra um número fixo de posições à frente no alfabeto. Por exemplo, com um deslocamento de 5, `a` seria substituído por `f`, `b` por `g` e assim por diante.

Para implementar essa cifra, você precisará criar uma nova versão do seu alfabeto que comece na posição indicada pelo deslocamento. Como você aprendeu em uma lição anterior, você pode extrair parte de uma string usando fatiamento de string:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Crie uma variável chamada `shifted_alphabet` e use a sintaxe de slicing para atribuir a ela a parte de `alphabet` que começa no índice de `shift`. Então, chame a função embutida `print()` para imprimir `shifted_alphabet` no terminal e veja o resultado.

# --hints--

Você deve declarar uma nova variável chamada `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

Você deve atribuir `alphabet[shift:]` a `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

Você deve chamar a função `print()` passando `shifted_alphabet` como argumento.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
