---
id: 680a3d4c065557260de2ff33
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Como você deve se lembrar das lições anteriores, em Python, você declara uma variável escrevendo o nome da variável no lado esquerdo do operador de atribuição (`=`) e o valor a ser atribuído no lado direito:

```py
variable_name = value
```

Crie uma variável chamada `shift` e atribua o valor `5` à sua nova variável.

# --hints--

Você deve declarar uma variável chamada `shift`. Preste atenção para colocar o nome da variável no início da linha.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

Você deve atribuir o valor `5` à sua variável `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
