---
id: 680a41ef5bfdc059cfd9d4b5
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Declare outra variável chamada `alphabet` e atribua a string `abcdefghijklmnopqrstuvwxyz` a essa variável.

# --hints--

Você deve declarar uma variável chamada `alphabet`. Certifique-se de colocar o nome da variável no início da linha.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

Você deve atribuir a string `abcdefghijklmnopqrstuvwxyz` à sua variável `alphabet`. Lembre-se de usar aspas simples ou duplas para delimitar a string.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
