---
id: 6818eb3f98c0bca5af7c8d03
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Agora seu código é reutilizável. No entanto, a função `caesar` imprime uma mensagem no terminal e retorna `None` por padrão. Para provar isso, imprima `encrypted_text` no final do seu código. 

# --hints--

Você deve imprimir `encrypted_text` fora do corpo da função.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
