---
id: 6818fcb250f34e62cab32c39
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Você pode ter notado que, embora a mensagem esteja criptografada, os caracteres maiúsculos foram deixados intactos. Isso acontece porque a tabela de tradução não inclui letras maiúsculas.

Para corrigir, você precisará modificar as strings passadas para sua chamada `str.maketrans()` usando o método `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Atualize sua chamada `str.maketrans()` concatenando a cada argumento a versão maiúscula do próprio argumento.

# --hints--

Você deve modificar os argumentos da sua chamada `str.maketrans()` concatenando a cada argumento a versão maiúscula do próprio argumento.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
