---
id: 6818ed1adc249bb91d373729
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Nas lições anteriores, você aprendeu que a declaração `return` é usada para retornar um valor de uma função, para que você possa usá-lo em outra parte do seu código:

```py
def spam():
    return 'Spam!'
```

Remova a chamada `print(encrypted_text)` da sua função. Então, exclua a declaração da variável `encrypted_text` e retorne diretamente `text.translate(translation_table)` da sua função.

# --hints--

Você não deve ter `print(encrypted_text)` dentro da função `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

Você não deve ter uma variável `encrypted_text` dentro da função `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Sua função `caesar` deve retornar `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)

```
