---
id: 6818e9c7c8dc5694e8f64fc3
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Agora, teste a função `caesar` chamando-a com a string `freeCodeCamp` e o número `3` como argumentos.Atribua a chamada da função a uma variável chamada `encrypted_text`.

# --hints--

Você deve ter uma variável chamada `encrypted_text` no escopo global.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

Você deve atribuir `caesar('freeCodeCamp', 3)` à variável `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--


--fcc-editable-region--
```
