---
id: 6818b80519e6c6c47a8bb6e8
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Como você aprendeu em uma lição anterior, funções podem ter parâmetros, que são variáveis que podem ser referenciadas dentro da função. Aqui está uma função `sum` com dois parâmetros, `num1` e `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

A mensagem a ser criptografada e o deslocamento ainda estão codificados diretamente na sua função. Dê à sua função dois parâmetros chamados `text` e `shift` e exclua as declarações das variáveis `text` e `shift` do corpo da função `caesar`.

# --hints--

Sua função `caesar` deve ter dois parâmetros, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

Você não deve ter uma variável `text` declarada dentro do corpo da função `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

Você não deve ter uma variável `shift` declarada dentro do corpo da função `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
