---
id: 6819be9055828d65c813eec5
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

A função `isinstance()` retorna `True` se seu primeiro argumento for uma instância do segundo argumento e `False` caso contrário:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

Substitua a condição atual do seu `if` por uma chamada `isinstance()`.Passe `shift` como o primeiro argumento e `int` como o segundo argumento. 

# --hints--

Sua declaração `if` deve usar `isinstance(shift, int)` como condição.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
    if True:
        return 'Shift must be an integer value.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
