---
id: 6819c5c9faf2548e556163d5
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Como você se lembra das lições anteriores, o operador lógico unário `not` nega uma expressão:

```py
print(not True) # False
print(not False) # True
```

Use o operador `not` para corrigir a condição da declaração `if` para que sua função retorne `Shift must be an integer value.` quando `shift` *não* for um inteiro.

# --hints--

Você deve usar o operador `not` para negar `isinstance(shift, int)` na sua declaração `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
