---
id: 6819cf38880d8cc9b9f6af7a
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Python permite que você especifique um valor padrão para os parâmetros em uma função, criando uma função que pode ser chamada com menos ou nenhum argumento. Aqui está como criar uma função com um parâmetro `name` que tem um valor padrão:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Adicione um terceiro parâmetro chamado `encrypt` à sua função e dê a ele um valor padrão de `True`. 

# --hints--

Você deve adicionar um terceiro parâmetro chamado `encrypt` com um valor padrão de `True` à sua função. Não modifique os parâmetros existentes.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
