---
id: 6819ca464db219ab86f37dc4
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Como você já verificou, o deslocamento passado para criptografar o texto deve ser positivo.Não pode exceder `25` (o último índice de `alphabet`) porém.

Adicione uma segunda condição à declaração `if` que verifica se `shift` é maior que `25`.Lembre-se de que a operação lógica OU em Python é implementada através do operador `or`.

Além disso, atualize a mensagem retornada para `Shift must be an integer between 1 and 25.`.

# --hints--

Você deve usar o operador `or` para adicionar uma segunda condição após `shift < 1` verificando se `shift` é maior que `25`.Você não deve modificar a condição existente.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Sua segunda declaração `if` deve retornar `Shift must be an integer between 1 and 25.`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)


encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
