---
id: 6819bc0637b80256a33adccc
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Agora que você implementou as funcionalidades básicas do `cipher`, é hora de adicionar algumas validações. Para isso, você precisará de uma declaração `if`. Aqui está um lembrete da sintaxe para uma declaração `if`:

```py
if condition:
    # code to run when condition is true
```

No início do corpo da sua função, crie uma declaração `if`. Por enquanto, use `True` como condição e dentro do corpo da declaração `if` retorne a string `Shift must be an integer value.`

# --hints--

Você deve ter uma declaração `if` dentro da sua função `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Sua declaração `if` deve usar `True` como sua condição.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Você deve retornar a string `Shift must be an integer value.` do seu comando `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
