---
id: 681a2a6cf95e82ca30866364
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Declare duas funções chamadas `encrypt` e `decrypt`, ambas com os parâmetros `text` e `shift`.

Você vai usar `encrypt` para criptografia e `decrypt` para descriptografia, dando a cada ação um nome descritivo.

De ambas as funções novas, retorne uma chamada `caesar` que passe `text` e `shift`. Em `decrypt`, também passe `False` como terceiro argumento.

# --hints--

Você deve ter uma função chamada `encrypt` com dois parâmetros, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Sua função `encrypt` deve retornar `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

Você deve ter uma função chamada `decrypt` com dois parâmetros, `text` e `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Sua função `decrypt` deve retornar `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
