---
id: 681a1d60e037127e7c8e0230
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Você vai usar o parâmetro adicionado na etapa anterior para determinar se a função deve criptografar o texto passado para ela (comportamento padrão, `encrypt=True`), ou se deve descriptografar uma mensagem criptografada.

Crie uma declaração `if` que verifique se `encrypt` não é verdadeiro. Dentro da nova declaração `if`, defina `shift` como `- shift`. Isso é necessário para permitir que o deslocamento ocorra na direção oposta em relação ao processo de criptografia.

# --hints--

Você deve definir `shift` como `- shift` quando `encrypt` não for verdadeiro.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
