---
id: 63ec3287b182ec0efe8a3135
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

このテーブルの残りの製品を（順番に）`products` 配列に `new Dessert(...)` を使って追加してください。製品ごとに `id` をインクリメントし、カウントアップしてください。

| `name`                         | `price` | `category`  |
| ------------------------------ | ------- | ----------- |
| `French Macaron`               | `3.99`  | `Macaron`   |
| `Pumpkin Cupcake`              | `3.99`  | `Cupcake`   |
| `Chocolate Cupcake`            | `5.99`  | `Cupcake`   |
| `Chocolate Pretzels (4 Pack)`  | `10.99` | `Pretzel`   |
| `Strawberry Ice Cream`         | `2.99`  | `Ice Cream` |
| `Chocolate Macarons (4 Pack)`  | `9.99`  | `Macaron`   |
| `Strawberry Pretzel`           | `4.99`  | `Pretzel`   |
| `Butter Pecan Ice Cream`       | `2.99`  | `Ice Cream` |
| `Rocky Road Ice Cream`         | `2.99`  | `Ice Cream` |
| `Vanilla Macarons (5 Pack)`    | `11.99` | `Macaron`   |
| `Lemon Cupcakes (4 Pack)`      | `12.99` | `Cupcake`   |

# --hints--

2番目の製品には `id` プロパティが数値 `2` に設定されている必要があります。

```js
assert.strictEqual(products[1].id, 2);
```

2番目の製品には `name` プロパティが `French Macaron` に設定されている必要があります。

```js
assert.strictEqual(products[1].name, 'French Macaron');
```

2番目の製品には `price` プロパティが数値 `3.99` に設定されている必要があります。

```js
assert.strictEqual(products[1].price, 3.99);
```

2番目の製品には `category` プロパティが `Macaron` に設定されている必要があります。

```js
assert.strictEqual(products[1].category, 'Macaron');
```

3番目の製品には `id` プロパティが数値 `3` に設定されている必要があります。

```js
assert.strictEqual(products[2].id, 3);
```

3番目の製品には `name` プロパティが `Pumpkin Cupcake` に設定されている必要があります。

```js
assert.strictEqual(products[2].name, 'Pumpkin Cupcake');
```

3番目の製品には `price` プロパティが数値 `3.99` に設定されている必要があります。

```js
assert.strictEqual(products[2].price, 3.99);
```

3番目の製品には `category` プロパティが `Cupcake` に設定されている必要があります。

```js
assert.strictEqual(products[2].category, 'Cupcake');
```

4番目の製品には `id` プロパティが数値 `4` に設定されている必要があります。

```js
assert.strictEqual(products[3].id, 4);
```

4番目の製品には `name` プロパティが `Chocolate Cupcake` に設定されている必要があります。

```js
assert.strictEqual(products[3].name, 'Chocolate Cupcake');
```

4番目の製品には `price` プロパティが数値 `5.99` に設定されている必要があります。

```js
assert.strictEqual(products[3].price, 5.99);
```

4番目の製品には `category` プロパティが `Cupcake` に設定されている必要があります。

```js
assert.strictEqual(products[3].category, 'Cupcake');
```

5番目の製品には `id` プロパティが数値 `5` に設定されている必要があります。

```js
assert.strictEqual(products[4].id, 5);
```

5番目の製品には `name` プロパティが `Chocolate Pretzels (4 Pack)` に設定されている必要があります。

```js
assert.strictEqual(products[4].name, 'Chocolate Pretzels (4 Pack)');
```

5番目の製品には `price` プロパティが数値 `10.99` に設定されている必要があります。

```js
assert.strictEqual(products[4].price, 10.99);
```

5番目の製品には `category` プロパティが `Pretzel` に設定されている必要があります。

```js
assert.strictEqual(products[4].category, 'Pretzel');
```

6番目の製品には `id` プロパティが数値 `6` に設定されている必要があります。

```js
assert.strictEqual(products[5].id, 6);
```

6番目の製品には `name` プロパティが `Strawberry Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[5].name, 'Strawberry Ice Cream');
```

6番目の製品には `price` プロパティが数値 `2.99` に設定されている必要があります。

```js
assert.strictEqual(products[5].price, 2.99);
```

6番目の製品には `category` プロパティが `Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[5].category, 'Ice Cream');
```

7番目の製品には `id` プロパティが数値 `7` に設定されている必要があります。

```js
assert.strictEqual(products[6].id, 7);
```

7番目の製品には `name` プロパティが `Chocolate Macarons (4 Pack)` に設定されている必要があります。

```js
assert.strictEqual(products[6].name, 'Chocolate Macarons (4 Pack)');
```

7番目の製品には `price` プロパティが数値 `9.99` に設定されている必要があります。

```js
assert.strictEqual(products[6].price, 9.99);
```

7番目の製品には `category` プロパティが `Macaron` に設定されている必要があります。

```js
assert.strictEqual(products[6].category, 'Macaron');
```

8番目の製品には `id` プロパティが数値 `8` に設定されている必要があります。

```js
assert.strictEqual(products[7].id, 8);
```

8番目の製品には `name` プロパティが `Strawberry Pretzel` に設定されている必要があります。

```js
assert.strictEqual(products[7].name, 'Strawberry Pretzel');
```

8番目の製品には `price` プロパティが数値 `4.99` に設定されている必要があります。

```js
assert.strictEqual(products[7].price, 4.99);
```

8番目の製品には `category` プロパティが `Pretzel` に設定されている必要があります。

```js
assert.strictEqual(products[7].category, 'Pretzel');
```

9番目の製品には `id` プロパティが数値 `9` に設定されている必要があります。

```js
assert.strictEqual(products[8].id, 9);
```

9番目の製品には `name` プロパティが `Butter Pecan Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[8].name, 'Butter Pecan Ice Cream');
```

9番目の製品には `price` プロパティが数値 `2.99` に設定されている必要があります。

```js
assert.strictEqual(products[8].price, 2.99);
```

9番目の製品には `category` プロパティが `Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[8].category, 'Ice Cream');
```

10番目の製品には `id` プロパティが数値 `10` に設定されている必要があります。

```js
assert.strictEqual(products[9].id, 10);
```

10番目の製品には `name` プロパティが `Rocky Road Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[9].name, 'Rocky Road Ice Cream');
```

10番目の製品には `price` プロパティが数値 `2.99` に設定されている必要があります。

```js
assert.strictEqual(products[9].price, 2.99);
```

10番目の製品には `category` プロパティが `Ice Cream` に設定されている必要があります。

```js
assert.strictEqual(products[9].category, 'Ice Cream');
```

11番目の製品には `id` プロパティが数値 `11` に設定されている必要があります。

```js
assert.strictEqual(products[10].id, 11);
```

11番目の製品には `name` プロパティが `Vanilla Macarons (5 Pack)` に設定されている必要があります。

```js
assert.strictEqual(products[10].name, 'Vanilla Macarons (5 Pack)');
```

11番目の製品には `price` プロパティが数値 `11.99` に設定されている必要があります。

```js
assert.strictEqual(products[10].price, 11.99);
```

11番目の製品には `category` プロパティが `Macaron` に設定されている必要があります。

```js
assert.strictEqual(products[10].category, 'Macaron');
```

12番目の製品には `id` プロパティが数値 `12` に設定されている必要があります。

```js
assert.strictEqual(products[11].id, 12);
```

12番目の製品には `name` プロパティが `Lemon Cupcakes (4 Pack)` に設定されている必要があります。

```js
assert.strictEqual(products[11].name, 'Lemon Cupcakes (4 Pack)');
```

12番目の製品には `price` プロパティが数値 `12.99` に設定されている必要があります。

```js
assert.strictEqual(products[11].price, 12.99);
```

12番目の製品には `category` プロパティが `Cupcake` に設定されている必要があります。

```js
assert.strictEqual(products[11].category, 'Cupcake');
```

`products` 配列には 12 個の値があり、それぞれが `Dessert` のインスタンスである必要があります。

```js
assert.strictEqual(products.length, 12);
products.forEach(p => assert.instanceOf(p, Dessert));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp Shopping Cart</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1 class="title">Desserts Page</h1>
    </header>
    <main>
      <button id="cart-btn" type="button" class="btn">
        <span id="show-hide-cart">Show</span> Cart
      </button>
      <div id="cart-container">
        <button class="btn" id="clear-cart-btn">Clear Cart</button>
        <div id="products-container"></div>
        <p>Total number of items: <span id="total-items">0</span></p>
        <p>Subtotal: <span id="subtotal">$0</span></p>
        <p>Taxes: <span id="taxes">$0</span></p>
        <p>Total: <span id="total">$0</span></p>
      </div>
      <div id="dessert-card-container"></div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.title {
  color: var(--light-grey);
  text-align: center;
  margin: 25px 0;
}

#dessert-card-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
}

.dessert-card {
  background-color: var(--light-grey);
  padding: 15px;
  text-align: center;
  border-radius: 15px;
  margin: 20px 10px;
}

.dessert-price {
  font-size: 1.2rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 100px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#cart-btn {
  position: absolute;
  top: 0;
  right: 0;
}

.add-to-cart-btn {
  margin: 30px auto 10px;
}

#cart-container {
  display: none;
  position: absolute;
  top: 60px;
  right: 0;
  background-color: var(--light-grey);
  width: 200px;
  height: 400px;
  border: 8px double var(--black);
  border-radius: 15px;
  text-align: center;
  font-size: 1.2rem;
  overflow-y: scroll;
}

.product {
  margin: 25px 0;
}

.product-count {
  display: inline-block;
  margin-right: 10px;
}

.product-category {
  margin: 10px 0;
}

@media (min-width: 768px) {
  #dessert-card-container {
    flex-direction: row;
  }

  .dessert-card {
    flex: 1 0 21%;
  }

  #cart-container {
    width: 300px;
  }
}
```

```js
const cartContainer = document.getElementById("cart-container");
const productsContainer = document.getElementById("products-container");
const dessertCards = document.getElementById("dessert-card-container");
const cartBtn = document.getElementById("cart-btn");
const clearCartBtn = document.getElementById("clear-cart-btn");
const totalNumberOfItems = document.getElementById("total-items");
const cartSubTotal = document.getElementById("subtotal");
const cartTaxes = document.getElementById("taxes");
const cartTotal = document.getElementById("total");
const showHideCartSpan = document.getElementById("show-hide-cart");
let isCartShowing = false;

class Dessert {
  constructor(id, name, price, category) {
    this.id = id;
    this.name = name;
    this.price = price;
    this.category = category;
  }
}

--fcc-editable-region--
const products = [
  new Dessert(1, "Vanilla Cupcakes (6 Pack)", 12.99, "Cupcake"),

];
--fcc-editable-region--
```
