---
id: 5a858944d96184f06fd60d61
title: 創建你的第一個 CSS Grid
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

將任何 HTML 元素的 `display` 屬性設定為 `grid`，即可將其轉換為網格容器。這讓你能使用所有其他與 CSS Grid 相關的屬性。

**注意：** 在 CSS Grid 中，父元素稱為 <dfn>容器</dfn>，其子元素稱為 <dfn>項目</dfn>。

# --instructions--

將具有 `container` 類別的 div 的顯示方式更改為 `grid`。

# --hints--

`container` 類別應該有一個 `display` 屬性，值為 `grid`。

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
