---
id: 5a9036ee38fddaf9a66b5d35
title: 使用 grid-column-gap 創建行間隙
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cVZ8vfD'
forumTopicId: 301124
dashedName: create-a-column-gap-using-grid-column-gap
---

# --description--

到目前為止，在你創建的格線中，行都緊密相連。有時你會想要在行之間留有間隙。要在行之間添加間隙，請使用 `grid-column-gap` 屬性，如下所示：

```css
grid-column-gap: 10px;
```

這會在我們所有的行之間產生 10px 的空白。

# --instructions--

為格線中的行添加 `20px` 間距。

# --hints--

`container` 類別應該有一個 `grid-column-gap` 屬性，其值為 `20px`。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-column-gap\s*?:\s*?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

   
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-column-gap: 20px;
  }
</style>
```
