---
id: 5a90375238fddaf9a66b5d3b
title: 使用 align-self 垂直對齊一個項目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzd4fz'
forumTopicId: 301123
dashedName: align-an-item-vertically-using-align-self
---

# --description--

正如你可以水平對齊一個項目一樣，也有方法可以垂直對齊一個項目。為了做到這一點，你可以在項目上使用 `align-self` 屬性。這個屬性接受與上一個挑戰中 `justify-self` 相同的所有值。

# --instructions--

將該項目與類別 `item3` 在 `end` 垂直對齊。

# --hints--

`item3` 類別應該有一個 `align-self` 屬性，其值為 `end`。

```js
assert.match(code, /.item3\s*?{[\s\S]*align-self\s*?:\s*?end\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }

  .item3 {
    background: PaleTurquoise;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    align-self: end;
  }
</style>
```
