---
id: 5a9036ee38fddaf9a66b5d37
title: 更快地為 `grid-gap` 添加間隙
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` 是前兩個挑戰中 `grid-row-gap` 和 `grid-column-gap` 的簡寫屬性，更方便使用。如果 `grid-gap` 有一個值，它會在所有列和行之間創建間距。然而，如果有兩個值，它會使用第一個值設定列之間的間距，第二個值設定行之間的間距。

# --instructions--

使用 `grid-gap` 在列之間介紹 `10px` 間距，在行之間介紹 `20px` 間距。

# --hints--

`container` 類別應該有一個 `grid-gap` 屬性，為列之間添加 `10px` 的間隙，為行之間添加 `20px` 的間隙。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
