---
id: 5a90374338fddaf9a66b5d3a
title: 使用 justify-self 將項目水平對齊
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

在 CSS Grid 中，每個項目的內容位於一個稱為 <dfn>cell</dfn> 的框中。你可以使用網格項目的 `justify-self` 屬性水平對齊內容在其 cell 中的位置。預設情況下，這個屬性的值為 `stretch`，會使內容填滿整個 cell 的寬度。這個 CSS Grid 屬性也接受其他值：

`start`：將內容對齊至儲存格的左側，

`center`：將內容置中對齊於儲存格內，

`end`：將內容對齊至儲存格的右側。

# --instructions--

使用 `justify-self` 屬性將類別為 `item2` 的項目置中。

# --hints--

`item2` 類別應該有一個 `justify-self` 屬性，其值為 `center`。

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
