---
id: 5a90376038fddaf9a66b5d3c
title: 使用 justify-items 將所有項目水平對齊
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

有時你會希望你的 CSS Grid 中所有的項目都共享相同的對齊方式。你可以使用先前學過的屬性並分別對齊它們，或者你可以透過在你的網格容器上使用 `justify-items`，一次水平對齊所有項目。這個屬性可以接受你在前兩個挑戰中學到的所有相同值，不同之處在於它會將我們網格中的**所有**項目移動到所需的對齊位置。

# --instructions--

使用此屬性將所有項目水平置中。

# --hints--

`container` 類別應該有一個 `justify-items` 屬性，其值為 `center`。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
