---
id: 5a9036ee38fddaf9a66b5d34
title: 使用 CSS Grid 單位來改變行和列的大小
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

你可以在 CSS Grid 中使用絕對和相對單位，如 `px` 和 `em`，來定義列和行的大小。你也可以使用這些：

`fr`：將行或列設定為可用空間的分數，

`auto`：自動將行或列設為其內容的寬度或高度，

`%`：將行或列調整為其容器的百分比寬度。

這是用來在預覽中產生輸出的程式碼：

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

這段程式碼創建五個行。第一個行的寬度與其內容相同，第二個行為 50px，第三個行為其容器的 10％，至於最後兩個行；剩餘空間被分成三個部分，其中兩部分配置給第四個行，一部分配置給第五個行。

# --instructions--

建立一個具有三個行的網格，其寬度分別為：1fr、100px 和 2fr。

# --hints--

你的 `container` 類別應該有一個 `grid-template-columns` 屬性，該屬性包含三個行，寬度分別為 `1fr`、`100px` 和 `2fr`。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```
