---
id: 5a9036d038fddaf9a66b5d32
title: 為你的 `grid-template-columns` 添加行
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

僅僅建立一個網格元素並不會有太大進展。你還需要定義網格的結構。若要為網格添加一些行，請在網格容器上使用 `grid-template-columns` 屬性，如下所示：

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

這會為你的網格提供兩個寬度各為 50px 的行。給予 `grid-template-columns` 屬性的參數數量顯示網格中的行數，而每個參數的值顯示每個行的寬度。

# --instructions--

給網格容器三個寬度皆為 `100px` 的行。

# --hints--

`container` 類別應該有一個 `grid-template-columns` 屬性，設定為三個單位的 `100px`。

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```
