---
id: 5a94fe4469fb03452672e460
title: 使用 minmax 函式限制項目大小
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

還有另一個內建函式（程式）可與 `grid-template-columns` 和 `grid-template-rows` 一起使用，稱為 `minmax`。它用於在網格容器大小改變時限制項目的大小。為此，你需要為你的項目指定可接受的大小範圍。以下是一個範例：

```css
grid-template-columns: 100px minmax(50px, 200px);
```

在上述程式碼中，`grid-template-columns` 被設定為創建兩個行；第一個行寬度為 100px，第二個行的最小寬度為 50px，最大寬度為 200px。

# --instructions--

使用 `minmax` 函式，將 `repeat` 函式中的 `1fr` 替換為最小寬度為 `90px`、最大寬度為 `1fr` 的行大小，並調整預覽面板的大小以查看效果。

# --hints--

`container` 類別應該有一個 `grid-template-columns` 屬性，設定為重複 3 行，最小寬度為 `90px`，最大寬度為 `1fr`。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
