---
id: 5a94fe1369fb03452672e45d
title: 使用 `grid-area` 屬性將項目放置在網格區域中
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

在為你的網格容器建立區域模板後，如前一個挑戰所示，你可以透過取用你給予它的名稱，將項目放置在你的訂製區域中。為此，你可以在項目上使用 `grid-area` 屬性，如下所示：

```css
.item1 {
  grid-area: header;
}
```

這讓網格知道你想要將 `item1` 類別放在名為 `header` 的區域。在這種情況下，該項目將使用整個頂部列，因為整列都被命名為 `header` 區域。

# --instructions--

使用 `grid-area` 屬性將具有 `item5` 類別的元素放置在 `footer` 區域。

# --hints--

`item5` 類別應該有一個 `grid-area` 屬性，其值為 `footer`。

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
