---
id: 5a94fe8569fb03452672e464
title: 在格子中創建格子
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

將元素轉換為網格只會影響其直接子元素的行為。因此，透過將直接子元素轉換為網格，你就擁有了網格中的網格。

例如，透過設定具有 `item3` 類別的元素的 `display` 和 `grid-template-columns` 屬性，你可以在你的網格中創建一個網格。

# --instructions--

將具有 `item3` 類別的元素使用 `display` 和 `grid-template-columns` 轉換為具有兩個行，寬度分別為 `auto` 和 `1fr` 的網格。

# --hints--

`item3` 類別應該有一個 `grid-template-columns` 屬性，其值為 `auto` 和 `1fr`。

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

`item3` 類別應該有一個 `display` 屬性，值為 `grid`。

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
