---
id: 5a94fe2669fb03452672e45e
title: 使用 grid-area 而不建立 Areas 模板
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

你在上一個挑戰中學到的 `grid-area` 屬性可以用另一種方式使用。如果你的網格沒有可供參考的區域範本，你可以即時為一個項目創建一個區域，並像這樣放置：

```css
item1 {
  grid-area: 1/1/2/4;
}
```

這是使用你先前學習過的行號來定義此項目的區域位置。上面範例中的數字表述這些數值：

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

因此，範例中的條款將會佔用第 1 列與第 2 列之間的列，以及第 1 行與第 4 行之間的行。

# --instructions--

使用 `grid-area` 屬性，將具有 `item5` 類別的元素放置在第三和第四條水平線之間，以及第一和第四條垂直線之間。

# --hints--

`item5` 類別應該有一個 `grid-area` 屬性，使其填滿第三和第四條水平線，以及第一和第四條垂直線之間的整個區域。

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
