---
id: 5a94fe6269fb03452672e462
title: 使用 auto-fit 創建彈性版面配置
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c3dPph8'
forumTopicId: 301127
dashedName: create-flexible-layouts-using-auto-fit
---

# --description--

`auto-fit` 的運作方式幾乎與 `auto-fill` 相同。唯一的差別是當容器的大小超過所有項目總和的大小時，`auto-fill` 會持續插入空的列或行並將你的項目推到一旁，而 `auto-fit` 則會摺疊那些空的列或行並拉伸你的項目以適應容器的大小。

**注意：** 如果你的容器無法將所有項目放在同一列，它會將它們移到下一列。

# --instructions--

在第二個網格中，使用 `auto-fit` 和 `repeat` 來填滿具有最小寬度為 `60px` 且最大為 `1fr` 的行。然後調整預覽大小以查看差異。

# --hints--

`container2` 類別應該有一個帶有 `repeat` 和 `auto-fit` 的 `grid-template-columns` 屬性，該屬性將以最小寬度為 `60px`、最大為 `1fr` 的行填滿網格。

```js
assert.match(
  code,
  /.container2\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fit\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));
  }
  .container2 {
    grid-template-columns: repeat(auto-fit, minmax(60px, 1fr));
  }
</style>
```
