---
id: 5a94fe3669fb03452672e45f
title: 使用 `repeat` 函式減少重複
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

當你使用 `grid-template-columns` 和 `grid-template-rows` 來定義網格的結構時，你為你創建的每一列或每一列輸入了一個值。

假設你想要一個具有 100 列相同高度的網格。逐一插入 100 個值並不實際。幸運的是，有更好的方法－使用 `repeat` 函式來指定你想要重複的行或列數量，接著用逗號和你想要重複的值。

這是一個範例，將創建 100 列的網格，每列高度為 50px。

```css
grid-template-rows: repeat(100, 50px);
```

你也可以使用 repeat 函式重複多個值，並在定義網格結構時將該函式插入其他值中。以下是範例：

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

這表示：

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**注意：** `1fr 50px` 重複兩次，後面接著 20px。

# --instructions--

使用 `repeat` 來移除 `grid-template-columns` 屬性中的重複。

# --hints--

`container` 類別應該有一個 `grid-template-columns` 屬性，設定為重複 3 個寬度為 `1fr` 的行。

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
