---
id: 68e6698b29e562d21b1d260c
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop erstellen Sie eine Liste der wichtigsten Webbrowser. Der Boilerplate-Code für HTML wurde bereits für Sie bereitgestellt.

Beginnen Sie damit, eine Überschrift auf Ihrer Seite hinzuzufügen, die `List of Major Web Browsers` lautet, und verwenden Sie dafür ein `h1`-Element innerhalb des `body`-Elements.

# --hints--

Sie sollten ein `h1`-Element innerhalb des `body`-Elements haben.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Ihr `h1`-Element sollte den Text `List of Major Web Browsers` enthalten.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
