---
id: 68ecfa1003d88990bb96a33e
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, Ihren zweiten Browser zur Liste hinzuzufügen.

Fügen Sie unter Ihrem `Google Chrome`-Eintrag ein weiteres `dt`-Element hinzu, das den Text `Firefox` enthält.

Unter Ihrem `dt`-Element fügen Sie ein `dd`-Element mit folgendem Text hinzu:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

Sie sollten zwei `dt`-Elemente in Ihrem Code haben.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Ihr zweites `dt`-Element sollte den Text `Firefox` enthalten.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

Sie sollten zwei `dd`-Elemente in Ihrem Code haben.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Ihr zweites `dd`-Element sollte den Text `This is a free web browser developed by the Mozilla Corporation and first created in 2004.` enthalten.

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    <dl>
      <dt>Google Chrome</dt>
      <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        --fcc-editable-region--
      
        --fcc-editable-region--
    </dl>
  </body> 
</html>
```

