---
id: 68ecfa0f03d88990bb96a33c
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Wie Sie sich aus einer früheren Lektion erinnern, werden Definitionslisten verwendet, um Begriffe und Definitionen in einem organisierten und leicht lesbaren Format darzustellen.

Hier ist ein Beispiel:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Erstellen Sie unter dem `h1`-Element ein `dl`-Element. Dieses wird die Liste der Browser enthalten.

# --hints--

Sie sollten ein `dl`-Element haben.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Ihr `dl`-Element sollte unter dem `h1`-Element sein.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```

