---
id: 68ecfa1003d88990bb96a33d
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Beschreibungslisten bestehen aus einer Liste von Begriffen und Details. Das `dt`-Element stellt den Beschreibungsterm dar, während das `dd`-Element das Beschreibungsdetail darstellt.

Der erste Browser, den Sie zu Ihrer Beschreibungslist hinzufügen, wird Google Chrome sein.

Fügen Sie innerhalb Ihres `dl`-Elements ein `dt`-Element mit dem Text `Google Chrome` hinzu.

Unter Ihrem `dt`-Element fügen Sie ein `dd`-Element mit folgendem Text hinzu:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Sie sollten ein `dt`-Element innerhalb Ihres `dl`-Elements haben.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Ihr `dt`-Element sollte den Text `Google Chrome` enthalten.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Sie sollten ein `dd`-Element unter Ihrem `dt`-Element haben.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Ihr `dd`-Element sollte den Text `This is a free web browser developed by Google and first released in 2008.` enthalten.

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

