---
id: 69038dd51a41ae215fcc030a
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Der dritte Browser, den Sie zur Liste hinzufügen, wird der Safari-Webbrowser sein.

Fügen Sie ein weiteres `dt`-Element hinzu, das den Text `Safari` enthält.

Unter Ihrem `dt`-Element fügen Sie ein `dd`-Element mit folgendem Text hinzu:

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

Sie sollten drei `dt`-Elemente in Ihrem Code haben.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

Ihr drittes `dt`-Element sollte den Text `Safari` enthalten.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

Sie sollten drei `dd`-Elemente in Ihrem Code haben.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

Ihr drittes `dd`-Element sollte den Text `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.` enthalten.

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    <dl>
      <dt>Google Chrome</dt>
      <dd>This is a free web browser developed by Google and first released in 2008.</dd>

      <dt>Firefox</dt>
      <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        --fcc-editable-region--
      
        --fcc-editable-region--
    </dl>
  </body> 
</html>
```
