---
id: bad87fee1348bd9aedf08823
title: Ongeza ukingo hasi kwa kipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpyGs3'
forumTopicId: 16166
dashedName: add-a-negative-margin-to-an-element
---

# --description--

Ukingo wa kipengele `margin` unadhibiti kiasi cha nafasi kati ya ukingo wa kipengele `border` na vipengele vinavyouzunguka.

Ikiwa utaweka ukingo wa kipengele `margin` kuwa thamani hasi, kipengele hicho kitaongezeka ukubwa.

# --instructions--

Jaribu kuweka ukingo `margin` kuwa thamani hasi kama ile ya kisanduku cheusi.

Badilisha ukingo `margin` wa kisanduku cha buluu kuwa `-15px`, ili kijaze upana wote wa usawa wa kisanduku cha njano kinachokizunguka.

# --hints--

Darasa lako la `blue-box` linapaswa kuwapa vipengele ukingo `-15px` wa `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)["margin-top"]; 
assert.strictEqual(marginTop, "-15px"); 
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: -15px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
    margin-top: -15px;
  }
</style>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
