---
id: bad87fee1348bd9aedf08822
title: Badilisha ukingo wa kipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJarHW'
forumTopicId: 16654
dashedName: adjust-the-margin-of-an-element
---

# --description--

`margin` ya kipengele inadhibiti kiasi cha nafasi kati ya `border` ya kipengele na vipengele vinavyozunguka.

Hapa, tunaona kuwa kisanduku cha buluu na kisanduku cha nyekundu vimepangwa ndani ya kisanduku cha manjano. Kumbuka kuwa kisanduku cha nyekundu kina `margin` kubwa zaidi kuliko kisanduku cha buluu, na kufanya kionekane kidogo.

Unapoongeza `margin` ya kisanduku cha buluu, itaongeza umbali kati ya mpaka wake na vipengele vinavyozunguka.

# --instructions--

Badilisha `margin` ya kisanduku cha buluu ili ifanane na ile ya kisanduku cha nyekundu.

# --hints--

Darasa lako la `blue-box` linapaswa kutoa vipengele `20px` vya `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
    margin: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
