---
id: bad87fee1348bd9aedf04756
title: Ghairi mitindo katika CSS inayofuata
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDQug'
forumTopicId: 18253
dashedName: override-styles-in-subsequent-css
---

# --description--

Darasa letu la `pink-text` lilighairi tamko la CSS la kipengele chetu cha `body`!

Tumeonyesha tu kwamba madarasa yetu yata ghairi CSS ya kipengele cha `body`. Hivyo swali linalofuata kwa mantiki ni, tunawezaje ghairi darasa letu la `pink-text`?

# --instructions--

Tengeneza darasa la ziada la CSS linaloitwa `blue-text` linalompa kipengele rangi ya buluu. Hakikisha liko chini ya tamko la darasa lako la `pink-text`.

Tumia darasa la `blue-text` kwa kipengele chako cha `h1` pamoja na darasa lako la `pink-text`, na tuone ni lipi linaloshinda.

Kutumia sifa nyingi za darasa kwa kipengele cha HTML hufanywa kwa kuweka nafasi kati yao kama hii:

```html
class="class1 class2"
```

**Kumbuka:** Haijalishi ni kwa mpangilio gani madarasa yameorodheshwa ndani ya kipengele cha HTML.

Hata hivyo, mpangilio wa tamko la `class` katika sehemu ya `<style>` ndio muhimu. Tamko la pili daima huchukua kipaumbele juu ya la kwanza. Kwa sababu `.blue-text` limetangazwa la pili, linaghairi sifa za `.pink-text`.

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na darasa la `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Kipengele chako cha `h1` kinapaswa kuwa na darasa la `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Zote `blue-text` na `pink-text` zinapaswa kuwa sehemu ya kipengele kimoja cha `h1`.

```js
assert.isTrue(document.querySelector('.pink-text').classList.contains('blue-text'));
```

Kipengele chako cha `h1` kinapaswa kuwa na rangi ya buluu.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }

  .blue-text {
    color: blue;
  }  
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```
