---
id: bad87fee1348bd9aedf08719
title: Tumia msimbo mfupi wa hex
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpKAm'
forumTopicId: 18338
dashedName: use-abbreviated-hex-code
---

# --description--

Watu wengi huhisi kuzidiwa na uwezekano wa rangi zaidi ya milioni 16. Na ni vigumu kukumbuka msimbo wa hex. Kwa bahati nzuri, unaweza kuufupisha.

Kwa mfano, msimbo wa hex wa rangi nyekundu `#FF0000` unaweza kufupishwa kuwa `#F00`. Fomu hii iliyofupishwa hutoa tarakimu moja kwa nyekundu, tarakimu moja kwa kijani, na tarakimu moja kwa bluu.

Hii hupunguza jumla ya rangi zinazowezekana kuwa takriban 4,000. Lakini vivinjari vitatafsiri `#FF0000` na `#F00` kama rangi ile ile kabisa.

# --instructions--

Endelea, jaribu kutumia misimbo mifupi ya hex kuupaka rangi vipengele sahihi.

<table><tbody><tr><th>Rangi</th><th>Msimbo mfupi wa hex</th></tr><tr><td>Cyan</td><td><code>#0FF</code></td></tr><tr><td>Kijani</td><td><code>#0F0</code></td></tr><tr><td>Nyeekundu</td><td><code>#F00</code></td></tr><tr><td>Fuchsia</td><td><code>#F0F</code></td></tr></tbody></table>

# --hints--

Kipengele chako cha `h1` chenye maandishi `I am red!` kinapaswa kupewa `color` nyekundu.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Msimbo mfupi wa hex kwa rangi nyekundu unapaswa kutumika badala ya msimbo wa hex `#FF0000`.

```js
assert.match(code,/\.red-text\s*?{\s*?color\s*:\s*?#F00\s*?;?\s*?}/gi);
```

Kipengele chako cha `h1` chenye maandishi `I am green!` kinapaswa kupewa `color` kijani.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Msimbo mfupi wa hex kwa rangi kijani unapaswa kutumika badala ya msimbo wa hex `#00FF00`.

```js
assert.match(__helpers.removeCssComments(code), /\.green-text\s*?{\s*?color\s*:\s*?#0F0\s*?;?\s*?}/gi);
```

Kipengele chako cha `h1` chenye maandishi `I am cyan!` kinapaswa kupewa `color` cyan.

```js
const cyanText = document.querySelector('.cyan-text');
const color = window.getComputedStyle(cyanText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 255)');
```

Msimbo mfupi wa hex kwa rangi cyan unapaswa kutumika badala ya msimbo wa hex `#00FFFF`.

```js
assert.match(__helpers.removeCssComments(code), /\.cyan-text\s*?{\s*?color\s*:\s*?#0FF\s*?;?\s*?}/gi);
```

Kipengele chako cha `h1` chenye maandishi `I am fuchsia!` kinapaswa kupewa `color` fuchsia.

```js
const fuchsiaText = document.querySelector('.fuchsia-text');
const color = window.getComputedStyle(fuchsiaText)['color']; 
assert.strictEqual(color,'rgb(255, 0, 255)');
```

Msimbo mfupi wa hex kwa rangi fuchsia unapaswa kutumika badala ya msimbo wa hex `#FF00FF`.

```js
assert.match(__helpers.removeCssComments(code), /\.fuchsia-text\s*?{\s*?color\s*:\s*?#F0F\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: #000000;
  }
  .fuchsia-text {
    color: #000000;
  }
  .cyan-text {
    color: #000000;
  }
  .green-text {
    color: #000000;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #F00;
  }
  .fuchsia-text {
    color: #F0F;
  }
  .cyan-text {
    color: #0FF;
  }
  .green-text {
    color: #0F0;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="fuchsia-text">I am fuchsia!</h1>

<h1 class="cyan-text">I am cyan!</h1>

<h1 class="green-text">I am green!</h1>
```
