---
id: bad87fee1348bd9afdf08726
title: Tumia Noti ya Kuanzia Saa Kuelezea Ukingo wa Kipengee
challengeType: 0
videoUrl: 'https://scrimba.com/c/cnpybAd'
forumTopicId: 18345
dashedName: use-clockwise-notation-to-specify-the-margin-of-an-element
---

# --description--

Tujaribu tena, lakini kwa `margin` wakati huu.

Badala ya kuelezesha vigezo vya `margin-top`, `margin-right`, `margin-bottom`, na `margin-left` vya kipengee kimoja kimoja, unaweza kuviweka vyote katika mstari mmoja, kama hivi:

```css
margin: 10px 20px 10px 20px;
```

Thamani hizi nne zinafanya kazi kama saa: juu, kulia, chini, kushoto, na zitatoa matokeo sawa kabisa na kutumia maagizo ya ukingo wa upande maalum.

# --instructions--

Tumia noti ya kuanzia saa kumpa kipengee chenye darasa la `blue-box` ukingo wa `40px` upande wake wa juu na kushoto, lakini tu `20px` upande wake wa chini na kulia.

# --hints--

Darasa lako la `blue-box` linapaswa kumpa sehemu ya juu ya vipengee `40px` ya `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

Darasa lako la `blue-box` linapaswa kumpa sehemu ya kulia ya vipengee `20px` ya `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

Darasa lako la `blue-box` linapaswa kumpa sehemu ya chini ya vipengee `20px` ya `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

Darasa lako la `blue-box` linapaswa kumpa sehemu ya kushoto ya vipengee `40px` ya `margin`.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft, '40px');
```

Unapaswa kutumia noti ya kuanzia saa kuweka ukingo wa darasa la `blue-box`.

```js
const cssCode = __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(cssCode, /\.blue-box\s*{[\s\S]*margin[\s]*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
