---
id: bad87fee1248bd9aedf08824
title: Ongeza Vikingo Tofauti Kwenye Kila Upande wa Kipengee
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4RWh4'
forumTopicId: 16633
dashedName: add-different-margins-to-each-side-of-an-element
---

# --description--

Wakati mwingine utataka kubinafsisha kipengee ili kiwe na `margin` tofauti kwenye kila upande wake.

CSS inakuwezesha kudhibiti `margin` ya pande zote nne za kipengee kwa kutumia vigezo vya `margin-top`, `margin-right`, `margin-bottom`, na `margin-left`.

# --instructions--

Toa `margin` ya `40px` kwenye sehemu ya juu na kushoto ya kisanduku cha buluu, lakini tu `20px` kwenye sehemu ya chini na kulia.

# --hints--

Darasa lako la `blue-box` linapaswa kutoa `40px` ya `margin` kwenye sehemu ya juu ya vipengee.

```js
const blueBox = document.querySelector('.blue-box');
const marginTop = window.getComputedStyle(blueBox)['margin-top'];
assert.strictEqual(marginTop, '40px');
```

Darasa lako la `blue-box` linapaswa kutoa `20px` ya `margin` kwenye sehemu ya kulia ya vipengee.

```js
const blueBox = document.querySelector('.blue-box');
const marginRight = window.getComputedStyle(blueBox)['margin-right'];
assert.strictEqual(marginRight, '20px');
```

Darasa lako la `blue-box` linapaswa kutoa `20px` ya `margin` kwenye sehemu ya chini ya vipengee.

```js
const blueBox = document.querySelector('.blue-box');
const marginBottom = window.getComputedStyle(blueBox)['margin-bottom'];
assert.strictEqual(marginBottom, '20px');
```

Darasa lako la `blue-box` linapaswa kutoa `40px` ya `margin` kwenye sehemu ya kushoto ya vipengee.

```js
const blueBox = document.querySelector('.blue-box');
const marginLeft = window.getComputedStyle(blueBox)['margin-left'];
assert.strictEqual(marginLeft,'40px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    margin-top: 40px;
    margin-right: 20px;
    margin-bottom: 20px;
    margin-left: 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
