---
id: bad87fee1348bd9aedf08803
title: Badilisha rangi ya maandishi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkVmSm'
forumTopicId: 16775
dashedName: change-the-color-of-text
---

# --description--

Sasa tubadilishe rangi ya baadhi ya maandishi yetu.

Tunaweza kufanya hivi kwa kubadilisha `style` ya kipengele chako cha `h2`.

Sifa inayohusika na rangi ya maandishi ya kipengele ni sifa ya mtindo ya `color`.

Hivi ndivyo unavyoweza kuweka rangi ya maandishi ya kipengele chako cha `h2` kuwa buluu:

```html
<h2 style="color: blue;">CatPhotoApp</h2>
```

Kumbuka kuwa ni desturi nzuri kumaliza tamko za `style` ndani ya mstari kwa `;` .

# --instructions--

Badilisha mtindo wa kipengele chako cha `h2` ili rangi ya maandishi yake iwe nyekundu.

# --hints--

Kipengele chako cha `h2` kinapaswa kuwa na tamko la `style`.

```js
assert.exists(document.querySelector('h2').getAttribute('style'));
```

Kipengele chako cha `h2` kinapaswa kuwa na rangi iliyowekwa kuwa `red`.

```js
assert.strictEqual(document.querySelector('h2').style.color, 'red');
```

Tamko lako la `style` linapaswa kumalizika na `;`.

```js
assert.isTrue(document.querySelector('h2').getAttribute('style').endsWith(';'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2 style="color: red;">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
