---
id: bad87fee1348bd9aedf08721
title: Tumia msimbo wa hex kuchanganya rangi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cK89PhP'
forumTopicId: 18359
dashedName: use-hex-code-to-mix-colors
---

# --description--

Ili ukague, misimbo ya hex hutumia tarakimu 6 za hexadecimal kuwakilisha rangi, mbili kwa kila sehemu ya nyekundu (R), kijani (G), na buluu (B).

Kutoka kwa rangi hizi tatu safi (nyekundu, kijani, na buluu), tunaweza kubadilisha kiasi cha kila moja kuunda rangi zaidi ya milioni 16!

Kwa mfano, rangi ya chungwa ni nyekundu safi, iliyochanganywa na kijani kidogo, na haina buluu. Katika msimbo wa hex, hii inatafsiriwa kuwa `#FFA500`.

Tarakimu `0` ni nambari ya chini kabisa katika msimbo wa hex, na inawakilisha ukosefu kamili wa rangi.

Tarakimu `F` ni nambari ya juu kabisa katika msimbo wa hex, na inawakilisha mwangaza wa juu kabisa unaowezekana.

# --instructions--

Badilisha maneno ya rangi katika sehemu yetu ya `style` na misimbo yao sahihi ya hex.

<table><tbody><tr><th>Rangi</th><th>Msimbo wa hex</th></tr><tr><td>Dodger Blue</td><td><code>#1E90FF</code></td></tr><tr><td>Kijani</td><td><code>#00FF00</code></td></tr><tr><td>Chungwa</td><td><code>#FFA500</code></td></tr><tr><td>Nyekundu</td><td><code>#FF0000</code></td></tr></tbody></table>

# --hints--

Sehemu yako ya `h1` yenye maandishi `I am red!` inapaswa kupewa `color` nyekundu.

```js
const redText = document.querySelector('.red-text');
const color = window.getComputedStyle(redText)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Msimbo wa `hex code` kwa rangi nyekundu unapaswa kutumika badala ya neno `red`.

```js
assert.match(code, /\.red-text\s*?{\s*?color\s*:\s*?(#FF0000|#F00)\s*?;?\s*?}/gi);
```

Sehemu yako ya `h1` yenye maandishi `I am green!` inapaswa kupewa `color` kijani.

```js
const greenText = document.querySelector('.green-text');
const color = window.getComputedStyle(greenText)['color']; 
assert.strictEqual(color, 'rgb(0, 255, 0)');
```

Msimbo wa `hex code` kwa rangi kijani unapaswa kutumika badala ya neno `green`.

```js
assert.match(code, /\.green-text\s*?{\s*?color\s*:\s*?(#00FF00|#0F0)\s*?;?\s*?}/gi);
```

Sehemu yako ya `h1` yenye maandishi `I am dodger blue!` inapaswa kupewa `color` dodger blue.

```js
const blueText = document.querySelector('.dodger-blue-text');
const color = window.getComputedStyle(blueText)['color']; 
assert.strictEqual(color, 'rgb(30, 144, 255)');
```

Msimbo wa `hex code` kwa rangi dodger blue unapaswa kutumika badala ya neno `dodgerblue`.

```js
assert.match(code, /\.dodger-blue-text\s*?{\s*?color\s*:\s*?#1E90FF\s*?;?\s*?}/gi);
```

Sehemu yako ya `h1` yenye maandishi `I am orange!` inapaswa kupewa `color` chungwa.

```js
const orangeText = document.querySelector('.orange-text');
const color = window.getComputedStyle(orangeText)['color']; 

assert.strictEqual(color, 'rgb(255, 165, 0)');
```

Msimbo wa `hex code` kwa rangi chungwa unapaswa kutumika badala ya neno `orange`.

```js
assert.match(code, /\.orange-text\s*?{\s*?color\s*:\s*?#FFA500\s*?;?\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: black;
  }
  .green-text {
    color: black;
  }
  .dodger-blue-text {
    color: black;
  }
  .orange-text {
    color: black;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```

# --solutions--

```html
<style>
  .red-text {
    color: #FF0000;
  }
  .green-text {
    color: #00FF00;
  }
  .dodger-blue-text {
    color: #1E90FF;
  }
  .orange-text {
    color: #FFA500;
  }
</style>

<h1 class="red-text">I am red!</h1>

<h1 class="green-text">I am green!</h1>

<h1 class="dodger-blue-text">I am dodger blue!</h1>

<h1 class="orange-text">I am orange!</h1>
```
