---
id: bad87fee1348bd8aedf06756
title: Ghairi tamko za darasa la vitu kwa kupangilia sifa za kipengele cha id
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkpDhB'
forumTopicId: 18251
dashedName: override-class-declarations-by-styling-id-attributes
---

# --description--

Tumeonyesha tu kwamba vivinjari husoma CSS kutoka juu hadi chini kwa mpangilio wa tamko zao. Hii inamaanisha kwamba, wakati kuna mgongano, kivinjari kitaitumia tamko la CSS lililokuja la mwisho. Angalia kwamba hata kama tungeweka `blue-text` kabla ya `pink-text` katika madarasa ya kipengele cha `h1`, bado kingetazama mpangilio wa tamko na si mpangilio wa matumizi yao!

Lakini bado hatujamaliza. Kuna njia nyingine za kughairi CSS. Je, unakumbuka sifa za kipengele cha id?

Hebu tughairi madarasa yako ya `pink-text` na `blue-text`, na tufanye kipengele chako cha `h1` kuwa rangi ya chungwa, kwa kumpa kipengele cha `h1` id kisha tupangilie id hiyo.

# --instructions--

Mpe kipengele chako cha `h1` sifa ya kipengele cha `id` yenye thamani ya `orange-text`. Kumbuka, mitindo ya id inaonekana hivi:

```html
<h1 id="orange-text">
```

Acha madarasa ya `blue-text` na `pink-text` kwenye kipengele chako cha `h1`.

Tengeneza tamko la CSS kwa id yako ya `orange-text` katika kipengele chako cha `style`. Hapa kuna mfano wa jinsi inavyoonekana:

```css
#brown-text {
  color: brown;
}
```

**Kumbuka:** Haijalishi kama unatangaza CSS hii juu au chini ya darasa la `pink-text`, kwa sababu sifa ya kipengele cha `id` itakuwa na kipaumbele kila wakati.

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na darasa la `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Kipengele chako cha `h1` kinapaswa kuwa na darasa la `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Kipengele chako cha `h1` kinapaswa kuwa na id ya `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'),'orange-text');
```

Kipengele cha `h1` kinapaswa kuwa kimoja tu.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Id yako ya `orange-text` inapaswa kuwa na tamko la CSS.

```js
assert.match(__helpers.removeCssComments(code), /#orange-text\s*{/gi);
```

`h1` haipaswi kuwa na sifa za kipengele cha `style`.

```js
assert.notMatch(__helpers.removeHtmlComments(code), /<h1.*style.*>/gi);
```

Kipengele chako cha `h1` kinapaswa kuwa na rangi ya chungwa.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 165, 0)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
  #orange-text {
    color: orange;
  }  
</style>
<h1 id="orange-text"  class="pink-text blue-text">Hello World!</h1>
```
