---
id: bad87fee1348bd9aedf08826
title: Tumia Noti ya Kuanzia Saa Kuelezea Nafasi ya Ndani ya Kipengee
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB7mBS9'
forumTopicId: 18346
dashedName: use-clockwise-notation-to-specify-the-padding-of-an-element
---

# --description--

Badala ya kuelezea vigezo vya `padding-top`, `padding-right`, `padding-bottom`, na `padding-left` vya kipengee mmoja mmoja, unaweza kuviweka vyote katika mstari mmoja, kama hivi:

```css
padding: 10px 20px 10px 20px;
```

Thamani hizi nne hufanya kazi kama saa: juu, kulia, chini, kushoto, na zitatoa matokeo sawa kabisa na kutumia maagizo ya nafasi ya ndani ya upande maalum.

# --instructions--

Tumia noti ya kuanzia saa kutoa darasa la `.blue-box` nafasi ya ndani ya `padding` ya `40px` upande wake wa juu na kushoto, lakini tu `20px` upande wake wa chini na kulia.

# --hints--

Darasa lako la `blue-box` linapaswa kutoa kipengee upande wa juu nafasi ya ndani ya `40px` ya `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '40px');
```

Darasa lako la `blue-box` linapaswa kutoa kipengee upande wa kulia nafasi ya ndani ya `20px` ya `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingRight = window.getComputedStyle(blueBox)['padding-right'];
assert.strictEqual(paddingRight, '20px');
```

Darasa lako la `blue-box` linapaswa kutoa kipengee upande wa chini nafasi ya ndani ya `20px` ya `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingBottom = window.getComputedStyle(blueBox)['padding-bottom'];
assert.strictEqual(paddingBottom, '20px');
```

Darasa lako la `blue-box` linapaswa kutoa kipengee upande wa kushoto nafasi ya ndani ya `40px` ya `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingLeft = window.getComputedStyle(blueBox)['padding-left'];
assert.strictEqual(paddingLeft, '40px');
```

Unapaswa kutumia noti ya kuanzia saa kuweka nafasi ya ndani ya darasa la `blue-box`.

```js
const css =  __helpers.removeCssComments(document.querySelector('style:not(.fcc-hide-header)').textContent);
assert.match(css, /\.blue-box\s*{[\s\S]*padding\s*:\s*\d+px\s+\d+px\s+\d+px\s+\d+px(;\s*[^}]+\s*}|;?\s*})/);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px 40px 20px 40px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 40px 20px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
