---
id: bad87fee1348bd9aedf06756
title: Ghairi tamko za darasa la vitu kwa mitindo ndani ya mstari
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJDRha'
forumTopicId: 18252
dashedName: override-class-declarations-with-inline-styles
---

# --description--

Hivyo tumethibitisha kuwa tamko za id ghairi tamko za darasa la vitu, bila kujali wapi zimetangazwa katika CSS ya kipengee chako cha `style`.

Kuna njia nyingine za ghairi CSS. Je, unakumbuka mitindo ndani ya mstari?

# --instructions--

Tumia mtindo ndani ya mstari kujaribu kufanya kipengee chetu cha `h1` kuwa nyeupe. Kumbuka, mitindo ndani ya mstari inaonekana hivi:

```html
<h1 style="color: green;">
```

Acha madarasa ya `blue-text` na `pink-text` kwenye kipengee chako cha `h1`.

# --hints--

Kipengee chako cha `h1` kinapaswa kuwa na darasa la `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

Kipengee chako cha `h1` kinapaswa kuwa na darasa la `blue-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('blue-text'));
```

Kipengee chako cha `h1` kinapaswa kuwa na id ya `orange-text`.

```js
assert.strictEqual(document.querySelector('h1').getAttribute('id'), 'orange-text');
```

Kipengee chako cha `h1` kinapaswa kuwa na mtindo ndani ya mstari.

```js
assert.exists(document.querySelector('h1[style]'));
```

Kipengee chako cha `h1` kinapaswa kuwa nyeupe.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text">Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  #orange-text {
    color: orange;
  }
  .pink-text {
    color: pink;
  }
  .blue-text {
    color: blue;
  }
</style>
<h1 id="orange-text" class="pink-text blue-text" style="color: white">Hello World!</h1>
```
