---
id: bad87fee1348bd9aedf08736
title: Panga mtindo wa kipengele cha mwili wa HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB77PHW'
forumTopicId: 18313
dashedName: style-the-html-body-element
---

# --description--

Sasa tuanze upya na tuzungumzie urithi wa CSS.

Kila ukurasa wa HTML una kipengele cha `body`.

# --instructions--

Tunaweza kuthibitisha kuwa kipengele cha `body` kipo hapa kwa kumsabisha `background-color` ya rangi nyeusi.

Tunaweza kufanya hivi kwa kuongeza yafuatayo kwenye kipengele chetu cha `style`:

```css
body {
  background-color: black;
}
```

# --hints--

Kipengele chako cha `body` kinapaswa kuwa na `background-color` ya rangi nyeusi.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color'];

assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Sheria yako ya CSS inapaswa kuwa na muundo sahihi na mabano ya kufungua na kufunga.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

Sheria yako ya CSS inapaswa kumalizika na alama ya semikolon.

```js
assert.match(code, /<style>\s*body\s*\{\s*background.*\s*:\s*.*;\s*\}\s*<\/style>/i);
```

# --seed--

## --seed-contents--

```html
<style>

</style>
```

# --solutions--

```html
<style>
body {
  background-color: black;
}
</style>
```
