---
id: bad87fee1348bd9aecf08806
title: Tumia darasa la CSS kupamba kipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MvDtV'
forumTopicId: 18337
dashedName: use-a-css-class-to-style-an-element
---

# --description--

Madarasa ni mitindo inayoweza kutumika tena ambayo yanaweza kuongezwa kwa vipengele vya HTML.

Hapa kuna mfano wa tamko la darasa la CSS:

```html
<style>
  .blue-text {
    color: blue;
  }
</style>
```

Unaweza kuona kwamba tumetengeneza darasa la CSS linaloitwa `blue-text` ndani ya lebo ya `<style>`. Unaweza kutumia darasa kwa kipengele cha HTML kama hivi: `<h2 class="blue-text">CatPhotoApp</h2>`. Kumbuka kwamba katika kipengele chako cha CSS cha `style`, majina ya madarasa huanza na nukta. Katika sifa ya darasa ya vipengele vyako vya HTML, jina la darasa halijumuishi nukta.

# --instructions--

Ndani ya kipengele chako cha `style`, badilisha kichaguzi cha `h2` kuwa `.red-text` na sasisha thamani ya rangi kutoka `blue` hadi `red`.

Toa kipengele cha `h2` chenye thamani ya `class` kwa kipengele chako cha `red-text`.

# --hints--

Kipengele chako cha `h2` kinapaswa kuwa chepesi.

```js
const h2Element = document.querySelector('h2'); 
const color = window.getComputedStyle(h2Element)['color']; 
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Kipengele chako cha `h2` kinapaswa kuwa na darasa la `red-text`.

```js
assert.isTrue(document.querySelector('h2').classList.contains('red-text'));
```

Karatasi yako ya mitindo inapaswa kutangaza darasa la `red-text` na kuweka rangi yake kuwa `red`.

```js
assert.match(__helpers.removeCssComments(code), /\.red-text\s*\{\s*color\s*:\s*red;?\s*\}/g);
```

Hupaswi kutumia tamko la mtindo ndani ya mstari kama `style="color: red"` katika kipengele chako cha `h2`.

```js
assert.notExists(document.querySelector('h2').getAttribute('style'));
```

# --seed--

## --seed-contents--

```html
<style>
  h2 {
    color: blue;
  }
</style>

<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<style>
  .red-text {
    color: red;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
