---
id: bad82fee1322bd9aedf08721
title: Elewa vitengo kamili dhidi ya vitengo vinavyohusiana
challengeType: 0
videoUrl: 'https://scrimba.com/c/cN66JSL'
forumTopicId: 301089
dashedName: understand-absolute-versus-relative-units
---

# --description--

Changamoto kadhaa za mwisho zote ziliweka ukingo au nafasi ya ndani ya kipengee kwa kutumia pixels (`px`). Pixels ni aina ya kitengo cha urefu, ambacho ndicho kinachoeleza kivinjari jinsi ya kupima au kuweka nafasi ya kipengee. Mbali na `px`, CSS ina chaguzi kadhaa tofauti za vitengo vya urefu ambavyo unaweza kutumia.

Aina kuu mbili za vitengo vya urefu ni vitengo kamili na vitengo vinavyohusiana. Vitengo kamili vinahusishwa na vitengo halisi vya urefu. Kwa mfano, `in` na `mm` vinarejelea inchi na milimita, mtawalia. Vitengo kamili vya urefu vinakaribia kipimo halisi kwenye skrini, lakini kuna tofauti fulani kulingana na azimio la skrini.

Vitengo vinavyohusiana, kama `em` au `rem`, vinahusiana na thamani nyingine ya urefu. Kwa mfano, `em` inategemea ukubwa wa fonti ya kipengee. Ukitumia kuweka sifa ya `font-size` yenyewe, ni kulingana na `font-size` ya mzazi.

**Kumbuka:** Kuna chaguzi kadhaa za vitengo vinavyohusiana vinavyohusishwa na ukubwa wa eneo la kuona. Zinajadiliwa katika sehemu ya Kanuni za Ubunifu wa Tovuti Unayojibadilisha Kulingana na Kifaa.

# --instructions--

Ongeza sifa ya `padding` kwa kipengee chenye darasa la `red-box` na uweke kuwa `1.5em`.

# --hints--

Darasa lako la `red-box` linapaswa kuwa na sifa ya `padding`.

```js
const redBox =document.querySelector('.red-box'); 
const style = window.getComputedStyle(redBox); 

assert.notEqual(style['padding-top'], '0px');
assert.notEqual(style['padding-right'], '0px');
assert.notEqual(style['padding-bottom'], '0px');
assert.notEqual(style['padding-left'], '0px');
```

Darasa lako la `red-box` linapaswa kuwapa vipengee nafasi ya ndani ya `padding` ya 1.5em.

```js
assert.match(__helpers.removeCssComments(code), /\.red-box\s*?{[\s\S]*padding\s*:\s*?1\.5em/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;

  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 20px 40px 20px 40px;
  }

  .red-box {
    background-color: red;
    margin: 20px 40px 20px 40px;
    padding: 1.5em;
  }

  .green-box {
    background-color: green;
    margin: 20px 40px 20px 40px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box green-box">padding</h5>
</div>
```
