---
id: bad87fee1348bd9aede08718
title: Tumia thamani za RGB kupaka rangi vipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Njia nyingine unayoweza wakilisha rangi katika CSS ni kwa kutumia thamani za `RGB`.

Thamani ya `RGB` kwa rangi nyeusi inaonekana hivi:

```css
rgb(0, 0, 0)
```

Thamani ya `RGB` kwa rangi nyeupe inaonekana hivi:

```css
rgb(255, 255, 255)
```

Badala ya kutumia tarakimu sita za hexadecimal kama unavyofanya na msimbo wa hex, kwa `RGB` unabainisha mwangaza wa kila rangi kwa nambari kati ya 0 na 255.

Ukifanya hesabu, tarakimu mbili kwa rangi moja ni sawa na 16 mara 16, ambayo inatupa jumla ya thamani 256. Hivyo `RGB`, ambayo huanza kuhesabu kutoka sifuri, ina idadi sawa kabisa ya thamani zinazowezekana kama msimbo wa hex.

Hapa kuna mfano wa jinsi unavyobadilisha rangi ya nyuma ya `body` kuwa rangi ya chungwa kwa kutumia msimbo wake wa RGB.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Hebu tubadilishe msimbo wa hex katika rangi ya nyuma ya kipengele chetu cha `body` na thamani ya RGB kwa rangi nyeusi: `rgb(0, 0, 0)`

# --hints--

Kipengele chako cha `body` kinapaswa kuwa na rangi ya nyuma nyeusi.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Unapaswa kutumia `rgb` kumpa kipengele chako cha `body` rangi ya nyuma nyeusi.

```js
assert.match(__helpers.removeCssComments(code), /rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
