---
id: bad87fee1348bd9aedf08726
title: Tumia msimbo wa hex kwa rangi maalum
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

Je, ulijua kuna njia nyingine za wakilisha rangi katika CSS? Moja ya njia hizi huitwa msimbo wa hexadecimal, au msimbo wa hex kwa kifupi.

Kawaida tunatumia <dfn>tarakimu za desimali</dfn>, au nambari za msingi wa 10, ambazo hutumia alama 0 hadi 9 kwa kila tarakimu. <dfn>Hexadecimal</dfn> (au <dfn>hex</dfn>) ni nambari za msingi wa 16. Hii ina maana hutumia alama kumi na sita tofauti. Kama tarakimu za desimali, alama 0-9 zinawakilisha thamani kutoka sifuri hadi tisa. Kisha A,B,C,D,E,F zinawakilisha thamani kutoka kumi hadi kumi na tano. Kwa jumla, 0 hadi F zinaweza wakilisha tarakimu katika hexadecimal, zikitupa jumla ya thamani 16 zinazowezekana. Unaweza kupata taarifa zaidi kuhusu <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">nambari za hexadecimal hapa</a>.

Katika CSS, tunaweza kutumia tarakimu 6 za hexadecimal kuwakilisha rangi, mbili kwa kila sehemu ya nyekundu (R), kijani (G), na bluu (B). Kwa mfano, `#000000` ni nyeusi na pia ni thamani ya chini kabisa inayowezekana. Unaweza kupata taarifa zaidi kuhusu <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">mfumo wa rangi wa RGB hapa</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Badilisha neno `black` katika rangi ya nyuma ya sehemu yetu ya `body` na uwakilishe kwa msimbo wake wa hex, `#000000`.

# --hints--

Sehemu yako ya `body` inapaswa kuwa na `background-color` ya rangi nyeusi.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Msimbo wa hex kwa rangi nyeusi unapaswa kutumika badala ya neno `black`.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
