---
id: bad87fee1348bd9aedf08808
title: Eleza Jinsi Fonti Zinavyopaswa Kupungua Ubora
challengeType: 0
videoUrl: 'https://scrimba.com/c/cpVKBfQ'
forumTopicId: 18304
dashedName: specify-how-fonts-should-degrade
---

# --description--

Kuna fonti kadhaa za msingi zinazopatikana katika vivinjari vyote. Familia hizi za fonti za jumla ni pamoja na `monospace`, `serif` na `sans-serif`.

Wakati fonti moja haipatikani, unaweza kusema kwa kivinjari "kupunguza ubora" hadi fonti nyingine.

Kwa mfano, kama ungetaka kipengele kitumie fonti ya `Helvetica`, lakini kupunguza ubora hadi fonti ya `sans-serif` wakati `Helvetica` haipatikani, utaeleza hivyo kama ifuatavyo:

```css
p {
  font-family: Helvetica, sans-serif;
}
```

Majina ya familia ya fonti ya jumla hayazingatii herufi kubwa au ndogo. Pia, hayahitaji nukuu kwa sababu ni maneno muhimu ya CSS.

# --instructions--

Kuanza, tumia fonti ya `monospace` kwa kipengele cha `h2`, ili sasa kiwe na fonti mbili - `Lobster` na `monospace`.

Katika changamoto iliyopita, uli-import fonti ya `Lobster` kwa kutumia lebo ya `link`. Sasa toa maoni kwenye import hiyo ya fonti ya `Lobster` (ukitumia maoni ya HTML uliyojifunza awali) kutoka Google Fonts ili isipatikane tena. Angalia jinsi kipengele chako cha `h2` kinavyopungua ubora hadi fonti ya `monospace`.

**Kumbuka:** Ikiwa una fonti ya `Lobster` imewekwa kwenye kompyuta yako, hutagundua kupungua ubora kwa sababu kivinjari chako kinaweza kupata fonti hiyo.

# --hints--

Kipengele chako cha h2 kinapaswa kutumia fonti ya `Lobster`.

```js
const h2Element = document.querySelector('h2');
const fontFamily = window.getComputedStyle(h2Element)['font-family']; 
assert.match(fontFamily, /^"?lobster/i);
```

Kipengele chako cha h2 kinapaswa kupungua ubora hadi fonti ya `monospace` wakati `Lobster` haipatikani.

```js
assert.match(__helpers.removeCssComments(code), /\s*h2\s*\{\s*font-family\s*\:\s*(\'|"|)Lobster\1\s*,\s*monospace\s*;?\s*\}/gi);
```

Unapaswa kutoa maoni kwenye wito wako kwa Google kwa fonti ya `Lobster` kwa kuweka `<!--` mbele yake.

```js
assert.match(code, /<!--[^fc]/gi);
```

Unapaswa kufunga maoni yako kwa kuongeza `-->`.

```js
assert.match(code, /[^fc]-->/gi);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<!--<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">-->
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, monospace;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>
  
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
