---
id: bad88fee1348bd9aedf08825
title: Rekebisha Nafasi ya Ndani ya Kipengee
challengeType: 0
videoUrl: 'https://scrimba.com/c/cED8ZC2'
forumTopicId: 301083
dashedName: adjust-the-padding-of-an-element
---

# --description--

Sasa tuache Kidokezo cha Picha za Paka kwa muda kidogo na tujifunze zaidi kuhusu mtindo wa HTML.

Labda tayari umeona hili, lakini vipengee vyote vya HTML kwa msingi ni mstatili mdogo.

Sifa tatu muhimu zinadhibiti nafasi inayozunguka kila kipengee cha HTML: `padding`, `border`, na `margin`.

`padding` ya kipengee hudhibiti kiasi cha nafasi kati ya maudhui ya kipengee na `border` yake.

Hapa, tunaona kuwa kisanduku cha buluu na kisanduku cha nyekundu vimepangwa ndani ya kisanduku cha manjano. Kumbuka kuwa kisanduku cha nyekundu kina `padding` zaidi kuliko kisanduku cha buluu.

Unapoongeza `padding` ya kisanduku cha buluu, itapanua umbali (`padding`) kati ya maandishi na mpaka unaouzunguka.

# --instructions--

Badilisha `padding` ya kisanduku chako cha buluu ili ifanane na ile ya kisanduku chako cha nyekundu.

# --hints--

Darasa lako la `blue-box` linapaswa kutoa vipengee `20px` vya `padding`.

```js
const blueBox = document.querySelector('.blue-box');
const paddingTop = window.getComputedStyle(blueBox)['padding-top'];
assert.strictEqual(paddingTop, '20px');
```

# --seed--

## --seed-contents--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 10px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```

# --solutions--

```html
<style>
  .injected-text {
    margin-bottom: -25px;
    text-align: center;
  }

  .box {
    border-style: solid;
    border-color: black;
    border-width: 5px;
    text-align: center;
  }

  .yellow-box {
    background-color: yellow;
    padding: 10px;
  }

  .red-box {
    background-color: crimson;
    color: #fff;
    padding: 20px;
  }

  .blue-box {
    background-color: blue;
    color: #fff;
    padding: 20px;
  }
</style>
<h5 class="injected-text">margin</h5>

<div class="box yellow-box">
  <h5 class="box red-box">padding</h5>
  <h5 class="box blue-box">padding</h5>
</div>
```
