---
id: bad87fee1348bd9aedf08756
title: Panga mtindo mmoja kuwa muhimu zaidi kuliko mwingine
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ8wnHv'
forumTopicId: 18258
dashedName: prioritize-one-style-over-another
---

# --description--

Wakati mwingine vipengele vyako vya HTML vitapokea mitindo mingi inayopingana.

Kwa mfano, kipengele chako cha `h1` hakiwezi kuwa kijani na waridi kwa wakati mmoja.

Tuchunguze kinachotokea tunapounda darasa la vitu linalofanya maandishi kuwa waridi, kisha kulitumia kwenye kipengele. Je, darasa letu lita *ghairi* sifa ya `body` ya CSS ya kipengele cha `color: green;`?

# --instructions--

Unda darasa la vitu la CSS linaloitwa `pink-text` linalompa kipengele rangi ya waridi.

Mpa kipengele chako cha `h1` darasa la vitu la `pink-text`.

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na darasa la vitu `pink-text`.

```js
assert.isTrue(document.querySelector('h1').classList.contains('pink-text'));
```

`<style>` yako inapaswa kuwa na darasa la vitu la CSS la `pink-text` linalobadilisha `color`.

```js
assert.match(__helpers.removeCssComments(code), /\.pink-text\s*\{\s*color\s*:\s*.+\s*;?\s*\}/g);
```

Kipengele chako cha `h1` kinapaswa kuwa waridi.

```js
const h1Element = document.querySelector('h1');
const color = window.getComputedStyle(h1Element)['color']; 
assert.strictEqual(color, 'rgb(255, 192, 203)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
</style>
<h1>Hello World!</h1>
```

# --solutions--

```html
<style>
  body {
    background-color: black;
    font-family: monospace;
    color: green;
  }
  .pink-text {
    color: pink;
  }
</style>
<h1 class="pink-text">Hello World!</h1>
```
