---
id: bad87fee1348bd9aedf08807
title: Import fonti ya Google
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9MRsJ'
forumTopicId: 18200
dashedName: import-a-google-font
---

# --description--

Mbali na kubainisha fonti za kawaida zinazopatikana katika mifumo mingi ya uendeshaji, tunaweza pia kubainisha fonti za mtandao zisizo za kawaida, za kipekee kwa matumizi katika tovuti yetu. Kuna vyanzo vingi vya fonti za mtandao kwenye Intaneti. Kwa mfano huu tutazingatia maktaba ya fonti za Google.

Google Fonts ni maktaba ya bure ya fonti za mtandao ambayo unaweza kutumia katika CSS yako kwa kurejelea URL ya fonti hiyo.

Hivyo basi, tuendelee na ku-import na kutumia fonti ya Google (kumbuka kwamba kama Google imezuiwa katika nchi yako, utahitaji kuruka changamoto hii).

Ili ku-import fonti ya Google, unaweza kunakili URL ya fonti kutoka maktaba ya Google Fonts kisha kuibandika katika HTML yako. Kwa changamoto hii, tutaleta fonti ya `Lobster`. Ili kufanya hivyo, nakili kipande cha msimbo kilicho hapa chini na ubandike juu ya mhariri wako wa msimbo (kabla ya kipengele cha `style` kufunguliwa):

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
```

Sasa unaweza kutumia fonti ya `Lobster` katika CSS yako kwa kutumia `Lobster` kama FAMILY_NAME kama katika mfano ufuatao:

```css
font-family: FAMILY_NAME, GENERIC_NAME;
```

GENERIC_NAME ni hiari, na ni fonti ya njia mbadala ikiwa fonti nyingine iliyobainishwa haitapatikana. Hii itajadiliwa katika changamoto inayofuata.

Majina ya familia ni nyeti kwa herufi kubwa na ndogo na yanahitaji kuwekwa ndani ya nukuu ikiwa kuna nafasi katika jina. Kwa mfano, unahitaji nukuu kutumia fonti ya `"Open Sans"`, lakini siyo kutumia fonti ya `Lobster`.

# --instructions--

Import fonti ya `Lobster` kwenye ukurasa wako wa mtandao. Kisha, tumia kichaguzi cha kipengele kuweka `Lobster` kama fonti ya `font-family` kwa kipengele chako cha `h2`.

# --hints--

Unapaswa ku-import fonti ya `Lobster`.

```js
assert.exists(document.querySelector('link[href*="googleapis" i]'));
```

Kipengele chako cha `h2` kinapaswa kutumia fonti ya `Lobster`.

```js
const h2 = document.querySelector('h2'); 
const fontFamily = window.getComputedStyle(h2)['font-family']; 
assert.match(fontFamily, /lobster/i);
```

Unapaswa kutumia tu kichaguzi cha kipengele cha `h2` kubadilisha fonti.

```js
assert.match(__helpers.removeHtmlComments(code), /\s*[^\.]h2\s*\{\s*font-family\s*:\s*('|"|)Lobster\1\s*(,\s*('|"|)[a-z -]+\3\s*)?(;\s*\}|\})/gi);
```

Kipengele chako cha `p` kinapaswa bado kutumia fonti ya `monospace`.

```js
const paragraphElement = document.querySelector('p');
const fontFamily = window.getComputedStyle(paragraphElement)['font-family']; 
assert.match(fontFamily, /monospace/i);
```

# --seed--

## --seed-contents--

```html
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  p {
    font-size: 16px;
    font-family: monospace;
  }

  h2 {
    font-family: Lobster;
  }
</style>

<h2 class="red-text">CatPhotoApp</h2>
<main>
  <p class="red-text">Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <div>
    <p>Things cats love:</p>
    <ul>
      <li>catnip</li>
      <li>laser pointers</li>
      <li>lasagna</li>
    </ul>
    <p>Top 3 things cats hate:</p>
    <ol>
      <li>flea treatment</li>
      <li>thunder</li>
      <li>other cats</li>
    </ol>
  </div>

  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor" checked> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label><input type="checkbox" name="personality" checked> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
