---
id: 699b68f94a0d154c2a7a5ce0
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Agora que o objeto `profile` está completo, é hora de criar um novo objeto que representa os papéis dos usuários.

Crie uma nova variável chamada `userRoles` e atribua a ela um objeto. Seu objeto deve ter as seguintes propriedades e valores:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Você deve ter uma variável chamada `userRoles`.

```js
assert.isDefined(userRoles);
```

Sua variável `userRoles` deve ser um objeto.

```js
assert.isObject(userRoles);
```

Seu objeto `userRoles` deve ter uma propriedade `admin`.

```js
assert.property(userRoles, "admin");
```

Sua propriedade `admin` deve ter um valor `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Seu objeto `userRoles` deve ter uma propriedade `editor`.

```js
assert.property(userRoles, "editor");
```

Sua propriedade `editor` deve ter um valor `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Seu objeto `userRoles` deve ter uma propriedade `viewer`.

```js
assert.property(userRoles, "viewer");
```

Sua propriedade `viewer` deve ter um valor `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
