---
id: 699b4e284291adcfcb90df47
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai praticar trabalhar com tipos primitivos assim como tipos de objeto e array construindo um perfil de usuário com tipagem segura.

Comece criando uma variável chamada `profile` e atribua a ela um objeto. Seu objeto deve ter estas três propriedades e valores:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Você deve ter uma variável chamada `profile`.

```js
assert.exists(profile);
```

Sua variável `profile` deve ser um objeto.

```js
assert.isObject(profile);
```

Você deve ter uma propriedade `username` no seu objeto `profile`.

```js
assert.property(profile, "username");
```

Sua propriedade `username` deve ter um valor `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Seu objeto `profile` deve ter uma propriedade `age`.

```js
assert.property(profile, "age");
```

Sua propriedade `age` deve ter um valor `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Seu objeto `profile` deve ter uma propriedade `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Sua propriedade `isLoggedIn` deve ter um valor `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
