---
id: 699b6dd1a2961bd38491ef2e
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

O objeto `profile` tinha chaves e tipos rígidos. Se você tentasse adicionar uma propriedade que não estivesse definida ou atribuir um tipo incompatível, o TypeScript geraria um erro.

Seria bom se o objeto `userRoles` não fosse tão rígido e tivesse chaves mais flexíveis.

Como você lembra das lições anteriores, você pode alcançar esse resultado usando o tipo `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

Neste exemplo `studentGrades`, o TypeScript não precisa conhecer todas as matérias antecipadamente. Isso permite que você adicione mais matérias como `Art` ou `Music` depois. O TypeScript só se importa que os valores para cada chave sejam números.

Atualize seu objeto `userRoles` adicionando um tipo `Record` inline. Defina-o como `<string, string>` para que cada chave possa ser qualquer string e cada valor deve ser uma string.

# --hints--

Seu objeto `userRoles` deve ter um tipo `Record` inline definido como `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
