---
id: 699b70b10ae344c802d106b0
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Agora é hora de adicionar mais algumas propriedades ao objeto `userRoles`.

Adicione uma propriedade `moderator` com o valor `"medium-access"`.

Depois, logo abaixo, adicione uma propriedade `guest` com o valor `3`.

# --hints--

Seu objeto `userRoles` deve ter uma propriedade `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Sua propriedade `moderator` deve ter o valor `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Seu objeto `userRoles` deve ter uma propriedade `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

Sua propriedade `guest` deve ter o valor numérico `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
