---
id: 699b5b7cb68b6d35afca61b7
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Embora seja possível continuar adicionando propriedades assim ao objeto `profile`, seria bom restringir os tipos de propriedades que devem ser permitidos. Seria melhor definir a forma do objeto `profile` antecipadamente.

Nas lições anteriores, você aprendeu a trabalhar com tipos de objeto assim:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Neste exemplo, `composer` é um objeto com um tipo de objeto inline. Esse objeto tem três propriedades obrigatórias: `name`, `instrument` e `isActive`. As duas primeiras propriedades devem ser do tipo `string` e `isActive` deve ser do tipo `boolean`.

Atualize seu objeto `profile` existente para incluir um tipo de objeto inline. `username` deve ser tipado como `string`, `age` deve ser tipado como `number` e `isLoggedIn` deve ser tipado como `boolean`. 

# --hints--

Seu objeto `profile` deve ter um tipo de objeto inline com uma propriedade `username` definida como `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Seu objeto `profile` deve ter um tipo de objeto inline com uma propriedade `age` definida como `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Seu objeto `profile` deve ter um tipo de objeto inline com uma propriedade `isLoggedIn` definida como `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
