---
id: 699b71d98c820ac7ab9eea2a
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Se você abrir o console novamente, verá uma nova mensagem de erro:

```md
Type 'number' is not assignable to type 'string'.
```

Isso está acontecendo porque você está tentando adicionar uma propriedade `guest` com um valor numérico.

Para remover o erro do TypeScript, altere o valor da propriedade `guest` para `"read-only"`.

Com essa última alteração, sua oficina de perfil de usuário está completa!

# --hints--

Sua propriedade `guest` para o objeto `userRoles` deve ter o valor `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
