---
id: 699b61ca8c9cb9bc1968364e
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Se você abrir o console, deve ver a seguinte mensagem de erro:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

No JavaScript puro, você podia adicionar quaisquer propriedades e valores que quisesse ao objeto `profile`. Mas no TypeScript, o compilador verifica se o objeto tem apenas as propriedades que você definiu (`username`, `age` e `isLoggedIn` neste caso). Como `mood` não foi incluído no tipo, o TypeScript o trata como uma propriedade inválida e gera um erro.

Remova o `mood: null` do seu código.

# --hints--

Você não deve mais ter `mood: null` no seu código.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
