---
id: 699b64096301fd75e3efeb1b
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Se você abrir o console novamente, verá uma mensagem de erro diferente:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Como você lembra das lições anteriores, é possível marcar propriedades como opcionais usando o símbolo `?` ao lado do nome da propriedade assim:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

Neste exemplo, `numberOfAlbums` está marcado como opcional, então o TypeScript não gerará um erro neste caso.

Atualize sua propriedade `bio` para que ela seja marcada como opcional.

# --hints--

Sua propriedade `bio` deve estar marcada como opcional.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
