---
id: 699b5370a43e299b4673bcaa
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Seria útil ver as propriedades do objeto `profile` no console enquanto você o constrói.

Adicione um `console.log` que registre `profile` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar `profile` no console.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
