---
id: 699b65caf17a8e8d340af27d
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Agora Ă© hora de adicionar a Ășltima propriedade ao objeto `profile`.

Nas liĂ§Ă”es anteriores, vocĂȘ aprendeu a trabalhar com tipos array. VocĂȘ aprendeu sobre diferentes tipos de array como `string[]` ou `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

Neste exemplo, o array `scores` aceita apenas nĂșmeros. Se vocĂȘ tentar adicionar outro tipo como um `string` ou `null`, o TypeScript gerarĂĄ um erro.

Atualize seu objeto `profile` para adicionar uma propriedade `programmingLanguages`, tipada como um `string[]`. Dentro do literal do objeto, defina essa propriedade como `["JavaScript", "Python", "C++"]`.

# --hints--

VocĂȘ deve ter uma propriedade `programmingLanguages` dentro do seu objeto `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Sua propriedade `programmingLanguages` deve ter o valor `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Sua propriedade `programmingLanguages` deve ser tipada como um `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
