---
id: 699b62c99ec8c63470313365
title: passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora é hora de adicionar outra propriedade ao seu objeto `profile`.

Comece atualizando o tipo do objeto inline para incluir uma propriedade `bio` com o tipo definido como `string`.

# --hints--

Você deve atualizar o tipo do objeto inline para incluir uma propriedade `bio` definida como `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
