---
id: 67051431a73f1ca25d9a6f25
title: بناء لوحة أفكار مشروع
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب تعريف ثابت كائن باسم `projectStatus` يحتوي على المفاتيح الثلاثة: `PENDING`، `SUCCESS`، و`FAILURE`. يجب تعيين كل حالة إلى كائن يحتوي على مفتاح `description` بالقيم التالية: `Pending Execution`، `Executed Successfully`، و`Execution Failed`، على التوالي.
2. يجب تعريف صنف باسم `ProjectIdea` يحتوي على `constructor` يأخذ `title` و`description` كمعلمات. ابدأ خاصيتي `title` و`description` بالمعلمات المقدمة. يجب أن يحتوي الصنف أيضًا على خاصية باسم `status` تُعيّن افتراضيًا إلى القيمة `projectStatus.PENDING`.
3. يجب تعريف طريقة باسم `updateProjectStatus` داخل صنف `ProjectIdea`. يجب أن تقبل هذه الطريقة معلمة `newStatus` وتُحدث خاصية `status` بالقيمة المعطاة.
4. يجب تعريف صنف `ProjectIdeaBoard` يحتوي على `constructor` يقبل `title` ويُهيئ مصفوفة فارغة باسم `ideas` لتخزين مثيلات صنف `ProjectIdea`.
5. يجب تعريف طريقة باسم `pin` داخل صنف `ProjectIdeaBoard` تقبل مثيلًا من صنف `ProjectIdea` وتُضيف هذا المثيل إلى مصفوفة `ideas`.
6. يجب تعريف طريقة باسم `unpin` داخل صنف `ProjectIdeaBoard`. يجب أن تقبل هذه الطريقة مثيلًا من صنف `ProjectIdea` وتُزيله من مصفوفة `ideas`.
7. يجب تعريف طريقة باسم `count` تُرجع عدد أفكار المشاريع في مصفوفة `ProjectIdeaBoard` المعطاة.
8. يجب تعريف طريقة باسم `formatToString` تُرجع أسماء أفكار المشاريع، وصفها، وحالتها بالتنسيق التالي:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

يجب تعريف ثابت كائن باسم `projectStatus` يحتوي على حالات المشاريع الصحيحة ووصفها.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

يجب أن يكون لديك صنف `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

يجب أن يبدأ صنف `ProjectIdea` خاصيتي `title` و`description` بناءً على المعلمات الممررة. كما يجب تعيين الحالة إلى `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

يجب أن يحتوي صنف `ProjectIdea` على طريقة `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

يجب أن يكون لديك صنف `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

يجب أن يبدأ صنف `ProjectIdeaBoard` خاصية `title` بناءً على المعلمة الممررة ويُهيئ مصفوفة فارغة باسم `ideas` لتخزين مثيلات صنف `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

يجب أن يحتوي صنف `ProjectIdeaBoard` على طريقة `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

يجب أن يحتوي صنف `ProjectIdeaBoard` على طريقة `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

يجب أن يحتوي صنف `ProjectIdeaBoard` على طريقة `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

يجب أن يحتوي صنف `ProjectIdeaBoard` على طريقة `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` يجب أن تُرجع `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

استدعاء `updateProjectStatus(projectStatus.SUCCESS)` على `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` يجب أن يُحدث الحالة إلى `{ description: 'Executed Successfully' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

استدعاء `updateProjectStatus(projectStatus.FAILURE)` على `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` يجب أن يُحدث الحالة إلى `{ description: 'Execution Failed' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

استدعاء `updateProjectStatus(projectStatus.SUCCESS)` على `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` يجب أن يُحدث الحالة إلى `{ description: 'Executed Successfully' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

يجب أن تكون قادرًا على تثبيت كائن `ProjectIdea` على `ProjectIdeaBoard` باستخدام طريقة `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

يجب أن تكون قادرًا على إزالة تثبيت كائن `ProjectIdea` من `ProjectIdeaBoard` باستخدام طريقة `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

عندما يكون `new ProjectIdeaBoard("Empty Board")` فارغًا، يجب أن تُرجع `emptyBoard.formatToString()` النص `Empty Board has 0 idea(s)\n`.

```js
let emptyBoard = new ProjectIdeaBoard("Empty Board");

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

عندما تثبت `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` على `new ProjectIdeaBoard("Tech Projects Board")`، يجب أن تُرجع `techProjects.formatToString()` النص `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
