---
id: 6995d0325fd7da40854fac40
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이전 수업에서 이진 탐색과 선형 탐색의 차이점을 배웠습니다.

이번 워크숍에서는 이진 탐색 알고리즘을 직접 만들어 보면서 배운 내용을 실습합니다.

먼저, `binarySearch` 함수를 `searchList`과 `value` 매개변수로 만드세요.

`searchList`은 함수가 탐색할 요소들의 배열이고, `value`는 함수가 `searchList` 안에서 찾으려는 대상입니다.

# --hints--

`binarySearch` 함수를 만들어야 합니다.

```js
assert.isFunction(binarySearch);
```

`binarySearch` 함수는 `searchList`과 `value` 매개변수를 가져야 합니다.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
