---
id: 6996d0a2c2ac55041145a13f
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

두 번째 함수 호출에 `3, 5, 4` 값이 목록에 있는 것을 보셨나요? 이것은 이진 탐색이 의도한 대로 작동하고 있음을 나타냅니다.

이렇게 된 과정은 다음과 같습니다: 알고리즘이 먼저 `3`을 초기 목록의 중간으로 확인했습니다. `4`가 `3`보다 크기 때문에 탐색이 오른쪽 절반으로 이동했습니다. 그다음 새로운 중간인 `5`를 조사했습니다. `4`가 `5`보다 작아서 탐색이 왼쪽으로 이동했고, 결국 `4`을 최종 범위의 중간으로 확인했습니다.

함수를 다시 테스트하려면 `[1, 3, 5, 9, 14, 22], 10`로 호출하고 콘솔에 기록하세요. 이 경우 값이 발견되지 않을 상황입니다.

# --hints--

함수를 `binarySearch([1, 3, 5, 9, 14, 22], 10)`로 호출하고 콘솔에 기록해야 합니다.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,3,5,9,14,22],10))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
