---
id: 6996c60281feca6718cc15c1
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`binarySearch` 함수 안에 `pathToTarget` 변수를 만들고, 우선 빈 배열로 초기화하세요.

# --hints--

`pathToTarget` 함수 안에 빈 배열로 설정된 `[]` 변수가 있어야 합니다 `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
