---
id: 6996c60281feca6718cc15c1
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`binarySearch` 함수 안에 `pathToTarget` 변수를 만들고, 우선 빈 배열로 초기화하세요.

# --hints--

`binarySearch` 함수 안에 `pathToTarget` 변수를 선언하고 빈 배열로 초기화하세요.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
