---
id: 6996c8f94fe6b250f11a2a1e
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 `valueAtMiddle`를 `pathToTarget` 배열에 포함하여 대상인지 여부와 상관없이 탐색 중에 거친 단계를 추적하세요. `push()` 메서드를 사용하면 됩니다.

# --hints--

`push()` 메서드를 사용해 `valueAtMiddle`를 `pathToTarget` 배열에 추가하세요.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
