---
id: 6996c95e5bede03bdf2937a2
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

다음으로 할 일은 대상 `value`가 중간에 있는지 확인하는 조건을 만드는 것입니다.

`while` 반복문 안에 대상 `if`가 `value`와 같은지 확인하는 `valueAtMiddle` 문을 만드세요. 같으면 `pathToTarget` 변수를 반환하세요.

# --hints--

조건 `if`를 가진 `value === valueAtMiddle` 문을 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

`pathToTarget` 문 안에서 `if` 변수를 반환해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
