---
id: 6996c796385b6ea4e222d5e8
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

루프를 시작하려면 현재 탐색 공간의 중간 인덱스를 찾아야 합니다.

중간 인덱스는 `low`와 `high` 변수의 평균을 계산해서 구할 수 있습니다. 두 값을 더한 후 평균을 구할 때는 항상 정수 인덱스여야 하므로 내림 나눗셈(floor division)을 사용해야 합니다. 이렇게 하면 결과가 항상 내림된 정수가 됩니다.

`while` 루프 안에서 `mid`와 `low`의 평균으로 설정된 `high` 변수를 만드세요.

# --hints--

while 루프 안에서 `mid`로 설정된 `Math.floor((low + high) / 2)` 변수를 만들어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
