---
id: 6996c8375f4301004459ee85
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

`mid` 인덱스를 계산했으니, 이제 `searchList`에서 해당 인덱스에 있는 값을 실제로 가져와야 합니다. 이 값이 `value` 매개변수, 즉 검색하려는 대상과 비교할 값입니다.

`valueAtMiddle` 변수를 만들고 `searchList`에서 `mid` 인덱스에 위치한 요소를 할당하세요.

# --hints--

`valueAtMiddle` 변수를 만들고 `searchList[mid]`로 설정하세요.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
