---
id: 6996ca21a74104721f2c7e5d
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

초기 단계에서 알고리즘의 흐름을 이해하려면 테스트해 봐야 합니다.

그렇게 하려면 먼저 반복문을 빠져나와야 합니다. 현재 구현은 한 번만 반복하도록 되어 있어서 조건이 충족되지 않으면 무한 반복이 발생하기 때문입니다.

`if` 문장 바로 뒤에 `break` 키워드를 사용해 `while` 반복문을 빠져나오세요. 그리고 `while` 반복문 뒤에는 값을 찾지 못했음을 나타내기 위해 빈 배열을 반환하세요.

# --hints--

`break` 문장 뒤에 `while` 반복문을 빠져나오기 위해 `if` 키워드를 사용하세요.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

`while` 반복문 뒤에는 빈 배열을 반환해야 합니다.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
