---
id: 6996d1032feed88e2957dc8e
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

마지막 호출이 빈 배열을 출력하는 것을 볼 수 있습니다. 검색한 배열에 값이 없다는 것을 더 명확히 하기 위해 `return` 문을 업데이트하여 두 요소를 포함하는 배열을 반환하세요: 빈 배열과 문자열 `Value not found`입니다.

# --hints--

함수가 첫 번째 요소로 `[]`를, 두 번째 요소로 `"Value not found"`를 가진 배열을 반환하도록 `return` 문을 업데이트하세요.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const validExpressions = [
  "[[], 'Value not found']",
  "[[], \"Value not found\"]",
  "[[], `Value not found`]"
];
assert.isTrue(validExpressions.some(expr => binarySearch?.hasReturn(expr) || binarySearch?.hasReturn("(" + expr + ")")));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
