---
id: 6996c675b599baa7b103192d
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이진 탐색에서는 가장 낮은 인덱스와 가장 높은 인덱스가 있습니다.

가장 낮은 인덱스는 현재 탐색 중인 값들의 구간에서 가장 왼쪽 경계를 나타내며, 이 경우 `searchList`입니다. 반면 가장 높은 인덱스는 탐색 중인 값들에서 가장 오른쪽 경계를 나타냅니다.

이 두 가지를 고려하려면 검색 중인 전체 배열을 생각해야 합니다. 그래서 `binarySearch` 함수 안에 `low` 변수를 값 `0`으로 선언하고, 검색 중인 배열의 마지막 인덱스를 반영하는 값을 가진 `high` 변수를 선언하세요.

# --hints--

`binarySearch` 함수 안에 `low`라는 이름의 변수를 값 `0`으로 선언하세요.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

`binarySearch` 함수 안에 `high`라는 이름의 변수를 `searchList.length - 1`로 선언하세요.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
