---
id: 6996c7363a4a437839f1f15d
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

검색 범위를 정의했으니, 이제 이진 탐색 알고리즘을 수행할 반복문을 만드세요.

이진 탐색은 검색 공간을 반복해서 좁혀 나가면서 작동합니다. 이 과정은 검사할 유효한 요소 범위가 있을 때 계속됩니다.

`while` 반복문을 사용해 `low` 포인터가 `high` 포인터보다 작거나 같을 때 계속 실행되도록 표현할 수 있습니다.

만약 `low`가 `high`보다 커지면, 검색 공간이 비어 있다는 뜻이며 `value`가 배열에 없다는 의미입니다.

`binarySearch` 함수 안에 `while` 반복문을 만들고, `low`가 high보다 작거나 같은지 확인하는 조건을 넣으세요.

# --hints--

`while`이 `low`보다 작거나 같을 때 실행되는 `high` 반복문을 만들어야 합니다.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
