---
id: 6996ccadf9048c79b857bb15
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

첫 번째 호출이 `[3]`를 반환하는 것을 볼 수 있습니다. 이는 `3`이 `[1, 2, 3, 4, 5]`에서 중간값이기 때문입니다.

반면 두 번째 호출은 빈 배열을 반환합니다. 이는 중간값을 확인한 후, 현재 루프가 `low` 또는 `high` 변수를 업데이트하여 검색 범위를 좁히는 로직이 없기 때문입니다. 단순히 중간값을 한 번 확인하고 추가 검색 없이 종료합니다.

이진 검색이 계속해서 검색 범위를 좁힐 수 있도록, `else if`가 `value`보다 큰지 확인하는 `valueAtMiddle` 문을 추가하세요.

# --hints--

조건 `else if`을 가진 `value > valueAtMiddle` 문을 추가해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
