---
id: 6996cb5869acf58de751ae81
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

함수가 `searchList`라는 배열과 찾고자 하는 `value`을 받는다는 것을 기억하세요.

이제 `binarySearch([1, 2, 3, 4, 5], 3)`와 `binarySearch([1, 2, 3, 4, 5, 9], 4)`으로 함수를 호출하고 두 문장을 모두 콘솔에 출력하세요.

# --hints--

`binarySearch([1, 2, 3, 4, 5], 3)`로 함수를 호출하고 콘솔에 출력해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

`binarySearch([1, 2, 3, 4, 5, 9], 4)`로 함수를 호출하고 콘솔에 출력해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
