---
id: 6996c675b599baa7b103192d
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이진 탐색에서는 가장 낮은 인덱스와 가장 높은 인덱스가 있습니다.

가장 낮은 인덱스는 현재 탐색 중인 값들의 구간에서 가장 왼쪽 경계를 나타내며, 이 경우 `searchList`입니다. 반면 가장 높은 인덱스는 탐색 중인 값들에서 가장 오른쪽 경계를 나타냅니다.

이 두 가지를 고려하려면 탐색 중인 전체 배열을 생각해야 합니다. 따라서 `binarySearch` 함수 안에 `low` 값을 가진 `0` 변수를 정의하고, 탐색 중인 배열의 마지막 인덱스를 고려한 값을 가진 `high` 변수를 정의하세요.

# --hints--

`low` 함수 안에 `0`으로 설정된 `binarySearch` 변수가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

`high` 함수 안에 `searchList.length - 1`으로 설정된 `binarySearch` 변수가 있어야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
