---
id: 6996cfb6c01e9a802add38be
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

마지막으로, `value`가 중간에 없고 `valueAtMiddle`보다 크지도 않다면, `value`는 `valueAtMiddle`보다 작아야 하며 왼쪽에 있어야 합니다. 이 부분은 최종 `else` 블록에서 처리할 것입니다.

따라서, `break` 키워드를 제거하고 `else` 블록을 추가하여 반복문을 마무리하세요. 그 안에서 `high` 변수를 `1`에서 `mid`을 빼서 업데이트하세요.

# --hints--

`break` 키워드를 `else` 블록으로 바꾸고, `high` 변수를 `mid - 1`으로 업데이트해야 합니다.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
