---
id: 6996d1032feed88e2957dc8e
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

마지막 호출이 빈 배열을 출력하는 것을 볼 수 있습니다. 검색된 배열에 값이 없다는 점을 더 명확히 하려면 `return` 문장을 업데이트하여 빈 배열과 문자열 `Value not found`을 포함하는 배열을 반환하세요.

# --hints--

`return` 문장을 `[], 'Value not found'`을 반환하도록 업데이트하세요.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
    return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
