---
id: af7588ade1100bde429baf20
title: 建置一個遺漏字母偵測器
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `fearNotLetter` 的函式（程式）。
1. `fearNotLetter` 函式（程式）應該接受一個引數：一個表述按字母順序排列且可能缺少一個字母的字串。
1. 該函式（程式）應該在傳入的字母範圍中找到缺失的字母並傳回它。
1. 如果所有字母都存在於該範圍內，該函式應該傳回 `undefined`。

# --hints--

你應該有一個 `fearNotLetter` 函式（程式）。

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` 應該傳回字串 `d`。

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` 應該傳回字串 `i`。

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` 應該傳回字串 `u`。

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` 應該傳回字串 `e`。

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` 應該傳回 `undefined`。

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
