---
id: af7588ade1100bde429baf20
title: Construye un Detector de Letras Faltantes
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `fearNotLetter`.
1. La función `fearNotLetter` debe aceptar un argumento: una cadena que representa un rango de letras en orden alfabético que puede tener una letra faltante.
1. La función debe encontrar la letra faltante en el rango de letras pasado y devolverla.
1. Si todas las letras están presentes en el rango, la función debería devolver `undefined`.

# --hints--

Deberías tener una función `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` debe devolver la cadena `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` debe devolver la cadena `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` debe devolver la cadena `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` debe devolver la cadena `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` debe devolver `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
