---
id: 68e6a99a0fd7386fb66e156f
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、ゲーム設定パネルを作成しながらチェックボックスのスタイリングを練習します。

すべてのHTMLボイラープレート（`DOCTYPE`、`html`、`head`、`body`）は用意されています。

`body`の中に、`div`要素を作成し、その中に`h1`要素をネストしてください。`h1`要素にはテキスト`Game Settings`を設定してください。

# --hints--

`div`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

`h1`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

`h1`要素は`div`要素の中にあるべきです。

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

`h1`要素にはテキスト`Game Settings`があるべきです。

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
