---
id: 68e9a6201e2efa142d62bfe1
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

次は、チェックボックスの可視性を高めるためにサイズを大きくします。

まず、`input` のセレクターを設定しますが、特に `[type="checkbox"]` をターゲットにしてください。

セレクター内で、`width` と `height` に `20px` を設定します。これで以前より大きくなります。

そして、ラベルに設定した `cursor` の設定に合わせて、`cursor` を追加し、値に `pointer` を代入してください。これでチェックボックスにカーソルを合わせるとポインターが表示されます。

# --hints--

CSSに `input[type="checkbox"]` のセレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

`width` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

`width` プロパティに `20px` の値を設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

`height` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

`height` プロパティに `20px` の値を設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

`cursor` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

`cursor` プロパティに `pointer` の値を設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
