---
id: 68e9a834dfef752dc35f6f90
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

ユーザーがチェックボックスをクリックしたときに色が変わるようになったので、チェックボックスが選択されたことがもう少しわかりやすくなるようにスタイルを適用しましょう。

CSSの疑似クラスは連結できます。複数の条件を満たす要素をターゲットにするために、疑似クラスや疑似要素を続けて追加できます。

`input[type="checkbox"]` の新しいセレクターを作成してください。`:checked` 疑似クラスの後に`::after` 疑似要素を追加します。これにより、チェックボックスがチェックされたときに視覚的なインジケーターを追加できます。

次に、`content` を`"✓"` の値で追加してください。

# --hints--

`input[type="checkbox"]:checked::after` の疑似クラスセレクターと疑似要素セレクターを組み合わせた型セレクターを作成しているはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

`content` の疑似クラスセレクターと疑似要素セレクターを組み合わせた型セレクターに`input[type="checkbox"]:checked::after` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

`content` プロパティに`"✓"` の値を設定しているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```
