---
id: 68e9a7bd60c971278d2d4ffc
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

前のレッスンで、CSSの疑似クラスと疑似要素について学びました。今回はその知識を活かして、疑似クラスセレクターと組み合わせた型セレクターを作成します。

まず、型セレクターの形式を`input[type="checkbox"]`に設定し、その末尾に`:checked`を追加してください。

次に、`background-color`に`#f1be32`の値を設定します。

最後に、`border-color`に`#e2a60d`の値を設定してください。

# --hints--

`input[type="checkbox"]:checked`の型セレクターと疑似クラスセレクターの組み合わせができているはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

`background-color`プロパティを、`input[type="checkbox"]:checked`の型セレクターと疑似クラスセレクターの組み合わせに追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

`background-color`プロパティに`#f1be32`の値が設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

`border-color`プロパティを、`input[type="checkbox"]:checked`の型セレクターと疑似クラスセレクターの組み合わせに追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

`border-color`プロパティに`#e2a60d`の値が設定されているはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
