---
id: 68e9a4a47d8a9802cba886b7
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

body にスタイルを適用したので、次は Game Settings パネル用のコンテナカードを作成します。

まず、`settings-card` 要素に `div` のクラスを追加してください。

# --hints--

`div` には `settings-card` のクラスを付ける必要があります。

```js
assert.strictEqual(document.querySelector("div")?.className, "settings-card");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <div>
--fcc-editable-region--
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}
```
