---
id: 68e9a3194160f8f5620ddf21
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 要素の直後に、4つの `label` 要素を作成してください。

各 `label` 要素の中に、`input` 属性が `type` に設定された `"checkbox"` 要素を作成してください。

各 `input` 要素の後に、`label` 要素に次のテキストを順番に設定してください：`Sound Effects`、`Background Music`、`Hard Mode`、`Haptic Feedback`。

# --hints--

4つの `label` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

最初の `label` 要素には `input` 要素があるはずです。

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

最初の `input` 要素には `type` 属性が `checkbox` に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

最初の `label` 要素にはテキスト `Sound Effects` があるはずです。

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

`Sound Effects` のテキストは `input` 要素の後に来るはずです。

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

2番目の `label` 要素には `input` 要素があるはずです。

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

2番目の `input` 要素には `type` 属性が `checkbox` に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

2番目の `label` 要素にはテキスト `Background Music` があるはずです。

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

`Background Music` のテキストは `input` 要素の後に来るはずです。

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

3番目の `label` 要素には `input` 要素があるはずです。

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

3番目の `input` 要素には `type` 属性が `checkbox` に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

3番目の `label` 要素にはテキスト `Hard Mode` があるはずです。

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

`Hard Mode` のテキストは `input` 要素の後に来るはずです。

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

4番目の `label` 要素には `input` 要素があるはずです。

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

4番目の `input` 要素には `type` 属性が `checkbox` に設定されているはずです。

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

4番目の `label` 要素にはテキスト `Haptic Feedback` があるはずです。

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

`Haptic Feedback` のテキストは `input` 要素の後に来るはずです。

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
