---
id: 68e9a6bf1a74721a4e43e217
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

これからブラウザーがデフォルトで適用するチェックボックスを削除します。`input[type='checkbox']` セレクターの中に、値が `appearance` の `none` を追加してください。

`appearance` プロパティを `none` に設定すると、ブラウザーがチェックボックスに適用する見た目がクリアされ、好きなスタイルを表示できるようになります。

その後、チェックボックスが見えなくなるので、`border` の太さ、`2px` のスタイル、そして `solid` の16進コードを持つ `#f1be32` を設定してください。

# --hints--

`appearance` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

`appearance` プロパティに `none` の値があるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

`border` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

`border` プロパティに `2px solid #f1be32` の値があるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
