---
id: 68ee86e3ee07b01626a392cc
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`.settings-card` セレクター内で `margin` プロパティを `auto` に設定してください。

`margin` プロパティを `auto` に設定すると、要素のマージンが自動的に調整され、コンテナ内の残りのスペースが均等に分配されます。これは、ブロックレベル要素を水平方向に中央揃えにする際によく使われます。

最後に、`text-align` プロパティを `left` の値で設定してください。これにより、テキストなどのインラインコンテンツがそのコンテナの左側に揃います。

# --hints--

`margin` クラスセレクターに `settings-card` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

`margin` プロパティを `auto` の値で設定しているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

`text-align` クラスセレクターに `settings-card` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

`text-align` プロパティを `left` の値で設定しているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
