---
id: 68f1560aff1214103ab6b59b
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

これからCSSスタイリングでページを整えていきます。まずは`body`セレクターを作成してください。

`body`に`height`プロパティを値`100vh`で、さらに`background-color`プロパティを値`#f0f0f0`で設定してください。

`height`の`100vh`は、`body`がブラウザのビューポートの高さ全体を占めるようにし、薄いグレーの背景色がページの控えめなベースになります。

最後に、`text-align`プロパティを値`center`で設定してください。これにより、子要素が独自の整列でオーバーライドしない限り、ページ内のすべてのインラインコンテンツが中央に配置されます。

# --hints--

`body`型セレクターを使うべきです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

`height`型セレクター内に`body`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

`height`の値を持つ`100vh`プロパティがあるべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

`background-color`型セレクター内に`body`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

`background-color`の値を持つ`#f0f0f0`プロパティがあるべきです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

`text-align`型セレクター内に`body`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

`text-align`の値を持つ`center`プロパティがあるべきです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
