---
id: 68f15ef0446464325ce1aa45
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

チェックボックスとラベルの間にスペースが必要です。まず、CSSで`label`のセレクターを使い、`display`プロパティに`block`を設定してください。

次に、要素間の縦方向のスペースを追加するために`margin`プロパティに`8px auto`を設定します。これにより、各要素の上下に均一なスペースができ、コンテンツの区切りがはっきりして読みやすくなります。

最後に、`cursor`要素に`pointer`プロパティを`label`に設定してください。これにより、ユーザーが`label`にカーソルを合わせたときに手のアイコンに変わり、クリック可能であることを示してユーザー体験が向上します。

# --hints--

CSSに`label`の型セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

`display`セレクターに`label`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

`display`の値を持つ`block`プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

`margin`セレクターに`label`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

`margin`の値を持つ`8px auto`プロパティがあるはずです。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

`cursor`セレクターに`label`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

`cursor`の値を持つ`pointer`プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
