---
id: 68f15cd397503429f0345262
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`input[type="checkbox"]:checked::after` セレクターの宣言で、`display` を `block` に設定してください。`display` プロパティを `block` に設定すると、その要素はブロックレベル要素になり、コンテナの幅いっぱいを占めて新しい行から始まるため、幅、高さ、間隔をより簡単に制御できます。

次に、`text-align` を `center` に設定してください。これにより、インラインコンテンツ（テキストやインライン要素など）がブロック内で水平方向に中央揃えになります。ブロックは引き続きコンテナの幅いっぱいを占めますが、その中のすべてが中央に揃います。

# --hints--

疑似クラスセレクターと疑似要素セレクターが `display` の複合型セレクターに、`input[type="checkbox"]:checked::after` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

`display` プロパティを値 `block` で設定しているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

疑似クラスセレクターと疑似要素セレクターが `text-align` の複合型セレクターに、`input[type="checkbox"]:checked::after` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

`text-align` プロパティを値 `center` で設定しているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
