---
id: 68f29dbf9125476b4808d6fb
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

これで`settings-card`コンテナのスタイルを設定する時間です。`settings-card`のクラスセレクターを作成してください。

ここにコンテナのすべてのフォーマットを記述します。コンテナの全体のサイズを定義するために、`max-width`を`250px`に設定してください。

次に、コンテンツとコンテナのボーダーの間にスペースを作るために、`padding`を`20px`に設定してください。

その後、`border-radius`を`10px`に設定して、角を丸くしてください。

そして、`box-shadow`の値を持つ`0 2px 6px rgba(0,0,0,0.2)`を設定してください。これにより、コンテナに微妙な「浮き上がった」見た目ができて、奥行きが生まれます。

# --hints--

CSSに`settings-card`のクラスセレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

`max-width`クラスセレクターに`.settings-card`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

値が`max-width`の`250px`プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

`padding`クラスセレクターに`.settings-card`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

値が`padding`の`20px`プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

`border-radius`クラスセレクターに`.settings-card`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

値が`border-radius`の`10px`プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

`box-shadow`クラスセレクターに`.settings-card`プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

値が`box-shadow`の`0 2px 6px rgba(0,0,0,0.2)`があるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
