---
id: 68f29753fd5524530c6e65c8
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

次に、値が `vertical-align` の `middle` プロパティを設定してください。

`vertical-align` プロパティは、`inline` または `inline-block` 要素が周囲のテキストや他のインライン要素と垂直方向にどのように揃うかを制御します。これは、画像やアイコン、行内のテキストなどの垂直位置を調整するためによく使われます。

その後、`margin` プロパティを `15px` にしてチェックボックスを完成させてください。

# --hints--

`vertical-align` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

値が `vertical-align` の `middle` プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

`margin-right` セレクターに `input[type="checkbox"]` プロパティを追加してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

値が `margin-right` の `15px` プロパティがあるはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
