---
id: 6823e637c1c0ed56f781b4fc
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

最后，JavaScript 能使你的网页可交互——它能让你设定当用户点击按钮、提交表单或执行其他操作时网页应如何响应。

在本工作坊的最后一步中，要在页面中添加另一个 `h3` 和 `p` 元素，用于描述 JavaScript。

首先，添加一个 `h3` 元素，文本为：

```md
Introduction to JavaScript
```

然后，在 `h3` 元素下面，添加一个 `p` 元素，文本为：

```md
JavaScript adds interactivity to a webpage
```

# --hints--

你应该为页面添加第三个 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

你的 `h3` 元素的文本应该是 `Introduction to JavaScript`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

你应该为页面添加第四个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

你的 `p` 元素应该包含文本 `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
