---
id: 6823d9ac8bdc3853df65a1ff
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

HTML 中有六个标题元素： `h1` 到 `h6`。 它们用于表明网页各部分的重要性，其中 `h1` 最为重要，而 `h6` 则最为次要。

在你的 `p` 元素下方，添加一个 `h3` 标题，文本为：

```md
Introduction to HTML
```

# --hints--

你的 `h3` 元素应该有一个开始标签 `<h3>`。

```js
assert.exists(document.querySelector("h3"));
```

你的 `h3` 元素应该有一个结束标签 `</h3>`。

```js
assert.match(code, /<\/h3\s*\>/);
```

你的 `h3` 元素的文本应该是 `Introduction to HTML`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
