---
id: 6823e036ea4b71553558c01b
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

你会注意到，每个标题在预览中看起来都有点不同。 段落也是如此。

接下来，添加另一个 `p` 元素来显示以下文本：

```md
HTML represents the content and structure of a webpage
```

# --hints--

你应该有了第二个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你的 `p` 元素应该包含文本 `HTML represents the content and structure of a webpage` 。 你可以从说明中复制文本。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
