---
id: 6823ac607bfdbc46331b2559
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 代表超文本标记语言。 它是定义网页结构和内容的代码。 这是你的代码编辑器，你将在这里编写 HTML。

在编辑器中找到第 1 行并输入以下文本：

```md
Welcome to freeCodeCamp
```

完成后，点击“检查你的代码”按钮查看是否正确。

# --hints--

在你的编辑器中应包含文本 `Welcome to freeCodeCamp`。 仔细检查拼写。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
