---
id: 6823e169fda14755fbf00445
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

你逐渐掌握要领了。 接下来，在编辑器底部添加另一个 `h3` 元素，文本为：

```md
Introduction to CSS
```

# --hints--

你应该为页面添加第二个 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你的 `h3` 元素的文本应该是 `Introduction to CSS`。 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```
