---
id: 6823d6244511f252c8300eed
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

当你需要将一个段落添加到网页上时，你可以像这样使用 `p` 元素：

```html
<p>This is a paragraph element.</p>
```

将 `Learn the skills to become a full-stack developer` 转换为段落元素。

# --hints--

你的 `p` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素应有一个闭合标签。闭合标签的语法是：`</elementName>`。

```js
assert.match(code, /<\/p\s*\>/);
```

你的 `p` 元素的文本应为 `Learn the skills to become a full-stack developer`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

你的 `p` 元素应该在 `h2` 元素的下方。 

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
