---
id: 682ba2318000b62f179bdf04
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML 由元素组成。 你将使用的第一个元素是 `h1` 元素：

```html
<h1>Welcome to freeCodeCamp</h1>
```

它以一个开始标签（`<h1>`）开始，以一个结束标签（`</h1>`）结束，在标签之间包含了将要显示的文本。

在你的 `Welcome to freeCodeCamp` 文本前面添加一个开始标签，后面添加一个结束标签，将其转化为一个 `h1` 元素。

# --hints--

你的 `h1` 元素应该有一个开始标签 `<h1>`。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素应该有一个结束标签 `</h1>`。

```js
assert.match(code, /<\/h1\s*\>/);
```

你的 `h1` 元素看起来应该像这样： `<h1>Welcome to freeCodeCamp</h1>`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
