---
id: 682cd20b883fc7b25eb539c6
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在之前的两行文本下面，添加：

```md
Learn the skills to become a full-stack developer
```

# --hints--

你应该在编辑器中有文本 `Learn the skills to become a full-stack developer`。请仔细查看拼写。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```
