---
id: 682cd206883fc7b25eb539c5
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

你会注意到，你在编辑器中输入的 HTML 会在预览中显示。 在本工作坊中，你将为本课程大纲网页编写一小部分 HTML 代码。

在你的 `h1` 元素下方的空行中输入以下内容：

```md
Full-Stack Curriculum
```

# --hints--

你的编辑器中应该有文本 `Full-Stack Curriculum`。请仔细查看拼写。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
