---
id: 683921c4769fd23dbadec2fe
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

HTML 定义了网页的结构和内容，CSS 则用于添加样式——比如颜色、字体、间距和布局。

在你刚刚添加的 `h3` 的下面，添加另一个段落元素，文本为：

```md
CSS is used to style a webpage
```

# --hints--

你应该为页面添加第二个 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你的 `h3` 元素的文本应该是 `Introduction to CSS`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

你应该为页面添加第三个 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

你的 `p` 元素应该包含文本 `CSS is used to style a webpage`。 你可以从说明中复制文本。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
