---
id: afcc8d540bea9ea2669306b6
title: 构建一个重复字串的函数
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

在本实验中，你将创建一个函数，将给定的字串重复指定的次数。为本实验目的，请*勿*使用内置的 `.repeat()` 方法。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `repeatStringNumTimes` 的函数，该函数接受两个参数：一个字串和一个数字。
2. 该函数应返回重复指定次数的字串。
3. 如果数字小于或等于零，函数应返回空字串。

# --hints--

你应该创建一个名为 `repeatStringNumTimes` 的函数。

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` 应该接受两个参数。

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

函数 `repeatStringNumTimes` 应始终返回一个字串。

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` 应返回 `***`。

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` 应返回 `abcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` 应返回 `abcabcabcabc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` 应返回 `abc`。

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` 应返回 `********`。

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` 应返回 `""`。

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` 应返回 `""`。

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

不应使用内置的 `repeat()` 方法。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
