---
id: 587d824a367417b2b2512c45
title: Bodi ya Ujumbe Isiyojulikana
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

Jenga programu kamili ya JavaScript inayofanana kiutendaji na hii: <a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>.

Kufanya kazi kwenye mradi huu kutahusisha kuandika msimbo wako kwa kutumia mojawapo ya njia zifuatazo:

-   Nakili <a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">hii repo ya GitHub</a> na ukamilishe mradi wako kwa ndani.
-   Tumia mtengenezaji wa tovuti unayempenda kukamilisha mradi. Hakikisha unajumuisha mafaili yote kutoka kwenye repo yetu ya GitHub.

# --instructions--

1.  Weka `NODE_ENV` kujaribu bila nukuu unapokuwa tayari kuandika majaribio na DB kwa mfuatano wa muunganisho wa hifadhidata zako (katika `.env`)
2.  Inashauriwa kuunda controllers/mshughulikiaji na kushughulikia njia za `routes/api.js`
3.  Utaongeza vipengele vyovyote vya usalama kwenye `server.js`

Andika majaribio yafuatayo katika `tests/2_functional-tests.js`:

-   Kuunda nyuzi mpya: ombi la POST kwa `/api/threads/{board}`
-   Kutazama nyuzi 10 za hivi karibuni na majibu 3 kila moja: ombi la GET kwa `/api/threads/{board}`
-   Kufuta nyuzi kwa nenosiri lisilo sahihi: ombi la DELETE kwa `/api/threads/{board}` na `delete_password` isiyo halali
-   Kufuta nyuzi kwa nenosiri sahihi: ombi la DELETE kwa `/api/threads/{board}` na `delete_password` halali
-   Kuripoti nyuzi: ombi la PUT kwa `/api/threads/{board}`
-   Kuunda jibu jipya: ombi la POST kwa `/api/replies/{board}`
-   Kutazama nyuzi moja na majibu yote: ombi la GET kwa `/api/replies/{board}`
-   Kufuta jibu kwa nenosiri lisilo sahihi: ombi la DELETE kwa `/api/replies/{board}` na `delete_password` isiyo halali
-   Kufuta jibu kwa nenosiri sahihi: ombi la DELETE kwa `/api/replies/{board}` na `delete_password` halali
-   Kuripoti jibu: ombi la PUT kwa `/api/replies/{board}`

# --hints--

Unaweza kutumia mradi wako mwenyewe, si lazima URL ya mfano.

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Ruhusu tovuti yako ipakuliwe kwenye iFrame tu kwenye kurasa zako mwenyewe.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

Usiruhusu DNS prefetching.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

Ruhusu tovuti yako itume referrer kwa kurasa zako mwenyewe tu.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

Unaweza kutuma ombi la POST kwa `/api/threads/{board}` ukiwa na data ya fomu ikijumuisha `text` na `delete_password`. Rekodi iliyohifadhiwa kwenye hifadhidata itakuwa na angalau sehemu za `_id`, `text`, `created_on` (tarehe na saa), `bumped_on` (tarehe na saa, huanza sawa na `created_on`), `reported` (thamani ya kweli au siyo kweli), `delete_password`, na `replies` (safu ya data).

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la POST kwa `/api/replies/{board}` ukiwa na data ya fomu ikijumuisha `text`, `delete_password`, na `thread_id`. Hii itasasisha tarehe ya `bumped_on` kuwa tarehe ya maoni. Katika safu ya `replies` ya nyuzi, kitu kitatunzwa chenye angalau mali za `_id`, `text`, `created_on`, `delete_password`, na `reported`.

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la GET kwa `/api/threads/{board}`. Itarudisha safu ya nyuzi 10 za hivi karibuni zilizopandishwa kwenye bodi na majibu 3 ya hivi karibuni kwa kila moja. Sehemu za `reported` na `delete_password` hazitatumwa kwa mteja.

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la GET kwa `/api/replies/{board}?thread_id={thread_id}`. Itarudisha nyuzi nzima na majibu yote, pia ikiondoa sehemu zile zile kutoka kwa mteja kama jaribio la awali.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la DELETE kwa `/api/threads/{board}` na kupitisha `thread_id` na `delete_password` kufuta nyuzi. Itarudisha mfuatano wa herufi `incorrect password` au `success`.

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la DELETE kwa `/api/replies/{board}` na kupitisha `thread_id`, `reply_id`, na `delete_password`. Itarudisha mfuatano wa herufi `incorrect password` au `success`. Kwa mafanikio, maandishi ya `reply_id` yatabadilishwa kuwa `[deleted]`.

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la PUT kwa `/api/threads/{board}` na kupitisha `thread_id`. Itarudisha mfuatano wa herufi `reported`. Thamani ya `reported` ya `thread_id` itabadilishwa kuwa `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Unaweza kutuma ombi la PUT kwa `/api/replies/{board}` na kupitisha `thread_id` na `reply_id`. Itarudisha mfuatano wa herufi `reported`. Thamani ya `reported` ya `reply_id` itabadilishwa kuwa `true`.

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

Majaribio yote 10 ya utendaji yamekamilika na kupita.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```
