---
id: 579e2a2c335b9d72dd32e05c
title: 实现 Slice 和 Splice 算法
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

在本实验中，你需要创建一个算法来合并两个数组。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 创建一个接受两个数组和一个索引的 `frankenSplice` 函数。
1. 将第一个数组的每个元素按顺序拷贝到第二个数组中，从给定的索引开始，并返回结果数组。
1. 作为输入参数的两个数组在函数执行前后应保持不变。

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` 应返回 `[4, 1, 2, 3, 5]`。

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` 应返回 `["a", 1, 2, "b"]`。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` 应返回 `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`。

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

第一个数组中的所有元素应按原始顺序添加到第二个数组中。`frankenSplice([1, 2, 3, 4], [], 0)` 应该返回 `[1, 2, 3, 4]`。

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

函数运行后，第一个数组应保持不变。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

函数运行后，第二个数组应保持不变。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
