---
id: 68c2fb153760d075321c5da3
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Camperbot ha creado una página de blog sobre un viaje de codificación, pero parece que la página tiene algunos problemas de accesibilidad y malas prácticas. Tu tarea en este taller es corregir estos problemas para Camperbot.

El primer conjunto de errores tiene que ver con el uso de encabezados. En una lección anterior, aprendiste que la mejor práctica es usar solo un `h1` por página. Pero parece que Camperbot está usando varios `h1`. 

Deja el `<h1>Welcome to Camperbot's Blog</h1>` tal cual y cambia los otros elementos `h1` a elementos `h2`.

# --hints--

No debes alterar el `<h1>Welcome to Camperbot's Blog</h1>`.

```js
const heading = document.querySelectorAll("h1")[0];
assert.equal(heading?.textContent, "Welcome to Camperbot's Blog");
```

Debes cambiar el elemento `h1` con el texto `My Journey Learning to Code` a un elemento `h2`.

```js
const firstPost = document.getElementById("post1");
assert.equal(firstPost?.tagName, "H2");
assert.equal(firstPost?.textContent, "My Journey Learning to Code");
```

Debes cambiar el elemento `h1` con el texto `Accessibility Matters` a un elemento `h2`.

```js
const secondPost = document.getElementById("post2");
assert.equal(secondPost?.tagName, "H2");
assert.equal(secondPost?.textContent, "Accessibility Matters");
```

Debes cambiar el elemento `h1` con el texto `What's Next?` a un elemento `h2`.

```js
const thirdPost = document.getElementById("post3");
assert.equal(thirdPost?.tagName, "H2");
assert.equal(thirdPost?.textContent, "What's Next?");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Camperbot's Blog</title>
</head>
<body>
  <h1>Welcome to Camperbot's Blog</h1>

  <h2>Navigation</h2>
  <ul>
    <li><a href="#post1">My Journey</a></li>
    <li><a href="#post2">Accessibility</a></li>
    <li><a href="#post3">Next Steps</a></li>
  </ul>
--fcc-editable-region--
  <div>
    <h1 id="post1">My Journey Learning to Code</h1>
    <p>I started learning to code a few months ago and it's been a wild ride!</p>
    
    <h4>Early Challenges</h4>
    <p>At first, syntax was really confusing.</p>
    
    <h4>Breakthroughs</h4>
    <p>Eventually things started to click.</p>
  </div>

  <div>
    <h1 id="post2">Accessibility Matters</h1>
    <p>Today I learned that not everyone uses the web the same way I do.</p>

    <h5>Screen Readers</h5>
    <p>These tools help visually impaired users browse websites.</p>
  </div>

  <div>
    <h1 id="post3">What's Next?</h1>
    <p>I'm excited to dive into JavaScript and build interactive features!</p>

    <h3>Coming soon: My first JavaScript project!</h3>
    <p>Stay tuned for some exciting interactive blog features.</p>
  </div>
--fcc-editable-region--
  <h3>Contact Me</h3>
  <p>Email me at camperbot@blog.io</p>
</body>
</html>
```
