---
id: 68e76b1ce8216249d04df670
title: بناء صنف كوكب
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب عليك إنشاء صنف باسم `Planet`.
1. يجب أن يحتوي صنف `Planet` على طريقة `__init__` التي:
   - تحتوي على أربع معلمات: `self`، و`name`، و`planet_type`، و`star`.
   - ترفع استثناء `TypeError` مع الرسالة `name, planet type, and star must be strings` إذا كان أي من المعطيات الممررة ليس من نوع سلسلة نصية.
   - ترفع استثناء `ValueError` مع الرسالة `name, planet_type, and star must be non-empty strings` إذا كان أي من المعطيات الممررة سلسلة نصية فارغة.
   - تعيّن القيم الممررة إلى خصائص الكائن `name`، و`planet_type`، و`star`.
1. يجب أن يحتوي صنف `Planet` على طريقة `orbit` التي تُرجع سلسلة نصية بالشكل `{name} is orbiting around {star}...`.
1. يجب أن يحتوي صنف `Planet` على طريقة `__str__` التي تُرجع سلسلة نصية بالشكل `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. يجب عليك إنشاء ثلاث نسخ من صنف `Planet` باسم `planet_1`، و`planet_2`، و`planet_3`.
1. يجب طباعة كل كائن كوكب لرؤية ناتج طريقة `__str__`.
1. يجب استدعاء طريقة `orbit` على كل كائن كوكب وطباعة النتيجة.

# --hints--

يجب عليك إنشاء صنف باسم `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

يجب أن يحتوي صنف `Planet` على طريقة `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

يجب أن تحتوي طريقة `__init__` على أربع معلمات: `self`، و`name`، و`planet_type`، و`star`، بهذا الترتيب.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `TypeError` مع الرسالة `name, planet type, and star must be strings` إذا لم يكن المعطى الأول من نوع سلسلة نصية.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `TypeError` مع الرسالة `name, planet type, and star must be strings` إذا لم يكن المعطى الثاني من نوع سلسلة نصية.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `TypeError` مع الرسالة `name, planet type, and star must be strings` إذا لم يكن المعطى الثالث من نوع سلسلة نصية.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `ValueError` مع الرسالة `name, planet_type, and star must be non-empty strings` إذا كان المعطى الأول سلسلة نصية فارغة.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `ValueError` مع الرسالة `name, planet_type, and star must be non-empty strings` إذا كان المعطى الثاني سلسلة نصية فارغة.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب أن ترفع طريقة `__init__` استثناء `ValueError` مع الرسالة `name, planet_type, and star must be non-empty strings` إذا كان المعطى الثالث سلسلة نصية فارغة.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

يجب تعيين `name` إلى `self.name` في طريقة `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

يجب تعيين `planet_type` إلى `self.planet_type` في طريقة `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

يجب تعيين `star` إلى `self.star` في طريقة `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

يجب أن يحتوي صنف `Planet` على طريقة `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

يجب أن تُرجع طريقة `orbit` سلسلة نصية بالشكل `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

يجب أن يحتوي صنف `Planet` على طريقة `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

يجب أن تُرجع طريقة `__str__` سلسلة نصية بالشكل `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

يجب عليك إنشاء ثلاث نسخ من صنف `Planet` باسم `planet_1`، و`planet_2`، و`planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

يجب طباعة كل كائن كوكب لرؤية ناتج طريقة `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

يجب استدعاء طريقة `orbit` على كل كائن كوكب وطباعة النتيجة.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
