---
id: 655cd899f8de09431eabb40c
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

병합 정렬은 분할 정복 원리를 사용해 데이터 모음을 정렬하는 알고리즘입니다. 즉, 모음을 더 작은 하위 부분으로 "분할"하고, 하위 부분을 독립적으로 정렬해 "정복"한 다음, 정렬된 하위 부분을 병합합니다.

`merge_sort`라는 이름의 함수를 선언하고 매개변수로 `array`을 받는 것부터 시작하세요. 이 함수는 숫자 목록을 정렬하는 작업을 처리합니다. 함수 본문에서 `pass` 키워드를 사용하세요.

# --hints--

`merge_sort`라는 이름의 함수를 매개변수 하나 `array`와 함께 선언하세요. 함수 본문에서 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
