---
id: 6564aee9c077774ea49c3faf
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

병합 정렬 알고리즘은 세 가지 작업을 수행합니다:

- 정렬되지 않은 항목 수열을 하위 부분으로 나누기
- 하위 부분의 항목을 정렬하기
- 정렬된 하위 부분을 병합하기

위 작업은 하위 부분이 완전한 정렬된 수열로 병합될 때까지 재귀적으로 진행됩니다. 수열을 나누는 것부터 시작하세요.

먼저 `pass` 키워드를 변수 `middle_point`로 바꾸세요. 정수 나눗셈 연산자(`//`)를 사용해 `array` 리스트의 길이를 2로 나누고 결과를 새 `middle_point` 변수에 할당하세요. 코드를 들여쓰기 하는 것을 잊지 마세요.

# --hints--

함수 본문에 `pass`가 없어야 합니다.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

`middle_point` 변수를 선언하고 `len(array) // 2`을 할당해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
