---
id: 656639771fed09b5c0e8fe71
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`array` 리스트를 두 개의 별도 리스트로 나눈 후, 각 리스트를 계속 나누어 모든 요소가 각각의 리스트에 단독으로 있게 하세요. 숫자 하나만 있는 리스트는 항상 정렬된 상태입니다.

이를 위해 함수 안에서 재귀적으로 `merge_sort`을 호출하고, 호출 시 인수로 `left_part`를 전달하세요.

# --hints--

함수 본문 맨 아래에서 `merge_sort`을 `left_part` 인수와 함께 호출해야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
