---
id: 656680b0fc79f2c38a34d90e
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 리스트(`left_part`와 `right_part`)를 정렬하고 원래 리스트에 병합할 시간입니다.

두 리스트의 요소를 비교해서 더 작은 요소를 메인 리스트에 병합하면 됩니다. `left_part`와 `right_part`의 모든 인덱스에 대해 이 비교를 수행할 것입니다.

세 개의 변수를 만드세요: `left_array_index`, `right_array_index`, `sorted_index` 그리고 이들의 값을 `0`로 설정하세요. 이 변수들은 정렬 과정에서 각 인덱스를 추적하는 데 도움이 됩니다.

# --hints--

`left_array_index` 함수 안에 `merge_sort`라는 변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

`left_array_index`의 값은 `0`이어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

`right_array_index` 함수 안에 `merge_sort`라는 변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

`right_array_index`의 값은 `0`이어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

`sorted_index` 함수 안에 `merge_sort`라는 변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

`sorted_index`의 값은 `0`이어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
