---
id: 656707e11e671ed54c96f7ec
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`if` 조건이 `True`으로 평가되면, `left_part` 목록에 있는 요소가 `right_part` 목록에서 비교 대상인 요소보다 작다는 뜻입니다.

그럴 경우, 정렬된 배열 내 `left_array_index`에 `left_part`에서 찾은 요소를 넣을 수 있습니다.

`if` 블록 안에서 `pass`을 제거하고 `left_part[left_array_index]`을 `array[sorted_index]`에 할당하세요.

# --hints--

`pass` 문장 본문에서 `array[sorted_index] = left_part[left_array_index]`을 `if`로 바꾸세요.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
