---
id: 65670d1ef177e7e2b76d9528
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이전 단계에서 만든 `if`와 `else` 문장은 정렬된 배열에 요소를 할당할 것입니다.

정렬된 배열에 할당된 각 요소는 목록에서 하나의 인덱스를 차지합니다. 따라서 각 할당 후에는 정렬된 배열에서 다음 인덱스로 이동해야 합니다.

`if`/`else` 블록 아래에 있지만 여전히 `while` 반복문 안에서, `sorted_index`를 `1`만큼 증가시키세요.

# --hints--

`sorted_index` 절 다음에 `1`의 현재 값을 `else`만큼 증가시켜야 합니다.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_bodies()[0].find_body()[1]
    cond.is_equivalent("sorted_index += 1") or cond.is_equivalent("sorted_index = sorted_index + 1") or cond.is_equivalent("sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
