---
id: 656710d1e0ec62253426db24
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

생성한 `while` 반복문은 `left_part`에서 하나의 요소를 `right_part`의 다른 요소와 비교한 다음, 더 작은 요소를 주 `array` 목록에 추가합니다.

비교할 요소가 더 이상 없을 때까지 이 작업을 계속합니다. 하지만 `left_part`에는 여전히 요소가 남아 있고 `right_part`에는 없을 수도 있고, 그 반대일 수도 있습니다.

`while`이 `left_array_index`보다 작을 때 실행되는 또 다른 `len(left_part)` 반복문을 만드세요. 다음 단계에서 `left_part`에 남아 있는 요소를 `array` 목록에 복사하는 데 사용할 것입니다.

# --hints--

`while` 반복문의 조건은 `left_array_index < len(left_part)`이어야 합니다. 반복문 본문에 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

`pass` 반복문 본문에 `while` 키워드가 있어야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--
    
--fcc-editable-region--
```
