---
id: 656709e50ed928da35260276
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이전 단계에서, 더 작은 요소였기 때문에 `left_part`의 요소를 `array` 목록에 할당했습니다. 하지만 항상 그런 것은 아닙니다. 비교하는 경우에 따라 오른쪽 요소가 더 작을 수도 있습니다.

`else`의 요소가 `left_part`의 요소보다 작지 않을 때 실행할 `right_part` 절을 만드세요.

`else` 블록 안에서 `right_part[right_array_index]`를 `array[sorted_index]`에 할당하세요.

# --hints--

`array[sorted_index] = right_part[right_array_index]` 절 본문에 `else`가 있어야 합니다.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
