---
id: 656702f8b4cbd8cbf0a433c6
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

다음으로, `left_part`와 `right_part`에 요소가 남아 있는 동안 계속되는 반복문을 만들어야 합니다.

이를 위해 두 가지 조건을 가진 `while` 반복문을 만드세요: 하나는 `left_array_index`가 `left_part`의 길이보다 작은지 확인하는 조건이고, 다른 하나는 `right_array_index`가 `right_part`의 길이보다 작은지 확인하는 조건입니다.

# --hints--

`while` 반복문은 `left_array_index`가 `left_part`의 길이보다 작은지, 그리고 `right_array_index`가 `right_part`의 길이보다 작은지 확인해야 합니다.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

반복문 본문에 `pass` 키워드를 사용하세요.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
