---
id: 6569b20f829b7e69d43c232a
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`while` 반복문에서 마지막으로 해야 할 일은 정렬된 배열에서 다음 인덱스로 이동하는 것입니다.

`1` 값에 `sorted_index`을 더하세요.

# --hints--

`sorted_index` 반복문 내에서 현재 `while` 값을 1만큼 증가시켜야 합니다. 왼쪽 부분 요소를 정렬된 배열에 할당한 후에 증가시켜야 한다는 점을 기억하세요.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
