---
id: 6569b19d31a09b65b4bc390a
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`while` 반복문 안에서 `left_array_index` 값을 1만큼 증가시키세요.

# --hints--

`left_array_index` 반복문 안에서 현재 `while` 값을 1만큼 증가시켜야 합니다. 증가 작업은 왼쪽 부분 요소를 정렬된 배열에 할당한 후에 해야 한다는 점을 기억하세요.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index = left_array_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "left_array_index = 1 + left_array_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
