---
id: 6569b4e0bd29d17d4b53b16c
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

이제 `while`에 대해 동일한 `right_part` 반복문 로직을 복제할 차례입니다.

`while`가 `right_array_index`보다 작을 때 실행되는 `len(right_part)` 반복문을 만드세요. `while` 반복문 내에서 `right_part[right_array_index]`에 `array[sorted_index]`을 할당하세요. 마지막에 `right_array_index`과 `sorted_index`를 `1`만큼 증가시키세요.

# --hints--

`while` 반복문은 `right_array_index`가 `len(right_part)`보다 작은지 확인해야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

반복문 본문에서 `right_part[right_array_index]`에 `array[sorted_index]`를 할당해야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

`right_array_index` 반복문 내에서 `1`을 `while`만큼 증가시켜야 합니다. 오른쪽 부분 요소를 정렬된 배열에 할당한 후에 증가시켜야 한다는 점을 유의하세요.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

`sorted_index` 반복문 내에서 `1`의 현재 값을 `while`만큼 증가시켜야 합니다. 오른쪽 부분 요소를 정렬된 배열에 할당한 후에 증가시켜야 한다는 점을 유의하세요.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
