---
id: 6577254891048c8f2c19e961
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

함수 본문 맨 아래에서 `merge_sort()` 함수를 다시 호출하세요. 이번에는 함수 호출에 `right_part`을 인수로 전달하세요.

# --hints--

함수 본문 맨 아래에서 `merge_sort()`를 인수 `right_part`와 함께 호출해야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
