---
id: 657823a9f4f372518614c8b7
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

여전히 `else` 블록 안에서 `right_array_index`을 `1`만큼 증가시키세요.

# --hints--

정렬된 배열에 요소를 삽입한 후 `1`의 현재 값에 `right_array_index`을 더해야 합니다.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[1]

    cond.is_equivalent("right_array_index += 1") or cond.is_equivalent("right_array_index = right_array_index + 1") or cond.is_equivalent("right_array_index = 1 + right_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        
--fcc-editable-region--
        else:
            array[sorted_index] = right_part[right_array_index]

--fcc-editable-region--
```
