---
id: 657f3a661730891aa64f3568
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

슬라이스 구문을 사용해 `array`의 오른쪽 절반을 추출하고 `right_part`이라는 변수에 할당하세요.

# --hints--

`right_part`라는 변수가 있어야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

슬라이스 구문을 사용해 배열의 오른쪽 절반을 추출하고 `right_part` 변수에 할당해야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
