---
id: 6720b7f3155a9278d6b6a8b8
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在是時候為你的賀卡設計樣式。

為你的 `body` 元素添加一個選擇器，然後： 

- 將 `font-family` 更改為先是 `Arial`，接著是泛型的 `sans-serif`，
- 在所有邊緣給予 `40px` 的內距，
- 將 `text-align` 屬性設為 `center`。

# --hints--

你應該有一個 `body` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

你應該在 `body` 選擇器中將 `font-family` 設定為 `"Arial", sans-serif`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

你應該在 `body` 選擇器中將 `padding` 設為 `40px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

你應該在 `body` 選擇器中將 `text-align` 設定為 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
