---
id: 6720b3191bbafd269bae0ae8
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

創建一個具有 `id` 為 `greeting-card` 且 `class` 為 `card` 的 `div` 元素。

在 `div` 元素內，為你的 `h1` 添加文字 `Happy Birthday!`。然後添加一個段落元素，並為它添加名為 `message` 的 `class`，以及文字 `Wishing you all the happiness and joy on your special day!`。

# --hints--

你應該有一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

`div` 元素應該有一個 `id` 為 `greeting-card`。

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

`div` 元素應該有一個 `class` 為 `card`。

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

你應該在 `div` 元素內有一個 `h1` 元素。

```js
assert.exists(document.querySelector("div > h1"));
```

`h1` 元素應該包含文字 `Happy Birthday!`。

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

你應該在 `div` 裡面有一個 `p` 元素。

```js
assert.exists(document.querySelector("div > p"));
```

`p` 元素應該有一個 `class` 為 `message`。

```js
assert.exists(document.querySelector("div > p.message"));
```

`p` 元素應該包含文字 `Wishing you all the happiness and joy on your special day!`。

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
