---
id: 6720b1f2680ec706ed534acf
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將透過設計賀卡來練習使用不同的偽類別和偽元素。已為你提供了 HTML 樣板程式碼。 

開始工作坊時，請連結 `styles.css` 檔案。

# --hints--

你的 `link` 元素應該位於你的 `head` 元素內。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
