---
id: 6729005c87b6320356ac1140
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

`.card` 元素需要更多樣式：為你的 `.card` 添加 `border-radius` 為 `10px`，以及 `box-shadow`，值為 `0 4px 8px gray`。

# --hints--

`.card` 選擇器應該有一個 `border-radius` 屬性，值為 `10px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("border-radius"), "10px");
```

`.card` 選擇器應該有一個屬性 `box-shadow`，其值為 `0 4px 8px gray`。

```js
assert.strictEqual(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("box-shadow"),
  "gray 0px 4px 8px" // this is the value outputted by getPropertyValue
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
