---
id: 672b4434dd6d6a1edc405419
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

在 `.card-links` 元素內添加兩個 `a` 元素。

第一個應該有 `Send Card` 的文字、`send-link` 的類別和 `#send` 的 `href`。

第二個應該有 `Share on Social Media` 的文字、`share-link` 的 `class` 和 `#share` 的 `href`。

# --hints--

`.card-links` 內的第一個 `a` 元素應該有 `Send Card` 的文字。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

第一個 `a` 元素應該有一個值為 `send-link` 的 `class`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

第一個 `a` 元素應該有一個值為 `#send` 的 `href`。

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

`.card-links` 內的第二個 `a` 元素應該有 `Share on Social Media` 的文字。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

第二個 `a` 元素應該有一個 `class`，其值為 `share-link`。

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

第二個 `a` 元素應該有一個值為 `#share` 的 `href`。

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
