---
id: 672b42241c0c6515f26786f4
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

現在在 `.message` 元素下方添加一個新的 `div`。新的 `div` 應該有一個 `class` 屬性為 `card-links`。

# --hints--

`#greeting-card` 內的最後一個元素應該是 `div`。

```js
assert.strictEqual(document.querySelector("#greeting-card").lastElementChild.tagName, "DIV");
```

新的 `div` 應該有一個 `class` 為 `card-links`。

```js
assert.strictEqual(document.querySelector("#greeting-card").lastElementChild.className, "card-links");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      --fcc-editable-region--
      
      --fcc-editable-region--
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
