---
id: 672e410ddd98b409594890b6
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

新增兩個 `section` 元素，依序排列。第一個應該有一個 `id` 為 `send`，第二個應該有一個 `id` 為 `share`。

# --hints--

你應該有兩個 `section` 元素。

```js
assert.lengthOf(document.querySelectorAll("section"), 2);
```

第一個 `section` 應該有一個 `id` 為 `send`。

```js
assert.strictEqual(document.querySelectorAll("section")[0].id, "send");
```

第二個 `section` 應該有一個 `id` 為 `share`。

```js
assert.strictEqual(document.querySelectorAll("section")[1].id, "share");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
