---
id: 6733352ce2a5642827b4f7f0
title: 第 18 步
challengeType: 0
dashedName: step-18
---

# --description--

為你的 `.card-links` 選擇器添加，並將 `margin-top` 屬性設為 `20px`。

你可以為元素添加 `display: flex` 來設定使用 flexbox，你將在課程後面學習更多關於 flexbox 的內容，你可以將這視為一個小預覽。

為了讓兩個聯結周圍有相同的間距，為你的 `display` 屬性添加 `flex`，並將 `justify-content` 設定為 `space-around`。

# --hints--

你應該有一個 `.card-links` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

`.card-links` 選擇器應該設定 `margin-top` 屬性為 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

`.card-links` 選擇器應該有一個設定為 `flex` 的 `display` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

`.card-links` 選擇器應該設定 `justify-content` 屬性為 `space-around`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
