---
id: 67333c34e357d34620217615
title: 第 25 步
challengeType: 0
dashedName: step-25
---

# --description--

為 `section` 元素創建選擇器，並給予它們：

- 設定為 `20px auto` 的 `margin` 屬性，
- 設定為 `600px` 的 `max-width`。
- 設定為 `whitesmoke` 的 `background-color` 屬性
- 設定為 `20px` 的 `padding` 屬性
- 設定為 `10px` 的 `border-radius` 屬性

# --hints--

你應該有一個 `section` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("section"));
```

`section` 選擇器應該設定 `margin` 屬性為 `20px auto`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("section").getPropertyValue("margin"), ["20px auto", "20px auto 20px", "20px auto 20px auto"]);
```

`section` 選擇器應該設定 `max-width` 屬性為 `600px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("section").getPropertyValue("max-width"), "600px");
```

`section` 選擇器應該設定 `background-color` 屬性為 `whitesmoke`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("section")?.getPropertyValue("background-color").toLowerCase(), ["whitesmoke", "#f5f5f5", "rgb(245, 245, 245)", "hsl(0, 0%, 96%)"]);
```

`section` 選擇器應該設定 `padding` 屬性為 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("section").getPropertyValue("padding"), "20px");
```

`section` 選擇器應該設定 `border-radius` 屬性為 `10px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("section").getPropertyValue("border-radius"), "10px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

.card-links {
  margin-top: 20px;
  display: flex;
  justify-content: space-around;
}

.card-links a {
  text-decoration: none;
  font-size: 1em;
  padding: 10px 20px;
  border-radius: 5px;
  color: white;
  background-color: midnightblue;
  transition: background-color 0.3s ease;
}

.card-links a:hover {
  background-color: orangered;
}

.card-links a:active {
  background-color: midnightblue;
}

.card-links a:focus {
  outline: 2px solid yellow;
}

.card-links a:visited {
  color: crimson;
}

--fcc-editable-region--

--fcc-editable-region--
```
