---
id: 67333082ab35e414eb4ea70d
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

當 `a` 元素被暫留時，背景的顏色會過渡到不同的顏色。你可以使用 `transition` 屬性來調整該過渡的方式：

```css
a {
  transition: color 1s linear;
}
```

`transition` 屬性接受的值依序為，過渡所套用的屬性、過渡的持續時間，然後是時間函數。

如果有多個屬性具有過渡效果，你可以用逗號分隔為每個屬性寫入值：

```css
p {
  transition: property1 0.1s, property2 0.6s linear;
}
```

如果省略了某個值，例如第一個屬性的時間，則會套用預設值。

為你的 `.card` 選擇器添加 `transition: transform 0.3s, background-color 0.3s ease`。

試試看，暫留過渡已完成。

# --hints--

`.card` 選擇器應該有一個屬性 `transition` 設定為 `transform 0.3s, background-color 0.3s ease`。

```js
assert.oneOf(
  new __helpers.CSSHelp(document).getStyle(".card")?.getPropertyValue("transition"),
  [
    "transform 0.3s ease 0s, background-color 0.3s ease 0s", // this is what comes out of `getPropertyValue`
    "transform 0.3s, background-color 0.3s" // different environment got this instead
  ]
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
--fcc-editable-region--
  
--fcc-editable-region--
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

```
