---
id: 61537bb9b1a29430ac15ad38
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل عنصر `head`، أضف عنصر `title` مع النص `Photo Gallery`، وعنصر `link` لإضافة ملف `styles.css` إلى الصفحة.

# --hints--

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

يجب أن يحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لملف المشروع عنوان `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

تذكر، أن حالة الأحرف والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
