---
id: 61537a8054753e2f1f2a1574
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `head`، أضف وسم `meta` مع تعيين الخاصية `name` إلى `viewport` والخاصية `content` إلى `width=device-width, initial-scale=1`.

أضف أيضًا وسم `meta` مع تعيين الخاصية `charset` إلى `UTF-8`.

# --hints--

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

يجب أن يحتوي أحد عناصر `meta` على الخاصية `name` المعينة إلى `viewport`، والخاصية `content` المعينة إلى `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي عنصر `meta` الآخر على الخاصية `charset` المعينة إلى `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
