---
id: 61537c9eecea6a335db6da79
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أسفل عنصر `.header`، أنشئ عنصر `div` جديد وعيّن له `class` بقيمة `gallery`. سيعمل هذا العنصر `div` كحاوية لصور المعرض.

داخل عنصر `.gallery` هذا، أنشئ تسعة عناصر `img`.

# --hints--

يجب أن تنشئ عنصر `div` داخل عنصر `body` الخاص بك.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

يجب أن يحتوي عنصر `div` الجديد على `class` بقيمة `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

يجب أن يأتي عنصر `div` الجديد بعد عنصر `.header` الخاص بك.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

يجب أن يحتوي عنصر `.gallery` على تسعة عناصر `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
