---
id: 6153893900438b4643590367
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

قم بمحاذاة نص `.header` في الوسط. اجعل النص بأحرف كبيرة باستخدام خاصية `text-transform` مع القيمة `uppercase`.

أضف له حشو `32px` من جميع الجهات. اضبط الخلفية على `#0a0a23` والنص على اللون `#fff`.

أضف `border-bottom` بقيمة `4px solid #fdb347`.

# --hints--

يجب أن يكون لديك محدِّد `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

يجب أن يحتوي محدِّد `.header` على خاصية `text-align` بقيمة `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

يجب أن يحتوي محدِّد `.header` على خاصية `text-transform` بقيمة `uppercase`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

يجب أن يحتوي محدِّد `.header` على خاصية `padding` بقيمة `32px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

يجب أن يحتوي محدِّد `.header` على خاصية `background-color` بقيمة `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

يجب أن يحتوي محدِّد `.header` على خاصية `color` بقيمة `#fff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

يجب أن يحتوي محدِّد `.header` على خاصية `border-bottom` بقيمة `4px solid #fdb347`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
