---
id: 61537485c4f2a624f18d7794
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ بقالب HTML الجاهز القياسي الخاص بك. أضف إعلان `DOCTYPE`، وعنصر `html`، وعنصر `head`، وعنصر `body`.

أضف الخاصية `lang` إلى وسم البداية `<html>` مع تعيين القيمة إلى `en`.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب أن تضع مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب أن تعرف `html` كنوع المستند.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم بداية.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يحتوي وسم البداية `<html>` على الخاصية `lang` مع القيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك وسم بداية `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم بداية `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عنصرا `head` و`body` أخوين (متجاورين).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
