---
id: 6494da0daf5df5197963671d
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

لاحظ كيف تمتد حدود الصورة الزرقاء خارج حدود معرض الصور الحمراء. هذا بسبب الطريقة التي تحسب بها المتصفحات حجم عناصر الحاوية.

تُستخدم خاصية `box-sizing` لضبط هذا السلوك. افتراضيًا، يُستخدم نموذج `content-box`. مع هذا النموذج، عندما يكون للعنصر عرض محدد، يُحسب هذا العرض بناءً فقط على محتوى العنصر. تُضاف قيم الحشو والحدود إلى العرض الكلي، فيتوسع العنصر لاستيعاب هذه القيم.

جرّب تعيين `box-sizing` إلى `content-box` صراحةً باستخدام المحدد العالمي `*`. في هذه المرحلة، لن ترى أي تغييرات، لأنك تستخدم القيمة الافتراضية.

# --hints--

يجب أن يكون لديك محدد `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

يجب أن يحتوي محدد `*` على خاصية `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

يجب أن تكون خاصية `box-sizing` في محدد `*` مضبوطة على `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
