---
id: 64dd076a1b33c86d84b4232f
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

لرؤية حجم العناصر بشكل أفضل، إضافة حدود يمكن أن تكون مفيدة.

أعطِ عنصر `.gallery` خاصية `width` بقيمة `50%` وحدودًا (`border`) بقيمة `5px solid red`.

ثم أعطِ عناصر `img` خاصية `width` بقيمة `100%`، و`padding` بقيمة `5px`، وحدودًا (`border`) بقيمة `5px solid blue`.

# --hints--

يجب أن يكون لديك محدِّد `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

يجب أن يحتوي محدِّد `.gallery` على خاصية `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

يجب أن تكون خاصية `width` في محدِّد `.gallery` مضبوطة على `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

يجب أن يحتوي محدِّد `.gallery` على خاصية `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

يجب أن تكون خاصية `border` في محدِّد `.gallery` مضبوطة على `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

يجب أن يكون لديك محدِّد `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

يجب أن يحتوي محدِّد `img` على خاصية `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

يجب أن تكون خاصية `width` في محدِّد `img` مضبوطة على `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

يجب أن يحتوي محدِّد `img` على خاصية `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

يجب أن تكون خاصية `padding` في محدِّد `img` مضبوطة على `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

يجب أن يحتوي محدِّد `img` على خاصية `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

يجب أن تكون خاصية `border` في محدِّد `img` مضبوطة على `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
