---
id: 67116c175a77cb64c5c10d49
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na vitendakazi vya ngazi ya juu kama `map`, `filter` na `reduce`.

Katika warsha hii, utajenga programu ya msimamizi wa maktaba ambayo itakupa fursa ya kufanya mazoezi ya kutumia vitendakazi mbalimbali vya ngazi ya juu.

Ili kuanza, tengeneza kigezo kinachoitwa `library` na ukaweke thamani ya safu ya data tupu.

# --hints--

Unapaswa kutengeneza safu ya data `library`.

```js
assert.isArray(library)
```

Safu yako ya data `library` inapaswa kuwa tupu.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
