---
id: 67116d7584d0b469b14579bf
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Ndani ya safu ya data `library`, tengeneza kitu chenye sifa na thamani zifuatazo:

| Sifa | Thamani   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Unapaswa kuwa na kitu ndani ya safu ya data `library`.

```js
assert.isObject(library[0])
```

Kitu chako kinapaswa kuwa na sifa ya `title`.

```js
assert.property(library[0], "title")
```

Thamani ya sifa yako ya `title` inapaswa kuwa `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Kitu chako kinapaswa kuwa na sifa ya `author`.

```js
assert.property(library[0], "author")
```

Sifa yako ya `author` inapaswa kuwa na thamani `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Kitu chako kinapaswa kuwa na sifa ya `about`.

```js
assert.property(library[0], "about")
```

Sifa yako ya `about` inapaswa kuwa na thamani `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Kitu chako kinapaswa kuwa na sifa ya `pages`.

```js
assert.property(library[0], "pages")
```

Thamani ya sifa yako ya `pages` inapaswa kuwa nambari.

```js
assert.isNumber(library[0]?.pages)
```

Sifa yako ya `pages` inapaswa kuwekwa kuwa `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
