---
id: 67116d7584d0b469b14579c0
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Unda kitu kingine ndani ya safu ya data `library` chenye sifa na thamani zifuatazo:

| Sifa | Thamani   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Unapaswa kuwa na kitu kama kipengee cha pili cha safu ya data `library`.

```js
assert.isObject(library[1]);
```

Kitu chako kinapaswa kuwa na sifa ya `title`.

```js
assert.property(library[1], "title")
```

Sifa yako ya `title` inapaswa kuwa na `"Atomic Habits"` kama thamani yake.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Kitu chako kinapaswa kuwa na sifa ya `author`.

```js
assert.property(library[1], "author")
```

Sifa ya `author` inapaswa kuwa na `"James Clear"` kama thamani yake.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Kitu chako kinapaswa kuwa na sifa ya `about`.

```js
assert.property(library[1], "about")
```

Sifa ya `about` inapaswa kuwa na thamani `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Kitu chako kinapaswa kuwa na sifa ya `pages`.

```js
assert.property(library[1], "pages")
```

Thamani ya sifa yako ya `pages` inapaswa kuwa nambari.

```js
assert.isNumber(library[1]?.pages)
```

Sifa yako ya `pages` inapaswa kuwekwa kuwa `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
