---
id: 587d8255367417b2b2512c74
title: 創建優先權佇列類別
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

在這個挑戰中你將會建立一個優先權佇列。優先權佇列是一種特殊型別的佇列，其中項目可能具有額外的資訊來指定它們的優先權。這可以簡單地用整數來表述。項目的優先權會改寫項目出列的順序。如果一個具有較高優先權的項目在較低優先權的項目之後被加入佇列，則該較高優先權的項目會在所有其他項目之前被出列。

例如，假設我們有一個包含三個項目的優先權佇列：

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

這裡第二個值（一個整數）表述條款優先權。如果我們以優先權 `1`（假設較低的優先權具有優先）將 `['human', 1]` 加入佇列，它將成為第一個被取出的條款。群集將會是這樣：

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

我們已經開始在程式碼編輯器中撰寫 `PriorityQueue`。你需要為添加具有優先權的項目添加一個 `enqueue` 方法，為移除並傳回項目添加一個 `dequeue` 方法，為傳回佇列中項目數量添加一個 `size` 方法，為傳回佇列前端元素添加一個 `front` 方法，最後添加一個 `isEmpty` 方法，當佇列為空時會傳回 `true`，否則傳回 `false`。

`enqueue` 應該接受上述格式的條款（`['human', 1]`），其中 `1` 表示優先權。`dequeue` 和 `front` 應該只傳回條款的名稱，而非其優先權。

# --hints--

你的 `PriorityQueue` 類別應該有一個 `enqueue` 方法。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

你的 `PriorityQueue` 類別應該有一個 `dequeue` 方法。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

你的 `PriorityQueue` 類別應該有一個 `size` 方法。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

你的 `PriorityQueue` 類別應該有一個 `front` 方法。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

你的 `PriorityQueue` 類別應該有一個 `isEmpty` 方法。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

你的 `PriorityQueue` 類別應該在項目被加入佇列和移出佇列時，透過 `size` 方法正確追蹤目前的項目數量。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

`front` 方法應該在項目被加入佇列和移出佇列時，傳回佇列前端的正確項目。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

當佇列為空時，`isEmpty` 方法應該傳回 `true`。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

優先權佇列應該先傳回優先權較高的項目，再傳回優先權較低的項目，否則以先進先出順序傳回項目。

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```
