---
id: 587d8250367417b2b2512c60
title: 創建一個佇列類別
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

像堆疊一樣，佇列是一個元素的群集。但與堆疊不同，佇列遵循 FIFO（先進先出）原則。添加到佇列的元素會被推到佇列的尾端或末端，且只有佇列前端的元素可以被移除。

我們可以使用陣列來表述佇列，但就像堆疊一樣，我們想要限制對我們的佇列的控制元件數量。

佇列類別的兩個主要方法是 `enqueue` 方法和 `dequeue` 方法。`enqueue` 方法將元素推入佇列的尾端，`dequeue` 方法則移除並傳回佇列前端的元素。其他有用的方法有 `front`、`size` 和 `isEmpty` 方法。

# --instructions--

撰寫一個 `enqueue` 方法，將元素推入佇列尾端，一個 `dequeue` 方法，移除並傳回前端元素，一個 `front` 方法，讓我們查看前端元素，一個 `size` 方法，顯示長度，以及一個 `isEmpty` 方法，用來檢查佇列是否為空。

# --hints--

你的 `Queue` 類別應該有一個 `enqueue` 方法。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

你的 `Queue` 類別應該有一個 `dequeue` 方法。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

你的 `Queue` 類別應該有一個 `front` 方法。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

你的 `Queue` 類別應該有一個 `size` 方法。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

你的 `Queue` 類別應該有一個 `isEmpty` 方法。

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

`dequeue` 方法應該移除並傳回佇列的前端元素

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

`front` 方法應該傳回佇列前端元素的返回值

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

`size` 方法應該傳回佇列的長度

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

如果佇列中有元素，`isEmpty` 方法應該傳回 `false`。

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```
