---
id: 587d8254367417b2b2512c70
title: 在 ES6 中創建並添加到 Sets
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

既然你已經完成了 ES5，接下來你將在 ES6 中執行類似的操作。這將容易許多。ES6 內建了資料結構 `Set`，因此你之前手寫的許多操作現在都已包含在內。讓我們來看看：

要創建一個新的空集合：

```js
var set = new Set();
```

你可以創建一組帶有值：

```js
var set = new Set(1);
```

你可以用陣列創建一個集合：

```js
var set = new Set([1, 2, 3]);
```

一旦你創建了一個集合，你就可以使用 `add` 方法添加你想要的值：

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

提醒你，集合是一種不能包含重複值的資料結構：

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set contains [1, 2, 3] only
```

# --instructions--

對於這個練習，傳回一個包含以下值的集合：`1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

你的 `Set` 應該只包含值 `1, 2, 3, Taco, Cat, Awesome`。

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```
