---
id: 587d8255367417b2b2512c73
title: 使用 Spread 和 Notes 進行 ES5 `Set()` 整合
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

你還記得 ES6 展開運算子 `...` 嗎？

`...` 可以接受 ES6 中的可迭代物件並將它們轉換成陣列。

讓我們創建一個 Set，並查看 spread 函式（程式）。

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

在這個練習中，我們將傳遞一個 Set 物件給 `checkSet` 函式。它應該傳回一個包含 Set 值的陣列。

現在你已成功學會如何使用 ES6 的 `Set()` 物件，做得好！

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` 應該傳回 `[1, 2, 3, 4, 5, 6, 7]`。

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```
