---
id: 587d825a367417b2b2512c88
title: 反轉雙向鏈結串列
challengeType: 1
forumTopicId: 301714
dashedName: reverse-a-doubly-linked-list
---

# --description--

讓我們為我們的雙向鏈結串列再創建一個名為 reverse 的方法，該方法會就地反轉列表。當方法執行完畢後，head 應該指向先前的 tail，而 tail 應該指向先前的 head。現在，如果我們從 head 巡訪到 tail，應該會以與原始列表相反的順序遇到節點。嘗試反轉空列表應該傳回 null。

# --hints--

`DoublyLinkedList` 資料結構應該存在。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` 應該有一個名為 `reverse` 的方法。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.reverse == undefined) {
      return false;
    }
    return typeof test.reverse == 'function';
  })()
);
```

反轉空的列表應該傳回 `null`。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.reverse() == null;
  })()
);
```

`reverse` 方法應該反轉列表。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    var n1 = new Node(58, null);
    var n2 = new Node(61, n1);
    var n3 = new Node(32, n2);
    var n4 = new Node(95, n3);
    var n5 = new Node(41, n4);
    n1.next = n2;
    n2.next = n3;
    n3.next = n4;
    n4.next = n5;
    test.head = n1;
    test.tail = n5;
    test.reverse();

    var result = [];
    var current = test.head;
    while (current) {
      result.push(current.data);
      current = current.next;
    }

    return result.join('') == '4195326158';
  })()
);
```

當 `list` 被反轉時，`next` 和 `previous` 的引用應該被正確維護。

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    var n1 = new Node(11, null);
    var n2 = new Node(22, n1);
    var n3 = new Node(33, n2);
    var n4 = new Node(44, n3);
    var n5 = new Node(55, n4);
    n1.next = n2;
    n2.next = n3;
    n3.next = n4;
    n4.next = n5;
    test.head = n1;
    test.tail = n5;
    test.reverse();

    var result = [];
    var current = test.tail;
    while (current) {
      result.push(current.data);
      current = current.prev;
    }

    return result.join('') == '1122334455';
  })()
);
```

# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
  var Node = function(data, prev) {
    this.data = data;
    this.prev = prev;
    this.next = null;
  };
  var DoublyLinkedList = function() {
    this.head = null;
    this.tail = null;

    this.reverse = function() {
      if (!this.head || !this.head.next) {
        return this.head
      }

      let tail;
      let temp;
      let current = this.head;
      while(current !== null) {
        if(!tail) tail = current;
        temp = current.prev;
        current.prev = current.next;
        current.next = temp;
        current = current.prev;
      }

      this.head = temp.prev;
      this.tail = tail
    }
  };
```
