---
id: 587d8253367417b2b2512c6a
title: 型別陣列
challengeType: 1
forumTopicId: 301716
dashedName: typed-arrays
---

# --description--

陣列是可以容納許多不同元素的 JavaScript 物件。

```js
var complexArr = [1, 5, "2", "Word", {"name": "James"}];
```

基本上在背景中發生的是你的瀏覽器會自動為該陣列分配適當的記憶體空間。如果你添加或移除資料，它也會根據需要進行變更。

然而，在高效率和不同元素型別的世界中，有時你需要更具體地指定給陣列多少記憶體。

<dfn>Typed arrays</dfn> 是這個問題的解決方案。你現在可以指定要給陣列多少記憶體。以下是可用的不同陣列型別的基本概述，以及該陣列中每個元素的位元組大小。

<table><tbody><tr><th>Type</th><th>Each element size in bytes</th></tr><tr><td><code>Int8Array</code></td><td>1</td></tr><tr><td><code>Uint8Array</code></td><td>1</td></tr><tr><td><code>Uint8ClampedArray</code></td><td>1</td></tr><tr><td><code>Int16Array</code></td><td>2</td></tr><tr><td><code>Uint16Array</code></td><td>2</td></tr><tr><td><code>Int32Array</code></td><td>4</td></tr><tr><td><code>Uint32Array</code></td><td>4</td></tr><tr><td><code>Float32Array</code></td><td>4</td></tr><tr><td><code>Float64Array</code></td><td>8</td></tr></tbody></table>

創建這種類型的陣列有兩種方法。一種方法是直接創建。以下是如何創建長度為 3 的 `Int16Array`。

```js
var i8 = new Int16Array(3);
console.log(i8);
// Returns [0, 0, 0]
```

你也可以創建一個 <dfn>buffer</dfn> 來指定你想讓陣列佔用多少資料（以位元組為單位）。**注意**要使用 buffer 創建型別陣列，你需要指定的位元組數必須是上述位元組數的倍數。

```js
// Create same Int16Array array differently
var byteSize = 6; // Needs to be multiple of 2
var buffer = new ArrayBuffer(byteSize);
var i8View = new Int16Array(buffer);
buffer.byteLength; // Returns 6
i8View.byteLength; // Returns 6
console.log(i8View); // Returns [0, 0, 0]
```

<dfn>Buffers</dfn> 是僅用於攜帶資料的一般用途物件。你無法正常存取它們。若要存取它們，你需要先創建一個 <dfn>view</dfn>。

```js
i8View[0] = 42;
console.log(i8View); // Returns [42, 0, 0]
```

**注意** 具型別陣列沒有傳統陣列所具有的某些方法，例如 `.pop()` 或 `.push()`。具型別陣列也會讓檢查某物是否為陣列的 `Array.isArray()` 回傳失敗。雖然較簡單，但這對於較不複雜的 JavaScript 引擎來說，實作它們可能是一種優勢。

# --instructions--

首先創建一個 64 位元組的 `buffer`。然後創建一個名為 `i32View` 的 `Int32Array` 型別陣列來檢視它。

# --hints--

你的 `buffer` 應該是 64 位元組 大小。

```js
assert(buffer.byteLength === 64);
```

你的 `i32View` 視圖應該是 64 位元組大小。

```js
assert(i32View.byteLength === 64);
```

你的 `i32View` 視圖應該有 16 個元素長。

```js
assert(i32View.length === 16);
```

# --seed--

## --seed-contents--

```js
var buffer;
var i32View;
```

# --solutions--

```js
var buffer = new ArrayBuffer(64);
var i32View = new Int32Array(buffer);
```
