---
id: 587d825b367417b2b2512c8d
title: 創建一個 ES6 JavaScript Map
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

新版的 JavaScript 提供了內建的 `Map` 物件，該物件提供了我們在上一個挑戰中手動撰寫的許多功能。這個 `Map` 物件雖然與一般的 JavaScript 物件相似，但提供了一些普通物件所缺乏的有用功能。 例如，ES6 的 `Map` 會追蹤加入其中的項目的插入順序。以下是其方法的更完整概述： 

- `.has(key)` 根據鍵的存在傳回 true 或 false
- `.get(key)` 傳回與鍵相關的值
- `.set(key, value)` 設定一個新的 key、value 配對。
- `.delete(key)` 移除一個鍵、值配對
- `.clear()` 會移除所有鍵、值配對。
- `.keys()` 傳回一個新的 map 迭代器物件，該物件包含所有以插入順序排列的鍵。
- `.values()` 傳回一個新的 map 迭代器物件，該物件包含所有依插入順序排列的值。
- `.entries()` 傳回一個新的 map 迭代器物件，該物件包含所有依插入順序排列的鍵、值配對。

# --instructions--

定義一個 JavaScript Map 物件並指定給一個名為 myMap 的變數。為它添加鍵、值配對 `freeCodeCamp`、`Awesome!`。

# --hints--

`myMap` 物件應該存在。

```js
assert(typeof myMap === 'object');
```

`myMap` 應該包含鍵值對 `freeCodeCamp`、`Awesome!`。

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```
