---
id: 587d8251367417b2b2512c63
title: 從鏈結串列中移除元素
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

下一個任何鏈結串列實作都需要的重要方法是 `remove` 方法。此方法應該接受我們想要移除的元素作為引數，然後搜尋列表以找到並移除包含該元素的節點。

每當我們從鏈結串列中移除節點時，重要的是不要意外地使串列的其餘部分成為孤立節點。回想每個節點的 `next` 屬性都指向串列中緊接其後的節點。如果我們要移除中間的元素，例如，我們會想確保從該元素前一個節點的 `next` 屬性連接到中間元素的 `next` 屬性（也就是串列中的下一個節點！）

這聽起來可能真的很混亂，所以讓我們回到 conga 線的例子，這樣我們就有一個良好的概念模型。想像你自己在 conga 線中，而你正前方的人離開了隊伍。剛離開隊伍的人不再把手放在隊伍中的任何人身上——而你也不再把手放在離開的人身上。你向前跨步，並把手放在你看到的下一個人身上。

如果我們想移除的元素是 `head` 元素，我們會將 `head` 重新指定為鏈結串列的第二個節點。

# --instructions--

撰寫一個 `remove` 方法，該方法接受一個元素並將其從鏈結串列中移除。

**注意：** 每次從鏈結串列中移除一個元素時，`length` 應該減少 1。

# --hints--

你的 `LinkedList` 類別應該有一個 `remove` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

當第一個節點被移除時，你的 `remove` 方法應該將 `head` 指定到第二個節點。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

你的 `remove` 方法應該在每移除一個節點時將鏈結串列的 `length` 減少 1。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

你的 `remove` 方法應該將被移除節點的前一個節點的引用重新指定到被移除節點的 `next` 引用。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

如果元素不存在於鏈結串列中，你的 `remove` 方法不應該改變鏈結串列。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```
