---
id: 587d8250367417b2b2512c5f
title: 創建一個 Stack 類別
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

在上一節中，我們討論了什麼是堆疊以及如何使用陣列來表現堆疊。在本節中，我們將創建我們自己的堆疊類別。雖然你可以使用陣列來創建堆疊，但有時限制我們對堆疊的控制量是最好的。除了 `push` 和 `pop` 方法外，堆疊還有其他有用的方法。讓我們為堆疊類別添加 `peek`、`isEmpty` 和 `clear` 方法。

# --instructions--

撰寫一個 `push` 方法，將元素推入堆疊頂端，一個 `pop` 方法，移除並傳回堆疊頂端的元素，一個 `peek` 方法，查看堆疊頂端的元素，一個 `isEmpty` 方法，檢查堆疊是否為空，以及一個 `clear` 方法，移除堆疊中的所有元素。通常堆疊沒有這個，但我們已經添加了一個 `print` 輔助方法，會在主控台列印群集。

# --hints--

你的 `Stack` 類別應該有一個 `push` 方法。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

你的 `Stack` 類別應該有一個 `pop` 方法。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

你的 `Stack` 類別應該有一個 `peek` 方法。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

你的 `Stack` 類別應該有一個 `isEmpty` 方法。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

你的 `Stack` 類別應該有一個 `clear` 方法。

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

`peek` 方法應該傳回堆疊的頂端元素

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

`pop` 方法應該移除並傳回堆疊頂端的元素

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

如果堆疊不包含任何元素，`isEmpty` 方法應該傳回 true。

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

`clear` 方法應該從堆疊中移除所有元素。

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```
