---
id: 587d8255367417b2b2512c72
title: 在 ES6 `Set` 上使用 `.has` 和 `.size`
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

讓我們來看看 ES6 `Set` 物件上可用的 `.has` 和 `.size` 方法。

首先，創建一個 ES6 Set

```js
var set = new Set([1,2,3]);
```

.has 方法會檢查該值是否包含在集合中。

```js
var hasTwo = set.has(2);
```

.size 方法會傳回一個整數，表述 Set 的大小

```js
var howBig = set.size;
```

# --instructions--

在這個練習中，我們將傳遞一個陣列和一個值給 `checkSet()` 函式（程式）。你的函式（程式）應該從陣列引數創建一個 ES6 集合。找出該集合是否包含值引數。找出集合的大小。並將這兩個值傳回一個陣列中。

# --hints--

`checkSet([4, 5, 6], 3)` 應該傳回 [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
