---
id: 587d8251367417b2b2512c62
title: 創建鏈結串列類別
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

讓我們創建一個 `linked list` 類別。每個鏈結串列都應該從幾個基本屬性開始：一個 `head`（你的列表中的第一個項目）和一個 `length`（你的列表中的項目數量）。有時你會看到包含 `tail` 作為列表最後一個元素的鏈結串列實作品，但目前我們只使用這兩個。每當我們為鏈結串列添加一個元素時，我們的 `length` 屬性應該增加一。

我們會想要有一種方法將項目添加到我們的鏈結串列，所以我們會想要創建的第一個方法是 `add` 方法。

如果我們的列表是空的，為我們的鏈結串列添加一個元素就非常簡單：我們只需將該元素包裝在 `Node` 類別中，然後將該節點指定給我們鏈結串列的 `head`。

但是如果我們的列表已經有一個或多個成員呢？我們如何將一個元素添加到列表中？回想一下，鏈結串列中的每個節點都有一個 `next` 屬性。要將節點添加到列表中，請找到列表中的最後一個節點，並將該最後節點的 `next` 屬性指向我們的新節點。（提示：當節點的 `next` 屬性為 `null` 時，你就知道已經到達鏈結串列的末端。）

# --instructions--

撰寫一個 add 方法，將你推入鏈結串列的第一個節點指定給 `head`；之後，每當新增節點時，每個節點都應由前一個節點的 `next` 屬性參考。

注意

每當將元素添加到鏈結串列時，你的列表的 `length` 應該增加 1。

# --hints--

你的 `LinkedList` 類別應該有一個 `add` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

你的 `LinkedList` 類別應該將 `head` 指定為新增的第一個節點。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

你在 `LinkedList` 類別中的前一個 `node` 應該要有對最新創建的 `node` 的引用。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

你的 `LinkedList` 類別的 `size` 應該與鏈結串列中節點的數量相等。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```
