---
id: 587d8251367417b2b2512c64
title: 在鏈結串列中搜尋
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

讓我們為我們的鏈結串列類別添加更多有用的方法。如果我們能像 `Stack` 和 `Queue` 類別一樣判斷我們的列表是否為空，豈不是很有用？

我們也應該能夠在我們的鏈結串列中找到特定的元素。遍歷資料結構是你會想要多加練習的事情！讓我們創建一個 `indexOf` 方法，它接受一個 `element` 作為引數，並傳回該元素在鏈結串列中的 `index`。如果在鏈結串列中找不到該元素，則傳回 `-1`。

我們也來實作一個相反的行為：一個 `elementAt` 方法，它接受一個 `index` 作為引數並傳回指定 `index` 的 `element`。如果找不到 `element`，則傳回 `undefined`。

# --instructions--

撰寫一個 `isEmpty` 方法來檢查鏈結串列是否為空，一個 `indexOf` 方法來傳回給定元素的 `index`，以及一個 `elementAt` 方法來傳回給定 `index` 的 `element`。

# --hints--

你的 `LinkedList` 類別應該有一個 `isEmpty` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

當鏈結串列中至少有一個元素時，你的 `isEmpty` 方法應該傳回 `false`。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

當鏈結串列中沒有元素時，你的 `isEmpty` 方法應該傳回 `true`。

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

你的 `LinkedList` 類別應該有一個 `indexOf` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

你的 `indexOf` 方法應該傳回在鏈結串列中找到的指定元素的索引。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

如果在鏈結串列中找不到給定的元素，你的 `indexOf` 方法應該傳回 `-1`。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

你的 `LinkedList` 類別應該有一個 `elementAt` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

你的 `elementAt` 方法應該傳回在鏈結串列中指定索引處找到的元素。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

如果在鏈結串列中未找到指定索引的元素，你的 `elementAt` 方法應該傳回 `undefined`。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```
