---
id: 587d8251367417b2b2512c61
title: 操作鏈結串列中的節點
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

你在計算機科學中會遇到的另一個常見資料結構是 <dfn>鏈結串列</dfn>。鏈結串列是一種線性資料群集，由稱為「節點」的資料元素組成，每個節點都指向下一個節點。鏈結串列中的每個 <dfn>節點</dfn> 包含兩個關鍵資訊：`element` 本身，以及指向下一個 `node` 的引用。

想像你正站在康茄舞隊伍中。你雙手放在隊伍中下一個人的身上，而你身後的人雙手放在你身上。你可以看到正前方的人，但他們擋住了你對前面其他人的視線。節點就像康茄舞隊伍中的一個人：他們知道自己是誰，且只能看到下一個人，但不會知道前面或後面其他人的存在。

# --instructions--

在我們的程式碼編輯器中，我們已經創建了兩個節點，`Kitten` 和 `Puppy`，並且我們已經手動將 `Kitten` 節點連接到 `Puppy` 節點。

創建一個 `Cat` 和 `Dog` 節點，並手動將它們添加到該線路中。

# --hints--

你的 `Puppy` 節點應該有一個對 `Cat` 節點的引用。

```js
assert(Puppy.next.element === 'Cat');
```

你的 `Cat` 節點應該有一個對 `Dog` 節點的引用。

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```
