---
id: 587d8251367417b2b2512c65
title: 透過索引從鏈結串列中移除元素
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

在我們繼續探討另一個資料結構之前，讓我們再練習幾個關於鏈結串列的最後位元。

讓我們撰寫一個 `removeAt` 方法，用來移除指定 `index` 的 `element`。該方法應該被稱為 `removeAt(index)`。要移除某個 `index` 的 `element`，我們需要在沿著鏈結串列移動時持續計數每個節點。

一種常見的技術用於迭代鏈結串列的元素，涉及一個 <dfn>「runner」</dfn>，或稱為哨兵，它「指向」你的程式碼正在比較的節點。在我們的例子中，從我們列表的 `head` 開始，我們使用一個從 `0` 開始的 `currentIndex` 變數。`currentIndex` 應該在每經過一個節點時遞增一。

就像我們的 `remove(element)` 方法，<a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">我們在先前的課程中涵蓋過</a>，當我們在 `removeAt(index)` 方法中移除節點時，需要小心不要讓串列的其餘部分成為孤兒。我們透過確保對被移除節點有引用的節點，改為引用下一個節點，來保持節點的連續性。

# --instructions--

撰寫一個 `removeAt(index)` 方法，該方法會移除並傳回位於指定 `index` 的節點。如果給定的 `index` 為負數，或大於或相等於鏈結串列的 `length`，該方法應傳回 `null`。

**注意：** 請記得保持 `currentIndex` 的計數。

# --hints--

你的 `LinkedList` 類別應該有一個 `removeAt` 方法。

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

你的 `removeAt` 方法應該將鏈結串列的 `length` 減少 1。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

你的 `removeAt` 方法應該從鏈結串列中移除指定索引處的元素。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

當鏈結串列中只有一個元素時，你的 `removeAt` 方法應該移除並傳回指定索引處的元素，並減少鏈結串列的長度。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

你的 `removeAt` 方法應該傳回被移除節點的元素。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

如果給定的索引小於 `0`，你的 `removeAt` 方法應該傳回 `null`，且鏈結串列不應該改變。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

如果給定的索引大於或相等於 `length`，你的 `removeAt` 方法應該傳回 `null`，且鏈結串列不應該改變。

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```
