---
id: 587d8255367417b2b2512c72
title: 在 ES6 集合上使用 .has 和 .size
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Let's look at the .has and .size methods available on the ES6 Set object.

首先，創建一個 ES6 集合

```js
var set = new Set([1,2,3]);
```

.has 方法將檢查該值是否包含在集合中。

```js
var hasTwo = set.has(2);
```

.size 方法將返回一個表示集合大小的整數

```js
var howBig = set.size;
```

# --instructions--

在本練習中，我們將向 checkSet() 函數傳遞一個數組和一個值。 您的函數應該從數組參數創建一個 ES6 集合。 查找該集合是否包含 value 參數。 求集合的大小。 並在數組中返回這兩個值。

# --hints--

`checkSet([4, 5, 6], 3)` 應該返回 [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```
