---
id: 587d825b367417b2b2512c8c
title: 使用最小堆積實作堆積排序
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

現在我們可以新增和移除元素，讓我們來看看堆積可以用於哪些應用。堆積通常用來實作優先權佇列，因為它們總是將最大或最小值的項目儲存在第一個位置。此外，它們也用來實作一種稱為堆積排序的演算法。我們將在這裡看到如何做到這一點。堆積排序使用的是最小堆積，與最大堆積相反。最小堆積總是將最小值的元素儲存在根節點位置。

堆積排序的運作方式是先取得一個未排序的陣列，將陣列中的每個條款添加到一個最小堆積中，然後再從最小堆積中取出每個條款到一個新的陣列。最小堆積結構確保新的陣列會包含原始條款，且順序由小到大。這是其中一個最高效的排序演算法，平均和最壞情況的效率為 O(nlog(n))。

# --instructions--

讓我們使用最小堆積來實作堆積排序。你可以自由調整你的最大堆積程式碼。在物件 `MinHeap` 中創建 `insert`、`remove` 和 `sort` 方法。`sort` 方法應該傳回一個包含最小堆積中所有元素，並從小到大排序的陣列。

# --hints--

應該存在 `MinHeap` 資料結構。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` 應該有一個名為 `insert` 的方法。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` 應該有一個名為 `remove` 的方法。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` 應該有一個名為 `sort` 的方法。

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

`sort` 方法應該傳回一個陣列，該陣列包含所有已加入最小堆積的項目，且依排序順序排列。

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```
