---
id: 587d7fb7367417b2b2512c0a
title: Viele Datensätze mit model.create() erstellen
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

Manchmal müssen Sie viele Instanzen Ihrer Modelle erstellen, z. B. beim Befüllen einer Datenbank mit Anfangsdaten. `Model.create()` nimmt als erstes Argument ein Array von Objekten wie `[{name: 'John', ...}, {...}, ...]` und speichert sie alle in der Datenbank.

# --instructions--

Ändern Sie die `createManyPeople`-Funktion so, dass viele Personen mit `Model.create()` und dem Argument `arrayOfPeople` erstellt werden.

**Hinweis:** Sie können das Modell wiederverwenden, das Sie in der vorherigen Aufgabe instanziiert haben.

# --hints--

Das gleichzeitige Erstellen vieler Datenbankeinträge sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```

