---
id: 587d7fb7367417b2b2512c0b
title: "Verwenden Sie model.find(), um Ihre Datenbank zu durchsuchen"
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

In seiner einfachsten Anwendung akzeptiert `Model.find()` als erstes Argument ein Abfragedokument (ein JSON-Objekt) und dann einen Callback. Es gibt ein Array mit Übereinstimmungen zurück. Es unterstützt eine äußerst breite Palette von Suchoptionen. Lesen Sie mehr in der Dokumentation.

# --instructions--

Ändern Sie die `findPeopleByName`-Funktion so, dass alle Personen mit einem bestimmten Namen gefunden werden, indem Sie <code>Model.find() -\> [Person]</code> verwenden.

Verwenden Sie das Funktionsargument `personName` als Suchschlüssel.

# --hints--

Alle Elemente zu finden, die einem Kriterium entsprechen, sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```

