---
id: 587d7fb7367417b2b2512c0d
title: "Verwenden Sie model.findById(), um Ihre Datenbank nach _id zu durchsuchen"
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Beim Speichern eines Dokuments fügt MongoDB automatisch das Feld `_id` hinzu und setzt es auf einen eindeutigen alphanumerischen Schlüssel. Die Suche nach `_id` ist eine äußerst häufige Operation, daher stellt Mongoose eine spezielle Methode dafür bereit.

# --instructions--

Ändern Sie die `findPersonById`, um die einzige Person mit einer bestimmten `_id` zu finden, und verwenden Sie dabei `Model.findById() -> Person`. Verwenden Sie das Funktionsargument `personId` als Suchschlüssel.

# --hints--

Ein Element per Id zu finden sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

