---
id: 587d7fb6367417b2b2512c06
title: Mongoose installieren und einrichten
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

Bei diesen Aufgaben schreiben Sie Ihren Code mit einer der folgenden Methoden:

- Klonen Sie <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">dieses GitHub-Repository</a> und bearbeiten Sie die Aufgaben lokal.
- Verwenden Sie einen Website-Builder Ihrer Wahl, um das Projekt fertigzustellen. Stellen Sie sicher, dass Sie alle Dateien aus unserem GitHub-Repository einbinden.

In dieser Aufgabe richten Sie eine MongoDB-Atlas-Datenbank ein und importieren die erforderlichen Pakete, um eine Verbindung herzustellen.

Folgen Sie <a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">diesem Tutorial</a>, um eine gehostete Datenbank auf MongoDB Atlas einzurichten.

# --instructions--

`mongoose@^5.11.15` wurde zur `package.json`-Datei Ihres Projekts hinzugefügt. Fordern Sie zuerst mongoose als `mongoose` in `myApp.js` an. Erstellen Sie anschließend eine `.env`-Datei und fügen Sie eine `MONGO_URI`-Variable hinzu. Ihr Wert sollte die URI Ihrer MongoDB-Atlas-Datenbank sein. Umgeben Sie die URI unbedingt mit einfachen oder doppelten Anführungszeichen, und denken Sie daran, dass Sie in Umgebungsvariablen keine Leerzeichen um das `=` verwenden dürfen. Zum Beispiel `MONGO_URI='VALUE'`.

Wenn Sie fertig sind, verbinden Sie sich mit der Datenbank, indem Sie die `connect`-Methode in Ihrer `myApp.js`-Datei mit der folgenden Syntax aufrufen:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

Die Abhängigkeit "mongoose version ^5.11.15" sollte in package.json enthalten sein

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

"mongoose" sollte mit einer Datenbank verbunden sein

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```

