---
id: 587d7fb9367417b2b2512c12
title: Verkettung von Suchabfrage-Hilfsmitteln zur Eingrenzung der Suchergebnisse
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

Wenn Sie das Callback nicht als letztes Argument an `Model.find()` (oder an die anderen Suchmethoden) übergeben, wird die Abfrage nicht ausgeführt. Sie können die Abfrage in einer Variablen speichern, um sie später zu verwenden. Ein solches Objekt ermöglicht es Ihnen, eine Abfrage mit Verkettungssyntax aufzubauen. Die eigentliche Datenbanksuche wird ausgeführt, wenn Sie schließlich die Methode `.exec()` verketten. Sie müssen Ihr Callback immer an diese letzte Methode übergeben. Es gibt viele Abfrage-Hilfsmittel, hier verwenden wir die am häufigsten genutzten.

# --instructions--

Ändern Sie die Funktion `queryChain`, um Personen zu finden, die das in der Variablen `foodToSearch` angegebene Essen mögen. Sortieren Sie sie nach `name`, begrenzen Sie die Ergebnisse auf zwei Dokumente und blenden Sie deren Alter aus. Verketten Sie `.find()`, `.sort()`, `.limit()`, `.select()` und dann `.exec()`. Übergeben Sie das Callback `done(err, data)` an `exec()`.

# --hints--

Die Verkettung von Abfrage-Hilfsmitteln sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```

