---
id: 587d7fb8367417b2b2512c11
title: Viele Dokumente mit model.remove() löschen
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` ist nützlich, um alle Dokumente zu löschen, die bestimmten Kriterien entsprechen.

# --instructions--

Ändern Sie die `removeManyPeople`-Funktion so, dass alle Personen gelöscht werden, deren Name in der Variablen `nameToRemove` enthalten ist, und verwenden Sie dazu `Model.remove()`. Übergeben Sie es einem Abfragedokument mit dem gesetzten `name`-Feld und einem Callback.

**Hinweis:** Das `Model.remove()` gibt nicht das gelöschte Dokument zurück, sondern ein JSON-Objekt mit dem Ergebnis der Operation und der Anzahl der betroffenen Einträge. Vergessen Sie nicht, es an den `done()`-Callback zu übergeben, da wir es in den Tests verwenden.

# --hints--

Das gleichzeitige Löschen mehrerer Einträge sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

