---
id: 587d7fb6367417b2b2512c09
title: Einen Datensatz eines Modells erstellen und speichern
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

In dieser Aufgabe müssen Sie einen Datensatz eines Modells erstellen und speichern.

# --instructions--

Erstellen Sie innerhalb der `createAndSavePerson`-Funktion eine Dokumentinstanz mit dem `Person`-Modellkonstruktor, den Sie zuvor erstellt haben. Übergeben Sie dem Konstruktor ein Objekt mit den Feldern `name`, `age` und `favoriteFoods`. Deren Typen müssen mit denen im `personSchema` übereinstimmen. Rufen Sie dann die Methode `document.save()` auf der zurückgegebenen Dokumentinstanz auf. Übergeben Sie ihr ein Callback nach der Node-Konvention. Dies ist ein gängiges Muster; alle folgenden CRUD-Methoden nehmen eine Callback-Funktion wie diese als letztes Argument entgegen.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

Das Erstellen und Speichern eines Datenbankeintrags sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```

