---
id: 587d7fb7367417b2b2512c0c
title: "Verwenden Sie model.findOne(), um ein einzelnes passendes Dokument aus Ihrer Datenbank zurückzugeben"
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` verhält sich wie `Model.find()`, gibt aber nur ein Dokument zurück (kein Array), selbst wenn mehrere Elemente vorhanden sind. Es ist besonders nützlich, wenn Sie nach Eigenschaften suchen, die Sie als einzigartig deklariert haben.

# --instructions--

Ändern Sie die `findOneByFood`-Funktion so, dass sie nur eine Person findet, die ein bestimmtes Essen in den Favoriten hat, und verwenden Sie dabei `Model.findOne() -> Person`. Verwenden Sie das Funktionsargument `food` als Suchschlüssel.

# --hints--

Das Finden eines einzelnen Elements sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

