---
id: 587d7fb8367417b2b2512c10
title: Ein Dokument mit model.findByIdAndRemove löschen
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` und `findOneAndRemove` funktionieren wie die vorherigen Update-Methoden. Sie übergeben das gelöschte Dokument an die Datenbank. Wie üblich verwenden Sie das Funktionsargument `personId` als Suchschlüssel.

# --instructions--

Ändern Sie die `removeById`-Funktion, um eine Person anhand der `_id` der Person zu löschen. Sie sollten eine der Methoden `findByIdAndRemove()` oder `findOneAndRemove()` verwenden.

# --hints--

Das Löschen eines Elements sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```

