---
id: 587d7fb8367417b2b2512c0f
title: Neue Updates an einem Dokument mit model.findOneAndUpdate() durchführen
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Neuere Versionen von Mongoose bieten Methoden, um das Aktualisieren von Dokumenten zu vereinfachen. Einige fortgeschrittene Funktionen (z. B. pre/post Hooks, Validierung) verhalten sich bei diesem Ansatz anders, daher ist die klassische Methode in vielen Situationen weiterhin nützlich. `findByIdAndUpdate()` kann verwendet werden, wenn nach der ID gesucht wird.

# --instructions--

Ändern Sie die `findAndUpdate`-Funktion so, dass eine Person anhand von `Name` gefunden wird und das Alter der Person auf `20` gesetzt wird. Verwenden Sie den Funktionsparameter `personName` als Suchschlüssel.

**Hinweis:** Sie sollten das aktualisierte Dokument zurückgeben. Dafür müssen Sie das Optionsdokument `{ new: true }` als dritten Argument an `findOneAndUpdate()` übergeben. Standardmäßig geben diese Methoden das unveränderte Objekt zurück.

# --hints--

findOneAndUpdate eines Elements sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```

