---
id: 587d7fb8367417b2b2512c0e
title: 'Klassische Updates durchführen durch Suchen, Bearbeiten und Speichern'
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

Früher war dies der übliche Weg, wenn Sie ein Dokument bearbeiten und es anschließend irgendwie verwenden wollten (z. B. indem Sie es in einer Server-Antwort zurücksenden). Mongoose hat eine spezielle Update-Methode: `Model.update()`. Sie ist an den Low-Level-Mongo-Treiber gebunden. Sie kann viele Dokumente, die bestimmten Kriterien entsprechen, auf einmal bearbeiten, sendet aber nicht das aktualisierte Dokument zurück, sondern nur eine „Status“-Meldung. Außerdem erschwert sie die Validierung von Modellen, weil sie direkt den Mongo-Treiber aufruft.

# --instructions--

Ändern Sie die `findEditThenSave`-Funktion so, dass sie eine Person anhand von `_id` findet (verwenden Sie eine der oben genannten Methoden) mit dem Parameter `personId` als Suchschlüssel. Fügen Sie `"hamburger"` zur Liste der `favoriteFoods` der Person hinzu (Sie können `Array.push()` verwenden). Dann – innerhalb des Find-Callbacks – `save()` Sie das aktualisierte `Person`.

**Hinweis:** Das kann knifflig sein, wenn Sie in Ihrem Schema `favoriteFoods` als Array deklariert haben, ohne den Typ anzugeben (also `[String]`). In diesem Fall ist `favoriteFoods` standardmäßig vom Typ Mixed, und Sie müssen es manuell als bearbeitet markieren, indem Sie `document.markModified('edited-field')` verwenden. Siehe unseren <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">Mongoose-Artikel</a>.

# --hints--

Find-edit-update eines Elements sollte erfolgreich sein

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```

