---
id: 698c513508507ec8bb820227
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Para especificar el recurso multimedia para el audio, necesitarás agregar el atributo `src` al elemento `audio`.

Agrega el atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3`.

# --hints--

Tu elemento `audio` debe tener un atributo `src`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('src'));
```

Tu elemento `audio` debe tener un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3`.

```js
const audio = document.querySelector('audio');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio></audio>
--fcc-editable-region--

</body>
</html>
```
