---
id: 698c4fe24f1f1a17be74a23e
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

A continuación, crea un elemento `audio` debajo del elemento `p`. En los siguientes pasos, agregarás los atributos necesarios para que el elemento de audio reproduzca música. 

# --hints--

Debes tener un elemento `audio`.

```js
assert.exists(document.querySelector('audio'));
```

El nuevo elemento `audio` debe estar debajo del elemento `p`.

```js
assert.exists(document.querySelector('p + audio'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
