---
id: 698c4f4e965bd853ad5084e1
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del `h1`, agrega un elemento `h2` con el texto `Can't Stay Down`, este es el título de la primera canción.

Debajo del elemento `h2` agrega un elemento `p` con el texto `Artist: Quincy Larson`.

# --hints--

Deberías tener un elemento `h2`.

```js
assert.exists(document.querySelector('h2'));
```

Tu elemento `h2` debe estar debajo del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + h2'));
```

Tu elemento `h2` debe contener el texto `Can't Stay Down`.

```js
const h2 = document.querySelector('h2');
assert.strictEqual(h2?.textContent.trim(), "Can't Stay Down");
```

Debes tener un elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Tu elemento `p` debe contener el texto `Artist: Quincy Larson`.

```js
const p = document.querySelector('p');
assert.strictEqual(p?.textContent.trim(), 'Artist: Quincy Larson');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
