---
id: 698c4e69abe14a9746c72312
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, construirás un reproductor de música en HTML. La estructura básica de HTML ha sido proporcionada para ti.

Crea un elemento `h1` y asígnale el texto `freeCodeCamp Tunes`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe contener el texto `freeCodeCamp Tunes`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'freeCodeCamp Tunes');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
