---
id: 698c5710136539797f296098
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Debajo del elemento `p`, agrega un elemento `audio`, asígnale un atributo `src` con un valor de `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`, y los atributos `loop` y `controls`.

# --hints--

Debes tener un segundo elemento `audio`.

```js
assert.exists(document.querySelector('audio:nth-of-type(2)'));
```

El nuevo elemento `audio` debe estar debajo del elemento `p`.

```js
assert.exists(document.querySelector('p:nth-of-type(2) + audio:nth-of-type(2)'));
```

Tu elemento `audio` debe tener un atributo `src`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('src'));
```

Tu elemento `audio` debe tener un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.strictEqual(audio?.getAttribute('src'),
"https://cdn.freecodecamp.org/curriculum/js-music-player/cruising-for-a-musing.mp3");
```

Tu elemento `audio` debe tener el atributo `loop`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('loop'));
```

Tu elemento `audio` debe tener el atributo `controls`.

```js
const audio = document.querySelector('audio:nth-of-type(2)');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop controls></audio>
  
  <h2>Cruising for a Musing</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  
--fcc-editable-region--

</body>
</html>
```
