---
id: 698c50535b54b2d75fa985fb
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

El atributo `loop` reiniciará el audio una vez que la reproducción haya terminado. Piensa en un meme de internet que repite la reproducción. Omitir el atributo `loop` hará que el audio se reproduzca una vez.

El atributo `loop` es un atributo booleano y no necesita un valor.

Agrega el atributo `loop` al elemento `audio`.

# --hints--

Tu elemento `audio` debe tener el atributo `loop`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3"></audio>
--fcc-editable-region--

</body>
</html>
```
