---
id: 698c5229e47d57c6dfaa91bd
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

El atributo `controls` proporciona controles de reproducción que incluyen pausar/reanudar la reproducción, búsqueda y control de volumen para el elemento `audio`. 

El atributo `controls` es un atributo booleano y no necesita un valor.

Agrega el atributo `controls` al elemento `audio`.

Ahora deberías ver el elemento `audio` mostrado en la página.

# --hints--

Tu elemento `audio` debe tener el atributo `controls`.

```js
const audio = document.querySelector('audio');
assert.isTrue(audio?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Audio Element</title>
</head>
<body>
  <h1>freeCodeCamp Tunes</h1>

  <h2>Can't Stay Down</h2>
  <p>Artist: Quincy Larson</p>
  
--fcc-editable-region--
  <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/can't-stay-down.mp3" loop></audio>
--fcc-editable-region--

</body>
</html>
```
