---
id: 587d78af367417b2b2512b04
title: Produkt-Zielseite erstellen
challengeType: 25
dashedName: build-a-product-landing-page
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Ihre Produkt-Zielseite sollte ein `header`-Element mit einem entsprechenden `id="header"` haben.
2. Sie sollten ein Bild innerhalb des `header`-Elements mit einem entsprechenden `id="header-img"` haben (ein Logo wäre hier ein gutes Bild).
3. Innerhalb des `#header`-Elements sollten Sie ein `nav`-Element mit einem entsprechenden `id="nav-bar"` haben.
4. Sie sollten mindestens drei anklickbare Elemente innerhalb des `nav`-Elements haben, jeweils mit der Klasse `nav-link`.
5. Wenn Sie auf eine `.nav-link`-Schaltfläche im `nav`-Element klicken, sollten Sie zum entsprechenden Abschnitt der Zielseite gelangen.
6. Sie sollten ein eingebettetes Produktvideo mit `id="video"` haben.
7. Ihre Zielseite sollte ein `form`-Element mit einem entsprechenden `id="form"` haben.
8. Innerhalb des Formulars sollte es ein `input`-Feld mit `id="email"` geben, in das Sie eine E-Mail-Adresse eingeben können.
9. Das `#email`-Eingabefeld sollte einen Platzhaltertext haben, der den Nutzenden zeigt, wofür das Feld ist.
10. Das `#email`-Eingabefeld sollte HTML5-Validierung verwenden, um zu bestätigen, dass der eingegebene Text eine E-Mail-Adresse ist.
11. Innerhalb des Formulars sollte es eine Absende-`input` mit einem entsprechenden `id="submit"` geben.
12. Wenn Sie auf das `#submit`-Element klicken, sollte die E-Mail an eine statische Seite gesendet werden (verwenden Sie diese Beispiel-URL: `https://www.freecodecamp.org/email-submit`).
13. Die Navigationsleiste sollte immer oben im Viewport sein.
14. Ihre Produkt-Zielseite sollte mindestens eine Media Query enthalten.
15. Ihre Produkt-Zielseite sollte mindestens einmal CSS Flexbox verwenden.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verknüpfen und Ihr CSS anwenden.

# --hints--

Sie sollten ein `header`-Element mit einem `id` von `header` haben.

```js
const el = document.getElementById('header');
assert.exists(el); 
assert.strictEqual(el.tagName, 'HEADER'); 
```

Sie sollten ein `img`-Element mit einem `id` von `header-img` haben.

```js
const el = document.getElementById('header-img');
assert.exists(el); 
assert.strictEqual(el.tagName, 'IMG');
```

Ihr `#header-img` sollte ein Nachfahre des `#header` sein.

```js
const els = document.querySelectorAll('#header #header-img');
assert.isNotEmpty(els);
```

Ihr `#header-img` sollte ein `src`-Attribut haben.

```js
const el = document.getElementById('header-img');
assert.exists(el);
assert.exists(el.src); 
```

Der `#header-img`-Wert Ihres `src` sollte eine gültige URL sein (beginnt mit `http`).

```js
const el = document.getElementById('header-img');
assert.exists(el);
assert.isNotEmpty(el.getAttribute('src')); 
assert.match(el.src, /^http/);
```

Sie sollten ein `nav`-Element mit einem `id` von `nav-bar` haben.

```js
const el = document.getElementById('nav-bar');
assert.exists(el);
assert.strictEqual(el.tagName, "NAV");
```

Ihr `#nav-bar` sollte ein Nachfahre des `#header` sein.

```js
const els = document.querySelectorAll('#header #nav-bar');
assert.isNotEmpty(els);
```

Sie sollten mindestens 3 `.nav-link`-Elemente innerhalb des `#nav-bar` haben.

```js
const els = document.querySelectorAll('#nav-bar .nav-link');
assert.isAtLeast(els.length, 3); 
```

Jedes `.nav-link`-Element sollte ein `href`-Attribut haben.

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  assert.exists(el.href);
  assert.isNotEmpty(el.href);
})
assert.isNotEmpty(els); 
```

Jedes `.nav-link`-Element sollte mit einem entsprechenden Element auf der Zielseite verknüpft sein (hat ein `href` mit dem Wert der id eines anderen Elements, z. B. `#footer`).

```js
const els = document.querySelectorAll('.nav-link')
els.forEach(el => {
  const linkDestination = el.getAttribute('href').slice(1)
  assert.exists(document.getElementById(linkDestination));
})
assert.isNotEmpty(els); 
```

Sie sollten ein `video`- oder `iframe`-Element mit einem `id` von `video` haben.

```js
const el = document.getElementById('video');
assert.exists(el); 
assert.oneOf(el.tagName, ['VIDEO','IFRAME']);
```

Ihr `#video` sollte ein `src`-Attribut haben.

```js
let el = document.getElementById('video')
const sourceNode = el.children;
let sourceElement = null;
if (sourceNode.length) {
  sourceElement = [...video.children].filter(el => el.localName === 'source')[0];
}
if (sourceElement) {
  el = sourceElement;
}
assert.isTrue(el.hasAttribute('src'));
```

Sie sollten ein `form`-Element mit einem `id` von `form` haben.

```js
const el = document.getElementById('form');
assert.exists(el); 
assert.strictEqual(el.tagName, "FORM"); 
```

Sie sollten ein `input`-Element mit einem `id` von `email` haben.

```js
const el = document.getElementById('email');
assert.exists(el); 
assert.strictEqual(el.tagName, "INPUT"); 
```

Ihr `#email` sollte ein Nachfahre des `#form` sein.

```js
const els = document.querySelectorAll('#form #email')
assert.isNotEmpty(els);
```

Ihr `#email` sollte das `placeholder`-Attribut mit Platzhaltertext haben.

```js
const el = document.getElementById('email');
assert.exists(el); 
assert.exists(el.placeholder); 
assert.isNotEmpty(el.placeholder); 
```

Ihr `#email` sollte HTML5-Validierung verwenden, indem Sie sein `type` auf `email` setzen.

```js
const el = document.getElementById('email');
assert.exists(el); 
assert.strictEqual(el.type, "email");
```

Sie sollten ein `input`-Element mit einem `id` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.exists(el); 
assert.strictEqual(el.tagName, "INPUT");
```

Ihr `#submit` sollte ein Nachfahre des `#form` sein.

```js
const els = document.querySelectorAll('#form #submit');
assert.isNotEmpty(els);
```

Ihr `#submit` sollte ein `type` von `submit` haben.

```js
const el = document.getElementById('submit');
assert.exists(el); 
assert.strictEqual(el.type, 'submit');
```

Ihr `#form` sollte ein `action`-Attribut von `https://www.freecodecamp.org/email-submit` haben.

```js
const el = document.getElementById('form');
assert.exists(el); 
assert.strictEqual(el.action, 'https://www.freecodecamp.org/email-submit');
```

Ihr `#email` sollte ein `name`-Attribut von `email` haben.

```js
const el = document.getElementById('email');
assert.exists(el); 
assert.strictEqual(el.name, 'email');
```

Ihr `#nav-bar` sollte immer oben im Viewport sein.

```js
  const timeout = (milliseconds) => new Promise((resolve) => setTimeout(resolve, milliseconds));

  const header = document.getElementById('header');
  const headerChildren = header.children;
  const navbarCandidates = [header, ...headerChildren];
  
  // Return smallest top position of all navbar candidates
  const getNavbarPosition = (candidates = []) => {
    return candidates.reduce(
      (min, candidate) =>
        Math.min(min, Math.abs(candidate?.getBoundingClientRect().top)),
      Infinity
    );
  };
  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the viewport '
  );

  window.scroll(0, 500);
  await timeout(1);

  assert.approximately(
    getNavbarPosition(navbarCandidates),
    0,
    15,
    '#header or one of its children should be at the top of the ' +
      'viewport even after scrolling '
  );
    
  window.scroll(0, 0);
```

Ihre Produkt-Zielseite sollte mindestens eine Media Query verwenden.

```js
const htmlSourceAttr = Array.from(document.querySelectorAll('source')).map(el => el.getAttribute('media'))
const cssCheck = new __helpers.CSSHelp(document).getCSSRules('media');
assert.isTrue(cssCheck.length > 0 || htmlSourceAttr.length > 0);
```

Ihre Produkt-Zielseite sollte mindestens einmal CSS Flexbox verwenden.

```js
const hasFlex = (rule) => ["flex", "inline-flex"].includes(rule.style?.display)
const stylesheet = new __helpers.CSSHelp(document).getStyleSheet()
const cssRules = new __helpers.CSSHelp(document).styleSheetToCssRulesArray(stylesheet)
const mediaRules = new __helpers.CSSHelp(document).getCSSRules('media')
const usesFlex = cssRules.find(rule => hasFlex(rule))
const usesFlexMedia = mediaRules.find(mediaRule => {
  return [...mediaRule.cssRules].find(rule => hasFlex(rule))
})
assert.exists(usesFlex || usesFlexMedia);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Product Landing Page</title>
  </head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <link rel="stylesheet" href="styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.1/css/all.css"
        integrity="sha384-50oBUHEmvpQ+1lW4y57PTFmhCaXp0ML5d60M1M7uH2+nqUivzIebhndOJK28anvf" crossorigin="anonymous" />
</head>

<body>
    <div id="page-wrapper">
        <header id="header">
            <div class="logo">
                <img id="header-img"
                    src="https://cdn.freecodecamp.org/testable-projects-fcc/images/product-landing-page-logo.png"
                    alt="original trombones logo" />
            </div>

            <nav id="nav-bar">
                <ul>
                    <li><a class="nav-link" href="#features">Features</a></li>
                    <li><a class="nav-link" href="#how-it-works">How It Works</a></li>
                    <li><a class="nav-link" href="#pricing">Pricing</a></li>
                </ul>
            </nav>
        </header>

        <div class="container"></div>

        <section id="hero">
            <h2>Handcrafted, home-made masterpieces</h2>
            <form id="form" action="https://www.freecodecamp.org/email-submit">
                <input name="email" id="email" type="email" placeholder="Enter your email address" required />
                <input id="submit" type="submit" value="Get Started" class="btn" />
            </form>
        </section>

        <div class="container">
            <section id="features">
                <div class="grid">
                    <div class="icon"><i class="fa fa-3x fa-fire"></i></div>
                    <div class="desc">
                        <h2>Premium Materials</h2>
                        <p>
                            Our trombones use the shiniest brass which is sourced locally.
                            This will increase the longevity of your purchase.
                        </p>
                    </div>
                </div>
                <div class="grid">
                    <div class="icon"><i class="fa fa-3x fa-truck"></i></div>
                    <div class="desc">
                        <h2>Fast Shipping</h2>
                        <p>
                            We make sure you receive your trombone as soon as we have
                            finished making it. We also provide free returns if you are not
                            satisfied.
                        </p>
                    </div>
                </div>
                <div class="grid">
                    <div class="icon">
                        <i class="fa fa-3x fa-battery-full" aria-hidden="true"></i>
                    </div>
                    <div class="desc">
                        <h2>Quality Assurance</h2>
                        <p>
                            For every purchase you make, we will ensure there are no damages
                            or faults and we will check and test the pitch of your
                            instrument.
                        </p>
                    </div>
                </div>
            </section>
            <section id="how-it-works">
                <iframe id="video" height="315"
                    src="https://www.youtube-nocookie.com/embed/y8Yv4pnO7qc?rel=0&amp;controls=0&amp;showinfo=0"
                    frameborder="0" allowfullscreen></iframe>
            </section>
            <section id="pricing">
                <div class="product" id="tenor">
                    <div class="level">Tenor Trombone</div>
                    <h2>$600</h2>
                    <ol>
                        <li>Good for beginners</li>
                        <li>Excellent sound quality</li>
                        <li>Great for Jazz Bands</li>
                        <li>Nice and shiny</li>                        
                    </ol>
                    <button class="btn">Select</button>
                </div>
                <div class="product" id="bass">
                    <div class="level">Bass Trombone</div>
                    <h2>$900</h2>
                    <ol>
                        <li>Sound quality is unmatched</li>
                        <li>Best for professionals</li>
                        <li>Absolutely stunning</li>
                        <li>Durable and long lasting</li>
                    </ol>
                    <button class="btn">Select</button>
                </div>
                <div class="product" id="valve">
                    <div class="level">Valve Trombone</div>
                    <h2>$1200</h2>
                    <ol>
                        <li>Plays similar to a Trumpet</li>
                        <li>Great for Jazz Bands</li>
                        <li>Beautiful in sound and appearance</li>
                        <li>Just amazing</li>
                    </ol>
                    <button class="btn">Select</button>
                </div>
            </section>
            <footer>
                <ul>
                    <li><a href="#">Privacy</a></li>
                    <li><a href="#">Terms</a></li>
                    <li><a href="#">Contact</a></li>
                </ul>
                <span>Copyright 2016, Original Trombones</span>
            </footer>
        </div>
    </div>
</body>

</html>
```

```css
/** global element styling **/

@import 'https://fonts.googleapis.com/css?family=Lato:400,700';

* {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

body {
    background-color: #eee;
    font-family: 'Lato', sans-serif;
}

#page-wrapper {
    position: relative;
}

li {
    list-style: none;
}

a {
    color: #000;
    text-decoration: none;
}

/** global classes styling **/

.container {
    max-width: 1000px;
    width: 100%;
    margin: 0 auto;
}

.btn {
    padding: 0 20px;
    height: 40px;
    font-size: 1em;
    font-weight: 900;
    text-transform: uppercase;
    border: 3px black solid;
    border-radius: 2px;
    background: transparent;
    cursor: pointer;
}

.grid {
    display: flex;
}

header {
    position: fixed;
    top: 0;
    min-height: 75px;
    padding: 0px 20px;
    display: flex;
    justify-content: space-around;
    align-items: center;
    background-color: #eee;
}

@media (max-width: 600px) {
    header {
        flex-wrap: wrap;
    }
}

.logo {
    width: 60vw;
}

@media (max-width: 650px) {
    .logo {
        margin-top: 15px;
        width: 100%;
        position: relative;
    }
}

.logo>img {
    width: 100%;
    height: 100%;
    max-width: 300px;
    display: flex;
    justify-content: center;
    align-items: center;
    text-align: center;
    margin-left: 20px;
}

@media (max-width: 650px) {
    .logo>img {
        margin: 0 auto;
    }
}

nav {
    font-weight: 400;
}

@media (max-width: 650px) {
    nav {
        margin-top: 10px;
        width: 100%;
        display: flex;
        flex-direction: column;
        align-items: center;
        text-align: center;
        padding: 0 50px;
    }

    nav li {
        padding-bottom: 5px;
    }
}

nav>ul {
    width: 35vw;
    display: flex;
    flex-direction: row;
    justify-content: space-around;
}

@media (max-width: 650px) {
    nav>ul {
        flex-direction: column;
    }
}

#hero {
    display: flex;
    flex-direction: column;
    align-items: center;
    justify-content: center;
    text-align: center;
    height: 200px;
    margin-top: 50px;
}

#hero>h2 {
    margin-bottom: 20px;
    word-wrap: break-word;
}

#hero input[type='email'] {
    max-width: 275px;
    width: 100%;
    padding: 5px;
}

#hero input[type='submit'] {
    max-width: 150px;
    width: 100%;
    height: 30px;
    margin: 15px 0;
    border: 0;
    background-color: #f1c40f;
}

#hero input[type='submit']:hover {
    background-color: orange;
    transition: background-color 1s;
}

@media (max-width: 650px) {
    #hero {
        margin-top: 120px;
    }
}

#features {
    margin-top: 30px;
}

#features .icon {
    display: flex;
    align-items: center;
    justify-content: center;
    height: 125px;
    width: 20vw;
    color: darkorange;
}

@media (max-width: 550px) {
    #features .icon {
        display: none;
    }
}

#features .desc {
    display: flex;
    flex-direction: column;
    justify-content: center;
    height: 125px;
    width: 80vw;
    padding: 5px;
}

@media (max-width: 550px) {
    #features .desc {
        width: 100%;
        text-align: center;
        padding: 0;
        height: 150px;
    }
}

@media (max-width: 650px) {
    #features {
        margin-top: 0;
    }
}

#how-it-works {
    margin-top: 50px;
    display: flex;
    justify-content: center;
}

#how-it-works>iframe {
    max-width: 560px;
    width: 100%;
}

#pricing {
    margin-top: 60px;
    display: flex;
    flex-direction: row;
    justify-content: center;
}

.product {
    display: flex;
    flex-direction: column;
    align-items: center;
    text-align: center;
    width: calc(100% / 3);
    margin: 10px;
    border: 1px solid #000;
    border-radius: 3px;
}

.product>.level {
    background-color: #ddd;
    color: black;
    padding: 15px 0;
    width: 100%;
    text-transform: uppercase;
    font-weight: 700;
}

.product>h2 {
    margin-top: 15px;
}

.product>ol {
    margin: 15px 0;
}

.product>ol>li {
    padding: 5px 0;
}

.product>button {
    border: 0;
    margin: 15px 0;
    background-color: #f1c40f;
    font-weight: 400;
}

.product>button:hover {
    background-color: orange;
    transition: background-color 1s;
}

@media (max-width: 800px) {
    #pricing {
        flex-direction: column;
    }

    .product {
        max-width: 300px;
        width: 100%;
        margin: 0 auto;
        margin-bottom: 10px;
    }
}

footer {
    margin-top: 30px;
    background-color: #ddd;
    padding: 20px;
}

footer>ul {
    display: flex;
    justify-content: flex-end;
}

footer>ul>li {
    padding: 0 10px;
}

footer>span {
    margin-top: 5px;
    display: flex;
    justify-content: flex-end;
    font-size: 0.9em;
    color: #444;
}
```
