---
id: 612ea4c4993aba52ab4aa32e
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Ora è il momento di usare i pseudo-selettori che hai impostato prima. Per creare i tasti neri, aggiungi un nuovo selettore `.key.black--key::after`. Questo seleziona gli elementi con entrambe le classi `key` e `black--key` e prende il pseudo-elemento creato dopo questi elementi nell'HTML.

Nel nuovo selettore, imposta la proprietà `background-color` su `#1d1e22`. Imposta anche la proprietà `content` su `""`. Questo renderà i pseudo-elementi vuoti.

La proprietà `content` serve per impostare o sovrascrivere il contenuto di un elemento. Per impostazione predefinita, i pseudo-elementi creati dai selettori `::before` e `::after` sono vuoti, il che significa che non vengono renderizzati sulla pagina. Impostando la proprietà `content` su una stringa vuota `""`, ti assicuri che i pseudo-elementi vengano renderizzati, pur apparendo vuoti.

Se vuoi sperimentare, prova a rimuovere la proprietà `background-color` e a impostare valori diversi per la proprietà `content`, come `"♥"`. Ricordati di annullare queste modifiche quando hai finito così che i test passino.

# --hints--

Dovresti avere un selettore `.key.black--key::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Il tuo selettore `.key.black--key::after` dovrebbe avere una proprietà `background-color` impostata su `#1d1e22`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor, 'rgb(29, 30, 34)');
```

Il tuo selettore `.key.black--key::after` dovrebbe avere una proprietà `content` impostata su `""`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content, '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
