---
id: 612e89562043183c86df287c
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

I browser spesso applicano valori predefiniti di margine e spaziatura interna a elementi specifici. Per assicurarti che il layout del tuo pianoforte venga visualizzato correttamente, devi resettare il modello a riquadro.

Aggiungi un selettore `html` al tuo file CSS e imposta la proprietà `box-sizing` su `border-box`.

# --hints--

Dovresti avere un selettore `html`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Il tuo selettore `html` dovrebbe avere la proprietà `box-sizing` impostata su `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
