---
id: 612e7d1c29fb872d6384379c
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Annida un secondo `div` all'interno del tuo `div` esistente e imposta il `class` su `keys`.

# --hints--

Dovresti creare un secondo elemento `div`.

```js
const divDiv = document.querySelectorAll('div');
assert.lengthOf(divDiv, 2);
```

Il tuo nuovo elemento `div` dovrebbe essere all'interno del tuo elemento `div` esistente.

```js
const div = document.querySelector('div');
assert.lengthOf(div?.children, 1);
assert.equal(div?.children?.[0]?.localName, 'div');
```

Il tuo nuovo elemento `div` dovrebbe avere il `class` impostato su `keys`.

```js
const div = document.querySelector('div');
assert.equal(div?.children?.[0]?.className, 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
