---
id: 612e813b3ba67633222cbe54
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ricorda, un attributo `class` può contenere valori multipli. Per distinguere tra i tasti bianchi e neri, aggiungi un secondo valore `class` di `black--key` ai tuoi secondi, terzi, quinti, sesti e settimi elementi `.key`.

# --hints--

Il tuo secondo elemento `.key` dovrebbe avere anche un `class` di `black--key`.

```js
const key = document.querySelectorAll('.key')[1];
assert.isTrue(key?.className.includes('black--key'));
```

Il tuo terzo `.key` dovrebbe avere `black--key` nel `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert.isTrue(third?.classList?.contains('black--key'));
```

Il tuo quinto `.key` dovrebbe avere `black--key` nel `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert.isTrue(fifth?.classList?.contains('black--key'));
```

Il tuo sesto `.key` dovrebbe avere `black--key` nel `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert.isTrue(sixth?.classList.contains('black--key'));
```

Il tuo settimo `.key` dovrebbe avere `black--key` nel `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert.isTrue(seventh?.classList?.contains('black--key'));
```

Dovresti avere cinque elementi `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 5);
```

Dovresti avere sette elementi `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
