---
id: 612e804c54d5e7308d7ebe56
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

All'interno del tuo elemento `.keys`, aggiungi sette elementi `div`, ognuno con la classe `key`.

# --hints--

Dovresti creare sette nuovi elementi `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert.lengthOf(divDivDiv, 9);
```

I tuoi sette nuovi elementi `div` dovrebbero essere all'interno del tuo elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf([...keys?.children], 7);
[...keys?.children].forEach(
  child => assert.equal(child.tagName, 'DIV')
)
```

I tuoi sette nuovi elementi `div` dovrebbero avere tutti l'insieme `class` impostato su `key`.

```js
const keys = document.querySelector('.keys');
[...keys?.children].forEach(child => assert.isTrue(child?.classList.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
