---
id: 612e89d254fe5d3df7d6693d
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora che hai resettato il modello a riquadro html, devi applicarlo anche agli elementi al suo interno. Per farlo, dovrai impostare la proprietà `box-sizing` di tutti gli altri elementi su inherited, così da assicurarti che gli elementi selezionati adottino lo stesso valore del loro elemento genitore.

Dovrai anche selezionare i pseudo-elementi, che sono parole chiave speciali che seguono un selettore. I due pseudo-elementi che userai sono i pseudo-elementi `::before` e `::after`.

Il selettore `::before` crea un pseudo-elemento che è il primo figlio dell'elemento selezionato, mentre il selettore `::after` crea un pseudo-elemento che è l'ultimo figlio dell'elemento selezionato. Questi pseudo-elementi sono spesso usati per creare contenuti cosmetici, come vedrai più avanti in questo progetto.

Per ora, crea un selettore CSS per selezionare tutti gli elementi usando `*`, e includi i pseudo-elementi `::before` e `::after`. Imposta la proprietà `box-sizing` su `inherit`.

# --hints--

Dovresti avere un selettore `html``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('html'));
```

Il tuo selettore `html` dovrebbe avere la proprietà `box-sizing` impostata su `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing, 'border-box');
```

Dovresti avere un selettore `*, ::before, ::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Il tuo selettore `*, ::before, ::after` dovrebbe avere la proprietà `box-sizing` impostata su `inherit`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing, 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
