---
id: 612e9d142affc44a453655db
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

È il momento di definire lo stile delle stesse chiavi. Crea un selettore `class` per gli elementi `.key`. Imposta l'insieme `background-color` al valore `#ffffff`, la proprietà `position` su `relative`, la proprietà `width` su `41px` e la proprietà `height` su `175px`.

# --hints--

Dovresti avere un selettore `.key`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.key'));
```

Il tuo selettore `.key` dovrebbe avere una proprietà `background-color` impostata su `#ffffff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor, 'rgb(255, 255, 255)');
```

Il tuo selettore `.key` dovrebbe avere la proprietà `position` impostata su `relative`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.position,'relative');
```

Il tuo selettore `.key` dovrebbe avere una proprietà `width` impostata su `41px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.width, '41px');
```

Il tuo selettore `.key` dovrebbe avere una proprietà `height` impostata su `175px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.key')?.height, '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
