---
id: 612e8279827a28352ce83a72
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Ora, copia l'insieme di sette elementi `.key` e incolla due insiemi aggiuntivi all'interno del div `.keys`.

# --hints--

Dovresti avere 21 elementi `.key`.

```js
const keys = document.querySelectorAll('.key');
assert.lengthOf(keys, 21);
```

Dovresti avere 15 elementi `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert.lengthOf(blackKeys, 15);
```

Tutti gli elementi `.key` dovrebbero essere all'interno del tuo elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.key'), 21);
```

Tutti gli elementi `.black--key` dovrebbero essere all'interno del tuo elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert.lengthOf(keys?.querySelectorAll('.black--key'), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
