---
id: 612e78af05201622d4bab8aa
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Crea un elemento `div` all'interno del tuo elemento `body` con la proprietà `id` impostata su `piano`.

# --hints--

Dovresti creare un nuovo elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Il tuo `div` dovrebbe essere all'interno del tuo elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

Il tuo `div` dovrebbe avere la proprietà `id` impostata su `piano`.

```js
const div = document.querySelector('div');
assert.equal(div?.getAttribute('id'), 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
