---
id: 5c6c06847491271903d37cfd
title: Das value-Attribut bei Optionsschaltflächen und Kontrollkästchen verwenden
challengeType: 0
forumTopicId: 301099
dashedName: use-the-value-attribute-with-radio-buttons-and-checkboxes
---

# --description--

Wenn ein Formular abgesendet wird, werden die Daten an den Server gesendet und enthalten Einträge für die ausgewählten Optionen. Eingaben vom Typ `radio` und `checkbox` geben ihre Werte aus dem `value`-Attribut an.

Zum Beispiel:

```html
<label for="indoor">
  <input id="indoor" value="indoor" type="radio" name="indoor-outdoor">Indoor
</label>
<label for="outdoor">
  <input id="outdoor" value="outdoor" type="radio" name="indoor-outdoor">Outdoor
</label>
```

Hier haben Sie zwei `radio`-Eingaben. Wenn der Benutzer das Formular mit der ausgewählten `indoor`-Option absendet, enthält die Formulardatenzeile: `indoor-outdoor=indoor`. Diese stammt von den `name`- und `value`-Attributen der Eingabe „indoor“.

Wenn Sie das `value`-Attribut weglassen, verwendet die abgesendete Formulardaten den Standardwert, der `on` ist. In diesem Fall, wenn der Benutzer die Option „indoor“ anklickt und das Formular absendet, wäre die resultierende Formulardaten `indoor-outdoor=on`, was nicht nützlich ist. Daher muss das `value`-Attribut auf etwas gesetzt werden, um die Option zu identifizieren.

# --instructions--

Geben Sie jeder der vorhandenen `radio`- und `checkbox`-Eingaben das `value`-Attribut. Erstellen Sie keine neuen Optionsschaltflächen oder Kontrollkästchen. Verwenden Sie den Text der Eingabe-Bezeichnung in Kleinbuchstaben als Wert für das Attribut.

# --hints--

Eine Ihrer Optionsschaltflächen sollte das `value`-Attribut mit dem Wert `indoor` haben.

```js
const indoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "indoor");
assert.notEmpty(indoorRadioButton)
```

Eine Ihrer Optionsschaltflächen sollte das `value`-Attribut mit dem Wert `outdoor` haben.

```js
const outdoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "outdoor");
assert.notEmpty(outdoorRadioButton);
```

Eines Ihrer Kontrollkästchen sollte das `value`-Attribut mit dem Wert `loving` haben.

```js
const lovingCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "loving");
assert.notEmpty(lovingCheckbox); 
```

Eines Ihrer Kontrollkästchen sollte das `value`-Attribut mit dem Wert `lazy` haben.

```js
const lazyCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "lazy");
assert.notEmpty(lazyCheckbox); 
```

Eines Ihrer Kontrollkästchen sollte das `value`-Attribut mit dem Wert `energetic` haben.

```js
const energeticCheckbox =  [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "energetic");
assert.notEmpty(energeticCheckbox);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
