---
id: bad87fee1348bd9aecf08801
title: Einführung in HTML5-Elemente
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5 führt beschreibendere HTML-Tags ein. Dazu gehören `main`, `header`, `footer`, `nav`, `video`, `article`, `section` und andere.

Diese Tags geben Ihrer HTML eine beschreibende Struktur, machen Ihren HTML-Code leichter lesbar und helfen bei der Suchmaschinenoptimierung (SEO) sowie der Barrierefreiheit. Das `main`-HTML5-Tag hilft Suchmaschinen und anderen Entwicklern, den Hauptinhalt Ihrer Seite zu finden.

Beispielanwendung: ein `main`-Element mit zwei darin verschachtelten Kindelementen:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**Hinweis:** Viele der neuen HTML5-Tags und deren Vorteile werden im Abschnitt zur angewandten Barrierefreiheit behandelt.

# --instructions--

Erstellen Sie ein zweites `p`-Element mit folgendem Kitty-Ipsum-Text: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

Erstellen Sie dann ein `main`-Element und verschachteln Sie nur die zwei `p`-Elemente innerhalb des `main`-Elements.

# --hints--

Sie sollten 2 `p`-Elemente mit Kitty-Ipsum-Text haben.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

Jedes Ihrer `p`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Ihr `p`-Element sollte die ersten Wörter des bereitgestellten zusätzlichen `kitty ipsum`-Texts enthalten.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

Ihr Code sollte ein `main`-Element enthalten.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Das `main`-Element sollte zwei Absatz-Elemente als Kinder haben.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

Das öffnende `main`-Tag sollte vor dem ersten Absatz-Tag stehen.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

Das schließende `main`-Tag sollte nach dem zweiten schließenden Absatz-Tag stehen.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
