---
id: bad87fee1348bd9aedd08830
title: Absenden-Schaltfläche zu einem Formular hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cp2Nkhz'
forumTopicId: 16627
dashedName: add-a-submit-button-to-a-form
---

# --description--

Fügen wir Ihrem Formular eine `submit`-Schaltfläche hinzu. Wenn Sie auf diese Schaltfläche klicken, werden die Daten aus Ihrem Formular an die URL gesendet, die Sie im `action`-Attribut Ihres Formulars angegeben haben.

Hier ist ein Beispiel für eine Absenden-Schaltfläche:

```html
<button type="submit">this button submits the form</button>
```

# --instructions--

Fügen Sie als letztes Element Ihres `form`-Elements eine Schaltfläche mit dem Typ `submit` hinzu, und verwenden Sie `Submit` als Text.

# --hints--

Ihr `form` sollte ein `button` enthalten.

```js
const form = document.querySelector('form');
const children = form.querySelectorAll(`:scope ${"button"}`); 
assert.isNotEmpty(children); 
```

Ihre Absenden-Schaltfläche sollte das Attribut `type` mit dem Wert `submit` haben.

```js
assert.strictEqual(document.querySelector('button').getAttribute('type'), 'submit');
```

Ihre Absenden-Schaltfläche sollte nur den Text `Submit` enthalten.

```js
const text = document.querySelector('button').textContent;
assert.match(text,/^\s*submit\s*$/gi);
```

Ihr `button`-Element sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.strictEqual(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
