---
id: bad87fee1348bd9aede08830
title: Ein Formular-Element erstellen
challengeType: 0
forumTopicId: 16817
dashedName: create-a-form-element
---

# --description--

Sie können Webformulare erstellen, die tatsächlich Daten an einen Server absenden, und zwar mit reinem HTML. Das erreichen Sie, indem Sie ein `action`-Attribut an Ihrem `form`-Element angeben.

Zum Beispiel:

```html
<form action="url-where-you-want-to-submit-form-data">
  <input>
</form>
```

# --instructions--

Verschachteln Sie das vorhandene `input`-Element innerhalb eines `form`-Elements und weisen Sie dem `"https://www.freecatphotoapp.com/submit-cat-photo"`-Attribut des `action`-Elements den Wert `form` zu.

# --hints--

Das vorhandene `input`-Element sollte innerhalb eines `form`-Elements verschachtelt sein.

```js
const inputElem = document.querySelector('form input');
assert.strictEqual(inputElem.getAttribute('type'),'text'); 
assert.strictEqual(inputElem.getAttribute('placeholder'),'cat photo URL');
```

Ihr `form` sollte ein `action`-Attribut haben, das auf `https://www.freecatphotoapp.com/submit-cat-photo` gesetzt ist.

```js
const action = document.querySelector('form').getAttribute('action');
assert.match(action,/^https:\/\/(www\.)?freecatphotoapp\.com\/submit-cat-photo$/i);
```

Ihr `form`-Element sollte gut geöffnete und geschlossene Tags haben.

```js
assert.match(code,/<\/form>/g);
assert.match(code,/<form [^<]*>/g);
assert.strictEqual(code.match(/<\/form>/g).length,code.match(/<form [^<]*>/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
  </form>
</main>
```
