---
id: bad87fee1348bd9aede08817
title: Ein Ankerelement in einen Absatz verschachteln
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

Sie können Links innerhalb anderer Textelemente verschachteln.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

Lassen Sie uns das Beispiel aufschlüsseln. Normaler Text ist im `p`-Element eingeschlossen:

```html
<p> Here's a ... for you to follow. </p>
```

Als Nächstes folgt das *Ankerelement* `<a>` (das ein schließendes Tag `</a>` benötigt):  

```html
<a> ... </a>
```

`target` ist ein Attribut des Ankerelements, das angibt, wo der Link geöffnet wird. Der Wert `_blank` gibt an, den Link in einem neuen Tab zu öffnen. Das `href` ist ein Attribut des Ankerelements, das die URL-Adresse des Links enthält:  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

Der Text `link to www.freecodecamp.org` innerhalb des `a`-Elements wird <dfn>Ankertext</dfn> genannt und zeigt den Link zum Anklicken an:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

Die endgültige Ausgabe des Beispiels sieht so aus:  

Hier ist ein <a href="https://www.freecodecamp.org" target="_blank">Link zu www.freecodecamp.org</a>, dem Sie folgen können.

# --instructions--

Verschachteln Sie das vorhandene `a`-Element innerhalb eines neuen `p`-Elements. Erstellen Sie kein neues Ankerelement. Der neue Absatz sollte den Text `View more cat photos` enthalten, wobei `cat photos` ein Link ist und der Rest einfacher Text.

# --hints--

Sie sollten nur ein `a`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

Das `a`-Element sollte mit "`https://www.freecatphotoapp.com`" verknüpft sein.

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

Ihr `a`-Element sollte den Ankertext `cat photos` haben.

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

Sie sollten ein neues `p`-Element erstellen. Insgesamt sollten mindestens 3 `p`-Tags in Ihrem HTML-Code vorhanden sein.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

Ihr `a`-Element sollte innerhalb Ihres neuen `p`-Elements verschachtelt sein.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

Ihr `p`-Element sollte den Text `View more ` (mit einem Leerzeichen dahinter) enthalten.

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

Ihr `a`-Element sollte <em>nicht</em> den Text `View more` enthalten.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

Jedes Ihrer `p`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

Jedes Ihrer `a`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
