---
id: bad87fee1348bd9aedf08816
title: Verknüpfen mit externen Seiten mithilfe von Ankerelementen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

Sie können `a` (*Anker*)-Elemente verwenden, um auf Inhalte außerhalb Ihrer Webseite zu verlinken.

`a`-Elemente benötigen eine Ziel-Webadresse, die als `href`-Attribut bezeichnet wird. Sie benötigen außerdem einen Ankertext. Hier ist ein Beispiel:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

Dann zeigt Ihr Browser den Text `this links to freecodecamp.org` als Link an, auf den Sie klicken können. Und dieser Link führt Sie zur Webadresse `https://www.freecodecamp.org`.

# --instructions--

Erstellen Sie ein `a`-Element, das mit `https://www.freecatphotoapp.com` verknüpft ist und den Ankertext „cat photos“ hat.

# --hints--

Ihr `a`-Element sollte den Ankertext `cat photos` haben.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

Sie benötigen ein `a`-Element, das mit `https://www.freecatphotoapp.com` verknüpft ist.

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

Ihr `a`-Element sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
