---
id: bad88fee1348bd9aedf08816
title: Verknüpfen mit internen Abschnitten einer Seite mit Ankerelementen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

`a` (*Anker*) Elemente können auch verwendet werden, um interne Links zu erstellen, mit denen man zu verschiedenen Abschnitten innerhalb einer Webseite springen kann.

Um einen internen Link zu erstellen, weisen Sie dem `href` Attribut eines Links ein Hash-Symbol `#` plus den Wert des `id` Attributs für das Element zu, zu dem Sie intern verlinken möchten, normalerweise weiter unten auf der Seite. Dann müssen Sie dasselbe `id` Attribut dem Element hinzufügen, zu dem Sie verlinken. Ein `id` ist ein Attribut, das ein Element eindeutig beschreibt.

Unten sehen Sie ein Beispiel für einen internen Ankerlink und sein Ziel-Element:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

Wenn Benutzer auf den `Contacts` Link klicken, werden sie zum Abschnitt der Webseite mit der **Kontakte** Überschrift geführt.

# --instructions--

Ändern Sie Ihren externen Link in einen internen Link, indem Sie das `href` Attribut in `#footer` ändern und den Text von `cat photos` zu `Jump to Bottom` anpassen.

Entfernen Sie das `target="_blank"` Attribut aus dem Ankerelement, da dieses bewirkt, dass das verlinkte Dokument in einem neuen Fenster-Tab geöffnet wird.

Fügen Sie dann dem `id` Element am unteren Rand der Seite ein `footer` Attribut mit dem Wert `<footer>` hinzu.

# --hints--

Es sollte nur ein Ankerelement auf Ihrer Seite vorhanden sein.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Es sollte nur ein `footer` Tag auf Ihrer Seite vorhanden sein.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

Der `a` Tag sollte ein `href` Attribut mit dem Wert "#footer" haben.

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

Der `a` Tag sollte kein `target` Attribut haben.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

Der `a` Text sollte "Jump to Bottom" sein.

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

Der `footer` Tag sollte ein `id` Attribut mit dem Wert "footer" haben.

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
