---
id: bad87fee1348bd9aedf08801
title: Mit dem Absatz-Element informieren
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

Das `p`-Element ist das bevorzugte Element für Absatztexte auf Websites. `p` steht für „paragraph“ (Absatz).

Sie können ein Absatz-Element so erstellen:

```html
<p>I'm a p tag!</p>
```

# --instructions--

Erstellen Sie unter Ihrem `p`-Element ein `h2`-Element und geben Sie ihm den Text `Hello Paragraph`.

**Hinweis:** Nach Konvention werden alle HTML-Tags in Kleinbuchstaben geschrieben, zum Beispiel `<p></p>` und nicht `<P></P>`.

# --hints--

Ihr Code sollte ein gültiges `p`-Element enthalten.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

Ihr `p`-Element sollte den Text `Hello Paragraph` enthalten.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

Ihr `p`-Element sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
