---
id: bad87fee1348bd9aedc08830
title: "HTML5 verwenden, um ein Feld als Pflichtfeld festzulegen"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMd4EcQ'
forumTopicId: 18360
dashedName: use-html5-to-require-a-field
---

# --description--

Sie können bestimmte Formularfelder als Pflichtfelder festlegen, sodass Ihre Nutzer das Formular nicht absenden können, bevor sie diese ausgefüllt haben.

Wenn Sie zum Beispiel ein Texteingabefeld als Pflichtfeld definieren möchten, fügen Sie einfach das Attribut `required` in Ihr `input`-Element ein, so: `<input type="text" required>`

# --instructions--

Machen Sie Ihr Texteingabefeld `input` zu einem `required` Pflichtfeld, damit Ihre Nutzer das Formular nicht absenden können, ohne dieses Feld auszufüllen.

Versuchen Sie dann, das Formular abzusenden, ohne Text einzugeben. Sehen Sie, wie Ihr HTML5-Formular Sie darauf hinweist, dass das Feld erforderlich ist?

# --hints--

Ihr Texteingabe-`input`-Element sollte das `required`-Attribut haben.

```js
assert.isTrue(document.querySelector('input').required);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" required placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
