---
id: bad87fee1348bd9aedf08812
title: Bilder zu Ihrer Website hinzufügen
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

Sie können Bilder zu Ihrer Website hinzufügen, indem Sie das `img`-Element verwenden und mit dem `src`-Attribut auf die URL eines bestimmten Bildes verweisen.

Ein Beispiel dafür wäre:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

Beachten Sie, dass `img` ein Void-Element ist.

Alle `img`-Elemente **müssen** ein `alt`-Attribut haben. Der Text im `alt`-Attribut wird für Screenreader verwendet, um die Barrierefreiheit zu verbessern, und wird angezeigt, wenn das Bild nicht geladen werden kann.

**Hinweis:** Wenn das Bild rein dekorativ ist, ist die Verwendung eines leeren `alt`-Attributs eine bewährte Vorgehensweise.

Idealerweise sollte das `alt`-Attribut keine Sonderzeichen enthalten, es sei denn, sie sind notwendig.

Fügen wir unserem obigen `alt`-Beispiel ein `img`-Attribut hinzu:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

Versuchen wir, ein Bild zu unserer Website hinzuzufügen:

Fügen Sie innerhalb des vorhandenen `main`-Elements ein `img`-Element vor den vorhandenen `p`-Elementen ein.

Setzen Sie nun das `src`-Attribut so, dass es auf die URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg` zeigt.

Vergessen Sie abschließend nicht, Ihrem `img`-Element ein `alt`-Attribut mit passendem Text zu geben.

# --hints--

Ihre Seite sollte ein Bild-Element enthalten.

```js
assert.exists(document.querySelector('img'));
```

Ihr Bild sollte ein `src`-Attribut haben, das auf das Kätzchenbild zeigt.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

Das `alt`-Attribut Ihres Bild-Elements sollte nicht leer sein.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
