---
id: bad87fee1348bd9aedf08830
title: Platzhaltertext zu einem Textfeld hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

Platzhaltertext ist der Text, der in Ihrem `input` Element angezeigt wird, bevor der Benutzer eine Eingabe gemacht hat.

Sie können Platzhaltertext folgendermaßen erstellen:

```html
<input type="text" placeholder="this is placeholder text">
```

**Hinweis:** Denken Sie daran, dass `input` ein Void-Element ist.

# --instructions--

Setzen Sie den `placeholder`-Wert Ihres Text-`input` auf "cat photo URL".

# --hints--

Sie sollten ein `placeholder`-Attribut zum vorhandenen Text-`input`-Element hinzufügen.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

Sie sollten den Wert Ihres `placeholder`-Attributs auf `cat photo URL` setzen.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

Das fertige `input`-Element sollte kein schließendes Tag haben.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

Das fertige `input`-Element sollte eine gültige Syntax haben.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
