---
id: bad87fee1348bd9aedf08834
title: Ein Set von Optionsschaltflächen erstellen
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

Sie können <dfn>Optionsschaltflächen</dfn> für Fragen verwenden, bei denen der Benutzer nur eine Antwort aus mehreren Optionen geben soll.

Optionsschaltflächen sind eine Art von `input`.

Jede Ihrer Optionsschaltflächen kann in ihrem eigenen `label`-Element verschachtelt sein. Wenn Sie ein `input`-Element in ein `label`-Element einwickeln, wird die Optionsschaltflächen-Eingabe automatisch mit dem sie umgebenden Bezeichnungselement verknüpft.

Alle zusammengehörigen Optionsschaltflächen sollten dasselbe `name`-Attribut haben, um eine Optionsschaltflächengruppe zu erstellen. Durch das Erstellen einer solchen Gruppe wird beim Auswählen einer einzelnen Optionsschaltfläche automatisch die Auswahl der anderen Schaltflächen in derselben Gruppe aufgehoben, sodass nur eine Antwort vom Benutzer gegeben wird.

Hier ist ein Beispiel für eine Optionsschaltfläche:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

Es gilt als bewährte Vorgehensweise, ein `for`-Attribut auf dem `label`-Element zu setzen, mit einem Wert, der dem Wert des `id`-Attributs des `input`-Elements entspricht. Dies ermöglicht assistiven Technologien, eine verknüpfte Beziehung zwischen der Bezeichnung und dem zugehörigen `input`-Element herzustellen. Zum Beispiel:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

Wir können das `input`-Element auch innerhalb der `label`-Tags verschachteln:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

Fügen Sie Ihrem Formular ein Paar Optionsschaltflächen hinzu, jeweils in ihrem eigenen `label`-Element verschachtelt. Eine sollte die Option `indoor` haben, die andere die Option `outdoor`. Beide sollten das `name`-Attribut `indoor-outdoor` teilen, um eine Optionsschaltflächengruppe zu erstellen.

# --hints--

Ihre Seite sollte zwei `radio`-Schaltflächenelemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

Ihren Optionsschaltflächen sollte das `name`-Attribut `indoor-outdoor` zugewiesen werden.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

Jedes Ihrer zwei Optionsschaltflächenelemente sollte in seinem eigenen `label`-Element verschachtelt sein.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

Jedes Ihrer `label`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

Eine Ihrer Optionsschaltflächen sollte die Bezeichnung `indoor` haben.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

Eine Ihrer Optionsschaltflächen sollte die Bezeichnung `outdoor` haben.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

Jedes Ihrer Optionsschaltflächenelemente sollte innerhalb des `form`-Tags hinzugefügt werden.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
