---
id: bad87fee1348bd9aedf08820
title: Ein Bild in einen Link umwandeln
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cRdBnUr'
forumTopicId: 18327
dashedName: turn-an-image-into-a-link
---

# --description--

Sie können Elemente in Links verwandeln, indem Sie sie innerhalb eines `a`-Elements verschachteln.

Verschachteln Sie Ihr Bild innerhalb eines `a`-Elements. Hier ist ein Beispiel:

```html
<a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
```

Denken Sie daran, `#` als `href`-Eigenschaft Ihres `a`-Elements zu verwenden, um daraus einen toten Link zu machen.

# --instructions--

Platzieren Sie das vorhandene Bildelement innerhalb eines `a` (*Anker*)-Elements.

Wenn Sie das getan haben, fahren Sie mit dem Cursor über Ihr Bild. Der normale Zeiger Ihres Cursors sollte sich in den Link-Klick-Zeiger verwandeln. Das Foto ist jetzt ein Link.

# --hints--

Das vorhandene `img`-Element sollte innerhalb eines `a`-Elements verschachtelt sein.

```js
const anchor = document.querySelectorAll('a')[1];
const children = anchor.querySelectorAll("img");
assert.notEmpty(children);
```

Ihr `a`-Element sollte ein toter Link mit einem `href`-Attribut sein, das auf `#` gesetzt ist.

```js
const anchor = document.querySelectorAll('a')[1];
const parentHREF = anchor.querySelector("img").parentNode.getAttribute('href');
assert.match(parentHREF,new RegExp('#'));
```

Jedes Ihrer `a`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length, code.match(/<a/g).length)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
