---
id: bad87fee1348bd9aedf08817
title: Tote Links mit dem Rautezeichen erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

Manchmal möchten Sie `a`-Elemente auf Ihrer Website hinzufügen, bevor Sie wissen, wohin sie verlinken werden.

Das ist auch praktisch, wenn Sie das Verhalten eines Links mit `JavaScript` ändern, was wir später lernen werden.

# --instructions--

Der aktuelle Wert des `href`-Attributs ist ein Link, der auf "`https://www.freecatphotoapp.com`" zeigt. Ersetzen Sie den Wert des `href`-Attributs durch ein `#`, auch bekannt als Rautezeichen, um einen toten Link zu erstellen.

Zum Beispiel: `href="#"`

# --hints--

Ihr `a`-Element sollte ein toter Link sein, bei dem der Wert des `href`-Attributs auf "#" gesetzt ist.

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
