---
id: bad87fee1348bd9aedf08835
title: Ein Satz von Kontrollkästchen erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cqrkJsp'
forumTopicId: 16821
dashedName: create-a-set-of-checkboxes
---

# --description--

Formulare verwenden häufig <dfn>Kontrollkästchen</dfn> für Fragen, die mehr als eine Antwort zulassen.

Kontrollkästchen sind eine Art von `input`.

Jedes Ihrer Kontrollkästchen kann innerhalb eines eigenen `label`-Elements verschachtelt sein. Wenn Sie ein `input`-Element in ein `label`-Element einbetten, wird das Kontrollkästchen-Eingabeelement automatisch mit dem umgebenden Bezeichnungselement verknüpft.

Alle zugehörigen Kontrollkästchen-Eingaben sollten dasselbe `name`-Attribut haben.

Es gilt als bewährte Vorgehensweise, die Beziehung zwischen einem Kontrollkästchen-`input` und seiner entsprechenden `label` explizit zu definieren, indem Sie das `for`-Attribut des `label`-Elements so setzen, dass es mit dem `id`-Attribut des zugehörigen `input`-Elements übereinstimmt.

Hier ist ein Beispiel für ein Kontrollkästchen:

```html
<label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
```

# --instructions--

Fügen Sie Ihrem Formular einen Satz von drei Kontrollkästchen hinzu. Jedes Kontrollkästchen sollte innerhalb eines eigenen `label`-Elements verschachtelt sein. Alle drei sollten das `name`-Attribut `personality` gemeinsam haben.

# --hints--

Ihre Seite sollte drei Kontrollkästchen-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

Jedes Ihrer drei Kontrollkästchen-Elemente sollte in einem eigenen `label`-Element verschachtelt sein.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="checkbox"]:only-child'),3);
```

Stellen Sie sicher, dass jedes Ihrer `label`-Elemente ein schließendes Tag hat.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length)
```

Ihre Kontrollkästchen sollten das `name`-Attribut `personality` erhalten.

```js
assert.lengthOf([...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.name === "personality"),3);
```

Jedes Ihrer Kontrollkästchen sollte innerhalb des `form`-Tags hinzugefügt werden.

```js
assert.strictEqual(document.querySelector('label').parentNode.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="playful"><input id="playful" type="checkbox" name="personality">Playful</label>
    <label for="lazy"><input id="lazy" type="checkbox" 
name="personality">Lazy</label>
    <label for="evil"><input id="evil" type="checkbox" 
name="personality">Evil</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
