---
id: bad87fee1348bd9aedf08804
title: HTML auskommentieren
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

Denken Sie daran, dass Sie zum Starten eines Kommentars `<!--` verwenden müssen und zum Beenden eines Kommentars `-->`.

Hier müssen Sie den Kommentar beenden, bevor Ihr `h2`-Element beginnt.

# --instructions--

Kommentieren Sie Ihr `h1`-Element und Ihr `p`-Element aus, aber nicht Ihr `h2`-Element.

# --hints--

Ihr `h1`-Element sollte auskommentiert sein, damit es auf der Seite nicht sichtbar ist.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

Ihr `h2`-Element sollte nicht auskommentiert sein, damit es auf der Seite sichtbar ist.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

Ihr `p`-Element sollte auskommentiert sein, damit es auf der Seite nicht sichtbar ist.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

Jeder Ihrer Kommentare sollte mit `-->` geschlossen werden.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

Sie sollten die Reihenfolge der `h1`-, `h2`- oder `p`-Elemente im Code nicht ändern.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
