---
id: bad87fee1348bd9aedf08802
title: HTML-Kommentare entfernen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

Kommentieren ist eine Möglichkeit, anderen Entwicklern Kommentare im Code zu hinterlassen, ohne dass dies die Ausgabe beeinflusst, die dem Endnutzer angezeigt wird.

Kommentieren ist auch eine praktische Methode, um Code vorübergehend zu deaktivieren, ohne ihn vollständig löschen zu müssen.

Kommentare in HTML beginnen mit `<!--` und enden mit einem `-->`

# --instructions--

Entfernen Sie die Kommentare bei Ihren `h1`, `h2` und `p` Elementen.

# --hints--

Ihr `h1` Element sollte durch das Entfernen der Kommentare auf der Seite sichtbar sein.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

Ihr `h2` Element sollte durch das Entfernen der Kommentare auf der Seite sichtbar sein.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

Ihr `p` Element sollte durch das Entfernen der Kommentare auf der Seite sichtbar sein.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

Es sollten keine abschließenden Kommentar-Tags auf der Seite sichtbar sein (z. B. `-->`).

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
