---
id: bad87fee1348bd9aedf08833
title: Platzhaltertext in eine Lücke einfügen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

Webentwickler verwenden traditionell <dfn>lorem-ipsum-Text</dfn> als Platzhaltertext. Der lorem-ipsum-Text wird zufällig aus einem berühmten Abschnitt von Cicero aus dem antiken Rom entnommen.

Lorem-ipsum-Text wird seit dem 16. Jahrhundert von Setzern als Platzhaltertext verwendet, und diese Tradition setzt sich im Web fort.

Nun, 5 Jahrhunderte sind lang genug. Da wir eine CatPhotoApp bauen, verwenden wir stattdessen einen sogenannten „kitty ipsum“-Text.

# --instructions--

Ersetzen Sie den Text innerhalb Ihres `p`-Elements durch die ersten Wörter dieses kitty-ipsum-Texts: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

Ihr `p`-Element sollte die ersten Wörter des bereitgestellten „kitty ipsum“-Texts enthalten.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
