---
id: 698367b9d1c6914a22095aec
title: Implementieren Sie den Bubble-Sort-Algorithmus
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

Für dieses Labor implementieren Sie den Bubble-Sort-Algorithmus. Er beginnt am Anfang eines unsortierten Arrays und „blubbert“ unsortierte Werte zum Ende hinauf, indem er das Array so lange durchläuft, bis es vollständig sortiert ist. Dabei vergleicht er benachbarte Elemente und tauscht sie, wenn sie in der falschen Reihenfolge sind. Die Methode durchläuft das Array so lange, bis keine Vertauschungen mehr stattfinden, woraufhin das Array sortiert ist.

Diese Methode erfordert mehrere Durchläufe durch das Array und hat im Durchschnitt und im schlechtesten Fall eine quadratische Zeitkomplexität. Obwohl einfach, ist sie in den meisten Situationen meist unpraktisch.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Schreiben Sie eine Funktion `bubbleSort`, die ein Array ganzer Zahlen als Eingabe erhält und ein Array dieser Zahlen in aufsteigender Reihenfolge zurückgibt.

# --hints--

Sie sollten eine Funktion mit dem Namen `bubbleSort` haben.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` sollte ein sortiertes Array (von klein nach groß) zurückgeben.

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` sollte ein Array zurückgeben, das bis auf die Reihenfolge unverändert ist.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` sollte die eingebaute `.sort()`-Methode nicht verwenden.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
