---
id: 68e51d8f163d3f8b70f9c4e5
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على العمل مع الطريقتين `includes()` و `slice()`.

لبدء العمل، أنشئ متغيرًا باسم `fccSentence` وقم بتعيين السلسلة النصية `"freeCodeCamp is a great place to learn web development."` له.

# --hints--

يجب إعلان متغير باسم `fccSentence`.

```js
assert.exists(fccSentence);
```

يجب أن يكون المتغير `fccSentence` سلسلة نصية.

```js
assert.isString(fccSentence);
```

يجب تعيين السلسلة النصية `"freeCodeCamp is a great place to learn web development."` إلى المتغير `fccSentence`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
