---
id: 68e6172975008fea742a63d0
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

يمكنك استخدام `slice()` لاستخراج جزء من بداية سلسلة نصية.

أنشئ متغيرًا باسم `greetingWord`. ثم قم بتعيين نتيجة استخدام `message.slice(0, 7)` لاستخراج الكلمة الأولى `"Welcome"`.

ثم استخدم إما قالب نصي أو ربط سلاسل نصية لتسجيل الرسالة `The first word is "<greetingWord>".` في وحدة التحكم. استبدل `<greetingWord>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب إعلان متغير باسم `greetingWord`.

```js
assert.exists(greetingWord);
```

يجب أن يكون متغير `greetingWord` سلسلة نصية.

```js
assert.isString(greetingWord);
```

يجب استخدام `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

يجب تعيين نتيجة `message.slice(0, 7)` إلى متغير `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

يجب أن يساوي متغير `greetingWord` القيمة `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

يجب تسجيل الرسالة `The first word is "<greetingWord>".`، حيث يجب استبدال `<greetingWord>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
