---
id: 68e6172975008fea742a63ca
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

استخدم الآن إما قالب نصي أو ربط سلاسل نصية لتسجيل الرسالة `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` في وحدة التحكم. استبدل `<hasFreeCodeCamp>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن تسجل الرسالة `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`، حيث يجب استبدال `<hasFreeCodeCamp>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
