---
id: 68e6172975008fea742a63d1
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

يمكن لطريقة `slice()` استخدام مؤشرات سالبة للعد من نهاية السلسلة النصية. يشير `-1` إلى الحرف الأخير.

إليك مثالاً:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

أنشئ متغيرًا باسم `endPunctuation`. ثم عيّن له نتيجة استخدام `message.slice(-1)` لاستخراج الحرف الأخير.

ثم استخدم إما قالب نصي أو ربط سلاسل نصية لتسجيل الرسالة `The ending punctuation mark is a "<endPunctuation>"` في وحدة التحكم. استبدل `<endPunctuation>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب إعلان متغير باسم `endPunctuation`.

```js
assert.exists(endPunctuation);
```

يجب أن يكون متغير `endPunctuation` سلسلة نصية.

```js
assert.isString(endPunctuation);
```

يجب استخدام `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

يجب تعيين نتيجة `message.slice(-1)` إلى متغير `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

يجب أن يساوي متغير `endPunctuation` القيمة `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

يجب تسجيل الرسالة `The ending punctuation mark is a "<endPunctuation>"`، حيث يجب استبدال `<endPunctuation>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
