---
id: 68e6172975008fea742a63cf
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

حان الوقت الآن لتسجيل نتيجة استخدام الطريقة `slice()`.

استخدم إما قالب نصي أو ربط سلاسل نصية لتسجيل الرسالة `The word "<platform>" was sliced from the message.` في وحدة التحكم. استبدل `<platform>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن تسجل الرسالة `The word "<platform>" was sliced from the message.`، حيث يجب استبدال `<platform>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
