---
id: 68e6172875008fea742a63c9
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

كما تتذكر من الدروس السابقة، تتحقق الطريقة `includes()` مما إذا كانت السلسلة النصية تحتوي على سلسلة فرعية محددة وتُرجع `true` أو `false`.

إليك مثالًا يستخدم الطريقة `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

أنشئ متغيرًا باسم `hasFreeCodeCamp`. ثم عيّن له نتيجة استخدام الطريقة `includes()` على `fccSentence` للتحقق مما إذا كانت تحتوي على `"freeCodeCamp"`.

# --hints--

يجب إعلان متغير باسم `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

يجب استخدام `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

يجب تعيين نتيجة `fccSentence.includes("freeCodeCamp")` إلى متغير `hasFreeCodeCamp` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

يجب أن تكون نتيجة `fccSentence.includes("freeCodeCamp")` هي `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
