---
id: 68e6172975008fea742a63cc
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

طريقة `includes()` تميز بين الأحرف الكبيرة والصغيرة، مما يعني أنها تفرق بين الأحرف الكبيرة والأحرف الصغيرة.

أنشئ متغيرًا باسم `hasLowercaseFCC`. ثم عيّن له نتيجة استخدام طريقة `includes()` للتحقق مما إذا كانت `fccSentence` تحتوي على `"freecodecamp"` (كلها أحرف صغيرة).

ثم استخدم إما قالب نصي أو دمج سلاسل نصية لتسجيل الرسالة `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` في وحدة التحكم. استبدل `<hasLowercaseFCC>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب إعلان متغير باسم `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

يجب استخدام `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

يجب تعيين نتيجة `fccSentence.includes("freecodecamp")` إلى متغير `hasLowercaseFCC` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

يجب أن تكون نتيجة `fccSentence.includes("freecodecamp")` هي `false`.

```js
assert.isFalse(hasLowercaseFCC);
```

يجب تسجيل الرسالة التي تتضمن `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`، حيث يجب استبدال `<hasLowercaseFCC>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
