---
id: 68e6172975008fea742a63ce
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

تُستخرج طريقة `slice()` قسمًا من سلسلة نصية. تأخذ معلمتين: فهرس البداية (مشمولة) وفهرس النهاية (غير مشمول).

إليك مثالًا:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

تبدأ فهارس السلاسل النصية من 0. في سلسلة `message`، الحرف عند الفهرس 11 هو `f` (بداية "freeCodeCamp").

أنشئ متغيرًا باسم `platform`. ثم عيّن له نتيجة استخدام `message.slice(11, 23)` لاستخراج `"freeCodeCamp"`.

# --hints--

يجب إعلان متغير باسم `platform`.

```js
assert.exists(platform);
```

يجب أن يكون متغير `platform` سلسلة نصية.

```js
assert.isString(platform);
```

يجب استخدام `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

يجب تعيين نتيجة `message.slice(11, 23)` إلى متغير `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

يجب أن يساوي متغير `platform` القيمة `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
