---
id: 68e6172975008fea742a63cd
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

الآن ستستكشف طريقة `slice()` التي تستخرج جزءًا من سلسلة نصية.

أنشئ متغيرًا باسم `message` وقم بتعيين السلسلة النصية `"Welcome to freeCodeCamp!"` له.

ثم استخدم `console.log()` لتسجيل `"Here are some examples of the slice() method:"` في وحدة التحكم.

# --hints--

يجب إعلان متغير باسم `message`.

```js
assert.exists(message);
```

يجب أن يكون المتغير `message` سلسلة نصية.

```js
assert.isString(message);
```

يجب أن يكون المتغير `message` هو `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

يجب تسجيل `"Here are some examples of the slice() method:"` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
