---
id: 68e6172975008fea742a63cb
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

حان الوقت الآن لرؤية ما يحدث عندما لا يتم العثور على السلسلة الفرعية.

أنشئ متغيرًا باسم `hasJavaScript`. ثم عيّن له نتيجة استخدام الطريقة `includes()` للتحقق مما إذا كانت `fccSentence` تحتوي على `"JavaScript"`.

ثم استخدم إما قالب نصي أو ربط سلاسل نصية لتسجيل الرسالة `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` في وحدة التحكم. استبدل `<hasJavaScript>` بالقيمة الفعلية للمتغير.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب إعلان متغير باسم `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

يجب استخدام `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

يجب تعيين نتيجة `fccSentence.includes("JavaScript")` إلى متغير `hasJavaScript` الخاص بك.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

يجب أن تكون نتيجة `fccSentence.includes("JavaScript")` هي `false`.

```js
assert.isFalse(hasJavaScript);
```

يجب تسجيل الرسالة `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`، حيث يجب استبدال `<hasJavaScript>` بالقيمة الفعلية للمتغير.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
