---
id: 68ef6e42fa6e8a9a18a401fa
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

استخدم `console.log()` لطباعة `"Here are some examples of the includes() method:"` في وحدة التحكم.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

يجب أن تطبع `"Here are some examples of the includes() method:"` في وحدة التحكم.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
