---
id: 68c18b1d0fce081f85312d97
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

In una lezione precedente, hai imparato che un attributo con prefisso di un singolo underscore è destinato all'uso interno per convenzione.

Modifica entrambi gli attributi `name` e `level` in `_name` e `_level`, dato che non dovrebbero essere modificati dall'esterno della loro classe.

Nota che questo non impedisce di consultare o modificare l'attributo dall'esterno della classe. Inoltre, in Python c'è sempre un modo per consultare anche gli attributi privati (con prefisso di doppio underscore).

# --hints--

Il tuo metodo `__init__` dovrebbe impostare `self._name` su `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Il tuo metodo `__init__` dovrebbe impostare `self._level` su `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
