---
id: 68c15a2a1f0e70ca7154aba5
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop costruirai un sistema per tracciare gli stipendi dei dipendenti.

Inizia creando una classe chiamata `Employee`. Al suo interno crea un metodo `__init__` con i parametri `self`, `name` e `level`. All'interno del metodo `__init__` assegna `name` e `level` agli attributi dell'istanza con lo stesso nome.

# --hints--

Dovresti avere una classe chiamata `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

La tua classe `Employee` dovrebbe avere un metodo `__init__` con tre parametri `self`, `name` e `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Il tuo metodo `__init__` dovrebbe assegnare `self.name` a `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Il tuo metodo `__init__` dovrebbe assegnare `self.level` a `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
