---
id: 68c16924498728d259e4781c
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Ora crea un'istanza della classe `Employee` passando le stringhe `Charlie Brown` e `trainee`. Assegna l'istanza a una variabile chiamata `charlie_brown`.

# --hints--

Dovresti avere una variabile chiamata `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Dovresti assegnare `Employee('Charlie Brown', 'trainee')` a `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level


--fcc-editable-region--
```
