---
id: 68c27fa1697ceb7c578b2c9f
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Il decoratore `@property` si usa in Python per trasformare un metodo in una proprietà. Viene tipicamente usato per definire metodi getter, cioè metodi che servono a recuperare il valore di un attributo:

```py
class Person:
    def __init__(self, name):
        self._name = name

    @property
    def name(self):
        return self._name

p = Person('Alice')
print(p.name)  # Alice
```

Crea un metodo chiamato `name` con un parametro `self` e decoralo con `@property`. All’interno del metodo, restituisci `self._name`.
    
# --hints--

La tua classe `Employee` dovrebbe avere un metodo `name` con un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_args("self")`)) })
```

Il tuo metodo `name` dovrebbe essere decorato con `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_decorators("property")`)) })
```

Il tuo metodo `name` dovrebbe restituire `self._name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("name").has_return("self._name")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
