---
id: 68c2d2414f5e111c02f4db5e
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Il metodo `__repr__` viene chiamato automaticamente quando chiami la funzione `repr`. Per vederlo in azione, stampa il risultato della chiamata a `repr(charlie_brown)` alla fine del tuo codice.

# --hints--

Dovresti stampare `repr(charlie_brown)`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(repr(charlie_brown))")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)

--fcc-editable-region--
```
