---
id: 68c28b7cc3ab8f4daf144b5e
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Seguendo quanto fatto nei passaggi precedenti, crea un getter per l'attributo `_level`.

# --hints--

La tua classe `Employee` dovrebbe avere un metodo `level` con un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_args("self")`)) })
```

Il tuo metodo `level` dovrebbe essere decorato con `@property`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_decorators("property")`)) })
```

Il tuo metodo `level` dovrebbe restituire `self._level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("level").has_return("self._level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    
--fcc-editable-region--
charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)

```
