---
id: 68c281cecda94ba2c6838e17
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Ora che hai definito un getter per `name`, puoi consultare l'attributo `_name` tramite la proprietà `name`. Quindi stampa `charlie_brown.name` sulla console.

# --hints--

Dovresti stampare `charlie_brown.name` sulla console.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--
```
