---
id: 68c27eb292bfec6b9402d58f
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Aggiungi un metodo `__str__` alla classe `Employee`. Fallo restituire una f-string con il formato `name: level`, sostituendo `name` e `level` con i rispettivi attributi.

Dopodiché, stampa `charlie_brown` sulla console.

# --hints--

La tua classe `Employee` dovrebbe avere un metodo `__str__` con un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Il tuo metodo `__str__` dovrebbe restituire una f-string con il formato `name: level`, sostituendo `name` e `level` con i valori di `self._name` e `self._level`, rispettivamente.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Dovresti stampare `charlie_brown` sulla console.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--
```
