---
id: 68c28d659d345270d4a61084
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Il metodo `__repr__` è un metodo speciale che dovrebbe restituire una rappresentazione in stringa dell’oggetto che può essere usata per istanziarlo.

Per esempio, il metodo `__repr__` di `Employee('Charlie', 'developer')` dovrebbe restituire la stringa `Employee('Charlie', 'developer')`, che è la stessa stringa usata per creare l’oggetto.

Dai alla tua classe `Employee` un metodo `__repr__` con un parametro `self` e fallo restituire una stringa che può essere usata per istanziare l’oggetto.

# --hints--

La tua classe `Employee` dovrebbe avere un metodo `__repr__` con un parametro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Il tuo metodo `__repr__` dovrebbe restituire la stringa usata per istanziare l’oggetto.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
