---
id: 68c7f7012a700243eff1cbc0
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Il livello usato per istanziare un dipendente dovrebbe essere scelto tra livelli specifici. Userai `_base_salaries` per convalidare il livello e impostare lo stipendio corretto per il dipendente.

Dopo la tua istruzione `if` esistente, crea un'altra istruzione `if` che verifica se `level` non è in `Employee._base_salaries`.

All’interno della nuova istruzione `if`, genera una `ValueError` con il messaggio `Invalid value '{level}' for 'level' attribute.`, dove `{level}` dovrebbe essere sostituito dal valore dell’argomento `level`.

# --hints--

Quando `level` non è in `Employee._base_salaries`, dovresti generare una `ValueError` con il messaggio `Invalid value '{level}' for 'level' attribute.`, dove `{level}` dovrebbe essere sostituito dal valore dell’argomento `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Non dovresti generare alcuna eccezione quando `level` è in `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
