---
id: 68c809c990f253912a9e9209
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

Come hai imparato in una lezione precedente, un setter offre un modo per controllare come un attributo può essere modificato. Per assicurarti che `new_name` sia del tipo corretto, crea un'istruzione `if` che sollevi un `TypeError` con il messaggio `'name' must be a string.` quando `new_name` non è un'istanza di `str`.

# --hints--

Dovresti avere un'istruzione `if` all'interno del tuo setter `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].find_ifs()[0]`)) })
```

Dovresti sollevare un `TypeError` con il messaggio `'name' must be a string.` quando `new_name` non è un'istanza di `str`.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_names = [0, True, []]
  for i in non_names:
    try:
      emp.name = i
    except TypeError as e:
      assert str(e) == "'name' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_name"
`) })
```

Non dovresti sollevare alcuna eccezione quando `new_name` è una stringa.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.name = "Jack"
  except Exception:
    assert False, "Expected not to raise any exception with valid new_name"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
--fcc-editable-region--
        
--fcc-editable-region--
        self._name = new_name

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
