---
id: 68c9cab4b1118da59eecfc56
title: Passaggio 31
challengeType: 20
dashedName: step-31
---

# --description--

Dopo le istruzioni `if` esistenti, crea un'altra per sollevare un `ValueError` quando `new_level` è già il livello selezionato.

Nota che `_level` non esiste ancora durante l'inizializzazione, quindi usa `hasattr(self, '_level')` per verificare se esiste prima di confrontare. Questo evita un `AttributeError` quando l'oggetto viene creato per la prima volta.

Per il messaggio, usa `'{level}' is already the selected level.`, dove `{level}` dovrebbe essere sostituito dal livello corrente.

# --hints--

Dovresti avere una terza istruzione `if` all'interno del tuo setter `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Quando `new_level` è uguale a `self.level`, dovresti sollevare un `ValueError` con il messaggio `'{level}' is already the selected level.`, dove `{level}` dovrebbe essere sostituito dal livello corrente.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

La tua istruzione `if` dovrebbe usare `hasattr(self, '_level')` per verificare se `_level` esiste prima di confrontare.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
