---
id: 646c47867800472a4ed5d2ea
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Empieza con la estructura básica de HTML. Añade una referencia `DOCTYPE` de `html` y un elemento `html` con su atributo `lang` establecido en `en`. Además, agrega un `head` y un elemento `body` dentro del elemento `html`.

# --hints--

Debes tener la declaración `DOCTYPE` de `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Tu declaración `DOCTYPE` debe estar al comienzo de tu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener una etiqueta de apertura `html` con un atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Debes tener una etiqueta de cierre `html`.

```js
assert(code.match(/<\/html>/i));
```

Debes tener una etiqueta de apertura `head`.

```js
assert(code.match(/<head>/i));
```

Debes tener una etiqueta `head` de cierre.

```js
assert(code.match(/<\/head>/i));
```

Debes tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body>/i));
```

Debes tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body>/i));
```

Tu elemento `body` debe estar después del elemento `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Tus elementos `head` y `body` deben estar dentro del elemento `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
