---
id: 646c5e727e56e743c9aed4a1
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

La propiedad de posicionamiento en CSS te permite establecer cómo deseas que un elemento se posicione en el navegador. Tiene una propiedad `position` que se puede establecer en `static`, `absolute`, `relative`, `sticky` o `fixed`.

Una vez que estableces la propiedad `position` del elemento, puedes mover el elemento ajustando un valor en píxeles o un porcentaje para una o más de las propiedades `top`, `right`, `left` o `bottom`.

`static` es el posicionamiento predeterminado para todos los elementos. Si lo asignas a un elemento, no podrás moverlo con `top`, `right`, `left` o `bottom`.

Asigna a `.cat-head` una propiedad `position` de `static`, luego establece las propiedades `top` y `left` en `100px` cada una.


# --hints--

Tu selector `.cat-head` debe tener una propiedad `position` establecida en `static`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Tu selector `.cat-head` debe tener una propiedad `top` establecida en `100px`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Tu selector `.cat-head` debe tener una propiedad `left` establecida en `100px`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
