---
id: 646c48df8674cf2b91020ecc
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Agrega un elemento `link` dentro de tu elemento `head`. Para ese elemento `link`, establece el atributo `rel` en `stylesheet` y el atributo `href` en `./styles.css`.

# --hints--

Debes tener un elemento `link`.

```js
assert.match(code, /<link/)
```

El elemento `link` debe tener un atributo `rel` y un atributo `href` establecido en `stylesheet` y `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
