---
id: 6476fb7cbaafa36d65e9cf35
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` es un valor de la propiedad `position` que te permite hacer que un elemento quede fijo en la página sin importar a dónde el usuario se desplace en la página.

Tendrás que hacer más marcas para ver cómo funciona la posicionamiento `fixed`. En tu HTML, crea un elemento `div` con la clase `box`.

# --hints--

Debe tener un elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2);
```

Tu elemento `div` debe tener la clase `box`.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
