---
id: 6477062778c85972eb648030
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Ahora debes centrar la cabeza del gato. 

Dale al elemento `.cat-head` una propiedad `position` establecida en `absolute`. Establece un valor de `0` para las propiedades `right`, `left`, `top`, `bottom`, luego establece su propiedad `margin` en todos los lados a `auto`. Esa es una manera de centrar un elemento vertical y horizontalmente usando posicionamiento con CSS.

# --hints--

Tu selector `.cat-head` debe tener una propiedad `position` establecida en `absolute`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Tu selector `.cat-head` debe tener una propiedad `top` establecida en `0`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Tu selector `.cat-head` debe tener una propiedad `left` establecida en `0`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Tu selector `.cat-head` debe tener una propiedad `right` establecida en `0`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Tu selector `.cat-head` debe tener una propiedad `bottom` establecida en `0`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Tu selector `.cat-head` debe tener una propiedad `margin` establecida en `auto`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
