---
id: 64770351e8586671ec0911f0
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

El último valor de la propiedad position es `sticky`. La posición `sticky` es una estructura híbrida entre `relative` y `fixed`. Permite que un elemento **se adhiera** a una posición específica dentro de su elemento contenedor o la ventana gráfica, según la posición del desplazamiento.

Cambia el valor de la propiedad `position` de `.cat-head` a `sticky`, establece `top` en `0`, luego elimina `left` y su valor.

**Nota**: Para ver como `sticky` funciona, tienes que colocar un par de textos antes y despues de tu elemento `.cat-head` `div`. Si desplazas hacia abajo después de eso, verás que `.cat-head` se pega a la parte superior y permanece allí.

# --hints--

Tu selector `.cat-head` debe tener una propiedad `position` establecida en `sticky`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

Tu selector `.cat-head` debe tener una propiedad `top` establecida en `0`. Asegúrate de agregar un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

No debes tener la propiedad `left` y su valor en tu código.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
