---
id: 647d855ac12cd436059acd39
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Usa el selector universal para añadir `box-sizing: border-box;` a tu CSS. Esto asegura que los elementos incluyan el relleno y borde en su ancho y alto especificados.

# --hints--

Debes tener el selector universal (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

Tu selector universal debe tener una propiedad `box-sizing` establecida en `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
