---
id: 6140c7e645d8e905819f1dd4
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie mit dem standardmäßigen Boilerplate-Code. Fügen Sie Ihre `DOCTYPE`-Deklaration, Ihr `html`-Element mit der Sprache auf Englisch gesetzt, Ihre `head`- und `body`-Elemente hinzu.

Fügen Sie Ihr `meta`-Element für das korrekte `charset`, Ihr `title`-Element und ein `link`-Element für die `./styles.css`-Datei hinzu.

Setzen Sie das `title` auf `Ferris Wheel`.

# --hints--

Ihr Code sollte die `DOCTYPE`-Deklaration enthalten.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Sie sollten nach der `DOCTYPE`-Deklaration ein Leerzeichen einfügen.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Sie sollten den Dokumenttyp auf `html` festlegen.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Sie sollten die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Ihr `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/html\s*>/);
```

Ihre `DOCTYPE`-Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Sie sollten ein öffnendes `head`-Tag haben.

```js
assert.match(code, /<head\s*>/i);
```

Sie sollten ein schließendes `head`-Tag haben.

```js
assert.match(code, /<\/head\s*>/i);
```

Sie sollten ein öffnendes `body`-Tag haben.

```js
assert.match(code, /<body\s*>/i);
```

Sie sollten ein schließendes `body`-Tag haben.

```js
assert.match(code, /<\/body\s*>/i);
```

Ihr `body`-Element sollte das erste Element nach dem `head`-Element sein.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Ihre `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Ihr `head`-Element sollte das erste Element innerhalb Ihres `html`-Elements sein.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Ihr `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Ihr `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ihr Code sollte ein `meta`-Element haben.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ihr `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` haben.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Ihr `meta`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Ihr Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr Projekt sollte einen Titel von `Ferris Wheel` haben.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Denken Sie daran, dass Groß- und Kleinschreibung sowie Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Ihr Code sollte ein `link`-Element haben.

```js
assert.match(code, /<link/);
```

Sie sollten ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
