---
id: 6894c1b715b8d18a18ef3203
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个乐器库存来练习 Python 的类和对象。

要在 Python 中定义一个类，使用 `class` 关键字，后跟类名和冒号。

```python
class ClassName:
    # class contents go here
```

请记住，在 Python 中，类名通常使用 PascalCase，这意味着每个单词的首字母大写且单词之间没有空格。

首先创建一个名为 `MusicalInstrument` 的类。暂时将 `pass` 作为该类的体。

# --hints--

你应该创建一个名为 `MusicalInstrument` 的类。

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
