---
id: 6895e5337b1914050b0af7bd
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在，在 `instrument_1` 实例上调用 `get_fact()` 方法并打印结果。

由于 `get_fact()` 返回的是字串而不是直接打印，你需要使用 `print()` 来显示返回的字串。

# --hints--

你应该在 `instrument_1` 上调用 `get_fact()` 并打印结果。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
