---
id: 6895e4e023e07603bcef46f9
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

现在让我们使用你创建的方法。 

在你的 `instrument_1` 实例上调用 `play()` 方法。这将打印有关演奏双簧管的消息。

要在实例上调用方法，请使用点符号：

```python
instance_name.method_name()
```

# --hints--

你应该在 `instrument_1` 实例上调用 `play()` 方法。

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
