---
id: 6895dffc64e97df0e62e9ef9
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

方法是在类内部定义的函数，可以使用类的数据执行操作。以下是你如何创建方法：

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

在 Python 中，`__init__` 方法是一种特殊的方法，称为初始化方法，当你创建类的一个新建实例时会调用它。

在你的 `MusicalInstrument` 类中，添加一个带有三个参数的 `__init__` 方法：`self`、`name` 和 `instrument_type`。

`self` 参数是引用类实例的标准方式，并且作为类中所有实例方法的第一个参数是必需的。

# --hints--

你应该在你的 `MusicalInstrument` 类中定义一个 `__init__` 方法。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

你的 `__init__` 方法应有三个参数：`self`、`name` 和 `instrument_type`。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
