---
id: 6895e0bdb0348bf46952ab67
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在是创建你的 `MusicalInstrument` 类的实例的时候了。

要创建一个类的实例，你像调用函数一样调用类名，传递 `__init__` 方法所需的参数（除了 `self`，它由 Python 自动处理）。

```python
my_instance = ClassName(arg1, arg2)
```

当你创建此实例时，Python 会调用你的 `__init__` 方法并将这些值作为仅属于该特定实例的实例属性保存。

在类定义之外，为一个属于 `woodwind` 乐器的 `Oboe` 创建一个名为 `instrument_1` 的实例。

# --hints--

你应该创建一个名为 `instrument_1` 的 `MusicalInstrument` 类实例。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

你应该在创建 `instrument_1` 时传递 `Oboe` 和 `woodwind` 作为参数。

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
