---
id: 6895e4940e39950288e406c2
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

接下来，为你的类添加另一个名为 `get_fact` 的方法，该方法返回关于记录仪的一个事实。

此方法应仅有 `self` 参数，并返回一个 f-字串，内容为 `The [instrument name] is part of the [instrument type] family of instruments.`

与直接打印消息的 `play` 方法不同，此方法应返回一个可在代码其他地方使用的字串。

# --hints--

你应该定义一个名为 `get_fact` 且带有 `self` 参数的方法。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

你的 `get_fact` 方法应返回一个包含记录仪名称和类型的 f-string。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
