---
id: 6895e3f91c069e0068a42582
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

在你的 `play` 方法中，使用 `print()` 来显示有关演奏记录仪的消息。

使用 f-string 在消息中包含记录仪的名称：`The [instrument name] is fun to play!`。`[instrument name]` 应替换为记录仪的实际名称。

# --hints--

你的 `play` 方法应使用 f-string 来显示消息 `The [instrument name] is fun to play!`。使用 `self.name` 来访问记录仪名称。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
