---
id: 6895e197180c9af6988e1e19
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

为一个 `brass` 记录仪的 `Trumpet` 创建第二个名为 `instrument_2` 的实例。

注意你可以从同一个类模板创建多个实例。每个实例独立保存自己的属性值，不受其他实例影响。

# --hints--

你应该创建一个名为 `instrument_2` 的 `MusicalInstrument` 类实例。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

你在创建 `instrument_2` 时应传入 `Trumpet` 和 `brass` 作为参数。

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
