---
id: 6895e2becb7890fad818ad4d
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

既然你已经了解如何访问对象属性，删除现有的 `print` 调用。

# --hints--

你应该移除 `print(instrument_1.name)` 语句。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

你应该删除 `print(instrument_1.instrument_type)` 语句。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

你应该删除 `print(instrument_2.name)` 语句。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

你应该删除 `print(instrument_2.instrument_type)` 语句。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)

--fcc-editable-region--
```
