---
id: 68e4276483b0c71658aab1d1
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

如果你需要创建一个所有字母都是小写的新的字串，可以使用 `toLowerCase()` 方法。  

首先创建一个名为 `lowerCaseInput` 的变量，并将对 `cleanedInput` 变量使用 `toLowerCase()` 方法的结果赋值给它。 

然后记录字串 `"Result of using the toLowerCase() method:"`。

最后，为 `lowerCaseInput` 变量添加另一个 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个 `lowerCaseInput` 变量。

```js
assert.exists(lowerCaseInput);
```

你的 `lowerCaseInput` 变量应是一个字串。

```js
assert.isString(lowerCaseInput);
```

你应该将 `cleanedInput.toLowerCase()` 赋值给你的 `lowerCaseInput` 变量。

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

你应该将字串 `"Result of using the toLowerCase() method:"` 记录到控制台日志。

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

你应该将 `lowerCaseInput` 变量记录到控制台。

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
