---
id: 68e4276483b0c71658aab1d0
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

有时你需要创建一个不同大小写的新建字串。 

要创建一个所有字母都转换为大写的新建字串，你可以使用 `toUpperCase()` 方法。

这是一个示例：

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

创建一个名为 `upperCaseInput` 的变量，并将对 `cleanedInput` 变量使用 `toUpperCase()` 方法的结果赋值给它。

# --hints--

你应该有一个 `upperCaseInput` 变量。

```js
assert.exists(upperCaseInput);
```

你的 `upperCaseInput` 变量应是一个字串。

```js
assert.isString(upperCaseInput);
```

你应该将 `cleanedInput.toUpperCase()` 赋值给你的 `upperCaseInput` 变量。

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
