---
id: 68e4276483b0c71658aab1ce
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

有时你只想修剪字串开头的空白字符。这时，`trimStart()` 方法就非常有用。

这是一个示例：

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

请记住，字串是不可变的。此方法将创建一个带有结果的`新建`字串。不会修改原始字串。

创建一个变量 `trimmedStart` 并将使用 `trimStart()` 方法处理 `userInput` 变量的结果赋值给它。

# --hints--

你应该有一个 `trimmedStart` 变量。

```js
assert.exists(trimmedStart);
```

你的 `trimmedStart` 变量应是一个字串。

```js
assert.isString(trimmedStart);
```

你应该将 `userInput.trimStart()` 赋值给你的 `trimmedStart` 变量。

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
