---
id: 68e4276483b0c71658aab1cc
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

你的 `userInput` 值在字串的两端有一些多余的空格。要创建一个去除这些空白的新字串，你可以使用 `trim()` 方法。  

这是一个示例：

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

创建一个名为 `cleanedInput` 的变量，并将对 `userInput` 变量使用 `trim()` 方法的结果赋值给它。

# --hints--

你应该有一个 `cleanedInput` 变量。

```js
assert.exists(cleanedInput);
```

你的 `cleanedInput` 变量应是一个字串。

```js
assert.isString(cleanedInput);
```

你应该将 `userInput.trim()` 赋值给你的 `cleanedInput` 变量。

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
