---
id: 68e4276483b0c71658aab1cd
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

现在是查看结果的时候了。

首先为你的代码添加一个带有字串 `"Result of trimming whitespace from both ends:"` 的 `console.log()`。

然后在那个 `console.log()` 下面，添加另一个带有 `cleanedInput` 变量的 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"Result of trimming whitespace from both ends:"` 记录到控制台日志。

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

你应该将 `cleanedInput` 变量记录到控制台日志。

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
