---
id: 68e4276483b0c71658aab1d2
title: 步骤 13
challengeType: 1
dashedName: step-13
---

# --description--

有时你需要使用多个方法来创建一个新建字串。在接下来的几个步骤中，你将使用迄今为止学到的一些方法来创建一个使用驼峰式大小写命名约定的新建字串。

首先创建一个名为 `lowercaseWord` 的变量，并赋值为字串 `"camelcase"`。 

在 `lowercaseWord` 变量下面，创建另一个名为 `camelCasedVersion` 的变量，并暂时赋值为空字串。

在 `camelCasedVersion` 变量下，添加一个带有字串 `"Camel cased version:"` 的 `console.log()`，然后再添加一个用于 `camelCasedVersion` 变量的 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个名为 `lowercaseWord` 的变量。

```js
assert.exists(lowercaseWord);
```

你的 `lowercaseWord` 应该是一个字串。

```js
assert.isString(lowercaseWord);
```

你应该将字串 `"camelcase"` 赋值给你的 `lowercaseWord` 变量。

```js
assert.equal(lowercaseWord, "camelcase");
```

你应该有一个 `camelCasedVersion` 变量。

```js
assert.exists(camelCasedVersion);
```

你的 `camelCasedVersion` 变量应为字串。

```js
assert.isString(camelCasedVersion);
```

你应该为你的 `camelCasedVersion` 变量赋值一个空字串。

```js
assert.isEmpty(camelCasedVersion);
```

你应该将字串 `"Camel cased version:"` 记录到控制台日志。

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

你应该将 `camelCasedVersion` 变量记录到控制台。

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
