---
id: 68e4276483b0c71658aab1ca
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将练习使用各种用于格式化字串的字串方法。

首先创建一个名为 `userInput` 的变量，并赋值为 `"   Hello World!   "`。

# --hints--

你应该有一个 `userInput` 变量。

```js
assert.exists(userInput);
```

你的 `userInput` 变量应是一个字串。

```js
assert.isString(userInput);
```

你应该将 `"   Hello World!   "` 赋值给你的 `userInput` 变量。确保在 `H` 前有三个空格，在 `!` 后有三个空格。

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
