---
id: 68e4276483b0c71658aab1d3
title: 步骤14
challengeType: 1
dashedName: step-14
---

# --description--

正如你从之前的课程和研讨会中回忆的那样，你可以使用 `slice()` 方法提取字串的部分内容。

这是一个提醒：

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

将 `camelCasedVersion` 的值从空字串更改为对 `lowercaseWord` 变量使用 `slice()` 方法的结果。传入数字 `0` 和 `5`，它们表现为 `slice()` 方法的起始和结束索引。

现在你应该在控制台看到单词 `"camel"`。

# --hints--

你的 `camelCasedVersion` 变量不应是空字串。

```js
assert.isNotEmpty(camelCasedVersion);
```

你应该将 `lowercaseWord.slice(0, 5)` 赋值给 `camelCasedVersion` 变量。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

你不应该将字面字串 `"camel"` 赋值给 `camelCasedVersion`。请重新阅读说明以了解正确的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
