---
id: 68ed6079fe0e3e98b94a1b75
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

接下来，添加一个带有字串 `"Original input:"` 的 `console.log()`。 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"Original input:"` 记录到控制台日志。

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
