---
id: 68f1179eaea3bb59ba434386
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

要去除字串末尾的空白字符，你可以使用 `trimEnd()` 方法。

这是一个示例：

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

创建一个名为 `trimmedEnd` 的变量，并将对 `userInput` 变量使用 `trimEnd()` 方法的结果赋值给它。 

# --hints--

你应该有一个 `trimmedEnd` 变量。

```js
assert.exists(trimmedEnd);
```

你的 `trimmedEnd` 变量应是一个字串。

```js
assert.isString(trimmedEnd);
```

你应该将 `userInput.trimEnd()` 赋值给你的 `trimmedEnd` 变量。

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
