---
id: 68f3f5ca1d2a490c97f14f03
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

现在是查看在字串末尾修剪空白的结果的时候了。

在你的 `trimmedEnd` 变量下方，添加一个带有字串 `"After using the trimEnd() method, trailing spaces removed:"` 的 `console.log()`。 

然后在那个 `console.log()` 下面，添加另一个带有 `trimmedEnd` 变量的 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"After using the trimEnd() method, trailing spaces removed:"` 记录到控制台日志。

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

你应该将 `trimmedEnd` 变量记录到控制台日志。

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
