---
id: 68f3f3d65d51d80ac0b412fc
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

现在是查看在字串开头修剪空白结果的时候了。

在你的 `trimmedStart` 变量下，添加一个带有字串 `"After using the trimStart() method, leading spaces removed:"` 的 `console.log()`。 

然后在那个 `console.log()` 下面，添加另一个带有 `trimmedStart` 变量的 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"After using the trimStart() method, leading spaces removed:"` 记录到控制台日志。

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

你应该将 `trimmedStart` 变量记录到控制台日志。

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
