---
id: 68f3f890010c870e17bbe754
title: 步骤11
challengeType: 1
dashedName: step-11
---

# --description--

现在是查看结果的时候了。

在你的 `upperCaseInput` 变量下方，添加一个带有字串 `"Result of using the toUpperCase() method:"` 的 `console.log()`。 

然后在那个 `console.log()` 下面，添加另一个带有 `upperCaseInput` 变量的 `console.log()`。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将字串 `"Result of using the toUpperCase() method:"` 记录到控制台日志中。

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

你应该将 `upperCaseInput` 变量记录到控制台。

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
