---
id: 68f3eeaa65614107ecd5f1f2
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

在你的第一个 `console.log()` 下面，添加另一个 `console.log()` 以将 `userInput` 变量记录到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将 `userInput` 变量记录到控制台日志。

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
