---
id: 68f405c8c9df44132c8fe1bb
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

正如你从之前的课程中回忆的，驼峰式大小写命名规范以第一个单词的小写字母开头，后续每个单词以大写字母开头，后面跟着小写字母。

这里有一些提醒示例：

```md
freeCodeCamp
camelCase
learningIsFun
```

`lowercaseWord` 变量中的第二个单词是 `"case"`。要访问该单词中的 `c`，你可以使用 `lowercaseWord[5]`。

更新 `camelCasedVersion` 赋值，使用 `+` 操作符将 `lowercaseWord.slice(0, 5)` 与对 `lowercaseWord[5]` 使用正确的字串转大写方法的结果连接起来。

现在你应该能在控制台看到 `camelC`。

# --hints--

你应该将 `lowercaseWord.slice(0, 5)` 与 `lowercaseWord[5].toUpperCase()` 连接起来。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

你不应该将字面字符串 `"camelC"` 赋值给 `camelCasedVersion`。请重新阅读说明，了解正确的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
