---
id: 68f40925d810ae14b2d5a169
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

现在你的 `camelCasedVersion` 变量保存的是 `camelC` 的值。但期望的结果是 `camelCase`。 

要获取剩余的 `ase` 字母，你可以使用带有负索引值的 `slice()` 方法。查全率，使用负索引值会告诉 `slice()` 方法从末尾提取。 

这是一个提醒示例：

```js
const str = "abcd";
str.slice(-2); // cd
```

使用 `+` 操作符将 `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` 与对 `lowercaseWord` 使用 `slice()` 方法的结果连接起来。为 `slice()` 方法传入数字 `-3`。

经过最后的改变，你的工作室就完成了！

# --hints--

你应该将 `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` 与 `lowercaseWord.slice(-3)` 连接起来。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

你不应该将字面字串 `"camelCase"` 赋值给 `camelCasedVersion`。请重新阅读说明以获得正确的方法。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
