---
id: 60a3e3396c7b40068ad6998b
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

من المفيد أن تدفع الهوامش في اتجاه واحد.

في هذه الحالة، يدفع الهامش السفلي للعنصر `.one` العنصر `.two` إلى الأسفل بمقدار 20 بكسل.

في المحدِّد `.two`، استخدم الخاصية المختصرة `margin` لتعيين الهامش العلوي إلى `0`، والهامش الأفقي إلى `auto`، والهامش السفلي إلى `20px`. هذا سيزيل الهامش العلوي، ويُوسّط العنصر أفقيًا، ويضبط الهامش السفلي إلى 20 بكسل.

# --hints--

يجب تعيين الخاصية `margin` إلى `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

يجب أن يحتوي العنصر `.two` على قيمة `margin` تساوي `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
