---
id: 60a3e3396c7b40068ad69977
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

استخدم الحشو لضبط المسافة داخل العنصر.

في `.frame`، استخدم الخاصية المختصرة `padding` لزيادة المسافة بين عنصري `.frame` و `.canvas` بمقدار `50px`. ستزيد الخاصية المختصرة المسافة في الأعلى والأسفل واليسار واليمين من حدود العنصر واللوحة بداخله.

# --hints--

يجب تعيين الخاصية `padding` إلى `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

يجب أن يحتوي عنصر `.frame` على قيمة `padding` تساوي `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
