---
id: 60a3e3396c7b40068ad69970
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

قبل أن تبدأ في تنسيق الـ `div` الذي أضفته، يجب ربط CSS الخاص بك بـ HTML.

أضف عنصر `link` لربط ملف `styles.css` الخاص بك. عيّن الخاصية `href` إلى `styles.css`، وتذكر تعيين الخاصية `rel` إلى `stylesheet`.

# --hints--

يجب أن يحتوي شفرتك على عنصر `link`.

```js
assert(/<link/.test(code))
```

يجب أن يحتوي عنصر `link` الخاص بك على خاصية `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

يجب أن يحتوي عنصر `link` الخاص بك على خاصية `href` بقيمة `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
