---
id: 60a3e3396c7b40068ad6996a
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

يجب أن تكون الآن على دراية بالعناصر الأساسية التي يجب أن تحتويها صفحة HTML.

أعد إعداد شفرتك مع إعلان `DOCTYPE`، وعنصر `html` مع تعيين اللغة إلى الإنجليزية، وعنصر `head`، وعنصر `body`.

# --hints--

يجب أن تحتوي شفرتك على إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يكون لديك وسم افتتاحي `<html>` مع خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن تحتوي شفرتك على عنصر `head` داخل عنصر `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

يجب أن تحتوي شفرتك على عنصر `body` داخل عنصر `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

يجب أن يأتي عنصر `head` قبل عنصر `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
