---
id: 60a3e3396c7b40068ad69980
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

أضف عنصر `div` آخر بقيمة `class` تساوي `two` أسفل عنصر `one` الخاص بك مباشرة. سيكون هذا المستطيل الثاني.

# --hints--

لا تغير عنصر `.one` الموجود.

```js
assert(document.querySelectorAll('.one').length === 1);
```

يجب أن يكون لديك عنصر `div` ثاني داخل عنصر `.canvas` الخاص بك.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

يجب أن يحتوي عنصر `div` الثاني على قيمة `class` تساوي `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

يجب أن يأتي عنصر `.two` بعد عنصر `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
