---
id: 60a3e3396c7b40068ad69978
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

استخدم الهوامش لضبط المسافة خارج العنصر.

باستخدام خاصية `margin`، امنح عنصر `.frame` هامشًا عموديًا بقيمة `20px`، وهامشًا أفقيًا بقيمة `auto`. سيؤدي ذلك إلى تحريك الإطار لأسفل بمقدار 20 بكسل ومحاذاته أفقيًا في وسط الصفحة.

# --hints--

يجب ضبط خاصية `margin` على القيمة `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

يجب أن يحتوي عنصر `.frame` على قيمة `margin` تساوي `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
