---
id: 60a3e3396c7b40068ad69975
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

اكتب قاعدة جديدة باستخدام محدد الصنف `.frame`.

استخدم إعلان الاختصار `border` لإعطاء عنصر `.frame` حدودًا صلبة سوداء بعرض `50px`.

# --hints--

يجب أن يحتوي شفرتك على محدد `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

يجب تعيين الخاصية `border` إلى `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

يجب أن يحتوي عنصر `.frame` على `border` بقيمة `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
