---
id: 6823ac607bfdbc46331b2559
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 代表超文本標記語言。 它是定義網頁結構和內容的代碼。 這是你的代碼編輯器，你將在這裏編寫 HTML。

在編輯器中找到第 1 行並輸入以下文本：

```md
Welcome to freeCodeCamp
```

完成後，點擊“檢查你的代碼”按鈕查看是否正確。

# --hints--

在你的編輯器中應包含文本 `Welcome to freeCodeCamp`。 仔細檢查拼寫。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
