---
id: 6823e169fda14755fbf00445
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

你逐漸掌握要領了。 接下來，在編輯器底部添加另一個 `h3` 元素，文本爲：

```md
Introduction to CSS
```

# --hints--

你應該爲頁面添加第二個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你的 `h3` 元素的文本應該是 `Introduction to CSS`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```
