---
id: 6823e036ea4b71553558c01b
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

請注意，每個標題在預覽中看起來都有些不同。段落也是如此。

接著，添加另一個 `p` 元素來顯示以下文字：

```md
HTML represents the content and structure of a webpage
```

# --hints--

你應該有第二個 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

你的 `p` 元素應該包含文字 `HTML represents the content and structure of a webpage`。你可以從說明中複製該文字。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
