---
id: 6823d6244511f252c8300eed
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

當你需要為網頁添加段落時，你可以像這樣使用 `p` 元素：

```html
<p>This is a paragraph element.</p>
```

將 `Learn the skills to become a full-stack developer` 轉換成段落元素。

# --hints--

你的 `p` 元素應該有一個開啟標籤。開啟標籤具有此語法：`<elementName>`。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素應該有一個結束標籤。結束標籤具有這種語法：`</elementName>`。

```js
assert.match(code, /<\/p\s*\>/);
```

你的 `p` 元素的文字應該是 `Learn the skills to become a full-stack developer`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

你的 `p` 元素應該位於你的 `h2` 元素下方。 

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
