---
id: 6823e637c1c0ed56f781b4fc
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

最後，JavaScript 讓你的網頁具有互動性——它讓你能告訴網頁當有人點擊按鈕、提交表單或執行許多其他動作時該怎麼做。

在工作坊的最後一步，為頁面添加另一個描述 JavaScript 的 `h3` 和 `p` 元素。

首先，為你的文件添加一個帶有文字的 `h3` 元素：

```md
Introduction to JavaScript
```

然後，在該 `h3` 元素下方，添加一個包含以下文字的 `p` 元素：

```md
JavaScript adds interactivity to a webpage
```

# --hints--

你應該為你的頁面添加第三個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

你的 `h3` 元素的文字應該是 `Introduction to JavaScript`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

你應該為頁面添加第四個 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

你的 `p` 元素應該包含文字 `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
