---
id: 6823d6244511f252c8300eed
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

當你需要將一個段落添加到網頁上時，你可以像這樣使用 `p` 元素：

```html
<p>This is a paragraph element.</p>
```

把 `Learn the skills to become a full stack developer` 轉化成一個段落元素。

# --hints--

你的 `p` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector("p"));
```

你的 `p` 元素應該有一個結束標籤。 結束標籤的語法爲：`</elementName>`。

```js
assert.match(code, /<\/p\s*\>/);
```

你的 `p` 元素的文本應該是 `Learn the skills to become a full stack developer`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

你的 `p` 元素應該在 `h2` 元素的下方。

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full stack developer
--fcc-editable-region--
```
