---
id: 6823d9ac8bdc3853df65a1ff
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

HTML 中有六個標題元素：`h1` 到 `h6`。它們用來顯示你網頁中各區段的重要性，其中 `h1` 最重要，`h6` 最不重要。

在你的 `p` 元素下方，添加一個帶有文字的 `h3` 標題：

```md
Introduction to HTML
```

# --hints--

你的 `h3` 元素應該有一個開啟的 `<h3>` 頁籤。

```js
assert.exists(document.querySelector("h3"));
```

你的 `h3` 元素應該有一個閉合的 `</h3>` 頁籤。

```js
assert.match(code, /<\/h3\s*\>/);
```

你的 `h3` 元素的文字應該是 `Introduction to HTML`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
