---
id: 6823d9ac8bdc3853df65a1ff
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

HTML 中有六個標題元素： `h1` 到 `h6`。 它們用於表明網頁各部分的重要性，其中 `h1` 最爲重要，而 `h6` 則最爲次要。

在你的 `p` 元素下方，添加一個 `h3` 標題，文本爲：

```md
Introduction to HTML
```

# --hints--

你的 `h3` 元素應該有一個開始標籤 `<h3>`。

```js
assert.exists(document.querySelector("h3"));
```

你的 `h3` 元素應該有一個結束標籤 `</h3>`。

```js
assert.match(code, /<\/h3\s*\>/);
```

你的 `h3` 元素的文本應該是 `Introduction to HTML`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
