---
id: 6823ac607bfdbc46331b2559
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML 是超文字標記語言的縮寫。它是定義網頁結構和內容的程式碼。這是你的程式碼編輯器，你將在這裡撰寫 HTML。

在編輯器中找到第 1 行並輸入這段文字：

```md
Welcome to freeCodeCamp
```

完成後，點擊「check your code」按鈕以查看是否正確。

# --hints--

你應該在你的編輯器中有文字 `Welcome to freeCodeCamp`。請仔細檢查拼字。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
