---
id: 682ba2318000b62f179bdf04
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML 由元素組成。 你將使用的第一個元素是 `h1` 元素：

```html
<h1>Welcome to freeCodeCamp</h1>
```

它以一個開始標籤（`<h1>`）開始，以一個結束標籤（`</h1>`）結束，在標籤之間包含了將要顯示的文本。

在你的 `Welcome to freeCodeCamp` 文本前面添加一個開始標籤，後面添加一個結束標籤，將其轉化爲一個 `h1` 元素。

# --hints--

你的 `h1` 元素應該有一個開始標籤 `<h1>`。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素應該有一個結束標籤 `</h1>`。

```js
assert.match(code, /<\/h1\s*\>/);
```

你的 `h1` 元素看起來應該像這樣： `<h1>Welcome to freeCodeCamp</h1>`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
