---
id: 682ba2318000b62f179bdf04
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML 是由元素組成。你將使用的第一個元素是 `h1` 元素：

```html
<h1>Welcome to freeCodeCamp</h1>
```

它以開啟的頁籤 (`<h1>`) 開始，以關閉的頁籤 (`</h1>`) 結束，並在頁籤之間有它將顯示的文字。

將你的 `Welcome to freeCodeCamp` 文字轉換成 `h1` 元素，方法是在它前面添加開啟標籤，並在它後面添加關閉標籤。

# --hints--

你的 `h1` 元素應該有一個開啟的 `<h1>` 頁籤。

```js
assert.exists(document.querySelector("h1"));
```

你的 `h1` 元素應該有一個閉合的 `</h1>` 標籤。

```js
assert.match(code, /<\/h1\s*\>/);
```

你的 `h1` 元素應該看起來像這樣：`<h1>Welcome to freeCodeCamp</h1>`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
