---
id: 682cd206883fc7b25eb539c5
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

請注意，你在編輯器中撰寫的 HTML 會顯示在預覽中。在本工作坊中，你將撰寫部分課程網頁的 HTML。

在你的 `h1` 元素下方的空白行輸入以下內容：

```md
Full-Stack Curriculum
```

# --hints--

你應該在你的編輯器中有 `Full-Stack Curriculum` 這段文字。請再次確認拼寫。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
