---
id: 682cd206883fc7b25eb539c5
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

你會注意到，你在編輯器中輸入的 HTML 會在預覽中顯示。 在本工作坊中，你將爲本課程大綱網頁編寫一小部分 HTML 代碼。

在你的 `h1` 元素下方的空行中輸入以下內容：

```md
Full Stack Curriculum
```

# --hints--

在你的編輯器中應該包含文本 `Full Stack Curriculum`。 仔細檢查拼寫。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full\s*stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
