---
id: 683921c4769fd23dbadec2fe
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

當 HTML 定義網頁的結構與內容時，CSS 用於添加樣式——例如顏色、字型、間距和版面配置。

在你剛剛添加的 `h3` 下方，添加另一個段落元素，內容為：

```md
CSS is used to style a webpage
```

# --hints--

你應該為頁面添加第二個 `h3` 元素。

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

你的 `h3` 元素的文字應該是 `Introduction to CSS`。 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

你應該為頁面添加第三個 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

你的 `p` 元素應該包含文字 `CSS is used to style a webpage`。你可以從說明中複製該文字。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
