---
id: 6690be42942f421648b119eb
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

在 `p` 元素的下方添加一个 `img` 元素。 此时，没有图片会出现在浏览器中。

# --hints--

你的 `img` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素应该没有结束标签 `</img>`。

```js
assert.notMatch(code, /<\/img\>/);
```

应该只有一个 `img` 元素。 删除多余的。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

你的 `img` 元素应该在 `p` 元素的下方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


