---
id: 6690beec809b4f184cab6914
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

You can link to another page with the anchor (`a`) element.

这是一个链接到 `https://www.freecodecamp.org` 的示例：

```html
<a href="https://www.freecodecamp.org"></a>
```

Add an anchor element after the paragraph that links to `https://freecatphotoapp.com`. At this point, the link won’t show up in the preview.

# --hints--

你的锚点元素（`a`）应该有一个开始标签。 开始标签的语法是：`<elementName>`。

```js
assert.exists(document.querySelector('a'));
```

你的锚元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert.match(code, /<\/a\>/);
```

Your anchor (`a`) element should be below the `p` element. 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('A'));
```

你的锚点元素（`a`）缺少 `href` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert.isTrue(document.querySelector('a').hasAttribute('href'));
```

你的锚点元素（`a`）应该链接到 `https://freecatphotoapp.com`。 你可能省略了这个 URL 或者有拼写错误。

```js
assert.equal(
  document.querySelector('a').getAttribute('href'),
    'https://freecatphotoapp.com'
);
```

Although you have set the anchor ('a') element's `href` attribute to the correct link, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert.notMatch(
  code, /\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```


