---
id: 669842cee55b683aa13ee758
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

Within the second `section` element, add a new `h2` element with the text `Cat Lists`.

# --hints--

你的 `section` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

你的 `h2` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

你的 `h2` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

你的第二个 `h2` 元素应该在第二个 `section` 元素的结束标签的上方。 它的位置不正确。

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

第二个 `h2` 元素应该有文本 `Cat Lists`。 你可能忽略了文本或有拼写错误。

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


