---
id: 6698491eff5f2e43c20c6503
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

After the last `img` element, add a `figcaption` element with the text `Cats hate other cats.`

# --hints--

你的 `figcaption` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('figcaption'), 2);
```

你的 `figcaption` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert.lengthOf(code.match(/<\/figcaption\>/g), 2);
```

在第二个 `section` 元素的结束标签上方应该有一个 `figure` 元素。

```js
assert.equal(document.querySelectorAll('main > section')[1].lastElementChild.nodeName, 'FIGURE');
```

最后一个 `img` 元素应该嵌套在 `figure` 元素中。

```js
const catsImg = document.querySelectorAll('figure > img')[1];
assert.equal(
    catsImg?.getAttribute('src')?.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg'
);
```

你的 `figure` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert.lengthOf(document.querySelectorAll('figure'), 2);
```

你的 `figure` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert.lengthOf(code.match(/<\/figure\>/g), 2);
```

`figcaption` 元素应该嵌套在 `figure` 元素中。

```js
assert.lengthOf(document.querySelectorAll('figure > figcaption'), 2);
```

嵌套在 `figure` 元素中的 `figcaption` 元素应位于 `img` 元素下方。 你的 `img` 元素和 `figcaption` 元素的顺序错误。

```js
assert.equal(
  document.querySelectorAll('figcaption')[1].previousElementSibling.nodeName,
    'IMG'
);
```

`figcaption` 元素应包含文本 `Cats hate other cats.`。你遗漏了一个单词或有拼写错误。

```js
assert.match(
  document
    .querySelectorAll('figcaption')[1]
    .innerText.toLowerCase(),
    /Cats hate other cats\.?$/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```


