---
id: 6842a6cd9836f0114a5b7a8a
title: Erstellen Sie einen Zahlenmuster-Generator
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

In diesem Labor üben Sie die Grundlagen von Python, indem Sie eine kleine App erstellen, die ein Zahlenmuster erzeugt.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `number_pattern` definieren, die einen einzelnen Parameter `n` annimmt (der eine positive ganze Zahl darstellt).
1. `number_pattern` sollte eine `for`-Schleife verwenden.
1. `number_pattern(n)` sollte einen String zurückgeben, der alle ganzen Zahlen von 1 bis einschließlich `n` enthält, getrennt durch ein Leerzeichen. Zum Beispiel sollte `number_pattern(4)` den String `1 2 3 4` zurückgeben.
1. Wenn das an die Funktion übergebene Argument kein ganzzahliger Wert ist, sollte die Funktion `Argument must be an integer value.` zurückgeben.
1. Wenn das an die Funktion übergebene Argument kleiner als 1 ist, sollte die Funktion `Argument must be an integer greater than 0.` zurückgeben.

# --hints--

Sie sollten eine `number_pattern`-Funktion haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

Die `number_pattern`-Funktion sollte einen Parameter namens `n` haben.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` sollte `1 2 3 4` zurückgeben.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` sollte `1 2 3 4 5 6 7 8 9 10 11 12` zurückgeben.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` sollte für jede positive ganze Zahl eine durch Leerzeichen getrennte Liste von Zahlen zurückgeben.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` sollte `Argument must be an integer value.` zurückgeben, wenn ein Wert übergeben wird, der keine ganze Zahl ist.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` sollte `Argument must be an integer greater than 0.` zurückgeben, wenn eine nicht-positive ganze Zahl übergeben wird.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
