---
id: 60eebd07ea685b0e777b5583
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the Registration Form project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
