---
id: 60f0286404aefb0562a4fdf9
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

Add a `title` and `meta` element inside the `head` element. Give your project a title of `Registration Form`, and add the `charset` attribute with a value of `utf-8` to your `meta` element.

# --hints--

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 要素は `head` 要素内にある必要があります。

```js
assert.exists(document.querySelector('head > title'));
```

プロジェクトには `Registration Form` というタイトルが必要です。 Remember, the casing and spelling matters for the title.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

`head` 要素の中に新しい `meta` 要素を作成する必要があります。

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` 要素に値が `utf-8` の `charset` 属性を追加する必要があります。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

`meta` 要素は空要素なので、終了タグ `</meta>` はありません。

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
