---
id: 60f027099a15b00485563dd2
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

Below the `DOCTYPE`, add an `html` element with a `lang` attribute set to `en`, so that you have a place to start putting some code.

# --hints--

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 要素が開始タグを持ち、`lang` 属性が `en` という値に設定されている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`html` タグは正しい順序になっている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html` 要素は 1 つだけにしてください。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
