---
id: 62cc5b1779e4d313466f73c5
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`head` 要素内に `link` 要素をネストしてください。 `rel` 属性の値は `stylesheet`、`href` 属性の値は `styles.css` に設定してください。

# --hints--

コードには `link` 要素が必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグを変更しないでください。 終了タグを削除していないか確認してください。

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

`link` 要素が 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 要素は空要素なので、終了タグ `</link>` は必要ありません。

```js
assert.notMatch(code, /<\/link>/);
```

`link` 要素は `head` 要素内に置く必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素には `rel` 属性があり、その値が `stylesheet` に設定されている必要があります。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
