---
id: 5dc1798ff86c76b9248c6eb3
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

تُستخدم عناصر العناوين من `h1` إلى `h6` للدلالة على أهمية المحتوى الذي يليها. كلما كان الرقم أقل، زادت الأهمية، لذا فإن عناصر `h2` أقل أهمية من عناصر `h1`.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

استخدم عنصر `h1` واحدًا فقط في الصفحة وضع العناوين الأقل أهمية تحت العناوين الأعلى أهمية.

أسفل عنصر `h1`، أضف عنصر `h2` بالنص التالي:

`Cat Photos`

# --hints--

يجب أن يحتوي عنصر `h1` على وسم بداية. صيغ وأسلوب وسم البداية هو: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/h1\>/));
```

يجب أن يكون لديك عنصر `h1` واحد فقط. احذف العنصر الزائد.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

يجب أن يكون نص عنصر `h1` هو `CatPhotoApp`. إما أنك حذفت النص أو هناك خطأ مطبعي.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

يجب أن تكون وسوم `h2` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?H2>/);
```

يجب أن يحتوي عنصر `h2` على وسم بداية. صيغ وأسلوب وسم البداية هو: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

يجب أن يحتوي عنصر `h2` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/h2\>/));
```

يجب أن يكون نص عنصر `h2` هو `Cat Photos`. ضع فقط النص `Cat Photos` بين وسم البداية ووسم الإغلاق لعنصر `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

يجب أن يكون عنصر `h2` أسفل عنصر `h1`. عنصر `h1` أكثر أهمية ويجب أن يكون فوق عنصر `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```

