---
id: 5dc17d3bf86c76b9248c6eb4
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يُستخدم العنصر `p` لإنشاء فقرة نصية على مواقع الويب. أنشئ عنصر `p` أسفل عنصر `h2` وأضف إليه النص التالي:

`Everyone loves cute cats online!`

# --hints--

يجب أن يحتوي عنصر `p` على وسم بداية. صيغة وسم البداية هي: `<elementName>`.

```js
assert(document.querySelector('p'));
```

يجب أن تكون وسوم `p` بأحرف صغيرة. وفقًا للاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?P>/);
```

يجب أن يحتوي عنصر `p` على وسم إغلاق. يحتوي وسم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/p\>/));
```

يجب أن يكون نص عنصر `p` هو `Everyone loves cute cats online!`، إما أنك حذفت النص أو يوجد خطأ مطبعي.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

يجب أن يكون عنصر `p` أسفل عنصر `h2`. ترتيبهم الحالي خاطئ.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```

