---
id: 5dc174fcf86c76b9248c6eb2
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

عناصر HTML لها وسم بداية مثل `<h1>` ووسم إغلاق مثل `</h1>`. النص الذي يعرضه العنصر يكون بين وسم البداية ووسم الإغلاق.

غيّر نص عنصر `h1` أدناه من `Hello World` إلى `CatPhotoApp` وراقب التغيير في معاينة المتصفح.

عندما تنتهي، اضغط على زر "Check Your Code" لترى إن كان صحيحًا.

# --hints--

يجب أن يكون النص `CatPhotoApp` موجودًا في الشفرة. قد ترغب في التحقق من تهجئتك.

```js
assert(code.match(/catphotoapp/i));
```

يجب أن يحتوي عنصر `h1` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

يجب أن تكون وسوم `h1` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
assert.notMatch(code, /<\/?H1>/);
```

يجب أن يحتوي عنصر `h1` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/h1\>/));
```

لديك أكثر من عنصر `h1`. احذف عنصر `h1` الزائد.

```js
assert(document.querySelectorAll('h1').length === 1);
```

يجب أن يكون نص عنصر `h1` هو `CatPhotoApp`. إما أنك حذفت النص، أو هناك خطأ مطبعي، أو أنه ليس بين وسم البداية ووسم الإغلاق لعنصر `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

