---
id: 5dc2385ff86c76b9248c6eb7
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

يحتوي HTML5 على بعض العناصر التي تحدد مناطق محتوى مختلفة. تجعل هذه العناصر شفرة HTML أسهل في القراءة وتساعد في تحسين محركات البحث (SEO) وإمكانية الوصول.

يُستخدم العنصر `main` لتمثيل المحتوى الرئيسي لجسم مستند HTML. يجب أن يكون المحتوى داخل عنصر `main` فريدًا للمستند ولا يجب تكراره في أجزاء أخرى من المستند.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

حدد القسم الرئيسي لهذه الصفحة بإضافة وسم بداية `<main>` قبل عنصر `h1`، ووسم الإغلاق `</main>` بعد عنصر `p`.

# --hints--

يجب أن يحتوي عنصر `main` على وسم بداية. تحتوي وسوم البداية على هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('main'));
```

يجب أن تكون وسوم `main` بأحرف صغيرة. حسب الاصطلاح، تُكتب جميع وسوم HTML بأحرف صغيرة.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0];
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

يجب أن يحتوي عنصر `main` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/main\>/));
```

يجب أن يكون وسم بداية عنصر `main` أسفل وسم بداية عنصر `body`. ترتيب الوسوم لديك خاطئ.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

يجب أن يكون وسم بداية عنصر `main` أعلى من عنصر `h1`. ترتيب الوسوم لديك خاطئ.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

يجب أن يكون وسم إغلاق عنصر `main` أسفل عنصر `p`. ترتيب الوسوم لديك خاطئ.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
  
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

