---
id: 5dc23991f86c76b9248c6eb8
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

في الخطوة السابقة، وضعت عناصر `h1`، و`h2`، والتعليق، و`p` داخل عنصر `main`. هذا يُسمى *التسكين*. يجب وضع العناصر المسكنة بمقدار مسافتين إضافيتين إلى اليمين من العنصر الذي تم تسكينها فيه. تُسمى هذه المسافة بالمسافة البادئة وتُستخدم لجعل شفرة HTML أسهل في القراءة.

إليك مثالًا على التسكين والمسافة البادئة:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

عنصر `h1`، وعنصر `h2`، والتعليق مُزاحان بمقدار مسافتين أكثر من عنصر `main` في الشفرة أدناه. استخدم شريط المسافة في لوحة المفاتيح لإضافة مسافتين إضافيتين أمام عنصر `p` ليُزاح بشكل صحيح أيضًا.

# --hints--

يجب أن يحتوي شفرتك على عنصر `h2` بنص `Cat Photos`. ربما حذفته عن طريق الخطأ، أو أنه مفقود وسم البداية ووسم الإغلاق، أو تغير النص.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

لا يجب عليك إضافة عناصر `ul` أو `li` من المثال.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

لا تغير المسافة البادئة في السطر الذي يحتوي على عنصر `h2`. يجب أن يبدأ وسم البداية بمقدار 6 مسافات من بداية السطر. يمكنك إعادة تشغيل الخطوة لاستعادة المسافة البادئة الأصلية.

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

يجب أن تحتوي شفرتك على تعليق. لقد حذفت التعليق من خطوة سابقة.

```js
assert(code.match(/<!--.*-->/));
```

يجب أن يكون نص التعليق `TODO: Add link to cat photos`. لا تغير نص التعليق أو المسافات فيه.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

لا تغير المسافة البادئة في السطر الذي يحتوي على عنصر التعليق. يجب أن يبدأ وسم البداية بمقدار 6 مسافات من بداية السطر. يمكنك إعادة تشغيل الخطوة لاستعادة المسافة البادئة الأصلية.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

يجب أن تحتوي شفرتك على عنصر `p`. لقد حذفت عنصر `p` من خطوة سابقة.

```js
assert(document.querySelector('p'));
```

يجب أن يكون نص عنصر `p` هو `Everyone loves cute cats online!` لا تغير النص أو المسافات أو علامات الترقيم في عنصر `p`.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

يجب أن يكون وسم البداية لعنصر `p` بمسافة بادئة تطابق عناصر `h2` والتعليق. يجب أن يبدأ وسم البداية بمقدار 6 مسافات من بداية السطر.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

