---
id: 5dc24165f86c76b9248c6ebc
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

يجب أن تحتوي جميع عناصر `img` على خاصية `alt`. يُستخدم نص خاصية `alt` لقارئ الشاشة لتحسين إمكانية الوصول ويُعرض إذا فشلت الصورة في التحميل.

إليك مثالًا على عنصر `img` مع خاصية `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

داخل عنصر `img`، أضف خاصية `alt` بهذا النص:

`A cute orange cat lying on its back`

# --hints--

يجب أن يحتوي شفرتك على عنصر `img`. لقد أزلت عنصر `img` في خطوة سابقة.

```js
assert(document.querySelector('img'));
```

عنصر `img` الخاص بك لا يحتوي على خاصية `alt`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

قيمة خاصية `alt` في عنصر `img` الخاص بك مختلفة عن `A cute orange cat lying on its back`. تأكد من أن قيمة خاصية `alt` محاطة بعلامات اقتباس.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

