---
id: 5dc23f9bf86c76b9248c6eba
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

يمكنك إضافة صور إلى موقعك باستخدام عنصر `img`. عناصر `img` لها وسم بداية بدون وسم إغلاق. العنصر الذي لا يحتوي على وسم إغلاق يُعرف باسم <dfn>العنصر الفارغ</dfn>.

أضف عنصر `img` أسفل عنصر `p`. في هذه المرحلة، لن تظهر أي صورة في المتصفح.

# --hints--

يجب أن يحتوي عنصر `img` على وسم بداية. وسوم البداية لها هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('img'));
```

يجب ألا يحتوي عنصر `img` على وسم إغلاق `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

يجب أن يكون لديك عنصر `img` واحد فقط. احذف أي عناصر إضافية.

```js
assert(document.querySelectorAll('img').length === 1);
```

يجب أن يكون عنصر `img` أسفل عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

