---
id: 5ddbd81294d8ddc1510a8e56
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

يجب وضع نص الرابط بين وسم البداية ووسم الإغلاق لعنصر وسم الرابط (`a`).

إليك مثالًا على رابط يحتوي على النص `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

أضف نص وسم الرابط `cat photos` إلى عنصر وسم الرابط. سيصبح هذا نص الرابط.

# --hints--

يجب أن يحتوي عنصر وسم الرابط (`a`) على وسم بداية. صيغ أوسمة البداية تكون بهذا الشكل: `<elementName>`.

```js
assert(document.querySelector('a'));
```

يجب أن يحتوي عنصر وسم الرابط (`a`) على وسم إغلاق. أوسمة الإغلاق تحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/a\>/));
```

يجب أن يكون نص عنصر وسم الرابط (`a`) هو `cat photos`. تأكد من وضع نص الرابط بين وسم البداية ووسم الإغلاق لعنصر وسم الرابط (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

