---
id: 5dfa37b9eacea3f48c6300b0
title: الخطوة 22
challengeType: 0
dashedName: step-22
---

# --description--

لإنشاء قائمة غير مرتبة من العناصر، يمكنك استخدام العنصر `ul`.

بعد العنصر `h3` الذي يحتوي على النص `Things cats love:`، أضف عنصر قائمة غير مرتبة (`ul`). لاحظ أنه لن يظهر شيء في هذه المرحلة.

# --hints--

يجب أن يحتوي عنصر `ul` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

يجب أن يحتوي عنصر `ul` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/ul>/));
```

يجب أن يكون عنصر `ul` فوق وسم الإغلاق للعنصر `section` الثاني.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

