---
id: 5dfa3589eacea3f48c6300ae
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

داخل العنصر الثاني من نوع `section`، أضف عنصرًا جديدًا من نوع `h2` يحتوي على النص `Cat Lists`.

# --hints--

يجب أن يحتوي عنصر `section` على وسم بداية. صيغ وسوم البداية هي: `<elementName>`.

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

يجب أن يحتوي عنصر `h2` على وسم بداية. صيغ وسوم البداية هي: `<elementName>`.

```js
assert(document.querySelectorAll('h2').length === 2);
```

يجب أن يحتوي عنصر `h2` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

يجب أن يكون عنصر `h2` الثاني مباشرة فوق وسم الإغلاق للعنصر الثاني من نوع `section`. موقعه الحالي غير صحيح.

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

يجب أن يحتوي عنصر `h2` الثاني على النص `Cat Lists`. إما أنك حذفت النص أو هناك خطأ مطبعي.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

