---
id: 5dfa2407b521be39a3de7be1
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

لفتح الروابط في تبويب جديد، يمكنك استخدام الخاصية `target` على وسم الرابط (`a`).

تحدد الخاصية `target` مكان فتح المستند المرتبط. `target="_blank"` يفتح المستند المرتبط في تبويب أو نافذة جديدة.

إليك الصيغة النحوية الأساسية لوسم `a` مع خاصية `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

أضف خاصية `target` بقيمة `_blank` إلى وسم البداية لوسم الرابط (`a`) لكي يفتح الرابط في تبويب جديد.

# --hints--

يجب أن يحتوي وسم `p` على وسم رابط (`a`) مسكن مع النص `cat photos`. ربما حذفته أو به خطأ مطبعي.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

وسم الرابط (`a`) الخاص بك لا يحتوي على خاصية `target`. تحقق من وجود فراغ بعد اسم وسم البداية و/أو وجود فراغات قبل أسماء الخصائص كلها.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

يجب أن تكون قيمة الخاصية `target` هي `_blank`. إما أنك أغفلت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

