---
id: 5dfb655eeacea3f48c6300b3
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

عنصر `figure` يمثل محتوى مستقلًا وسيسمح لك بربط صورة مع تسمية توضيحية.

قم بتسكين الصورة التي أضفتها للتو داخل عنصر `figure`.

# --hints--

يجب أن يحتوي عنصر `figure` على وسم بداية. صيغ وسم البداية هي كالتالي: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

يجب أن يحتوي عنصر `figure` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/figure\>/));
```

يجب أن يكون هناك عنصر `figure` مباشرة فوق وسم إغلاق العنصر الثاني من نوع `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

يجب أن تكون صورة اللازانيا `img` مسكنة داخل عنصر `figure`.

```js
assert.strictEqual(document.querySelector('figure > img')?.src.toLowerCase(), 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
      
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
      
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

