---
id: 5dfb6a35eacea3f48c6300b4
title: خطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

يُستخدم عنصر تسمية توضيحية للرسم البياني (`figcaption`) لإضافة تسمية توضيحية تصف الصورة الموجودة داخل عنصر `figure`.

إليك مثالًا على عنصر `figcaption` مع التسمية التوضيحية `A cute cat`:

```html
<figure>
  <img src="image.jpg" alt="A description of the image">
  <figcaption>A cute cat</figcaption>
</figure>
```

بعد الصورة المسكنة في عنصر `figure`، أضف عنصر `figcaption` مع النص التالي:

`Cats love lasagna.`

# --hints--

يجب أن يحتوي عنصر `figcaption` على وسم بداية. تحتوي وسوم البداية على الصيغة النحوية التالية: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

يجب أن يحتوي عنصر `figcaption` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

يجب أن يكون عنصر `figcaption` مسكنًا داخل عنصر `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

يجب أن يكون عنصر الصورة `img` الخاص باللازانيا مسكنًا داخل عنصر `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

يجب أن يكون عنصر `figcaption` المسكن داخل عنصر `figure` أسفل عنصر `img`. لقد وضعتهم بالترتيب الخطأ.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

يجب أن يكون نص عنصر `figcaption` هو `Cats love lasagna.` لقد حذفت النص أو به خطأ مطبعي.

```js
assert(
  document.querySelector('figcaption').innerText.trim().match(/^Cats love lasagna\.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

