---
id: 5ef9b03c81a63668521804d3
title: الخطوة 30
challengeType: 0
dashedName: step-30
---

# --description--

بعد القائمة المرتبة، أضف عنصر `figure` آخر.

# --hints--

يجب أن يحتوي عنصر `figure` على وسم البداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelectorAll('figure').length === 2);
```

يجب أن يحتوي عنصر `figure` على وسم الإغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/figure>/g).length === 2);
```

يجب أن يكون هناك عنصر `figure` فوق وسم الإغلاق الخاص بالعنصر `section` الثاني مباشرة.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'FIGURE');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
--fcc-editable-region--
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

