---
id: 5ef9b03c81a63668521804df
title: الخطوة 48
challengeType: 0
dashedName: step-48
---

# --description--

تُستخدم الخاصية `id` لتحديد عناصر HTML معينة. يجب أن تكون قيمة كل خاصية `id` فريدة عن جميع قيم `id` الأخرى في الصفحة بأكملها.

إليك مثال على عنصر `input` يحتوي على خاصية `id`:

```html
<input id="email">
```

أضف خاصية `id` بالقيمة `indoor` إلى زر الاختيار. عندما تحتوي العناصر على خصائص متعددة، لا يهم ترتيب الخصائص.

# --hints--

يجب أن يبقى زر الاختيار بين وسم البداية ووسم الإغلاق لعنصر `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

يجب أن يحتوي زر الاختيار على خاصية `id`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل جميع أسماء الخصائص.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

يجب أن يحتوي زر الاختيار على خاصية `type`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل جميع أسماء الخصائص.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

يجب أن يحتوي زر الاختيار على خاصية `id` بالقيمة `indoor`. إما أنك حذفت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

