---
id: 5ef9b03c81a63668521804db
title: الخطوة 43
challengeType: 0
dashedName: step-43
---

# --description--

لمنع المستخدم من إرسال النموذج عندما تكون المعلومات المطلوبة مفقودة، يجب إضافة الخاصية `required` إلى عنصر `input`.

إليك مثالًا على حقل إدخال يحتوي على الخاصية `required`:

```html
<input type="text" name="firstName" required>
```

لا حاجة لتعيين قيمة للخاصية `required`. فقط أضف كلمة `required` إلى عنصر `input`، مع التأكد من وجود مسافة بينها وبين الخصائص الأخرى.

# --hints--

لقد حذفت عنصر `input` الخاص بك أو أن صياغته النحوية غير صحيحة. يجب أن تكون قيم جميع الخصائص محاطة بعلامات اقتباس.

```js
assert(document.querySelectorAll('input').length);
```

يجب أن يحتوي عنصر `form` على عنصر `input` فقط. احذف أي عناصر HTML إضافية أو نصوص داخل عنصر `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

يجب أن يحتوي عنصر `input` على الخاصية `required`. تذكر، فقط أضف كلمة `required` داخل وسم عنصر `input`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

لا يجب إعطاء قيمة للخاصية `required`.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

