---
id: 5ef9b03c81a63668521804dc
title: الخطوة 46
challengeType: 0
dashedName: step-46
---

# --description--

يمكنك استخدام أزرار الاختيار للأسئلة التي تريد فيها إجابة واحدة فقط من بين عدة خيارات.

إليك مثالًا على زر اختيار مع النص مضبوطًا على `cat`:

```html
<input type="radio"> cat
```

تذكر أن عنصر `input` هو عنصر فارغ.

قبل مُدخَل النص، أضف زر اختيار مع النص مضبوطًا على:

`Indoor`

# --hints--

يجب أن تنشئ عنصر `input` لزر الاختيار الخاص بك. تحقق من الصيغة النحوية.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

يجب أن يحتوي عنصر `input` على وسم بداية، ولكن ليس وسم إغلاق.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

يجب أن تكون قد أضفت عنصر `input` واحدًا فقط لزر الاختيار الخاص بك. احذف أي إضافات زائدة.

```js
assert(document.querySelectorAll('form > input'), 2);
```

يجب أن يكون عنصر `input` الجديد فوق عنصر `input` الموجود مع خاصية `type` المضبوطة على `text`. لقد وضعتهم بترتيب خاطئ.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

عنصر `input` الجديد لا يحتوي على خاصية `type`. تحقق من وجود مسافة بعد اسم وسم البداية.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

يجب أن يحتوي عنصر `input` الجديد على خاصية `type` واحدة فقط. احذف أي إضافات زائدة.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

يجب أن يحتوي عنصر `input` الجديد على خاصية `type` بقيمة `radio`. إما أنك حذفت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

على الرغم من أنك ضبطت خاصية `type` لعنصر `input` الجديد على `radio`، يُنصح دائمًا بإحاطة قيمة الخاصية بعلامات اقتباس.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

يجب أن يكون نص ` Indoor` الخاص بزر `radio` موجودًا بعده وليس قبله.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

يجب أن يكون النص ` Indoor` موجودًا مباشرة إلى يمين زر `radio` الخاص بك. إما أنك حذفت النص أو هناك خطأ مطبعي.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

