---
id: 5ef9b03c81a63668521804dd
title: الخطوة 47
challengeType: 0
dashedName: step-47
---

# --description--

تُستخدم عناصر `label` للمساعدة في ربط النص الخاص بعنصر `input` مع عنصر `input` نفسه (خاصة لتقنيات المساعدة مثل قارئ الشاشة).

إليك مثالًا على عنصر `label` مع زر `radio`:

```html
<label><input type="radio"> cat</label>
```

في المثال، النقر على كلمة `"cat"` سيؤدي أيضًا إلى اختيار زر `radio`.

قم بتسكين زر `radio` داخل عنصر `label`.

# --hints--

يجب التأكد من بقاء زر الاختيار موجودًا.

```js
assert(document.querySelector('input[type="radio"]'));
```

يجب أن يكون النص ` Indoor` موجودًا مباشرة إلى يمين زر `radio` الخاص بك. تأكد من وجود مسافة بين العنصر والنص. إما أنك حذفت النص أو هناك خطأ مطبعي.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

يجب أن يحتوي عنصر `label` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('label'));
```

يجب أن يحتوي عنصر `label` على وسم الإغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/label\>/));
```

يجب أن يكون زر الاختيار ونصه كلهم بين وسم البداية ووسم الإغلاق لعنصر `label`.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

