---
id: 5ef9b03c81a63668521804e2
title: الخطوة 56
challengeType: 0
dashedName: step-56
---

# --description--

تستخدم النماذج عادة مربعات اختيار لأسئلة قد يكون لها أكثر من إجابة. عنصر `input` مع خاصية `type` مضبوطة على `checkbox` يُنشئ مربع اختيار.

تحت عنصر `legend` الذي أضفته للتو، أضف `input` مع خاصية `type` مضبوطة على `checkbox` والنص الخاص به مضبوط على:

` Loving`

# --hints--

يجب أن يحتوي عنصر `input` الخاص بمربع الاختيار على وسم بداية، ولكن لا يحتوي على وسم إغلاق.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

يجب أن تضيف عنصر مُدخَل واحد فقط لمربع الاختيار. احذف أي عناصر إضافية.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

يجب أن يكون عنصر `input` الجديد تحت عنصر `legend` الذي يحتوي على النص `What's your cat's personality?`. لقد وضعتهم بترتيب خاطئ.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

عنصر `input` الجديد لا يحتوي على خاصية `type`. تحقق من وجود مسافة بعد اسم وسم البداية.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

يجب أن يحتوي عنصر `input` الجديد على خاصية `type` بقيمة `checkbox`. إما أنك حذفت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

على الرغم من أنك ضبطت خاصية `type` لعنصر `input` الجديد على `checkbox`، يُنصح دائمًا بإحاطة قيمة الخاصية بعلامات اقتباس.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

يجب أن يكون النص ` Loving` موجودًا مباشرة إلى يمين مربع الاختيار. تأكد من وجود مسافة بين العنصر والنص. إما أنك حذفت النص أو به خطأ مطبعي.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

