---
id: 5ef9b03c81a63668521804d6
title: الخطوة 37
challengeType: 0
dashedName: step-37
---

# --description--

الآن أضف نموذج ويب لجمع المعلومات من المستخدمين.

يُستخدم عنصر `form` للحصول على معلومات من المستخدم مثل اسمه، بريده الإلكتروني، وتفاصيل أخرى.

بعد عنوان `Cat Form`، أضف عنصر `form`.

# --hints--

يجب أن يحتوي عنصر `form` على وسم البداية ووسم الإغلاق. قد يكون أحد الوسمين مفقودًا أو كلاهما، أو في ترتيب خاطئ.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

وسوم عنصر `form` ليست في الترتيب الصحيح.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

يجب أن يكون عنصر `form` المسكن داخل آخر عنصر `section` أسفل عنصر `h2`. لديك عنصر `h2` وعنصر `form` بترتيب خاطئ.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

يجب ألا يحتوي عنصر `form` على محتوى. احذف أي عناصر HTML أو نصوص بين وسوم عنصر `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

