---
id: 5ef9b03c81a63668521804d2
title: الخطوة 29
challengeType: 0
dashedName: step-29
---

# --description--

شفرة القائمة المرتبة (`ol`) مشابهة للقائمة غير المرتبة، لكن عناصر القائمة في القائمة المرتبة تُرقم عند العرض.

أسفل عنصر `h3`، أضف قائمة مرتبة تحتوي على هذه العناصر الثلاثة:

`flea treatment`
`thunder`
`other cats`

# --hints--

يجب أن يحتوي عنصر `ol` على وسم البداية. صيغ وسوم البداية هي: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

يجب أن يحتوي عنصر `ol` على وسم الإغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/ol>/));
```

يجب أن يكون عنصر `ol` فوق وسم الإغلاق للعنصر الثاني من نوع `section`. ترتيبهم لديك خاطئ.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

يجب أن تكون عناصر `li` الثلاثة مسكنة داخل عنصر `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

يجب أن تحتوي على ثلاثة عناصر `li` بالنصوص `flea treatment` و `thunder` و `other cats` بأي ترتيب.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

يجب أن يكون لديك عنصر `ol` واحد فقط.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

