---
id: 5ef9b03c81a63668521804d8
title: الخطوة 39
challengeType: 0
dashedName: step-39
---

# --description--

يسمح لك عنصر `input` بعدة طرق لجمع البيانات من نموذج الويب. مثل عناصر `img`، عناصر `input` هي <dfn>عنصر فارغ</dfn> ولا تحتاج إلى وسم إغلاق.

قم بتسكين عنصر `input` داخل عنصر `form`.

# --hints--

يجب أن يحتوي عنصر `form` على وسم بداية ووسم إغلاق بالترتيب الصحيح. قد يكون أحد الوسمين أو كلاهما مفقودًا، أو موجودًا بترتيب خاطئ.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

يجب أن يحتوي وسم البداية لعنصر `form` على خاصية `action` فقط. احذف أي شيء آخر قد كتبته فيه.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

يجب أن تنشئ عنصر `input`. تحقق من الصيغة النحوية.

```js
assert(document.querySelector('input'));
```

يجب أن يحتوي عنصر `input` على وسم بداية فقط، وليس وسم إغلاق.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

يجب أن يكون عنصر `input` مسكنًا داخل عنصر `form`.

```js
assert(document.querySelector('form > input'));
```

يجب أن يحتوي عنصر `form` على عنصر `input` فقط. احذف أي عناصر HTML أو نصوص بين وسمي عنصر `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

