---
id: 5ef9b03c81a63668521804d7
title: الخطوة 38
challengeType: 0
dashedName: step-38
---

# --description--

تشير الخاصية `action` إلى المكان الذي يجب إرسال بيانات النموذج إليه.

إليك مثال على عنصر `form` يحتوي على خاصية `action`:

```html
<form action="/submit-url"></form>
```

في المثال، تخبر الخاصية `action="/submit-url"` المتصفح أن بيانات النموذج يجب أن تُرسل إلى المسار `/submit-url`.

أضف خاصية `action` بالقيمة `https://freecatphotoapp.com/submit-cat-photo` إلى عنصر `form`.

# --hints--

يجب أن يحتوي عنصر `form` على وسم بداية ووسم إغلاق بالترتيب الصحيح. قد يكون أحد الوسمين مفقودًا أو كلاهما، أو قد يكونان بترتيب خاطئ.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

يجب أن يكون عنصر `form` المتسكن داخل آخر عنصر `section` أسفل عنصر `h2`. ترتيب عنصر `h2` و`form` لديك خاطئ.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

يجب ألا يحتوي عنصر `form` على محتوى. احذف أي عناصر HTML أو نصوص بين وسمي عنصر `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

لا يحتوي عنصر `form` على خاصية `action`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء الخصائص كلها.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

يجب أن يحتوي عنصر `form` على خاصية `action` بالقيمة `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

على الرغم من أنك ضبطت الخاصية `action` على عنوان URL الصحيح، يُنصح دائمًا بوضع قيمة الخاصية بين علامات اقتباس.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

