---
id: 5ef9b03c81a63668521804ea
title: الخطوة 67
challengeType: 0
dashedName: step-67
---

# --description--

لاحظ أن كل ما أضفته إلى الصفحة حتى الآن داخل عنصر `body`. كل عناصر محتوى الصفحة التي يجب عرضها على الصفحة توضع داخل عنصر `body`. مع ذلك، معلومات مهمة أخرى توضع داخل عنصر `head`.

يُستخدم عنصر `head` لاحتواء البيانات الوصفية عن الوثيقة، مثل عنوانها، وروابط أوراق الأنماط، والسكريبتات. البيانات الوصفية هي معلومات عن الصفحة لا تُعرض مباشرة على الصفحة.

أضف عنصر `head` فوق عنصر `body`.

# --hints--

لقد حذفت إما عنصر `body` أو أنه يفتقد وسم البداية أو وسم الإغلاق.

```js
assert(document.querySelector('body') && code.match(/<\/body>/));
```

يجب أن يحتوي عنصر `head` على وسم بداية. أوسمة البداية لها الصيغة النحوية التالية: `<elementName>`.

```js
assert(code.match(/\<head\>/));
```

يجب أن يحتوي عنصر `head` على وسم إغلاق. أوسمة الإغلاق تحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/\<\/head\>/));
```

يجب أن يكون عنصر `head` فوق وسم بداية عنصر `body`. لقد وضعته في مكان آخر.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/\<\/head\>\<body\>/));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<html>

  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```

