---
id: 5ef9b03c81a63668521804d9
title: خطوة 42
challengeType: 0
dashedName: step-42
---

# --description--

يُستخدم نص العنصر النائب لإعطاء الناس تلميحًا عن نوع المعلومات التي يجب إدخالها في المُدخَل.

إليك مثالًا على عنصر `input` مع تعيين عنصر نائب إلى `e.g. Jane Doe`:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

أضف نص العنصر النائب `cat photo URL` إلى عنصر `input` الخاص بك.

# --hints--

لقد حذفت إما عنصر `input` الخاص بك أو أن صياغته النحوية غير صحيحة. يجب أن تكون قيم جميع الخصائص محاطة بعلامات اقتباس.

```js
assert(document.querySelectorAll('input').length);
```

يجب أن يحتوي عنصر `form` الخاص بك على عنصر `input` فقط. احذف أي عناصر HTML إضافية أو نصوص داخل عنصر `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

عنصر `input` الخاص بك لا يحتوي على خاصية `placeholder`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

يجب أن يحتوي عنصر `input` الخاص بك على خاصية `placeholder` بقيمة `cat photo URL`. إما أنك أغفلت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

على الرغم من أنك عينت خاصية `placeholder` لعنصر `input` إلى `cat photo URL`، يُنصح دائمًا بإحاطة قيمة الخاصية بعلامات اقتباس.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

