---
id: 5ef9b03c81a63668521804e1
title: الخطوة 52
challengeType: 0
dashedName: step-52
---

# --description--

يُستخدم عنصر `fieldset` لتجميع المُدخَلات والمُسميات المرتبطة معًا في نموذج ويب. عناصر `fieldset` هي <dfn>عناصر على مستوى الكتلة</dfn>، مما يعني أنها تظهر في سطر جديد.

قم بتسكين زري الاختيار `Indoor` و`Outdoor` داخل عنصر `fieldset`، ولا تنسَ إزاحة زري الاختيار إلى الداخل.

# --hints--

يجب أن يبقى كلا زري الاختيار بين وسم البداية ووسم الإغلاق لعنصر `label`.

```js
const labelChildNodes = [...document.querySelectorAll('label')].map((node) => [...node.childNodes]);
assert(
  labelChildNodes.filter(
    childNodes =>
      childNodes.filter(node => node.nodeName === 'INPUT').length === 1
  ).length === 2
);
```

يجب أن يحتوي عنصر `fieldset` على وسم بداية. صيغ وسوم البداية هي كما يلي: `<elementName>`.

```js
assert(document.querySelector('fieldset'));
```

يجب أن يحتوي عنصر `fieldset` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/fieldset\>/));
```

يجب أن يكون كل من زري الاختيار والمُسميات المرتبطة بين وسمي البداية والإغلاق لعنصر `fieldset`.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(
  radioButtons.every((btn) => btn.parentNode.parentNode.nodeName === 'FIELDSET')
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
        
          <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
        
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

