---
id: 5ef9b03c81a63668521804d0
title: خطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

لتوضيح كلمة أو عبارة معينة، يمكنك استخدام العنصر `em`.

أبرز كلمة `love` في عنصر `figcaption` عن طريق تغليفها بعنصر التأكيد `em`.

# --hints--

يجب أن يحتوي عنصر التأكيد `em` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelector('em'));
```

يجب أن يحتوي عنصر التأكيد `em` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/em\>/));
```

لقد حذفت إما عنصر `figcaption` أو أنه يفتقد وسم بداية أو وسم إغلاق.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

يجب أن يحيط عنصر التأكيد `em` بالنص `love`. إما أنك حذفت النص أو يوجد خطأ مطبعي.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

يجب أن يكون نص `figcaption` هو `Cats love lasagna`. تحقق من الأخطاء المطبعية وأن المسافات اللازمة موجودة حول وسم البداية ووسم الإغلاق لعنصر `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

