---
id: 5efc4f528d6a74d05e68af74
title: الخطوة 58
challengeType: 0
dashedName: step-58
---

# --description--

هناك طريقة أخرى لربط نص عنصر `input` بالعنصر نفسه. يمكنك تسكين النص داخل عنصر `label` وإضافة خاصية `for` بقيمة مطابقة لخاصية `id` في عنصر `input`.

معطى عنصر `input` كما يلي:

```html
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

مثال على عنصر `label` مرتبط بهذا العنصر `input` هو:

```html
<label for="breakfast">Breakfast</label>
```

اربط النص `Loving` بمربع الاختيار عن طريق تسكين النص `Loving` فقط داخل عنصر `label` ومنحه خاصية `for` مناسبة.

# --hints--

تأكد من بقاء عنصر مربع الاختيار موجودًا، وأنك لم تُجرِ أي تغييرات عليه. يجب أن يبقى كما هو في الشفرة الأصلية.

```js
const checkboxElementRegex = /\s*<input\s+id\s*=\s*"loving"\s*type\s*=\s*"checkbox"\s*>/;
assert.match(code, checkboxElementRegex);
```

يجب أن يحتفظ مربع الاختيار بخاصية `id` بقيمة `loving`. المتوقع `--fcc-expected--`، لكن وُجد `--fcc-actual--`.

```js
const checkboxElementId = document.querySelector('input[type="checkbox"]')?.getAttribute('id');
assert.strictEqual(checkboxElementId, 'loving');
```

أضف عنصر `label` جديدًا بعد عنصر مربع الاختيار. تأكد من وجود وسم افتتاح ووسم إغلاق له.

```js
assert.lengthOf(document.querySelectorAll('label'), 3);
assert.lengthOf(code.match(/<\/label\>/g), 3);
```

يجب أن يكون عنصر `label` الخاص بك موجودًا بعد مربع الاختيار.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElement = checkboxElement?.nextElementSibling?.nodeName;
assert.strictEqual(checkboxSiblingElement, 'LABEL');
```

يجب أن يحتوي عنصر `label` على النص `Loving`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementText = checkboxElement?.nextElementSibling?.innerText;

assert.strictEqual(checkboxSiblingElementText, 'Loving');
```

يجب أن يكون نص المسمى `Loving` داخل عنصر `label` الجديد فقط. تأكد من عدم تكرار النص خارج عنصر `label`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingNode = checkboxElement?.nextSibling?.nodeValue ?? '';
const labelElement = checkboxElement?.nextElementSibling;
const labelSiblingNode = labelElement?.nextSibling?.nodeValue ?? '';

assert.notInclude(checkboxSiblingNode, 'Loving');
assert.notInclude(labelSiblingNode, 'Loving');
```

عنصر `label` الخاص بك لا يحتوي على خاصية `for`. تأكد من وجود مسافة بين اسم وسم `label` الافتتاحي وخاصية `for`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.isNotNull(checkboxSiblingElementAttr);
```

يجب أن يحتوي عنصر `label` على خاصية `for` بقيمة `loving`. المتوقع `--fcc-expected--`، لكن وُجد `--fcc-actual--`.

```js
const checkboxElement = document.querySelector('input[type="checkbox"]');
const checkboxSiblingElementAttr = checkboxElement?.nextElementSibling?.getAttribute('for');

assert.strictEqual(checkboxSiblingElementAttr, 'loving');
```

يجب أن تكون هناك مسافة بين مربع الاختيار وعنصر `label` الجديد.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

