---
id: 5efc575c8d6a74d05e68af77
title: خطوة 61
challengeType: 0
dashedName: step-61
---

# --description--

أضف مربع اختيار نهائي بعد السابق مع قيمة خاصية `id` تساوي `energetic`. يجب أن تكون خاصية `name` هي نفسها في مربع الاختيار السابق.

أضف أيضًا عنصر `label` إلى يمين مربع الاختيار الجديد مع النص `Energetic`. تأكد من ربط عنصر الـ `label` بمربع الاختيار الجديد.

# --hints--

تحتاج إلى إضافة مربع اختيار جديد.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

يجب أن يحتوي مربع الاختيار الجديد على خاصية `id` بقيمة `energetic` وخاصية `name` بقيمة `personality`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل جميع أسماء الخصائص.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'energetic' &&
      checkbox.getAttribute('name') === 'personality'
  )
);
```

يجب أن يكون مربع الاختيار الجديد بعد الثاني. ترتيبهم خاطئ لديك.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('lazy') < checkboxes.indexOf('energetic'));
```

على يمين مربع الاختيار الجديد، يجب أن يكون هناك عنصر `label` بالنص `Energetic`.

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[2]?.nextElementSibling;
assert.isOk(
  nextElementSibling?.nodeName === 'LABEL' &&
  nextElementSibling?.innerText?.trim().match(/^Energetic$/i)
);
```

يجب أن يحتوي الـ `label` الجديد على خاصية `for` بنفس قيمة خاصية `id` لمربع الاختيار الجديد. إما أنك حذفت القيمة أو بها خطأ مطبعي.

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[2]?.nextElementSibling?.getAttribute('for'),
  'energetic'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality"> <label for="lazy">Lazy</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

