---
id: 5f05a1d8e233dff4a68508d8
title: الخطوة 49
challengeType: 0
dashedName: step-49
---

# --description--

أنشئ زر اختيار آخر أسفل الأول. سكّنه داخل عنصر `label` ويكون نص الـ `label` هو `Outdoor`. أعط زر الاختيار خاصية `id` بقيمة `outdoor`.

# --hints--

يجب أن يكون فقط زر اختيار `Indoor` الأصلي من نوع `input` هو الذي يحمل خاصية `id` بقيمة `indoor`. يمكنك إعادة الخطوة لاستعادة شفرة HTML الأصلية إذا لزم الأمر.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

لا تعدل على زر اختيار `Indoor`. يمكنك إعادة الخطوة لاستعادة شفرة HTML الأصلية إذا لزم الأمر.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

أضف عنصر `input` جديد واحد فقط مسكن داخل عنصر `label` جديد. تأكد من أن `label` الجديد يحتوي على وسم افتتاح ووسم إغلاق.

```js
assert(document.querySelectorAll('label input').length === 2);
```

لا تضف أي خصائص إلى أي من وسم افتتاح `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

يجب أن يحتوي `input` الجديد على خاصية `id` بقيمة `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

يجب أن يحتوي `input` الجديد على خاصية `type` بقيمة `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

لا تضف أي عناصر جديدة غير عنصر `input` المسكن داخل `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

لا يجب أن يكون هناك نص على يسار `input` الجديد.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

يجب أن يكون نص المسمى لزر الاختيار الجديد هو بالضبط `Outdoor`. يشمل ذلك حالة الأحرف.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

يجب أن يكون زر الاختيار والمسمى الجديدان مباشرة أسفل/بعد زر الاختيار والمسمى `Indoor`. لا يجب أن يكون هناك أي وسم آخر بينهما.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

