---
id: 5f07fb1579dc934717801375
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

حان الوقت لإضافة قسم جديد. أضف عنصر `section` ثالثًا أسفل عنصر `section` الثاني.

# --hints--

يجب أن يحتوي عنصر `section` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

يجب إضافة وسم بداية `section` واحد فقط. احذف أي إضافات زائدة.

```js
assert(document.querySelectorAll('section').length === 3);
```

يجب أن يحتوي عنصر `section` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

يجب إضافة وسم إغلاق `section` واحد فقط. احذف أي إضافات زائدة.

```js
assert(code.match(/<\/section>/g).length === 3);
```

يجب أن تكون كل عناصر `section` بين وسم البداية ووسم الإغلاق لعنصر `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

يجب ألا يحتوي آخر عنصر `section` على محتوى. احذف أي عناصر HTML أو نصوص داخل عنصر `section`.

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

