---
id: 5f07c98cdb9413cbd4b16750
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

حان الوقت لإضافة قسم جديد. أضف عنصر `section` ثانيًا أسفل عنصر `section` الموجود.

# --hints--

يجب أن يحتوي عنصر `section` على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

يجب أن تضيف وسم بداية واحد فقط لعنصر `section`. احذف أي إضافات زائدة.

```js
assert(document.querySelectorAll('section').length === 2);
```

يجب أن يحتوي عنصر `section` على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

يجب أن تضيف وسم إغلاق واحد فقط لعنصر `section`. احذف أي إضافات زائدة.

```js
assert(code.match(/<\/section>/g).length === 2);
```

يجب ألا يكون عنصر `section` الثاني مسكنًا داخل عنصر `section` الأول.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

يجب أن يكون كلا عنصري `section` بين وسم البداية ووسم الإغلاق لعنصر `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

