---
id: 5f07be6ef7412fbad0c5626b
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

قبل إضافة أي محتوى جديد، يجب استخدام عنصر `section` لفصل محتوى صور القطط عن المحتوى المستقبلي.

يُستخدم عنصر `section` لتعريف أقسام في المستند، مثل الفصول، الأقسام العلوية، التذييلات، أو أي أقسام أخرى في المستند. هو عنصر دلالي يساعد في تحسين محركات البحث وإمكانية الوصول.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

ضع عناصر `h2` و `p` ووسم الرابط (`a`) داخل عنصر `section` كمسكن.

# --hints--

يجب أن يحتوي عنصر `section` على وسم بداية. صيغ وسوم البداية تكون على الشكل التالي: `<elementName>`.

```js
assert(document.querySelector('section'));
```

يجب أن يحتوي عنصر `section` على وسم إغلاق. تحتوي وسوم الإغلاق على `/` مباشرة بعد الحرف `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

يجب أن يكون عنصر `section` كاملاً بين وسوم البداية والإغلاق لعنصر `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

يجب أن تكون عناصر `h2` واثنان من عناصر `p` ووسم الرابط (`a`) الموجودة بين وسوم البداية والإغلاق لعنصر `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4);
```

لا يجب أن يكون عنصر `h1` مسكنًا داخل عنصر `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

