---
id: 5f0d4ab1b435f13ab6550052
title: الخطوة 54
challengeType: 0
dashedName: step-54
---

# --description--

بعد ذلك، ستضيف بعض عناصر `input` جديدة للنموذج، فأضف عنصر `fieldset` آخر مباشرةً تحت عنصر `fieldset` الحالي.

# --hints--

يجب أن يحتوي عنصر `fieldset` الجديد على وسم بداية. وسم البداية له هذه الصيغة النحوية: `<elementName>`.

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

يجب أن تضيف وسم بداية واحد فقط لعنصر `fieldset`. احذف أي إضافات زائدة.

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

يجب أن يحتوي عنصر `fieldset` الجديد على وسم إغلاق. وسم الإغلاق يحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

يجب أن تضيف وسم إغلاق واحد فقط لعنصر `fieldset`. احذف أي إضافات زائدة.

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

لا يجب أن يكون عنصر `fieldset` الثاني مسكنًا داخل عنصر `fieldset` الأول.

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

يجب أن يكون كلا عنصري `fieldset` فوق حقل النص. الترتيب غير صحيح.

```js
const formChildren = document.querySelector('form')?.children;
assert.isTrue(
  formChildren?.[0]?.nodeName === 'FIELDSET' &&
  formChildren?.[1]?.nodeName === 'FIELDSET' &&
  formChildren?.[2]?.nodeName === 'INPUT' &&
  formChildren?.[2]?.getAttribute('type') === 'text'
);
```

يجب أن يكون عنصر `fieldset` الجديد تحت عنصر `fieldset` الموجود. الترتيب لديك خاطئ.

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

