---
id: 5f0d48e7b435f13ab6550051
title: خطوة 53
challengeType: 0
dashedName: step-53
---

# --description--

عنصر `legend` يعمل كتسمية توضيحية للمحتوى داخل عنصر `fieldset`. يمنح المستخدمين سياقًا حول ما يجب عليهم إدخاله في ذلك الجزء من النموذج.

أضف عنصر `legend` بالنص `Is your cat an indoor or outdoor cat?` فوق كلا زري الاختيار.

# --hints--

يجب أن يحتوي عنصر `legend` على وسم بداية. صيغ أوسمة البداية هي كما يلي: `<elementName>`.

```js
assert(document.querySelector('legend'));
```

يجب أن يحتوي عنصر `legend` على وسم إغلاق. أوسمة الإغلاق تحتوي على `/` مباشرة بعد حرف `<`.

```js
assert(code.match(/<\/legend\>/));
```

يجب أن يكون عنصر `legend` هو العنصر الأول أسفل وسم البداية لعنصر `fieldset` وقبل وسم البداية لأول `label` لزر الاختيار. هو ليس في الموضع الصحيح.

```js
const fieldsetElem = document.querySelector('fieldset');
const fieldsetElemChildren = fieldsetElem.children;
assert(
  fieldsetElem.firstElementChild.nodeName === 'LEGEND' &&
    fieldsetElemChildren[1].nodeName === 'LABEL' &&
    fieldsetElemChildren[1].children[0].nodeName === 'INPUT' &&
    fieldsetElemChildren[1].children[0].id === 'indoor'
);
```

يجب أن يكون نص عنصر `legend` هو `Is your cat an indoor or outdoor cat?`. إما أنك حذفت النص، أو به خطأ مطبعي، أو أنه ليس بين وسم البداية ووسم الإغلاق لعنصر `legend`.

```js
const extraSpacesRemoved = document
  .querySelector('legend')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/Is your cat an indoor or outdoor cat\??$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

