---
id: 7cf9b03d81a65668421804c3
title: الخطوة 41
challengeType: 0
dashedName: step-41
---

# --description--

لكي يتم الوصول إلى بيانات النموذج من الموقع المحدد في خاصية `action`، يجب أن تعطي حقل النص خاصية `name` وتعيّن لها قيمة لتمثيل البيانات التي تُرسل.

إليك مثالًا على عنصر `input` مع خاصية `name`:

```html
<input type="text" name="name">
```

أضف خاصية `name` بالقيمة `catphotourl` إلى حقل النص الخاص بك.

# --hints--

لقد حذفت عنصر `input` الخاص بك أو أن صياغته غير صحيحة. يجب أن تكون قيم جميع الخصائص محاطة بعلامات اقتباس.

```js
assert(document.querySelectorAll('input').length);
```

يجب أن يحتوي نموذجك `form` على عنصر `input` فقط. احذف أي عناصر HTML إضافية أو نصوص داخل عنصر `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

عنصر `input` الخاص بك لا يحتوي على خاصية `name`. تحقق من وجود مسافة بعد اسم وسم البداية و/أو وجود مسافات قبل أسماء جميع الخصائص.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

يجب أن يحتوي عنصر `input` على خاصية `name` بالقيمة `catphotourl`. إما أنك أغفلت القيمة أو بها خطأ مطبعي. تذكر أن قيم الخصائص يجب أن تكون محاطة بعلامات اقتباس.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

على الرغم من أنك عينت خاصية `name` لعنصر `input` بالقيمة `catphotourl`، يُنصح دائمًا بإحاطة قيمة الخاصية بعلامات اقتباس.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

