---
id: ab6137d4e35944e21037b769
title: Criar um Conversor para Título em Maiúsculas
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

Neste laboratório você criará uma função que converte uma string para title case. Title case significa que a primeira letra de cada palavra é maiúscula e o restante da palavra está em minúsculas.

"Web Development Is Awesome" é um exemplo de uma string em formato de título.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função `titleCase` que recebe uma string como argumento.
1. A função `titleCase` deve retornar uma string com a primeira letra de cada palavra em maiúscula e o restante da palavra em minúscula.
1. `titleCase("I like to code")` deve retornar `"I Like To Code"`.
1. `titleCase("javaScript is fun")` deve retornar `"Javascript Is Fun"`.

# --hints--

Você deve criar uma função chamada `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` deve receber um único argumento.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` deve retornar uma string.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` deve retornar a string `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` deve retornar a string `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` deve retornar a string `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
