---
id: 69aaafeaa4037197c75587cd
title: 큐 구현하기
challengeType: 26
dashedName: implement-a-queue
---

# --description--

이 실습에서는 함수를 사용해 큐 자료 구조를 구현할 것입니다. 큐는 선입선출(FIFO) 자료 구조로, 요소가 큐의 뒤에 추가되고 앞에서 제거됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. 큐의 뒤에 요소를 추가하는 `enqueue` 함수가 있어야 합니다.
1. 큐의 앞 요소를 제거하고 반환하는 `dequeue` 함수가 있어야 합니다.
1. 큐의 앞 요소를 제거하지 않고 반환하는 `front` 함수가 있어야 합니다.
1. 큐에 있는 요소의 개수를 반환하는 `size` 함수가 있어야 합니다.
1. 큐가 비어 있으면 `true`를, 그렇지 않으면 `false`를 반환하는 `isEmpty` 함수가 있어야 합니다.

# --hints--

`enqueue` 함수가 있어야 합니다.

```js
assert.isFunction(enqueue);
```

`enqueue` 함수는 큐의 뒤에 요소를 추가해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`dequeue` 함수가 있어야 합니다.

```js
assert.isFunction(dequeue);
```

`dequeue` 함수는 큐의 앞 요소를 제거하고 반환해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

빈 큐에 호출하면 `dequeue`이 `undefined`를 반환해야 합니다.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` 함수는 `dequeue` 후에도 큐의 뒤에 계속 요소를 추가해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

서로 다른 큐 인스턴스는 상태를 공유하지 않아야 합니다.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

`front` 함수가 있어야 합니다.

```js
assert.isFunction(front);
```

`front` 함수는 큐의 앞 요소를 제거하지 않고 반환해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

빈 큐에 호출하면 `front`이 `undefined`를 반환해야 합니다.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

`size` 함수가 있어야 합니다.

```js
assert.isFunction(size);
```

`size` 함수는 큐에 있는 요소의 정확한 개수를 반환해야 합니다.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

`isEmpty` 함수가 있어야 합니다.

```js
assert.isFunction(isEmpty);
```

빈 큐에 대해 `isEmpty` 함수는 `true`를 반환해야 합니다.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

비어 있지 않은 큐에 대해 `isEmpty` 함수는 `false`를 반환해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

모든 요소가 큐에서 제거된 후 `isEmpty` 함수는 `true`를 반환해야 합니다.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
