---
id: 6a89c69c8dbfda74151fa0b6
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将创建一个使用递归的 `countup` 函数，用于创建一个从 `1` 到指定数字的数字列表。

首先定义一个名为 `countup` 的函数，该函数有一个名为 `number` 的参数。为使函数在为空时仍然有效，向函数体中添加 `pass` 关键字。

# --hints--

你应该有一个名为 `countup` 的函数，并带有一个名为 `number` 的参数。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

你的 `countup` 函数应包含 `pass` 关键字。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
