---
id: 6a89c69c91bc7365fbd4b733
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

为你的 `countup` 函数添加基本情形。用一个 `if` 语句替换 `pass`，该语句检查 `number` 是否小于 `1`。在 `if` 语句内部，返回一个空列表。

# --hints--

当 `number` 小于 `1` 时，你的 `countup` 函数应该返回一个空列表。

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

当 `number` 大于或相等于 `1` 时，你的 `countup` 函数不应该返回空列表。

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
