---
id: 6a89c69c99298438d6fbe60b
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

现在添加递归情形。在 `if` 语句之后，创建一个名为 `count_list` 的变量。将调用 `countup` 并传入 `number - 1` 的结果赋值给它。

每次调用使用更小的数字，因此递归会向基本情形推进。

# --hints--

你应该将 `countup(number - 1)` 赋值给一个名为 `count_list` 的变量。

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
