---
id: 6a89c69c6af682856558d78b
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

递归调用返回一个包含较小数字的列表。使用 `append()` 将 `number` 添加到 `count_list`，然后返回 `count_list`。

# --hints--

你应该将 `number` 追加到 `count_list`。

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

你的 `countup` 函数应该返回 `count_list`。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

当用 `3` 调用时，你的 `countup` 函数应该返回 `[1, 2, 3]`。

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
