---
id: 6a89c69c111694edc6d1c628
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

通过在 `print()` 调用中传入 `5` 来测试你的 `countup` 函数。输出应该是 `[1, 2, 3, 4, 5]`。

至此，countup 函数练习完成。

# --hints--

你应该在 `print()` 函数内调用 `countup(5)`。

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
