---
id: 698b42492a7f5f737895892e
title: بناء محول من قائمة مجاورة إلى مصفوفة مجاورة
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

في هذا المختبر، ستبني دالة تحول تمثيل قائمة مجاورة للرسم البياني إلى مصفوفة مجاورة. القائمة المجاورة هي كائن حيث يمثل كل مفتاح عقدة، والقيمة المقابلة هي مصفوفة من العقد التي ترتبط بها عقدة المفتاح. المصفوفة المجاورة هي مصفوفة ثنائية الأبعاد حيث تكون القيمة في الموقع `[i][j]` هي `1` إذا كان هناك حافة من العقدة `i` إلى العقدة `j`، و`0` خلاف ذلك.

على سبيل المثال، بالنظر إلى القائمة المجاورة:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

المصفوفة المجاورة المقابلة ستكون:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `adjacencyListToMatrix` لتحويل القائمة المجاورة إلى مصفوفة مجاورة.
2. يجب أن تأخذ الدالة كائنًا يمثل القائمة المجاورة لرسم بياني غير موزون (موجه أو غير موجه) كمعلمتها.
3. يجب أن تقوم الدالة بـ:
   - تحويل القائمة المجاورة إلى مصفوفة مجاورة.
   - طباعة كل صف في المصفوفة المجاورة.
   - إرجاع المصفوفة المجاورة.

على سبيل المثال، يجب أن تطبع الدالة `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})`:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

وترجع `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

يجب أن تعرف دالة باسم `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

يجب أن تحتوي الدالة `adjacencyListToMatrix` على معلمة واحدة.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

يجب أن تحدد الدالة عدد العقد بشكل صحيح من القائمة المجاورة.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

يجب أن تضبط الدالة قيم المصفوفة إلى `1` للحواف الموجودة بشكل صحيح.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

يجب أن تطبع الدالة كل صف من المصفوفة.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

يجب أن ترجع الدالة المصفوفة المجاورة.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

عند إعطاء القائمة المجاورة `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`، يجب أن ترجع الدالة `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

عند إعطاء القائمة المجاورة `{0: [1], 1: [0]}`، يجب أن ترجع الدالة `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

عند إعطاء القائمة المجاورة `{0: [], 1: [], 2: []}`، يجب أن ترجع الدالة `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
