---
id: 68da6138d70960c197ec5eec
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Antes de implementar el algoritmo principal, necesitas validar la entrada. El parámetro `pairs` debe ser un entero, ya que representa el número de pares de paréntesis que se van a generar.

Agrega una estructura `if` al inicio de tu función para verificar si `pairs` no es un entero. Usa la función `isinstance()` para eso.

Si la condición es verdadera, devuelve la cadena `The number of pairs should be an integer`.

# --hints--

Deberías tener una estructura `if` en tu función `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Tu sentencia `if` debe verificar si `pairs` no es un entero usando `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Debes devolver la cadena `The number of pairs should be an integer` cuando la entrada no sea un entero.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
