---
id: 68da6138d70960c197ec5eef
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Para el enfoque de búsqueda en anchura, usarás una cola para rastrear diferentes estados mientras estructuras las combinaciones de paréntesis. Cada estado estará representado como una tupla que contiene tres elementos:

- La cadena actual que se está construyendo
- El número de paréntesis de apertura usados hasta ahora
- El número de paréntesis de cierre usados hasta ahora

Crea una variable llamada `queue` e inicialízala con una lista que contiene una tupla: `('', 0, 0)`. Esto representa el estado inicial con una cadena vacía y cero paréntesis usados.

# --hints--

Debes declarar una variable llamada `queue` en tu función `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Debes inicializar `queue` con una lista que contenga la tupla `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
