---
id: 68da6138d70960c197ec5ef3
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora necesitas manejar la adición de paréntesis de cierre. La regla clave es que solo puedes agregar un paréntesis de cierre si mantiene el equilibrio, lo que significa que debe haber más paréntesis de apertura usados que de cierre.

Todavía dentro del bloque `else`, agrega una segunda sentencia `if` para verificar si `closes_used < opens_used`.

Si esta condición es verdadera, agrega otra nueva tupla a la `queue`: `(current + ')', opens_used, closes_used + 1)`. Esto representa el estado después de agregar un paréntesis de cierre.

# --hints--

Debes tener una segunda estructura `if` dentro de tu bloque `else`.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Tu segunda estructura `if` anidada debe verificar si `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Debes agregar `(current + ')', opens_used, closes_used + 1)` a la cola cuando la condición sea verdadera.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
