---
id: 68da6138d70960c197ec5eee
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

Ahora vas a configurar la estructura de datos para almacenar tus resultados. Crea una variable llamada `result` e inicialízala como una lista vacía. Esta lista almacenará todas las combinaciones válidas de paréntesis que generes.

Actualiza tu sentencia return para devolver `result` en lugar de una lista vacía.

# --hints--

Debes declarar una variable llamada `result` en tu función `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Debes inicializar `result` como una lista vacía.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("result").is_equivalent("result = []")
`) })
```

Debes retornar `result` al final de tu función.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
