---
id: 68da6138d70960c197ec5eed
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

A continuación, necesitas validar que el número de pares sea al menos uno, ya que no puedes generar combinaciones de paréntesis con cero o un número negativo de pares.

Agrega otra sentencia `if` para verificar si `pairs` es menor que `1`. Si esta condición es verdadera, devuelve la cadena `The number of pairs should be at least 1`.

# --hints--

Deberías tener una segunda estructura `if` en tu función `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Tu segunda estructura `if` debe verificar si `pairs` es menor que `1`.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Debes devolver la cadena `The number of pairs should be at least 1` cuando `pairs` sea menor que `1`.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
