---
id: 68da6138d70960c197ec5ef2
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Si la cadena actual aún no está completa, necesitas explorar los siguientes estados posibles. Añade una cláusula `else` a tu sentencia `if`.

Dentro del bloque `else`, manejarás la adición de paréntesis de apertura. Agrega una sentencia `if` para verificar si `opens_used < pairs`. Esto asegura que no uses más paréntesis de apertura de los permitidos.

Si esta condición es verdadera, agrega una nueva tupla a la `queue`: `(current + '(', opens_used + 1, closes_used)`. Esto representa el estado después de agregar un paréntesis de apertura.

# --hints--

Debes tener una cláusula `else` para tu sentencia `if` dentro del ciclo `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

Deberías tener una estructura `if` dentro del bloque `else` que verifique si `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Debes agregar `(current + '(', opens_used + 1, closes_used)` a la cola cuando la condición sea verdadera.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
