---
id: 68da604b8b6012a33e2344fb
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, implementarás una función que genera todas las combinaciones válidas de paréntesis usando un enfoque de búsqueda en anchura (BFS). Por ejemplo, las combinaciones válidas de dos pares de paréntesis son `(())` y `()()`.

Comienza creando una función llamada `gen_parentheses` con un solo parámetro `pairs`. Por ahora, devuelve una lista vacía desde la función.

# --hints--

Debes definir una función llamada `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Tu función `gen_parentheses` debe tener un solo parámetro llamado `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Tu función debería devolver una lista vacía.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
