---
id: 68da6138d70960c197ec5ef1
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

Ahora necesitas identificar cuándo has construido una combinación completa de paréntesis. Una combinación completa tiene una longitud igual al doble del número de pares (ya que cada par aporta un paréntesis de apertura y uno de cierre).

Dentro de tu ciclo `while`, agrega una estructura `if` para verificar si `len(current) == 2 * pairs`. Cuando esta condición sea verdadera, añade `current` a la lista `result`.

# --hints--

Debes tener una estructura `if` dentro de tu ciclo `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Tu sentencia `if` debería verificar si `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Debes agregar `current` a la lista `result` dentro de tu nueva estructura `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
