---
id: 68da6138d70960c197ec5ef0
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora implementarás el ciclo principal de BFS. Crea un ciclo `while` que continúe mientras la `queue` no esté vacía (mientras `queue` evalúe a `True`).

Dentro del ciclo, imprime `queue`.

# --hints--

Deberías tener un ciclo `while` en tu función `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Tu ciclo `while` debe tener la condición `queue`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Debes imprimir `queue` dentro de tu ciclo `while`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
