---
id: 68e00c5349cbddfef6dbc053
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Dentro de tu ciclo `while`, usa `queue.pop(0)` para eliminar y obtener el primer elemento de la cola. Esto implementa el comportamiento primero en entrar, primero en salir (FIFO) característico de BFS.

Desempaqueta esta tupla en tres variables: `current`, `opens_used` y `closes_used`.

# --hints--

Debes usar `queue.pop(0)` para obtener el primer elemento de la cola.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Debes desempaquetar el resultado en tres variables: `current`, `opens_used` y `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
