---
id: 699b62c99ec8c63470313365
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora es momento de agregar otra propiedad a tu objeto `profile`.

Comienza actualizando el tipo del objeto en línea para incluir una propiedad `bio` con un tipo establecido en `string`.

# --hints--

Debes actualizar el tipo del objeto en línea para incluir una propiedad `bio` establecida en `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
