---
id: 699b68f94a0d154c2a7a5ce0
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Ahora que el objeto `profile` está completo, es momento de crear un nuevo objeto que represente los roles de usuario.

Crea una nueva variable llamada `userRoles` y asígnale un objeto. Tu objeto debe tener las siguientes propiedades y valores:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

Deberías tener una variable llamada `userRoles`.

```js
assert.isDefined(userRoles);
```

Tu variable `userRoles` debe ser un objeto.

```js
assert.isObject(userRoles);
```

Tu objeto `userRoles` debe tener una propiedad `admin`.

```js
assert.property(userRoles, "admin");
```

Tu propiedad `admin` debe tener un valor `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Tu objeto `userRoles` debe tener una propiedad `editor`.

```js
assert.property(userRoles, "editor");
```

Tu propiedad `editor` debe tener un valor `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Tu objeto `userRoles` debe tener una propiedad `viewer`.

```js
assert.property(userRoles, "viewer");
```

Tu propiedad `viewer` debe tener un valor `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
