---
id: 699b65caf17a8e8d340af27d
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Ahora es momento de agregar la última propiedad al objeto `profile`.

En lecciones anteriores, aprendiste a trabajar con tipos de arreglos. Aprendiste sobre diferentes tipos de arreglos como `string[]` o `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

En este ejemplo, el arreglo `scores` solo acepta números. Si intentas agregar otro tipo como un `string` o `null`, TypeScript mostrará un error.

Actualiza tu objeto `profile` para agregar una propiedad `programmingLanguages`, tipada como un `string[]`. Dentro del literal del objeto, asigna a esta propiedad el valor `["JavaScript", "Python", "C++"]`.

# --hints--

Deberías tener una propiedad `programmingLanguages` dentro de tu objeto `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Tu propiedad `programmingLanguages` debería tener el valor `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Tu propiedad `programmingLanguages` debería estar tipada como un `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
