---
id: 699b5b7cb68b6d35afca61b7
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Aunque es posible seguir agregando propiedades así al objeto `profile`, sería bueno restringir los tipos de propiedades que deberían permitirse. Sería mejor definir la estructura del objeto `profile` de antemano.

En lecciones anteriores, aprendiste a trabajar con tipos de objetos así:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

En este ejemplo, `composer` es un objeto con un tipo de objeto en línea. Este objeto tiene tres propiedades obligatorias: `name`, `instrument` y `isActive`. Las dos primeras propiedades deben ser del tipo `string` y `isActive` debe ser del tipo `boolean`.

Actualiza tu objeto `profile` existente para incluir un tipo de objeto en línea. `username` debe tener el tipo `string`, `age` debe tener el tipo `number` y `isLoggedIn` debe tener el tipo `boolean`.

# --hints--

Tu objeto `profile` debe tener un tipo de objeto en línea con una propiedad `username` establecida en `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Tu objeto `profile` debe tener un tipo de objeto en línea con una propiedad `age` establecida en `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Tu objeto `profile` debe tener un tipo de objeto en línea con una propiedad `isLoggedIn` establecida en `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
