---
id: 699b70b10ae344c802d106b0
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Ahora es momento de agregar un par de propiedades más al objeto `userRoles`.

Agrega una propiedad `moderator` con un valor de `"medium-access"`.

Luego, debajo de esa, agrega una propiedad `guest` con un valor de `3`.

# --hints--

Tu objeto `userRoles` debe tener una propiedad `moderator`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

Tu propiedad `moderator` debe tener un valor de `"medium-access"`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

Tu objeto `userRoles` debe tener una propiedad `guest`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

Tu propiedad `guest` debe tener un valor numérico `3`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
