---
id: 699b64096301fd75e3efeb1b
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Si abres la consola de nuevo, verás un mensaje de error diferente:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Como recuerdas de lecciones anteriores, puedes marcar propiedades como opcionales usando el símbolo `?` junto al nombre de la propiedad así:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

En este ejemplo, `numberOfAlbums` está marcada como opcional, por lo que TypeScript no mostrará un error en este caso.

Actualiza tu propiedad `bio` para que esté marcada como opcional.

# --hints--

Tu propiedad `bio` debe estar marcada como opcional.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
