---
id: 699b4e284291adcfcb90df47
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, practicarás trabajar con tipos primitivos así como con tipos de objetos y arreglos construyendo un perfil de usuario con tipado seguro.

Comienza creando una variable llamada `profile` y asígnale un objeto. Tu objeto debe tener estas tres propiedades y valores:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

Debes tener una variable llamada `profile`.

```js
assert.exists(profile);
```

Tu variable `profile` debe ser un objeto.

```js
assert.isObject(profile);
```

Debes tener una propiedad `username` en tu objeto `profile`.

```js
assert.property(profile, "username");
```

Tu propiedad `username` debe tener un valor de `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Tu objeto `profile` debe tener una propiedad `age`.

```js
assert.property(profile, "age");
```

Tu propiedad `age` debe tener un valor de `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Tu objeto `profile` debe tener una propiedad `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Tu propiedad `isLoggedIn` debe tener un valor de `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
