---
id: 699b6dd1a2961bd38491ef2e
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

El objeto `profile` tenía claves y tipos estrictos. Si intentabas agregar una propiedad que no estaba definida o asignar un tipo incompatible, TypeScript lanzaba un error.

Sería bueno que el objeto `userRoles` no fuera tan estricto y tuviera claves más flexibles.

Como recuerdas de las lecciones anteriores, puedes lograr este resultado usando el tipo `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

En este ejemplo `studentGrades`, TypeScript no necesita conocer todas las materias desde el principio. Esto te permite agregar más materias como `Art` o `Music` más adelante. A TypeScript solo le importa que los valores de cada clave sean números.

Actualiza tu objeto `userRoles` agregando un tipo `Record` en línea. Establécelo como `<string, string>` para que cada clave pueda ser cualquier cadena y cada valor debe ser una cadena.

# --hints--

Tu objeto `userRoles` debe tener un tipo `Record` en línea establecido como `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
