---
id: 699b71d98c820ac7ab9eea2a
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

Si abres la consola de nuevo, verás un nuevo mensaje de error:

```md
Type 'number' is not assignable to type 'string'.
```

Esto ocurre porque estás intentando agregar una propiedad `guest` con un valor numérico.

Para eliminar el error de TypeScript, cambia el valor de la propiedad `guest` a `"read-only"`.

Con ese último cambio, ¡tu taller de perfil de usuario está completo!

# --hints--

Tu propiedad `guest` para el objeto `userRoles` debe tener el valor `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
