---
id: 699b5370a43e299b4673bcaa
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Sería útil ver las propiedades del objeto `profile` en la consola mientras lo construyes.

Agrega un `console.log` que registre `profile` en la consola.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Debes registrar `profile` en la consola.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
