---
id: 699b61ca8c9cb9bc1968364e
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Si abres la consola, deberías ver el siguiente mensaje de error:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

En JavaScript puro, podías agregar las propiedades y valores que quisieras al objeto `profile`. Pero en TypeScript, el compilador verifica que el objeto solo tenga las propiedades que definiste (`username`, `age` y `isLoggedIn` en este caso). Como `mood` no estaba incluida en el tipo, TypeScript la considera una propiedad inválida y lanza un error.

Elimina `mood: null` de tu código.

# --hints--

Ya no deberías tener `mood: null` en tu código.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
