---
id: 68e2c945cc1d8e778152be31
title: بناء واجهة لاعب
challengeType: 27
dashedName: build-a-player-interface
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف صنفًا مجردًا اسمه `Player` يرث من الصنف `abc.ABC`.
1. يجب أن يحتوي صنف `Player` على دالة `__init__` تقوم بتعيين:

- الخاصية `moves` إلى قائمة فارغة.
   - الخاصية `position` إلى `(0, 0)`.
   - الخاصية `path` إلى قائمة تحتوي على الموقع الابتدائي.

1. يجب أن يحتوي صنف `Player` على طريقة اسمها `make_move` التي:

- تستخدم `random.choice` للحصول على حركة عشوائية من الخاصية `moves` (المعرفة في الصنف الملموس).
   - تضيف قيم الحركة المختارة إلى الموقع الحالي وتحدث الخاصية `position`.
   - تضيف الصف المرتب الجديد `position` إلى الخاصية `path`.
   - تعيد الموقع الجديد `position`.

1. يجب أن يحتوي صنف `Player` على طريقة مجردة باسم `level_up` تُنفذ في الأصناف الملموسة.
1. عرّف صنفًا باسم `Pawn` يرث من صنف `Player`.
1. يجب أن يستخدم صنف `Pawn` الدالة `super()` لاستدعاء طريقة `__init__` الخاصة بالصنف الأب، ثم يعيّن الخاصية `moves` إلى قائمة من الصفوف المرتبة التي تمثل إحداثيات `x, y`.
1. يجب أن تمثل كل صف مرتب حركة بمقدار `1` وحدة في الاتجاهات التالية: أعلى، أسفل، يسار، يمين.
1. يجب أن ينفذ صنف `Pawn` طريقة `level_up` ملموسة تضيف المزيد من الحركات إلى الخاصية `moves`. يجب أن تمثل الحركات المضافة الحركات القطرية الأربعة (مثلًا، `1` وحدة إلى الأسفل زائد `1` وحدة إلى اليسار).

**ملاحظة:** يجب استيراد وحدات مكتبة النظام القياسية دون استخدام ألقاب. ستفشل الاختبارات المتعلقة بصنف `Player` حتى يصبح صنف `Pawn` قابلًا للإنشاء.

# --hints--

يجب أن يكون لديك صنف اسمه `Player`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Player")`) })
```

يجب أن يرث صنف `Player` من الصنف `ABC` في وحدة `abc`.

```js
({ test: () => runPython(`
from abc import ABC
import ast
cls = _Node(_code).find_class("Player")
assert cls.inherits_from("ABC") or cls.inherits_from("abc.ABC")
`) })
```

يجب أن يحتوي صنف `Player` على دالة `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("__init__")`) })
```

يجب أن تحتوي دالة `__init__` في صنف `Player` على معلمة واحدة هي `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("__init__").has_args("self")`) })
```

يجب أن تعين دالة `__init__` في صنف `Player` الخاصية `self.moves` إلى قائمة فارغة.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.moves = []") or init.has_stmt("self.moves = list()")
`) })
```

يجب أن تعين دالة `__init__` في صنف `Player` الخاصية `self.position` إلى `(0, 0)`.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.position = (0, 0)")
`) })
```

يجب أن تعين دالة `__init__` في صنف `Player` الخاصية `self.path` إلى قائمة تحتوي الموقع الابتدائي.

```js
({ test: () => runPython(`
init = _Node(_code).find_class("Player").find_function("__init__")
assert init.has_stmt("self.path = [self.position]")
`) })
```

يجب أن يحتوي صنف `Player` على طريقة `make_move`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("make_move")`) })
```

يجب أن تحتوي طريقة `make_move` في صنف `Player` على معلمة واحدة هي `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("make_move").has_args("self")`) })
```

يجب أن تستخدم طريقة `make_move` في صنف `Player` الدالة `random.choice` للحصول على حركة عشوائية من الخاصية `moves`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").block_has_call("choice", "make_move")`) })
```

يجب أن تحدث طريقة `make_move` في صنف `Player` الخاصية `position` بإضافة إحداثيات الحركة المختارة عشوائيًا إليها.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.position == (1, 2)
  p.make_move()
  assert p.position == (2, 4)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

يجب أن تضيف طريقة `make_move` في صنف `Player` الصف المرتب الجديد `position` إلى الخاصية `path`.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  p.make_move()
  assert p.path == [(0, 0), (1, 2)]
  p.make_move()
  assert p.path == [(0, 0), (1, 2), (2, 4)]

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

يجب أن تعيد طريقة `make_move` في صنف `Player` الخاصية `position` المحدثة.

```js
({ test: () => runPython(`
def test_func():
  try:
    p = Pawn()
  except TypeError:
    assert False, "Unable to instantiate Pawn"

  assert p.make_move() == (1, 2)

def custom_choice(seq):
  return (1, 2)

flag = None

try:
  original_choice = choice
  choice = custom_choice
  flag = "choice"
    
  test_func()

except NameError:
  try:
    original_choice = random.choice
    random.choice = custom_choice
    flag = "random.choice"

    test_func()

  except NameError:
    assert False, "missing import"
finally:
  if flag == "choice":
    choice = original_choice
  elif flag == "random.choice":
    random.choice = original_choice
`) })
```

يجب أن يحتوي صنف `Player` على طريقة `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").has_function("level_up")`) })
```

يجب أن تحتوي طريقة `level_up` في صنف `Player` على معلمة واحدة هي `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Player").find_function("level_up").has_args("self")`) })
```

يجب أن تكون طريقة `level_up` في صنف `Player` طريقة مجردة.

```js
({ test: () => runPython(`
target = _Node(_code).find_class("Player").find_function("level_up")
assert target.has_decorators("abstractmethod") or target.has_decorators("abc.abstractmethod")
`) })
```

يجب أن يكون صنف `Player` صنفًا مجردًا.

```js
({ test: () => runPython(`
try:
  Player()
except TypeError as e:
  assert str(e) == "Can't instantiate abstract class Player with abstract method level_up"
else:
  assert False, "Player class should not be instantiable"
`) })
```

يجب أن يكون لديك صنف اسمه `Pawn`.

```js
({ test: () => runPython(`assert _Node(_code).has_class("Pawn")`) })
```

يجب أن يرث صنف `Pawn` من صنف `Player`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").inherits_from("Player")`) })
```

يجب أن يحتوي صنف `Pawn` على دالة `__init__`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("__init__")`) })
```

يجب أن تحتوي دالة `__init__` في صنف `Pawn` على معلمة واحدة هي `self`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_args("self")`) })
```

يجب أن تستدعي دالة `__init__` في صنف `Pawn` دالة `__init__` الخاصة بالأصل باستخدام الدالة `super`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").find_function("__init__").has_stmt("super().__init__()")`) })
```

يجب أن تُعيّن طريقة `__init__` في `Pawn` الخاصية `moves` إلى قائمة من الصفوف المرتبة التي تمثل إحداثيات `x, y`، حيث يمثل كل صف مرتب حركة بمقدار `1` وحدة في الاتجاهات التالية: أعلى، أسفل، يسار، يمين.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]
assert isinstance(p.moves, list)
assert len(p.moves) == 4
assert all(move in p.moves for move in moves)
`) })
```

يجب أن يحتوي صنف `Pawn` على طريقة `level_up`.

```js
({ test: () => runPython(`assert _Node(_code).find_class("Pawn").has_function("level_up")`) })
```

يجب أن تضيف طريقة `level_up` في `Pawn` الحركات الأربعة القطرية بمقدار `1` وحدة إلى الخاصية `moves`.

```js
({ test: () => runPython(`
p = Pawn()
moves = [(0, 1), (1, 0), (-1, 0), (0, -1), (1, 1), (1, -1), (-1, 1), (-1, -1)]
p.level_up()
assert isinstance(p.moves, list)
assert len(p.moves) == 8
assert all(move in p.moves for move in moves)
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
from abc import ABC, abstractmethod
from random import choice


class Player(ABC):
    def __init__(self):
        self.moves = []
        self.position = (0, 0)
        self.path = [self.position]

    def make_move(self):
        x, y = self.position
        move = choice(self.moves)
        new_x = move[0] + x
        new_y = move[1] + y
        self.position = (new_x, new_y)
        self.path.append(self.position)
        return self.position

    @abstractmethod
    def level_up(self):
        pass


class Pawn(Player):
    def __init__(self):
        super().__init__()
        self.moves = [(0, 1), (1, 0), (-1, 0), (0, -1)]

    def level_up(self):
        self.moves.extend([(1, 1), (1, -1), (-1, 1), (-1, -1)])
```
