---
id: 5d5a813321b9e3db6c106a46
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript لغة قوية تتيح لك بناء مواقع ويب تفاعلية.

*ملاحظة*: لجميع المشاريع المتبقية في هذا المنهج، ستحتاج إلى مستوى أساسي من المعرفة بـ HTML وCSS. إذا كنت جديدًا على HTML وCSS، يرجى متابعة [شهادة تصميم الويب المتجاوب](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

للبدء، أنشئ قالب HTML جاهز قياسي يحتوي على `DOCTYPE` و`html` و`head` و`body`، ثم أضف وسم `meta` لـ `charset`. أضف عنصر `title` واستخدم النص `RPG - Dragon Repeller` له. أدرج وسم `link` لربط ملف الأنماط `styles.css`.

أخيرًا، أنشئ عنصر `div` مع تعيين `id` إلى `game` داخل عنصر `body`.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب تضمين مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب تعريف نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب إغلاق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم بداية. لا تنسَ خاصية `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون إعلان `DOCTYPE` في بداية ملف HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك وسم بداية `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم بداية `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عنصرا `head` و`body` متجاورين (أشقاء).

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

يجب أن تحتوي شفرتك على عنصر `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يحتوي عنصر `meta` على خاصية `charset` بالقيمة `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

يجب أن تحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لديك وسم إغلاق `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

يجب أن تحتوي شفرتك على عنصر `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

يجب أن يكون لديك عنصر `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

يجب أن يحتوي عنصر `div` على خاصية `id` بالقيمة `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

يجب أن يكون عنصر `div` داخل عنصر `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
