---
id: 62a1166ed9a56d439c0770e7
title: خطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

أنشئ ثلاثة عناصر `span` داخل عنصر `#stats`. أعط كل واحد منها الصنف `stat`. ثم أعطِ الأول النص `XP: 0`، والثاني النص `Health: 100`، والثالث النص `Gold: 50`.

# --hints--

يجب أن يكون لديك ثلاثة عناصر `span` داخل عنصر `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

يجب أن تعطي العناصر الثلاثة الجديدة `span` الصنف `stat`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

يجب أن يحتوي أول عنصر `.stat` على النص المقدم `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

يجب أن يحتوي ثاني عنصر `.stat` على النص المقدم `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

يجب أن يحتوي ثالث عنصر `.stat` على النص المقدم `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
