---
id: 62a115879a6d51422652cbfc
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

في لعبتك التمثيلية، سيتمكن المستخدمون من تتبع إحصائياتهم، وشراء الأسلحة، ومحاربة الوحوش. قبل أن تتابع الجزء التفاعلي من اللعبة باستخدام JavaScript، عليك أولًا إنشاء عناصر HTML التي ستعرض معلومات اللعبة.

أنشئ أربعة عناصر `div` داخل عنصر `#game` الخاص بك. امنحهم قيم `id` التالية بالترتيب: `stats`، `controls`، `monsterStats`، و`text`.

# --hints--

يجب أن تنشئ أربعة عناصر `div` جديدة.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

يجب أن تعطي أحد عناصر `div` الجديدة قيمة `id` تساوي `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

يجب أن تعطي أحد عناصر `div` الجديدة قيمة `id` تساوي `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

يجب أن تعطي أحد عناصر `div` الجديدة قيمة `id` تساوي `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

يجب أن تعطي أحد عناصر `div` الجديدة قيمة `id` تساوي `text`.

```js
assert.exists(document.querySelector('div#text'));
```

يجب أن تضع عناصر `div` الجديدة بالترتيب الصحيح.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

يجب أن تضع عناصر `div` الجديدة داخل عنصر `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
