---
id: 62a23c1d505bfa13747c8a9b
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

غلف الأرقام `0` و `100` و `50` بعناصر `span`، ثم غلف تلك العناصر الجديدة `span` بعناصر `strong`. بعد ذلك، أعطِ عناصر `span` الجديدة قيم `id` هي `xpText` و `healthText` و `goldText` على التوالي.

يجب أن يتبع جوابك هذا الهيكل الأساسي:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

أضف عنصر `strong` في أول عنصر `.stat` لديك.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

يجب أن يحتوي أول عنصر `strong` جديد على عنصر `span`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

يجب أن يحتوي أول عنصر `span` مسكن على الخاصية `id` بقيمة `xpText`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

يجب أن يغلف أول عنصر `span` النص `0`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

يجب أن يظل نص أول عنصر `.stat` هو `XP: 0`.

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

أضف عنصر `strong` في ثاني عنصر `.stat` لديك.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

يجب أن يحتوي ثاني عنصر `strong` جديد على عنصر `span`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

يجب أن يحتوي العنصر `span` المسكن الثاني على الخاصية `id` بقيمة `healthText`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

يجب أن يغلف ثاني عنصر `span` النص `100`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

يجب أن يظل نص ثاني عنصر `.stat` هو `Health: 100`.

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

أضف عنصر `strong` في ثالث عنصر `.stat` لديك.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

يجب أن يحتوي ثالث عنصر `strong` جديد على عنصر `span`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

يجب أن يحتوي العنصر `span` المسكن الثالث على الخاصية `id` بقيمة `goldText`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

يجب أن يغلف ثالث عنصر `span` النص `50`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

يجب أن يظل نص ثالث عنصر `.stat` هو `Gold: 50`.

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
