---
id: 62a2509ba163e020bb9d84ea
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

يمكنك الآن البدء في كتابة JavaScript. ابدأ بإنشاء عنصر `script`. يُستخدم هذا العنصر لتحميل JavaScript في ملف HTML الخاص بك.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

يجب أن يكون لديك عنصر `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

يجب أن يحتوي عنصر `script` على وسم البداية.

```js
assert.match(code, /<script\s*>/i);
```

يجب أن يحتوي عنصر `script` على وسم الإغلاق.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
