---
id: 62a255dae245b52317da824a
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

واحدة من أقوى الأدوات هي وحدة التحكم للمطوِّر. حسب متصفحك، قد تُفتح بالضغط على `F12` أو `Ctrl+Shift+I`. على نظام Mac، يمكنك الضغط على `Option + ⌘ + C` واختيار "Console". يمكنك أيضًا النقر على زر "Console" أعلى نافذة المعاينة لرؤية وحدة التحكم المدمجة لدينا.

تتضمن وحدة التحكم للمطوِّر الأخطاء البرمجية التي تُنتجها شفرتك، لكن يمكنك أيضًا استخدامها لرؤية قيم المتغيرات في شفرتك، وهذا مفيد في تصحيح الأخطاء.

أضف سطرًا `console.log("Hello World");` بين وسمَي `script`. ثم انقر على زر "Console" لفتح وحدة التحكم. يجب أن ترى النص `"Hello World"`.


# --hints--

يجب أن تضيف سطرًا `console.log("Hello World");` إلى شفرتك.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

يجب أن يكون سطر `console.log("Hello World");` بين وسمَي `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
