---
id: 62a23d1c5f1c93161f3582ae
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

مماثل لعنصر `#stats` الخاص بك، يحتاج عنصر `#monsterStats` إلى عنصرين `span`. امنحهما الصنف `stat` وأعطِ العنصر الأول النص `Monster Name: ` والعنصر الثاني النص `Health: `. بعد النص في كل منهما، أضف عنصر `strong` يحتوي على عنصر `span` فارغ مسكن بداخله. امنح عنصر `span` الداخلي الأول معرفًا `id` باسم `monsterName` وعنصر `span` الداخلي الثاني معرفًا `id` باسم `monsterHealth`.

# --hints--

يجب أن يحتوي عنصر `#monsterStats` على عنصرين `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

يجب أن يحتوي كلا العنصرين `span` الجديدين على قيمة `class` تساوي `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

يجب أن يحتوي أول عنصر `span` على عنصر `strong` مع عنصر `span` مسكن فارغ.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

يجب أن يحتوي ثاني عنصر `span` على عنصر `strong` مع عنصر `span` مسكن فارغ.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

يجب أن يحتوي أول عنصر `span` على النص `Monster Name: `. تأكد من أن التباعد صحيح.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

يجب أن يحتوي ثاني عنصر `span` على النص `Health: `. تأكد من أن التباعد صحيح.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

يجب أن تأتي عناصر `strong` و`span` بعد النص.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

يجب أن يحتوي أول عنصر `span` داخلي على `id` بقيمة `monsterName`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

يجب أن يحتوي العنصر الداخلي الثاني من نوع `span` على `id` بقيمة `monsterHealth`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
