---
id: 62a257659d0d1e2456f24ba2
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

قبل أن تبدأ بكتابة شفرة المشروع، انقلها إلى ملف خاص بها للحفاظ على التنظيم.

احذف سطر `console.log("Hello World");`. ثم أضف إلى عنصر `script` الفارغ الآن خاصية `src` بقيمة `./script.js`.

# --hints--

يجب ألا يحتوي شفرتك على سطر `console.log("Hello World");`.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

يجب أن يحتوي عنصر `script` على خاصية `src` بقيمة `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
