---
id: 62a3a7e4f1060e2fc5ffb34b
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

أنشئ متغيرًا آخر باسم `currentWeaponIndex` وضع قيمته على `0`.

# --hints--

يجب استخدام `let` لإعلان متغير باسم `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

يجب استخدام تنسيق سنام الجمل لتسمية متغيرك.

```js
assert.match(code, /currentWeaponIndex/);
```

يجب أن يكون متغير `currentWeaponIndex` مضبوطًا على `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

يجب تهيئة متغيرك إلى `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
