---
id: 62a3b41c9494f937560640ab
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

أعلن عن متغيرين إضافيين باسمَي `monsterHealth` و `inventory`.

بالنسبة لمتغير `inventory`، عيّن له قيمة مصفوفة تحتوي على السلسلة النصية `"stick"`.

تذكر أنك تعاملت مع المصفوفات في المشروع السابق بهذه الطريقة:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

يجب استخدام `let` لإعلان متغير `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

يجب استخدام تنسيق سنام الجمل لتسمية متغير `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

يجب ألا يكون لـ `monsterHealth` قيمة.

```js
assert.isUndefined(monsterHealth);
```

يجب ألا تعيّن قيمة لمتغير `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

يجب استخدام `let` لإعلان متغير `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

يجب أن يظل متغير `inventory` مصفوفة.

```js
assert.isArray(inventory);
```

يجب أن يحتوي متغير `inventory` على قيمة واحدة فقط.

```js
assert.lengthOf(inventory, 1);
```

يجب أن يتضمن متغير `inventory` السلسلة النصية `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
