---
id: 62a3a75d8466a12e009eff76
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

هيئ متغيرًا آخر باسم `health` بقيمة `100`، ومتغيرًا باسم `gold` بقيمة `50`.

# --hints--

يجب استخدام `let` لإعلان متغير باسم `health`.

```js
assert.match(code, /let\s+health/);
```

يجب تهيئة متغير باسم `health` بقيمة `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

يجب استخدام `let` لإعلان متغير باسم `gold`.

```js
assert.match(code, /let\s+gold/);
```

يجب تهيئة متغير باسم `gold` بقيمة `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

يجب أن تكون قيمة `health` هي `100`.

```js
assert.equal(health, 100);
```

يجب أن تكون قيمة `gold` هي `50`.

```js
assert.equal(gold, 50);
```

يجب أن تظل قيمة `xp` هي `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
