---
id: 62a3b79d520a7f3d0e25afd6
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

يتفاعل JavaScript مع HTML باستخدام <dfn>نموذج كائن المستند</dfn>، أو DOM. DOM هو شجرة من الكائنات تمثل HTML. يمكنك الوصول إلى HTML باستخدام كائن `document`، الذي يمثل مستند HTML كاملًا.

إحدى الطرق للعثور على عناصر محددة في HTML هي استخدام طريقة `querySelector()`. تأخذ طريقة `querySelector()` محدد CSS كمعلمة وتُرجع أول عنصر يطابق هذا المحدد. على سبيل المثال، للعثور على عنصر `<h1>` في HTML، تكتب:

```js
let h1 = document.querySelector("h1");
```

لاحظ أن `h1` هي سلسلة نصية وتطابق محدد CSS الذي ستستخدمه.

أنشئ متغيرًا باسم `button1` واستخدم `querySelector()` لتعيينه إلى العنصر الذي يحمل `id` بقيمة `button1`. تذكر أن محددات CSS للـ `id` تبدأ بـ `#`.

# --hints--

يجب استخدام `let` لإعلان متغير `button1`.

```js
assert.match(code, /let\s+button1/);
```

يجب استخدام `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

يجب استخدام المحدد `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
