---
id: 62a3a0a3c0a4b32915d26a6e
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

احذف سطر `console.log("Hello World");` لتبدأ بكتابة شفرة المشروع.

في المشروع السابق، تعلمت كيف تتعامل مع المتغيرات وكلمة `let` بهذه الطريقة:

```js
let camperbot;
```

كما تعلمت كيف تهيئ متغيرًا بقيمة بهذه الطريقة:

```js
let age = 32;
```

استخدم كلمة `let` لإعلان متغير اسمه `xp` وتعيين القيمة `0` له، وهو رقم.

# --hints--

يجب ألا يكون لديك سطر `console.log("Hello World");` في شفرتك.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

يجب استخدام كلمة `let` لإعلان متغيرك.

```js
assert.match(code, /let/);
```

يجب إعلان متغير اسمه `xp`.

```js
assert.match(code, /let\s+xp/);
```

يجب أن تكون قيمة `xp` هي `0`.

```js
assert.equal(xp, 0);
```

يجب تهيئة المتغير `xp` بالقيمة `0`.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
