---
id: 62a3bec30ea7f941412512dc
title: الخطوة 29
challengeType: 0
dashedName: step-29
---

# --description--

تمامًا كما فعلت مع الأزرار، أنشئ متغيرات لـ `id`s التالية واستخدم `querySelector()` لإعطائها العنصر كقيمة:

`text`، `xpText`، `healthText`، `goldText`، `monsterStats`، و `monsterName`.

تذكر إعلان هذه باستخدام الكلمة المفتاحية `const`، وسمِّ المتغيرات لتطابق الـ `id`s.

# --hints--

يجب إعلان متغير `text` باستخدام `const`.

```js
assert.match(code, /const\s+text/);
```

يجب أن يحتوي متغير `text` على قيمة عنصر `#text` الخاص بك.

```js
assert.match(code, /const\s+text\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#text\1\s*\)\s*/);
```

يجب إعلان متغير `xpText` باستخدام `const`.

```js
assert.match(code, /const\s+xpText/);
```

يجب أن يحتوي متغير `xpText` على قيمة عنصر `#xpText` الخاص بك.

```js
assert.match(code, /const\s+xpText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#xpText\1\s*\)\s*/);
```

يجب إعلان متغير `healthText` باستخدام `const`.

```js
assert.match(code, /const\s+healthText/);
```

يجب أن يحتوي متغير `healthText` على قيمة عنصر `#healthText` الخاص بك.

```js
assert.match(code, /const\s+healthText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#healthText\1\s*\)\s*/);
```

يجب إعلان متغير `goldText` باستخدام `const`.

```js
assert.match(code, /const\s+goldText/);
```

يجب أن يحتوي متغير `goldText` على قيمة عنصر `#goldText` الخاص بك.

```js
assert.match(code, /const\s+goldText\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#goldText\1\s*\)\s*/);
```

يجب إعلان متغير `monsterStats` باستخدام `const`.

```js
assert.match(code, /const\s+monsterStats/);
```

يجب أن يحتوي متغير `monsterStats` على قيمة عنصر `#monsterStats` الخاص بك.

```js
assert.match(code, /const\s+monsterStats\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterStats\1\s*\)\s*/);
```

يجب إعلان متغير `monsterName` باستخدام `const`.

```js
assert.match(code, /const\s+monsterName/);
```

يجب أن يحتوي متغير `monsterName` على قيمة عنصر `#monsterName` الخاص بك.

```js
assert.match(code, /const\s+monsterName\s*=\s*document\s*\.\s*querySelector\s*\(\s*('|")#monsterName\1\s*\)\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
