---
id: 62b46e3a8d4be31be5af793d
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

واجهنا مشكلة بسيطة. تحاول استعلام صفحتك عن عنصر زر، لكن وسم `script` موجود في `head` الخاص بـ HTML. هذا يعني أن شفرتك تُنفذ قبل أن ينتهي المتصفح من قراءة HTML، و`document.querySelector()` لن يرى الزر لأن المتصفح لم يعالجه بعد.

لإصلاح ذلك، حرّك وسم `script` خارج عنصر `head` وضعه في نهاية عنصر `body` (قبل وسم الإغلاق `</body>` مباشرة).

# --hints--

يجب ألا يكون وسم `script` داخل عنصر `head`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

يجب أن يأتي وسم `script` في نهاية عنصر `body`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
