---
id: a6e40f1041b06c996f7b2406
title: Tengeneza Kitafuta Kipengele cha Kwanza
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

Katika maabara hii, utaunda kitendakazi kinachopitia safu ya data na kurudisha kipengele cha kwanza kinachopita mtihani wa kitendakazi (pia kinachojulikana kama "mtihani wa kweli"). 

Kitendakazi kitapita kupitia safu ya data na kupima kila kipengele kwa kutumia kitendakazi cha mtihani kilichotolewa. Mwishoni, kitarudisha kipengele cha kwanza ambacho kitendakazi cha mtihani kitarudisha `true`.

Mfano:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `findElement` kinachokubali safu ya data na kitendakazi kama hoja.
2. Kitendakazi kinapaswa kurudisha kipengele cha kwanza katika safu ya data kinachopita mtihani wa kweli. Hii inamaanisha kwamba, kwa kuita kitendakazi kilichopitishwa `func`, ukitolea kipengele `x`, mtihani wa kweli unapitishwa ikiwa `func(x)` ni `true`.
3. Ikiwa hakuna kipengele kinachopita mtihani, kitendakazi kinapaswa kurudisha `undefined`.

# --hints--

Unapaswa kuwa na kitendakazi cha `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` inapaswa kurudisha `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` inapaswa kurudisha `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` inapaswa kurudisha `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` inapaswa kurudisha `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` inapaswa kurudisha `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` inapaswa kurudisha `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` inapaswa kurudisha `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
