---
id: 68e97fe79367ad7b5dd6c9cd
title: 第 4 步
challengeType: 0
dashedName: step-4
---

# --description--

在這個步驟中，添加一個 `h1` 元素，內容為 `XYZ Bookstore`。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素的文字應該是 `XYZ Bookstore`。請仔細檢查拼寫或大小寫錯誤。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

你的 `h1` 元素應該在 `body` 元素內。

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
