---
id: 68e97fe79367ad7b5dd6c9cd
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在此步驟中為你的文件添加一個帶有文字 `XYZ Bookstore` 的 `h1` 元素。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素的文字應該是 `XYZ Bookstore`。請仔細檢查拼寫或大小寫錯誤。

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

你的 `h1` 元素應該在 `body` 元素內。

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
</html>
```
