---
id: 68ebdbacdd3fa474132cc975
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `h1` 元素下方，添加一個 `p` 元素，內容為：`Browse our collection of amazing books!`。

# --hints--

你應該有一個 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素的文字應該是：`Browse our collection of amazing books!`。 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

你的 `p` 元素應該在你的 `h1` 元素下方。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  --fcc-editable-region--

  --fcc-editable-region--
</body>

</html>
```
