---
id: 68ec99e478211dc578699944
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

在第一個具有 `card` 類別的元素中，`h2` 元素下方添加一個帶有以下文字的 `p` 元素：

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

你應該在具有 `card` 類別的元素內巢狀一個 `p` 元素。

```js
assert.exists(document.querySelector('.card p'));
```

你的 `p` 元素的文字應該是 `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

你的 `p` 元素應該位於你的 `h2` 元素下方。

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
