---
id: 68eca6364c5616d393389a30
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

請記住，HTML 元素看起來像這樣：

```html
<element attribute="value">
    inner text
</element>
```

在具有類別 `card-container` 的元素下方，添加一個新的 `p` 元素，並包含此文字：

```md
Review your selections and continue to checkout.
```

在 `p` 元素下方，創建一個 `div` 元素，並將 `class` 屬性設為 `btn-container`。此容器將群組你的導覽按鈕元素。

# --hints--

新添加的 `p` 元素應該位於具有 `card-container` 類別的元素下方。

```js
assert.exists(document.querySelector('.card-container + p'));
```

你新的 `p` 元素的文字應該是 `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

你應該有一個 `div` 元素，其 `class` 為 `btn-container`。

```js
assert.exists(document.querySelector('div.btn-container'));
```

你的 `div` 元素具有 `btn-container` 的 `class`，應該放在新添加的 `p` 元素之後。

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>

--fcc-editable-region--
</body>

</html>
```
