---
id: 68eca5c0065c82d256d29ca3
title: 第 17 步
challengeType: 0
dashedName: step-17
---

# --description--

在第二張卡片內，為你的 `button` 元素添加 `class` 屬性並設為 `btn`，以及文字 `Buy Now`。

兩個 `button` 元素現在共用相同的 `class`，這表示它們可以一起一致地進行樣式設定。

# --hints--

你應該在第二個具有 `card` 類別的元素內巢狀一個 `button` 元素。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

你的第二張卡片的 `button` 元素應該有一個 `class` 屬性。

```js
const cards = document.querySelectorAll('.card');
assert.isTrue(cards[1]?.querySelector('button')?.hasAttribute('class'));
```

你的第二張卡片的 `button` 元素應該有一個值為 `btn` 的 `class`。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

你第二張卡片的 `button` 元素的文字應該是 `Buy Now`。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
  
    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
