---
id: 68eca2a795b333ca5fee30a8
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

現在創建第二個書籍卡片。為你的結構添加另一個 `div` 元素，並將 `class` 屬性設置為 `card`。注意你可以為多個元素重複使用相同的類別名稱，以套用一致的樣式。

# --hints--

你應該有兩個帶有 `class` 屬性且值為 `card` 的 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

具有 `card` 類別的第二個元素應該位於具有 `card-container` 類別的元素內。

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

具有 `card` 類別的兩個元素應該是兄弟元素。

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>

  --fcc-editable-region--
  </div>
</body>

</html>
```
