---
id: 68eca6e5a6759ed4ea0034dc
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

在具有 `btn-container` 類別的元素內，添加兩個 `button` 元素：

第一個按鈕：

- Id: `view-cart-btn`
- 類別：`btn`
- `View Cart`

第二個按鈕：

- 識別碼：`checkout-btn`
- 類別：`btn`
- 結帳


恭喜你！你已成功使用 div、類別 和 id 建立書店頁面的結構來組織你的內容。

# --hints--

你應該在 `.btn-container` 元素內巢狀兩個 `button` 元素。

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

在類別為 `btn-container` 的元素內，你的第一個按鈕應該有一個 `id` 屬性，值為 `view-cart-btn`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

在具有 `btn-container` 類別的元素內，你的第一個按鈕應該有一個 `class` 屬性，值為 `btn`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

你第一個按鈕的文字應該是 `View Cart`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

在具有 `btn-container` 類別的元素內，你的第二個按鈕應該有一個 `id` 屬性，值為 `checkout-btn`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

在具有 `class="btn-container"` 的元素內，你的第二個按鈕應該有一個 `class` 屬性，值為 `btn`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

你第二個按鈕的文字應該是 `Checkout`。

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
--fcc-editable-region--
  <p>Review your selections and continue to checkout.</p>
  <div class="btn-container">

  </div>
--fcc-editable-region--
</body>

</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>

  <div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
  </div>

  <p>Review your selections and continue to checkout.</p>

  <div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
  </div>
</body>

</html>
```
