---
id: 68eca5412b5dd9d06b3d6404
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

在第二個卡片的 `h2` 元素下方，添加一個帶有此文字的 `p` 元素：

```md
This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.
```

# --hints--

你應該在第二個具有 `card` 類別的元素內巢狀一個 `p` 元素。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

你的第二張卡片的 `p` 元素文字應該是 `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

你的第二張卡片的 `p` 元素應該位於你的 `h2` 元素下方。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
