---
id: 68ec98b38d83a3c28dd30efe
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

在第一個具有 `card` 類別的元素內，為你的元素添加一個 `h2` 元素，文字為 `Sally's SciFi Adventure`。

# --hints--

你應該在具有 `card` 類別的元素內巢狀一個 `h2` 元素。

```js
assert.exists(document.querySelector('.card h2'));
```

你的 `h2` 元素的文字應該是 `Sally's SciFi Adventure`。 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
