---
id: 68ec6cd5b7e8f5a8f7319e32
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

`div` 元素用作容器來群組其他 HTML 元素。當你想要群組將共用一組 CSS 樣式的 HTML 元素時，你主要會使用 `div` 元素。

在 `p` 元素下方，添加一個 `div` 元素。這個 `div` 將會是你的書籍卡片的容器。

**注意**：這個工作坊不會套用 CSS。類別和群組元素對 CSS 樣式很有用，但在這個工作坊中，它們僅用於結構和群組內容。你會在後面的模組學習樣式如何運作。

# --hints--

你應該有一個 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素應該位於你的 `p` 元素下方。

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
