---
id: 68ec6d9a315221aa31e54816
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

`class` 屬性用於識別一個或多個元素以進行樣式設定。與 `id` 屬性不同，類別名稱不需要唯一：多個元素可以共用相同的類別。

這裡有一個範例：

```html
<p class="example">example paragraph</p>
```

為你的 `div` 元素添加 `class` 屬性，並將其值設為 `card-container`。

# --hints--

你的 `div` 元素應該有一個 `class` 屬性。

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

你的 `div` 元素應該有一個 `class` 屬性，值為 `card-container`。

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
--fcc-editable-region-- 
  <div>

  </div>
--fcc-editable-region--
</body>

</html>
```
