---
id: 68ec9332a9b5b2b32487bd00
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

`id` 屬性為 HTML 元素添加唯一的識別符。每個 `id` 在一個頁面中應該是唯一的，且只能使用一次。

`id` 值不能包含空格，且只能包含字母、數字、底線和破折號。

這裡有一個範例：

```html
<p id="para">example paragraph</p>
```

為你的類別為 `card` 的元素添加一個 `id` 屬性，並將其值設為 `sally-adventure-book`。

# --hints--

你的具有 `card` 類別的元素應該有一個 `id` 屬性。

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

你的元素擁有 `card` 類別，應該有一個 `id`，其值為 `sally-adventure-book`。

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
