---
id: 68eca11e99e3c5c894ca9d69
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

`button` 元素用於在網頁上創建可點擊的按鈕。按鈕是使用者可以點擊以執行操作的互動元素。

在具有 `class` 為 `card` 的元素內新增一個 `button` 元素，為該按鈕設定 `class` 屬性為 `btn`，並設定文字為 `Buy Now`。

# --hints--

你應該在具有 `card` 類別的元素內巢狀一個 `button` 元素。

```js
assert.exists(document.querySelector('.card button'));
```

你的 `button` 元素應該有一個 `class` 屬性。

```js
assert.isTrue(document.querySelector('.card button')?.hasAttribute('class'));
```

你的 `button` 元素應該有一個 `class`，其值為 `btn`。

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

你的 `button` 元素的文字應該是 `Buy Now`。 

```js
assert.equal(document.querySelector('.card button')?.innerText.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
  --fcc-editable-region--
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
