---
id: 68eca4c4b0f952cf09fabe09
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

在具有 `card` 類別的第二個元素內，添加一個帶有文字 `Dave's Cooking Adventure` 的 `h2` 元素。

# --hints--

你應該在第二個具有 `card` 類別的元素內巢狀一個 `h2` 元素。

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

第二個具有 `card` 類別的元素內的 `h2` 元素應該有 `Dave's Cooking Adventure` 的文字。

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
  --fcc-editable-region--
    <div class="card" id="dave-cooking-book">

    </div>
  --fcc-editable-region--
  </div>
</body>

</html>
```
