---
id: 69134b6cb43aa5f254950a10
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將透過建立顯示不同書籍資訊的書籍卡片來建置一個書店頁面。你將練習使用 `div` 元素、類別和 ID 來組織內容。

從建立 HTML 樣板開始你的書店頁面。

為你的文件添加 `<!DOCTYPE html>` 宣告式以及 `html` 和 `head` 元素。

為你的 `html` 元素添加 `lang` 屬性並將其設為 `"en"`。

# --hints--

你的程式碼應該包含 `<!DOCTYPE html>` 宣告式。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

你的程式碼應該同時包含開啟和關閉的 `<html>` 頁籤。

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

你的 `html` 元素應該有一個 `lang` 屬性設定為 `"en"`。

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

你的程式碼應該同時包含開啟和關閉的 `<head>` 頁籤。

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

你的 `head` 元素應該巢狀在 `html` 元素內。

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
