---
id: 69134d14150d51f8066a3246
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在，改善你的 HTML 文件的結構，以確保你的頁面正確編碼。

在 `head` 元素內，添加 `<meta charset="UTF-8">` 元素。

最後，在 `head` 區段下方添加一個 `body` 元素。這裡將放置你所有可見的頁面內容。

# --hints--

你的程式碼應該包含一個 `<meta charset="UTF-8">` 元素。

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

你的程式碼應該包含開啟和關閉的 `<body>` 頁籤。

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

你的 `body` 元素應該位於 `head` 元素之下。

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
<head>
  <title>XYZ Bookstore Page</title>
</head>

</html>
--fcc-editable-region--
```
