---
id: 691f85085787a8c4ca589fbb
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

將 `title` 元素添加到你的 `head` 元素內。 

將頁面標題設定為 `XYZ Bookstore Page`。

# --hints--

你的程式碼應該同時包含開啟和關閉的 `<title>` 頁籤。 

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

你的標題應該是 `XYZ Bookstore Page`。

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

你的 `title` 元素應該巢狀在 `head` 元素內。

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
  --fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
<head>

</head>
  --fcc-editable-region-- 
</html>
```
