---
id: 67172b43f84bcd2dec238a3d
title: بناء منظم كتب
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يكون لديك مصفوفة من الكائنات باسم `books` حيث يحتوي كل كائن في المصفوفة على سلسلة نصية `title`، وسلسلة نصية أخرى `authorName`، ورقم `releaseYear`.

2. يجب أن تحتوي مصفوفة `books` على ثلاثة كائنات على الأقل.

3. يجب أن يكون لديك دالة استدعاء راجع باسم `sortByYear` تقبل كتابين كمعلمات لفرز المصفوفة.

4. يجب أن تُرجع الدالة `sortByYear` القيمة `-1` إذا كان `releaseYear` للكتاب الأول أصغر من ذلك الخاص بالكتاب الثاني.

5. يجب أن تُرجع الدالة `sortByYear` القيمة `1` إذا كان `releaseYear` للكتاب الأول أكبر من ذلك الخاص بالكتاب الثاني.

6. يجب أن تُرجع الدالة `sortByYear` القيمة `0` إذا كانت قيمتا `releaseYear` متساويتين.

7. يجب أن تقوم بتصفية الكتب التي كُتبت بعد سنة معينة مثل 1950 من مصفوفة `books` وحفظ المصفوفة المصفاة في مصفوفة جديدة باسم `filteredBooks`.

8. يجب أن ترتب الكتب في مصفوفة `filteredBooks` حسب `releaseYear` بترتيب تصاعدي. تعلمت في درس سابق أن طريقة `sort()` ترتب المصفوفة في مكانها، مما يعني أن مصفوفة `filteredBooks` ستتغير.

# --hints--

يجب أن يكون لديك دالة `sortByYear` في شفرتك.

```js
assert.isFunction(sortByYear);
```

يجب أن تأخذ دالة `sortByYear` معلمتين.

```js
assert.lengthOf(sortByYear, 2);
```

يجب أن تُرجع دالة `sortByYear` القيمة `-1` إذا كان `releaseYear` لكائن `book1` أصغر من ذلك الخاص بكائن `book2`، والقيمة `1` إذا كان `releaseYear` لكائن `book1` أكبر من ذلك الخاص بكائن `book2`، والقيمة `0` في جميع الحالات الأخرى.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

يجب أن يكون لديك مصفوفة `books` في شفرتك.

```js
assert.isArray(books)
```

يجب أن تحتوي مصفوفة `books` على ثلاثة كائنات كتب على الأقل.

```js
assert.isAtLeast(books.length, 3)
```

يجب أن تحتوي مصفوفة `books` على كائنات كل منها يحتوي على سلسلة نصية `title`، وسلسلة نصية أخرى `authorName`، ورقم `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

يجب أن يكون لديك مصفوفة `filteredBooks` في شفرتك.

```js
assert.isArray(filteredBooks)
```

يجب أن تحتوي مصفوفة `filteredBooks` على بعض الكتب التي تم تصفيتها من مصفوفة `books` وألا تكون فارغة.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

يجب أن تحتوي مصفوفة `filteredBooks` فقط على الكتب الموجودة في مصفوفة `books` الأصلية.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

يجب أن تتضمن مصفوفة `filteredBooks` فقط الكتب التي صدرت في أو قبل سنة محددة.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

يجب أن تستدعي دالة الترتيب `sort` ذات الرتبة الأعلى بتمرير دالة الاستدعاء الراجع `sortByYear` على مصفوفة `filteredBooks`.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
