---
id: 616965351e74d4689eb6de30
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

ウェブページには、`meta` 要素を複数追加できます。 各 `meta` 要素は、他の HTML 要素では表現できない、そのページ自体に関する情報を追加します。

`head` 要素内に `meta` 要素をもう 1 つ追加してください。 `name` 属性を `"viewport"` に、`content` 属性を `"width=device-width, initial-scale=1.0"` に設定してください。これにより、あらゆるデバイスでページが同じように表示されます。

# --hints--

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

新しく追加した `meta` 要素は空要素なので、終了タグ `</meta>` は必要ありません。

```js
assert.notMatch(code, /<\/meta>/i);
```


新しい `meta` 要素の `name` 属性には `"viewport"` を、`content` 属性には `"width=device-width, initial-scale=1.0"` を設定する必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

`meta` 要素は 2 つとも `head` 要素内に置く必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
