---
id: 61695d1fbc003856628bf561
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

The `charset` attribute specifies the character encoding used by the document. `utf-8` (Unicode Transformation Format – 8-bit) is a character encoding standard used for electronic communication.

`head` 要素内に `meta` 要素をネストし、`charset` 属性を `"utf-8"` に設定してください。

# --hints--

`meta` 要素が 1 つ必要です。

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

`meta` 要素は空要素なので、終了タグ `</meta>` は必要ありません。

```js
assert.notMatch(code, /<\/meta>/i);
```

`meta` 要素には値が `"utf-8"` に設定された `charset` 属性が必要です。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

`meta` 要素は `head` 要素内に置く必要があります。

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
