---
id: 61696ef7ac756c829f9e4048
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

このプロジェクトでは外部 CSS ファイルを使用してページのスタイルを設定します。 `styles.css` ファイルは既に作成してあります。 ですがそれを使えるようにするには、ページにリンクする必要があります。

`head` 要素内に `link` 要素を 1 つネストしてください。 その `rel` 属性を `"stylesheet"` に、`href` 属性を `"styles.css"` に設定してください。

# --hints--

`link` 要素が 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 要素は空要素なので、終了タグ `</link>` は必要ありません。

```js
assert.notMatch(code, /<\/link>/i);
```

`link` 要素には値が `"stylesheet"` の `rel` 属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には値が `"styles.css"` の `href` 属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

`link` 要素は `head` 要素内に置く必要があります。

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
