---
id: 66db529d37ad966480ebb633
title: Ein Programm zum Auswählen des Mittagessens erstellen
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

In diesem Labor erstellen Sie ein Programm, das bei der Verwaltung von Mittagsoptionen hilft. Sie arbeiten mit einem Array von Mittagessen, fügen Elemente zum Array hinzu und entfernen sie, und wählen zufällig eine Mittagsoption aus.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `lunches`-Variable erstellen und ihr ein leeres Array zuweisen, das zum Speichern der Mittagessen-Elemente verwendet wird.

2. Sie sollten eine Funktion `addLunchToEnd` erstellen, die ein Array als ersten Parameter und einen String als zweiten Parameter entgegennimmt. Die Funktion soll:
    - Den String ans Ende des Arrays anhängen.
    - Den String `[Lunch Item] added to the end of the lunch menu.` in der Konsole ausgeben, wobei `[Lunch Item]` der an die Funktion übergebene String ist.
    - Das aktualisierte Array zurückgeben.

3. Sie sollten eine Funktion `addLunchToStart` erstellen, die ein Array als ersten Parameter und einen String als zweiten Parameter entgegennimmt. Die Funktion soll:
    - Den String an den Anfang des Arrays einfügen.
    - Den String `[Lunch Item] added to the start of the lunch menu.` in der Konsole ausgeben, wobei `[Lunch Item]` der an die Funktion übergebene String ist.
    - Das aktualisierte Array zurückgeben.

4. Sie sollten eine Funktion `removeLastLunch` erstellen, die ein Array als Parameter entgegennimmt. Die Funktion soll:
   - Das letzte Element aus dem Array entfernen.
   - Wenn die Entfernung erfolgreich ist, den String `[Lunch Item] removed from the end of the lunch menu.` in der Konsole ausgeben, wobei `[Lunch Item]` das entfernte Element ist.
   - Wenn das Array leer ist, den String `"No lunches to remove."` in der Konsole ausgeben.
   - Das aktualisierte Array zurückgeben.

5. Sie sollten eine Funktion `removeFirstLunch` erstellen, die ein Array als Parameter entgegennimmt. Die Funktion soll:
   - Das erste Element aus dem Array entfernen.
   - Wenn die Entfernung erfolgreich ist, den String `[Lunch Item] removed from the start of the lunch menu.` in der Konsole ausgeben, wobei `[Lunch Item]` das entfernte Element ist.
   - Wenn das Array leer ist, den String `"No lunches to remove."` in der Konsole ausgeben.
   - Das aktualisierte Array zurückgeben.

6. Sie sollten eine Funktion `getRandomLunch` erstellen, die ein Array als Parameter entgegennimmt. Die Funktion soll:
   - Ein zufälliges Element aus dem Array auswählen.
   - Wenn erfolgreich, den String `Randomly selected lunch: [Lunch Item]` in der Konsole ausgeben, wobei `[Lunch Item]` ein zufälliges Element aus dem Array ist.
   - Wenn das Array leer ist, den String `"No lunches available."` in der Konsole ausgeben.

7. Sie sollten eine Funktion `showLunchMenu` erstellen, die ein Array als Parameter entgegennimmt und:
   - Wenn Elemente im Array sind, den String `Menu items: [Lunch Item], [Lunch Item]...` in der Konsole ausgeben, wobei jedes `[Lunch item]` eines der Elemente im Array in der Reihenfolge ist.
   - Wenn das Array leer ist, den String `"The menu is empty."` in der Konsole ausgeben.

# --hints--

Sie sollten eine Variable `lunches` deklarieren und ihr ein leeres Array zuweisen, um die Mittagessen-Elemente zu speichern.

```js
assert.isArray(lunches);
```

Sie sollten eine Funktion `addLunchToEnd` definieren.

```js
assert.isFunction(addLunchToEnd);
```

Die Funktion `addLunchToEnd` sollte zwei Parameter haben.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` sollte den String `"Tacos added to the end of the lunch menu."` in der Konsole ausgeben.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` sollte `["Pizza", "Tacos", "Burger"]` zurückgeben.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Sie sollten eine Funktion `addLunchToStart` definieren.

```js
assert.isFunction(addLunchToStart);
```

Die Funktion `addLunchToStart` sollte zwei Parameter haben.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` sollte den String `"Sushi added to the start of the lunch menu."` in der Konsole ausgeben.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` sollte `["Pizza", "Burger", "Sushi"]` zurückgeben.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Sie sollten eine Funktion `removeLastLunch` definieren.

```js
assert.isFunction(removeLastLunch)
```

Die Funktion `removeLastLunch` sollte einen Parameter haben.

```js
assert.lengthOf(removeLastLunch, 1);
```

Wenn das Eingabe-Array leer ist, sollte die Funktion `removeLastLunch` den String `"No lunches to remove."` in der Konsole ausgeben.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` sollte den String `"Toast removed from the end of the lunch menu."` in der Konsole ausgeben.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` sollte `["Sushi", "Pizza"]` zurückgeben.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Sie sollten eine Funktion `removeFirstLunch` definieren.

```js
assert.isFunction(removeFirstLunch);
```

Die Funktion `removeFirstLunch` sollte einen einzelnen Parameter haben.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Wenn das Eingabe-Array leer ist, sollte die Funktion `removeFirstLunch` den String `"No lunches to remove."` in der Konsole ausgeben.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` sollte den String `"Salad removed from the start of the lunch menu."` in der Konsole ausgeben.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` sollte `["Pizza", "Burger"]` zurückgeben.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch` und `removeFirstLunch` sollten dasselbe Array zurückgeben, das als Argument übergeben wurde, nachdem es aktualisiert wurde.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Sie sollten eine Funktion `getRandomLunch` definieren.

```js
assert.isFunction(getRandomLunch);
```

Die Funktion `getRandomLunch` sollte einen einzelnen Parameter haben.

```js
assert.lengthOf(getRandomLunch, 1);
```

Wenn das Eingabe-Array leer ist, sollte die Funktion `getRandomLunch` den String `"No lunches available."` in der Konsole ausgeben.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Wenn das Eingabe-Array nicht leer ist, sollte die Funktion `getRandomLunch` einen String im Format `Randomly selected lunch: [Lunch Item]` in der Konsole ausgeben.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

Die Funktion `getRandomLunch` sollte das übergebene Array als Argument nicht verändern.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Sie sollten eine Funktion `showLunchMenu` definieren.

```js
assert.isFunction(showLunchMenu);
```

Die Funktion `showLunchMenu` sollte einen einzelnen Parameter haben.

```js
assert.lengthOf(showLunchMenu, 1);
```

Wenn das Eingabe-Array leer ist, sollte die Funktion `showLunchMenu` den String `"The menu is empty."` in der Konsole ausgeben.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` sollte `"Menu items: Greens, Corns, Beans"` in der Konsole ausgeben.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` sollte `"Menu items: Pizza, Burger, Fries, Salad"` in der Konsole ausgeben.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
