---
id: 587d8249367417b2b2512c41
title: Convertitore da unità metriche a imperiali
challengeType: 4
forumTopicId: 301570
dashedName: metric-imperial-converter
---

# --description--

Crea un'app JavaScript full-stack che sia funzionalmente simile a questa: <a href="https://metric-imperial-converter.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://metric-imperial-converter.freecodecamp.rocks/</a>. Lavorare su questo progetto ti richiederà di scrivere il tuo codice usando uno dei seguenti metodi:

- Clona <a href="https://github.com/freeCodeCamp/boilerplate-project-metricimpconverter/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa il progetto localmente.
- Usa un site builder a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

**Nota:** I test di questo progetto non funzionano usando `glitch.com`.

# --instructions--

- Completa la logica di conversione necessaria in `/controllers/convertHandler.js`
- Completa le rotte necessarie in `/routes/api.js`
- Copia il file `sample.env` in `.env` e imposta le variabili in modo appropriato
- Per eseguire i test automaticamente, aggiungi `NODE_ENV=test` nel tuo file `.env`
- Per eseguire i test in console, usa il comando `npm run test`

Scrivi i seguenti test in `tests/1_unit-tests.js`:

- `convertHandler` dovrebbe leggere correttamente un input numerico intero.
- `convertHandler` dovrebbe leggere correttamente un input numerico decimale.
- `convertHandler` dovrebbe leggere correttamente un input frazionario.
- `convertHandler` dovrebbe leggere correttamente un input frazionario con decimale.
- `convertHandler` dovrebbe restituire correttamente un errore su una doppia frazione (cioè `3/2/3`).
- `convertHandler` dovrebbe impostare correttamente il valore numerico predefinito a `1` quando non viene fornito alcun input numerico.
- `convertHandler` dovrebbe leggere correttamente ogni unità di input valida.
- `convertHandler` dovrebbe restituire correttamente un errore per un'unità di input non valida.
- `convertHandler` dovrebbe restituire l'unità di ritorno corretta per ogni unità di input valida.
- `convertHandler` dovrebbe restituire correttamente la stringa scritta per ogni unità di input valida.
- `convertHandler` dovrebbe convertire correttamente `gal` in `L`.
- `convertHandler` dovrebbe convertire correttamente `L` in `gal`.
- `convertHandler` dovrebbe convertire correttamente `mi` in `km`.
- `convertHandler` dovrebbe convertire correttamente `km` in `mi`.
- `convertHandler` dovrebbe convertire correttamente `lbs` in `kg`.
- `convertHandler` dovrebbe convertire correttamente `kg` in `lbs`.

Scrivi i seguenti test in `tests/2_functional-tests.js`:

- Converti un input valido come `10L`: richiesta `GET` a `/api/convert`.
- Converti un input non valido come `32g`: richiesta `GET` a `/api/convert`.
- Converti un numero non valido come `3/7.2/4kg`: richiesta `GET` a `/api/convert`.
- Converti un numero E un'unità non validi come `3/7.2/4kilomegagram`: richiesta `GET` a `/api/convert`.
- Converti senza numero come `kg`: richiesta `GET` a `/api/convert`.

# --hints--

Puoi fornire il tuo progetto, non l'URL di esempio.

```js
  assert(
    !/.*\/metric-imperial-converter\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Puoi `GET` `/api/convert` con un singolo parametro contenente un numero e un'unità accettati e farlo convertire. (Suggerimento: dividi l'input cercando l'indice del primo carattere che segnerà l'inizio dell'unità)

```js

```

Puoi convertire `'gal'` in `'L'` e viceversa. (1 gal in 3.78541 L)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 3.78541);
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10gal');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 37.8541);
    assert.equal(data2.returnUnit, 'L');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.26417);
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10l');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 2.64172);
    assert.equal(data4.returnUnit, 'gal');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Puoi convertire `'lbs'` in `'kg'` e viceversa. (1 lbs in 0.453592 kg)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1lbs');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 0.45359);
    assert.equal(data1.returnUnit, 'kg');
    const response2 = await fetch(code + '/api/convert?input=10lbs');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 4.53592);
    assert.equal(data2.returnUnit, 'kg');
    const response3 = await fetch(code + '/api/convert?input=1kg');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 2.20462);
    assert.equal(data3.returnUnit, 'lbs');
    const response4 = await fetch(code + '/api/convert?input=10kg');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 22.04624);
    assert.equal(data4.returnUnit, 'lbs');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Puoi convertire `'mi'` in `'km'` e viceversa. (1 mi in 1.60934 km)

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.returnNum, 1.60934);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=10mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.returnNum, 16.0934);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(code + '/api/convert?input=1km');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.returnNum, 0.62137);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(code + '/api/convert?input=10km');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.returnNum, 6.21373);
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Tutte le unità in ingresso dovrebbero essere accettate sia in maiuscolo che minuscolo, ma dovrebbero essere restituite sia in `initUnit` che `returnUnit` in minuscolo, tranne per il litro, che dovrebbe essere rappresentato con `'L'` maiuscola.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=1gal');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.equal(data1.initUnit, 'gal');
    assert.equal(data1.returnUnit, 'L');
    const response2 = await fetch(code + '/api/convert?input=10L');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.equal(data2.initUnit, 'L');
    assert.equal(data2.returnUnit, 'gal');
    const response3 = await fetch(code + '/api/convert?input=1l');
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.equal(data3.initUnit, 'L');
    assert.equal(data3.returnUnit, 'gal');
    const response4 = await fetch(code + '/api/convert?input=10KM');
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.equal(data4.initUnit, 'km');
    assert.equal(data4.returnUnit, 'mi');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Se l'unità di misura non è valida, verrà restituito `'invalid unit'`.

```js
  try {
    const response = await fetch(code + '/api/convert?input=1min');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid unit' || data === 'invalid unit');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Se il numero non è valido, verrà restituito `'invalid number'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2gal'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(data.error === 'invalid number' || data === 'invalid number');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Se sia l'unità che il numero non sono validi, verrà restituito `'invalid number and unit'`.

```js
  try {
    const response = await fetch(
      code + '/api/convert?input=1//2min'
    );
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.text();
    assert(
      data.error === 'invalid number and unit' ||
        data === 'invalid number and unit'
    );
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Puoi usare frazioni, decimali o entrambi nel parametro (es. 5, 1/2, 2.5/6), ma se non viene fornito nulla, il valore predefinito sarà 1.

```js
  try {
    const response1 = await fetch(code + '/api/convert?input=mi');
    if (!response1.ok) {
      throw Error(await response1.text());
    }
    const data1 = await response1.json();
    assert.approximately(data1.initNum, 1, 0.001);
    assert.approximately(data1.returnNum, 1.60934, 0.001);
    assert.equal(data1.returnUnit, 'km');
    const response2 = await fetch(code + '/api/convert?input=1/5mi');
    if (!response2.ok) {
      throw Error(await response2.text());
    }
    const data2 = await response2.json();
    assert.approximately(data2.initNum, 1 / 5, 0.1);
    assert.approximately(data2.returnNum, 0.32187, 0.001);
    assert.equal(data2.returnUnit, 'km');
    const response3 = await fetch(
      code + '/api/convert?input=1.5/7km'
    );
    if (!response3.ok) {
      throw Error(await response3.text());
    }
    const data3 = await response3.json();
    assert.approximately(data3.initNum, 1.5 / 7, 0.001);
    assert.approximately(data3.returnNum, 0.13315, 0.001);
    assert.equal(data3.returnUnit, 'mi');
    const response4 = await fetch(
      code + '/api/convert?input=3/2.7km'
    );
    if (!response4.ok) {
      throw Error(await response4.text());
    }
    const data4 = await response4.json();
    assert.approximately(data4.initNum, 3 / 2.7, 0.001);
    assert.approximately(data4.returnNum, 0.69041, 0.001);
    assert.equal(data4.returnUnit, 'mi');
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Il tuo ritorno consisterà in `initNum`, `initUnit`, `returnNum`, `returnUnit` e `string` che scrivono le unità nel formato `'{initNum} {initUnitString} converts to {returnNum} {returnUnitString}'` con il risultato arrotondato a 5 decimali.

```js
  try {
    const response = await fetch(code + '/api/convert?input=2mi');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const data = await response.json();
    assert.equal(data.initNum, 2);
    assert.equal(data.initUnit, 'mi');
    assert.approximately(data.returnNum, 3.21868, 0.001);
    assert.equal(data.returnUnit, 'km', 'returnUnit did not match');
    assert.equal(data.string, '2 miles converts to 3.21868 kilometers');
  } catch (xhr) {
    throw new Error(xhr.responseText || xhr.message);
  }
```

Tutti i 16 test delle unità sono completi e superati.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const unitTests = getTests.filter(test => {
      return !!test.context.match(/Unit Tests/gi);
    });
    assert.isAtLeast(unitTests.length, 16, 'At least 16 tests passed');
    unitTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```

Tutti i 5 test funzionali sono completi e superati.

```js
  try {
    const response = await fetch(code + '/_api/get-tests');
    if (!response.ok) {
      throw Error(await response.text());
    }
    const getTests = await response.json();
    assert.isArray(getTests);
    const functTests = getTests.filter(test => {
      return !!test.context.match(/Functional Tests/gi);
    });
    assert.isAtLeast(functTests.length, 5, 'At least 5 tests passed');
    functTests.forEach(test => {
      assert.equal(test.state, 'passed', 'Tests in Passed State');
      assert.isAtLeast(
        test.assertions.length,
        1,
        'At least one assertion per test'
      );
    });
  } catch (err) {
    throw new Error(err.responseText || err.message);
  }
```
