---
id: 66a8347f97a33e36ffc81b9f
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have an opening `body` tag.

```js
assert.match(code, /<body>/i);
```

You should have a closing `body` tag.

```js
assert.match(code, /<\/body>/i);
```

你不應該改變你的 `head` 元素。 確認你沒有刪除結束標籤。

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

你的 `body` 元素應該在你的 `head` 元素之後。

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
