---
id: 699d20623e57e8ffee88545a
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

目录解析完成后，你可以按字段进行查找。`findByAuthor` 会过滤作者包含搜索词的条目。

创建一个名为 `findByAuthor` 的函数，它接受 `catalog` 和 `author` 作为参数。在函数内部，声明一个名为 `searchTerm` 的变量，并将 `author` 转为小写，以便后续进行不区分大小写的匹配。

还要声明一个名为 `results` 的变量，并赋值为空数组，用来存放匹配项。

# --hints--

你应该有一个名为 `findByAuthor` 的函数，带有 `catalog` 和 `author` 两个参数。

```js
const regex = __helpers.functionRegex('findByAuthor', ['catalog', 'author']);
assert.match(__helpers.removeJSComments(code), regex);
```

你应该声明一个名为 `searchTerm` 的变量，并赋值为 `author.toLowerCase()`。

```js
assert.match(findByAuthor.toString(), /searchTerm\s*=\s*author\.toLowerCase\s*\(\s*\)/);
```

你应该声明一个名为 `results` 的变量，并赋值为空数组。

```js
assert.match(findByAuthor.toString(), /(const|let|var)\s+results\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);
--fcc-editable-region--

--fcc-editable-region--
```
