---
id: 699d20623e57e8ffee885450
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

其余的目录卡片已经添加。注意接近末尾的三个条目缺少字段；你将在后续步骤中处理它们。

`parseCard` 会接受一个未经处理的目录字串并返回一个包含四个字段的结构化对象：`title`、`author`、`year` 和 `location`。现在，创建一个名为 `parseCard` 的函数，它有一个名为 `rawString` 的参数。

# --hints--

你应该有一个名为 `parseCard` 的函数。

```js
assert.isFunction(parseCard);
```

`parseCard` 应该是一个带有名为 `rawString` 参数的函数。

```js
const regex = __helpers.functionRegex('parseCard', ['rawString']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
