---
id: 699d20623e57e8ffee88545c
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

现在你将创建 `groupByDecade`，它按年代组织目录。每个密钥将是类似 `"1970s"` 的字串，每个值是该时期的书籍数组。

创建一个名为 `groupByDecade` 的函数，它接受 `catalog` 作为参数。在函数内部，声明 `const grouped = {}`。添加一个 `for` 循环来迭代 `catalog`，并在循环内声明 `const book = catalog[i]`。

# --hints--

你应该有一个名为 `groupByDecade` 的函数，带有一个 `catalog` 参数。

```js
const regex = __helpers.functionRegex('groupByDecade', ['catalog']);
assert.match(__helpers.removeJSComments(code), regex);
```

你应该声明一个名为 `grouped` 的变量，并将其设置为空对象。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+grouped\s*=\s*\{\s*\}/);
```

你应该使用 `for` 循环来迭代 `catalog`。

```js
assert.match(groupByDecade.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*catalog\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

你应该声明一个名为 `book` 的变量，并将其设置为当前目录条目。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+book\s*=\s*catalog\[/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

--fcc-editable-region--

--fcc-editable-region--
```
