---
id: 699d20623e57e8ffee885451
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

`parseCard` 函数将接受一个未经处理的目录字串并返回一个结构化的对象。在创建你的函数之前，先设置一种方法来查看它的输出。

调用 `parseCard` 并传入 `rawCatalogCards[2]` 作为参数，将结果赋值给名为 `cardResult` 的变量，并将 `cardResult` 记录到控制台。

目前 `parseCard` 没有体，所以它返回 `undefined`。随着你在接下来的步骤中为函数添加代码，日志输出将准确显示函数在每个阶段返回的内容。

# --hints--

你应该调用 `parseCard` 并传入 `rawCatalogCards[2]` 作为参数。

```js
assert.match(__helpers.removeJSComments(code), /parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)/);
```

你应该将调用 `parseCard` 并传入 `rawCatalogCards[2]` 的结果赋值给名为 `cardResult` 的变量。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+cardResult\s*=\s*parseCard\s*\(\s*rawCatalogCards\s*\[\s*2\s*\]\s*\)\s*;?/);
```

你应该将 `cardResult` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*cardResult\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {}

--fcc-editable-region--

--fcc-editable-region--
```
