---
id: 699d20623e57e8ffee885454
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

解析目录卡的第一步是分离四个字段。`.split()` 方法会在你提供的分隔符每次出现时，将字串拆分成一个子字串数组。

例如：

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

在 `parseCard` 内，调用 `.split("|")` 处理 `rawString` 并返回结果。完成这一步后，控制台应显示一个包含四个字串的数组。

**注意：** 你可能会注意到值周围有多余的空白。你将在下一步清理这些空白。

# --hints--

你应该在 `.split("|")` 函数内调用 `rawString` 处理 `parseCard`。

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` 应该返回一个数组。

```js
assert.isArray(parseCard("A | B | C | D"));
```

当使用 `"A | B | C | D"` 调用时，`parseCard` 应该返回一个包含四个元素的数组。

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
