---
id: 699d20623e57e8ffee885455
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

从控制台的输出可以看到，拆分后，每个部分仍然有多余的前导和/或尾随空白。要清理每个元素，你需要单独处理每一个。

这正是使用 `for` 循环的完美场景，它会对可迭代对象（例如数组）中的每个元素重复执行一段代码：

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i` 从 `0`（第一个索引）开始，循环在 `i` 小于数组长度时继续，`i++` 在每次迭代后将 `i` 增加一。

在 `parseCard` 内，声明一个名为 `trimmedParts` 的空数组。然后添加一个 `for` 循环来迭代 `parts`。循环体暂时留空，下一步你会填充它。

# --hints--

你应该声明一个名为 `trimmedParts` 的变量，并将其设置为空数组。

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

你应该使用 `for` 循环来迭代 `parts`。

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
