---
id: 699d20623e57e8ffee885459
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard` 已经经过测试并且工作正常。`cardResult` 变量不再需要。

首先删除 `const cardResult` 行和 `console.log(cardResult)` 行。然后创建一个名为 `parseCatalog` 的函数，该函数接受 `rawCards` 作为参数。现在，在函数内部声明一个空的数组，名为 `catalog`，并返回它。

在函数下面，调用 `parseCatalog(rawCatalogCards)`，将结果赋值给一个名为 `catalog` 的变量，并打印 `catalog.length`。

# --hints--

`parseCatalog` 应该是一个返回数组的函数。

```js
assert.isArray(parseCatalog(rawCatalogCards));
```

`catalog` 应该被定义并赋值为 `parseCatalog(rawCatalogCards)`。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+catalog\s*=\s*parseCatalog\s*\(/);
```

你应该删除 `cardResult` 变量和相关的 `console.log()` 调用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

你应该将 `catalog.length` 打印到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
