---
id: 699d1f8e51fcdc8230f779fb
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将构建一个遗产图书馆目录系统。你将解析未经处理的图书数据，查找和分组条目，验证数据质量，并将结果导出为不同格式。

每张目录卡以单个字串的形式保存，格式为 `"Title | Author | Year | Location"`，字段由管道字符（`|`）分隔。

你将在整个工作坊中使用这些字串组成的数组来构建和测试你的目录函数。创建一个名为 `rawCatalogCards` 的数组，并向其中添加以下两个字串：

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

你应该声明一个名为 `rawCatalogCards` 的变量。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` 应该是一个数组。

```js
assert.isArray(rawCatalogCards);
```

`rawCatalogCards` 的第一个元素应该是 `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`。

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

`rawCatalogCards` 的第二个元素应该是 `"The Shining | King, Stephen | 1977 | Shelf K1"`。

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
