---
id: 699d20623e57e8ffee885457
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

有些目录卡片缺少字段。例如，条目 16 的作者为空，因此 `parseCard` 当前为该字段返回空字串。

`||` 操作符提供了一个备用方案：当 `a || b` 为真值时，`a` 计算结果为 `a`，否则为 `b`。由于空字串是假值，`"" || "Unknown"` 计算结果为 `"Unknown"`。

更新返回语句中的 `title`、`author` 和 `location`，使用 `|| "Unknown"` 作为备用。保留 `year` 以供下一步使用。

# --hints--

当传入的字串作者字段为空时，返回对象的 `author` 属性应为 `"Unknown"`。

```js
assert.strictEqual(parseCard("X |  | 2000 | Y").author, "Unknown");
```

当传入的字串标题字段为空时，返回对象的 `title` 属性应为 `"Unknown"`。

```js
assert.strictEqual(parseCard(" | A | 2000 | Y").title, "Unknown");
```

当传入的字串位置字段为空时，返回对象的 `location` 属性应为 `"Unknown"`。

```js
assert.strictEqual(parseCard("X | A | 2000 | ").location, "Unknown");
```

你还不应该为 `year` 添加备用。

```js
assert.strictEqual(parseCard("X | A |  | Y").year, "");
```

当所有字段都存在时，`parseCard` 应该返回它们的实际值，而不是 `"Unknown"`。

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--
  return {
    title: title,
    author: author,
    year: year,
    location: location
  };
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
