---
id: a789b3483989747d63b0e427
title: Construa o Localizador do Maior Número
challengeType: 26
dashedName: build-the-largest-number-finder
---

# --description--

Neste laboratório, você construirá uma função que retorna um array contendo o maior número de cada sub-array fornecido.

Lembre-se, você pode iterar por um array com um simples loop `for` e acessar cada elemento com a sintaxe de array `arr[i]`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função `largestOfAll` que recebe um array de arrays como argumento.
2. A função deve retornar um array contendo o maior número de cada sub-array.

# --hints--

Você deve criar uma função chamada `largestOfAll`.

```js
assert.isFunction(largestOfAll);
```

`largestOfAll` deve receber um único parâmetro.

```js
assert.lengthOf(largestOfAll, 1);
```

`largestOfAll([[4, 5, 1, 3], [13, 27, 18, 26], [32, 35, 37, 39], [1000, 1001, 857, 1]])` deve retornar um array.

```js
assert.isArray(
  largestOfAll([
    [4, 5, 1, 3],
    [13, 27, 18, 26],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ])
);
```

`largestOfAll([[13, 27, 18, 26], [4, 5, 1, 3], [32, 35, 37, 39], [1000, 1001, 857, 1]])` deve retornar `[27, 5, 39, 1001]`.

```js
assert.deepEqual(
  largestOfAll([
    [13, 27, 18, 26],
    [4, 5, 1, 3],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ]),
  [27, 5, 39, 1001]
);
```

`largestOfAll([[4, 9, 1, 3], [13, 35, 18, 26], [32, 35, 97, 39], [1000000, 1001, 857, 1]])` deve retornar `[9, 35, 97, 1000000]`.

```js
assert.deepEqual(
  largestOfAll([
    [4, 9, 1, 3],
    [13, 35, 18, 26],
    [32, 35, 97, 39],
    [1000000, 1001, 857, 1]
  ]),
  [9, 35, 97, 1000000]
);
```

`largestOfAll([[17, 23, 25, 12], [25, 7, 34, 48], [4, -10, 18, 21], [-72, -3, -17, -10]])` deve retornar `[25, 48, 21, -3]`.

```js
assert.deepEqual(
  largestOfAll([
    [17, 23, 25, 12],
    [25, 7, 34, 48],
    [4, -10, 18, 21],
    [-72, -3, -17, -10]
  ]),
  [25, 48, 21, -3]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function largestOfAll(arr) {
  return arr.map(subArr => Math.max.apply(null, subArr));
}

largestOfAll([
  [4, 5, 1, 3],
  [13, 27, 18, 26],
  [32, 35, 37, 39],
  [1000, 1001, 857, 1]
]);
```
