---
id: 66c051d13a6a20255a963695
title: Jenga orodha ya kazi yenye mtindo
challengeType: 25
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

Katika maabara hii, utazoeza mitindo tofauti inayoweza kutumika kwa viungo vinapobofyanwa, kupewa makini, kubofyanwa, na kutembelewa.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na orodha isiyopangwa moja yenye darasa la `todo-list`.
2. Ndani ya orodha isiyopangwa, unapaswa kuwa na vitu vinne vya orodha.
3. Ndani ya kila kipengee cha orodha, kunapaswa kuwepo:

   - Kipengele cha `input` chenye aina ya `checkbox` na `id` kikiwa na thamani ya kipekee.
   - Kipengele cha `label` chenye sifa ya `for` iliyowekwa kwa thamani ya `input` ya kipengele cha `id` kinachohusiana.
   - Orodha isiyopangwa yenye darasa la `sub-item`.
   - Kipengee cha orodha chenye kipengele cha nanga ndani yake. Tagi ya nanga inapaswa kuwa na darasa la `sub-item-link`, thamani halali ya `href`, na thamani ya `target` inayofanya kiungo kufunguka kwenye kichupo kipya.

4. Vipengele vyako vya `a` visiwe na mapambo yoyote ya maandishi.
5. Unapaswa kuweka rangi ya maandishi ya viungo visivyotembelewa kwa rangi unayotaka.
6. Viungo vyako vinapotembelewa, rangi inapaswa kubadilika kuwa rangi nyingine unayotaka.
7. Viungo vyako vinapobofyanwa, rangi inapaswa kubadilika kuwa rangi nyingine unayotaka.
8. Viungo vyako vinapopewa makini, panapaswa kuwepo na mstari wa rangi unaozunguka kiungo.
9. Viungo vyako vinapobofyanwa, rangi inapaswa kubadilika kuwa rangi nyingine unayotaka.

**Kumbuka:** Hakikisha umeunganisha karatasi yako ya mitindo (stylesheet) katika HTML yako na kutumia CSS yako.

# --hints--

Unapaswa kuwa na orodha isiyopangwa moja yenye darasa la `todo-list`.

```js
assert.exists(document.querySelector('ul.todo-list'));
assert.lengthOf(document.querySelectorAll('ul.todo-list'), 1);
```

Unapaswa kuwa na vitu vinne vya orodha ndani ya orodha isiyopangwa.

```js
const li = document.querySelectorAll('ul.todo-list > li');
assert.lengthOf(li, 4);
```

`li` ndani ya `ul` yenye darasa la `todo-list` inapaswa kuwa na kipengele cha `input` chenye `type` ya `checkbox`.

```js
const checkboxes = document.querySelectorAll(
  'ul.todo-list li input[type="checkbox"]'
);
assert.lengthOf(checkboxes, 4);
```

`li` ndani ya `ul` yenye darasa la `todo-list` inapaswa kuwa na kipengele cha `label`.

```js
const labels = document.querySelectorAll(
  'ul.todo-list li label'
);
assert.lengthOf(labels, 4);
```

Vipengele vyote vya `input` vinapaswa kuwa na `id`.

```js
const inputs = document.querySelectorAll('input');
assert.isAbove(inputs.length, 0);
for (let input of inputs) {
  assert.isAbove(input?.id.length, 0);
}
```

Vipengele vyote vya `label` vinapaswa kuwa na sifa ya `for`.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.htmlFor.length, 0);
}
```

Vipengele vyote vya `label` vinapaswa kuwa na maandishi.

```js
const labels = document.querySelectorAll('label');
assert.isAbove(labels.length, 0);
for (let label of labels) {
  assert.isAbove(label.innerText.length, 0);
}
```

Sifa za `id` na `for` za vipengele vya `input` na `label` zinapaswa kuwa na thamani zinazolingana.

```js
const labels = document.querySelectorAll("label");
assert.isNotEmpty(labels);

document.querySelectorAll("input").forEach((input, index) => {
  assert.equal(input.id, labels[index].htmlFor);
});
```

Baada ya vipengele vya `label`, panapaswa kuwepo na orodha isiyopangwa yenye darasa la `sub-item`.

```js
const labelSbItems = document.querySelectorAll('ul.todo-list li label + ul.sub-item');
assert.lengthOf(labelSbItems, 4);
```

`li` ndani ya `ul` yenye darasa la `sub-item` inapaswa kuwa na kipengele cha nanga chenye darasa la `sub-item-link`.

```js
const allItems = document.querySelectorAll(
  'ul.todo-list li label + ul.sub-item li a.sub-item-link'
);
assert.lengthOf(allItems, 4);
```

Vipengele vyote vya nanga vinapaswa kuwa na sifa halali ya `href`.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.href.length, 0);
}
```

Vipengele vyote vya nanga vinapaswa kuwa na maandishi.

```js
const links = document.querySelectorAll('a');
assert.isAbove(links.length, 0);
for (let link of links) {
  assert.isAbove(link.innerText.length, 0);
}
```

Kila kipengele cha `a` kinapaswa kuwa na sifa ya `target` yenye thamani ya `_blank`.

```js
const anchors = document.querySelectorAll('a');
assert.isAbove(anchors.length, 0);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

Vipengele vyako vya `a` visiwe na mapambo yoyote ya maandishi.

```js
assert.isTrue(
  new __helpers.CSSHelp(document).getStyle('a.sub-item-link')
    ?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('a')?.textDecoration === 'none' ||
    new __helpers.CSSHelp(document).getStyle('.sub-item-link')
      ?.textDecoration === 'none'
);
```

Viungo visivyotembelewa vinapaswa kuwa na rangi ya maandishi iliyowekwa kwa rangi unayotaka.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:link')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:link')?.color
);
```

Viungo vinapaswa kubadilisha rangi vinapobofyanwa.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:hover')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:hover')?.color
);
```

Viungo vinapaswa kubadilisha rangi vinapobofyanwa.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:active')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:active')?.color
);
```

Viungo vinapaswa kuwa na mstari wa rangi wanapopewa makini.

```js
const css = new __helpers.CSSHelp(document);
assert.exists(
  css.getStyle('.sub-item-link:focus')?.outline ||
    css.getStyle('a:focus')?.outline ||
    css.getStyle('a.sub-item-link:focus')?.outline ||
    css.getStyle('.sub-item-link:focus')?.border ||
    css.getStyle('a:focus')?.border ||
    css.getStyle('a.sub-item-link:focus')?.border
);
```

Viungo vinapaswa kubadilisha rangi mara vinapotembelewa.

```js
assert.exists(
  new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a:visited')?.color ||
    new __helpers.CSSHelp(document).getStyle('a.sub-item-link:visited')?.color
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div class="container">
      <h1>My To-Do List</h1>
      <ul class="todo-list">
        <li>
          <input type="checkbox" id="task1" />
          <label for="task1">Explore gaming keyboards</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link">Store link</a>
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task2" />
          <label for="task2">Finish the report</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >Request Access</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task3" />
          <label for="task3">View Phone's Warranty</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Receipts</a
              >
            </li>
          </ul>
        </li>

        <li>
          <input type="checkbox" id="task4" />
          <label for="task4">Check Processor Specs</label>
          <ul class="sub-item">
            <li>
              <a href="#" target="_blank" class="sub-item-link"
                >View Model Number</a
              >
            </li>
          </ul>
        </li>
      </ul>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: #f0f0f0;
  margin: 0;
  padding: 20px;
}

.container {
  max-width: 400px;
  margin: 0 auto;
  background-color: #fff;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
  color: #333;
}

.todo-list {
  list-style-type: none;
  padding: 0;
}

.todo-list li {
  margin-bottom: 15px;
  line-height: 1.6;
  padding: 10px;
  background-color: #e5f1e5;
  border-radius: 10px;
  position: relative;
}

a.sub-item-link {
  margin-left: 10px;
  text-decoration: none;
}

.sub-item-link:link {
  color: blue;
}

.sub-item-link:visited {
  color: purple;
}

.sub-item-link:hover {
  color: red;
}

.sub-item-link:focus {
  outline: 2px solid blue;
  outline-offset: 2px;
}

.sub-item-link:active {
  color: green;
}

.sub-item {
  list-style-type: square;
}
```
