---
id: 690ddb5e6ee94eb73ed172b0
title: Ein Formular für Spenden debuggen
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Eine lokale Wohltätigkeitsorganisation hat eine Website mit einem Spendenformular erstellt, aber es gibt mehrere Probleme, die behoben werden müssen. Das Formular ist nicht barrierefrei und enthält einige HTML-Syntaxfehler.

Ihre Aufgabe ist es, alle Fehler zu beheben, damit das Formular korrekt funktioniert und für alle Benutzer zugänglich ist. Erfüllen Sie die Punkte in den User Stories unten und klicken Sie auf „Check Your Code“, um zu sehen, ob Sie alle Fehler behoben haben.

**User Stories:**

1. Die `input`-Elemente sind Void-Elemente und sollten keine schließenden Tags haben. Entfernen Sie alle `</input>`-schließenden Tags aus dem Formular.
2. Fügen Sie für jedes Eingabefeld im Formular `label`-Elemente hinzu, damit Benutzer wissen, wofür jedes Feld ist. Der Text der Bezeichnung sollte mit dem übereinstimmen, was derzeit neben jeder Eingabe steht.
3. Der `Email Address:`-Eingabetyp sollte ein `email` statt eines `text` sein.
4. Sie sollten jedes `label`-Element mit seinem entsprechenden `input`-Element über das `for`-Attribut am `label` und ein passendes `id`-Attribut am `input` verknüpfen.
5. Fügen Sie den Eingabefeldern für Text, E-Mail und Zahl das `required`-Attribut hinzu (aber nicht dem Kontrollkästchen oder der Absenden-Schaltfläche), um sicherzustellen, dass Benutzer die erforderlichen Informationen ausfüllen.

# --hints--

Sie sollten keine `</input>`-schließenden Tags in Ihrem Code haben.

```js
assert.notMatch(code, /<\/input>/);
```

Sie sollten genau fünf `input`-Elemente in Ihrem Formular haben.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Sie sollten genau vier `label`-Elemente in Ihrem Formular haben.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Ihr erstes `label` sollte den Text `Full Name:` haben.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Ihr erstes `input` sollte ein `required`-Attribut haben.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Ihr zweites `label` sollte den Text `Email Address:` haben.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Ihr zweites `input` sollte ein `required`-Attribut haben.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Ihr drittes `label` sollte den Text `Donation Amount ($):` haben.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Ihr drittes `input` sollte ein `required`-Attribut haben.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Ihr viertes `label` sollte den Text `Subscribe` haben.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Sie sollten jedes `input`-Element mit einem `label`-Element über die `for`- und `id`-Attribute verknüpfen.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Ihr `input` für die E-Mail-Adresse sollte vom Typ `email` sein.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Ihr Kontrollkästchen-`input` sollte kein `required`-Attribut haben.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Ihre Absenden-`button` sollte kein `required`-Attribut haben.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
