---
id: 69860157f7fce39f9df399a3
title: تنفيذ خوارزمية N-Queens
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

تطلب مشكلة N-Queens وضع N ملكة على رقعة شطرنج بحجم N×N بحيث لا تهاجم أي ملكتين بعضهما البعض (لا يتشاركان في صف أو عمود أو قطر).

على سبيل المثال، إذا كانت الرقعة 4×4، فإن أحد الترتيبات الصحيحة هو:

```md
[1, 3, 0, 2]
```

هذا يعني أنه في الصف 0، وُضعت الملكة في العمود 1؛ في الصف 1، وُضعت الملكة في العمود 3؛ في الصف 2، وُضعت الملكة في العمود 0؛ وفي الصف 3، وُضعت الملكة في العمود 2.

بصريًا، يبدو هذا الترتيب كما يلي:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

حيث تمثل `Q` الملكة وتمثل `.` مربعًا فارغًا.

في هذا المختبر، ستنفذ حل مشكلة N-Queens باستخدام طريقة البحث بالعمق أولاً.


**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي على دالة باسم `dfsNQueens`.
2. يجب أن تقبل الدالة معلمة واحدة فقط: عدد صحيح `n`.
3. إذا كان `n` أقل من `1`، يجب أن تُرجع الدالة مصفوفة فارغة (`[]`).
4. يجب أن تُرجع الدالة قائمة من الحلول؛ كل حل هو قائمة بطول `n`، حيث العنصر في الفهرس `i` هو فهرس العمود (مُرقم من صفر) للملكة في الصف `i`.

# --hints--

يجب أن يكون لديك دالة باسم `dfsNQueens` تأخذ معلمة واحدة.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

إذا كانت قيمة `n` أقل من `1`، يجب أن تُرجع الدالة مصفوفة فارغة.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

يجب أن تُرجع الدالة مصفوفة من الحلول، حيث كل حل هو مصفوفة بطول `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` يجب أن تُرجع `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` يجب أن تُرجع `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` يجب أن تُرجع `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` يجب أن تُرجع `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` يجب أن تُرجع `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` يجب أن تكون `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` يجب أن تكون `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

يجب أن تُرجع `dfsNQueens` النتيجة الصحيحة.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
