---
id: 613297a923965e0703b64796
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

你可能已經熟悉 `meta` 元素；它用於指定關於頁面的資訊，例如標題、描述、關鍵字和作者。

為你的頁面添加具有適當 `charset` 值的 `meta` 元素。

`charset` 屬性指定頁面的字元編碼，而現今大多數瀏覽器只支援 `UTF-8` 編碼。

# --hints--

你應該在 `head` 元素內創建一個 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你應該為 `meta` 頁籤設定 `UTF-8` 的 `charset`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
