---
id: 6133d11ef548f51f876149e3
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

導覽是存取性的核心部分，而螢幕閱讀器依賴你提供頁面的結構。這是透過語意化的 HTML 元素來完成的。

為你的頁面添加 `header` 和 `main` 元素。

`header` 元素將用於介紹頁面，以及提供導覽選單。

`main` 元素將包含你頁面的核心內容。

# --hints--

你應該為你的 `body` 添加一個 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你應該為你的 `body` 添加一個 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

`header` 元素應該位於 `main` 元素之前。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    
  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
