---
id: 6133acc353338c0bba9cb553
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

最後在 `head` 中，`title` 元素對於螢幕閱讀器理解頁面內容非常有用。此外，它是 *SEO* 的重要零件。

為你的頁面提供一個描述性且簡潔的 `title`。

# --hints--

你應該為你的 `head` 添加一個 `title` 元素。

```js
assert.exists(document.querySelector('head > title'));
```

你不應該讓 `title` 超過 60 個字元。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

嘗試讓你的 `title` 元素更具描述性。*提示：至少 15 個字元*

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
