---
id: 613e2546d0594208229ada50
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

在 `header` 中，透過巢狀一個 `img`、`h1` 和 `nav` 元素來提供頁面的背景關係。

`img` 應該指向 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`，具有 `id` 為 `logo`，並且有 `alt` 文字為 `freeCodeCamp`。

`h1` 應該包含文字 `HTML/CSS Quiz`。

# --hints--

你應該為你的 `header` 元素添加一個 `img` 元素。

```js
assert.exists(document.querySelector('header > img'));
```

你應該為你的 `header` 元素添加一個 `h1` 元素。

```js
assert.exists(document.querySelector('header > h1'));
```

你應該為你的 `header` 元素添加一個 `nav` 元素。

```js
assert.exists(document.querySelector('header > nav'));
```

你應該依序放置 `img`、`h1` 和 `nav` 元素。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

你應該為 `img` 元素指定一個 `src` 屬性，值為 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

你應該為 `img` 元素指定一個 `id` 屬性為 `logo`。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

你應該為 `img` 元素給予一個 `alt` 屬性，值為 `freeCodeCamp`。

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

你應該給 `h1` 元素添加文字 `HTML/CSS Quiz`。

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
