---
id: 613e275749ebd008e74bb62e
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

目前，`img` 正假設其預設大小，這太大了。CSS 有一個 `max` 函式（程式），它會傳回一組以逗號分隔的值中最大的值。<dfn>例如</dfn>：

```css
img {
  width: max(250px, 25vw);
}
```

在上述範例中，如果視窗寬度小於 1000 像素，影像的寬度將會是 250px。如果視窗寬度大於 1000 像素，影像的寬度將會是 25vw。這是因為 25vw 相等於視窗寬度的 25％。

使用其 `id` 作為選擇器縮放影像，並將 `width` 設定為 `10rem` 或 `18vw` 中的最大值。 

# --hints--

你應該使用 `#logo` 選擇器來標的 `img` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

你應該給 `img` 一個 `width`，值為 `max(10rem, 18vw)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
