---
id: 6141fabd6f75610664e908fd
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

由於這是一個測驗，你將需要一個表單讓使用者提交答案。你可以使用 `section` 元素語意上分隔表單內的內容。

在 `main` 元素內，創建一個包含三個巢狀 `section` 元素的表單。然後，使用正確的方法，讓表單提交到 `https://freecodecamp.org/practice-project/accessibility-quiz`。

# --hints--

你應該將 `form` 元素巢狀於你的 `main` 元素內。

```js
assert.exists(document.querySelector('main > form'));
```

你應該在你的 `form` 元素中巢狀三個 `section` 元素。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

你應該為 `form` 元素指定一個 `action` 屬性。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

你應該為 `action` 屬性指定 `https://freecodecamp.org/practice-project/accessibility-quiz` 的值。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

你應該為 `form` 元素指定一個 `method` 屬性。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

你應該為 `form` 元素指定一個值為 `post` 的 `method` 屬性。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
