---
id: 614202874ca576084fca625f
title: 第 16 步
challengeType: 0
dashedName: step-16
---

# --description--

每個 `region` 角色都需要一個標籤，這有助於螢幕閱讀器使用者了解該區域的用途。為你的區域添加標籤的一種方法是，在區域內新增標題元素，然後使用 `aria-labelledby` 屬性取用它。

將以下 `aria-labelledby` 屬性添加到 `section` 元素：

- `student-info`
- `html-questions`
- `css-questions`

然後，在每個 `section` 元素內，巢狀一個 `h2` 元素，其 `id` 與對應的 `aria-labelledby` 屬性相符。為每個 `h2` 提供適當的文字內容。

# --hints--

你應該為第一個 `section` 元素給予一個 `aria-labelledby` 屬性，值為 `student-info`。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

你應該為第二個 `section` 元素添加一個 `aria-labelledby` 屬性，值為 `html-questions`。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

你應該為第三個 `section` 元素添加一個 `aria-labelledby` 屬性，值為 `css-questions`。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

你應該在第一個 `section` 元素中巢狀一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

你應該在第二個 `section` 元素中巢狀一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

你應該在第三個 `section` 元素內巢狀一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

你應該為第一個 `h2` 元素給予一個 `id` 屬性，值為 `student-info`。

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

你應該為第二個 `h2` 元素指定一個 `id` 屬性，值為 `html-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

你應該為第三個 `h2` 元素指定一個 `id` 屬性，值為 `css-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

你應該為第一個 `h2` 元素提供適當的文字內容。*提示：我會選擇 `Student Info`*

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[0]?.textContent?.length, 1);
```

你應該為第二個 `h2` 元素提供適當的文字內容。*提示：我會選擇 `HTML`*

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[1]?.textContent?.length, 1);
```

你應該為第三個 `h2` 元素提供適當的文字內容。*提示：我會選擇 `CSS`*

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[2]?.textContent?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
