---
id: 614ccc21ea91ef1736b9b578
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

歡迎來到無障礙測試的第一部分。 隨着您成爲經驗豐富的HTML和CSS開發人員，我們已經爲您提供了基本的模版文件。

通過爲你的 `html` 元素添加一個 `lang` 屬性來開始這段無障礙之旅。 這將幫助屏幕閱讀器識別頁面的語言。

# --hints--

你應該給 `html` 元素一個 `lang` 屬性。 _提示：對於英文，你可以使用值 `en` 。_

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
