---
id: 614ccc21ea91ef1736b9b578
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

歡迎來到無障礙測驗的第一部分。隨著你成為經驗豐富的 HTML 和 CSS 開發者，我們已經為你準備了基本的樣板。

開始這段無障礙旅程，請為你的 `html` 元素提供 `lang` 屬性。這將協助螢幕閱讀器識別頁面的語言。

# --hints--

你應該為 `html` 元素指定一個 `lang` 屬性。*提示：你可以使用值 `en` 代表英文。*

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
