---
id: afcc8d540bea9ea2669306b6
title: Construir uma Função de Repetição de String
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

Neste laboratório, você criará uma função que repete uma string dada um número específico de vezes. Para o propósito deste laboratório, não use o método embutido `.repeat()`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `repeatStringNumTimes` que recebe dois parâmetros: uma string e um número.
2. A função deve retornar a string repetida o número especificado de vezes.
3. Se o número for menor ou igual a zero, a função deve retornar uma string vazia.

# --hints--

Você deve criar uma função chamada `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` deve receber dois parâmetros.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

A função `repeatStringNumTimes` deve sempre retornar uma string.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` deve retornar a string `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` deve retornar a string `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` deve retornar a `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` deve retornar a string `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` deve retornar a `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` deve retornar uma string vazia (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` deve retornar `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

O método integrado `repeat()` não deve ser usado.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
