---
id: 587d774c367417b2b2512a9d
title: Jua Wakati Maandishi ya Alt Yanapaswa Kuachwa Tupu
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

Katika changamoto ya mwisho, ulijifunza kwamba kujumuisha sifa ya kipengele ya `alt` wakati wa kutumia lebo za `img` ni lazima. Hata hivyo, wakati mwingine picha zimekusanywa pamoja na maelezo ya picha tayari kuyazungumzia, au zinatumika kwa ajili ya mapambo tu. Katika hali hizi, maandishi ya `alt` yanaweza kuonekana kuwa ya ziada au hayahitajiki.

Wakati picha tayari imefafanuliwa kwa maudhui ya maandishi au haiongezi maana kwenye ukurasa, `img` bado inahitaji sifa ya kipengele ya `alt`, lakini inaweza kuwekwa kuwa mfuatano wa herufi tupu. Hapa kuna mfano:

```html
<img src="visualDecoration.jpeg" alt="">
```

Picha za nyuma kawaida huangukia chini ya lebo ya 'mapambo' pia. Hata hivyo, mara nyingi hutumika kwa sheria za CSS, na kwa hivyo si sehemu ya alama za uandishi ambazo kisomaji cha skrini huchakata.

**Kumbuka:** Kwa picha zilizo na maelezo ya picha, bado unaweza kutaka kujumuisha maandishi ya `alt` kwani husaidia injini za utafutaji kuorodhesha maudhui ya picha.

# --instructions--

Camper Cat ameandika msimbo wa ukurasa wa muundo kwa sehemu ya blogu ya tovuti yake. Anapanga kuongeza mapumziko ya kuona kati ya makala zake mbili kwa picha ya mapambo ya upanga wa samurai. Ongeza sifa ya kipengele ya `alt` kwenye lebo ya `img` na uiweke kuwa mfuatano wa herufi tupu. (Kumbuka kwamba picha ya `src` haunganishi na faili halisi - usijali kwamba hakuna mapanga yanayoonekana kwenye onyesho.)

# --hints--

Leboa ya `img` inapaswa kuwa na sifa ya kipengele ya `alt`.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Sifa ya kipengele ya `alt` inapaswa kuwekwa kuwa mfuatano wa herufi tupu.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
