---
id: 587d774e367417b2b2512aa0
title: Weka maudhui ndani ya kipengele cha article
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` ni mojawapo ya vipengele vipya vya HTML5 vinavyoongeza maana ya kimaelezo katika alama zako za uandishi. `article` ni kipengele cha kugawanya sehemu na hutumika kufunika maudhui huru, yenye kujitegemea. Lebo hii inafaa sana kwa maingizo ya blogu, machapisho ya jukwaa, au makala za habari.

Kuhakikisha kama maudhui yanaweza kusimama peke yake kawaida ni uamuzi wa mtumiaji, lakini unaweza kutumia vipimo rahisi viwili. Jiulize kama ukiondoa muktadha wote unaozunguka, je, maudhui hayo bado yangeeleweka? Vivyo hivyo, kwa maandishi, je, maudhui yangedumu ikiwa yangekuwa katika mtiririko wa RSS?

Kumbuka kuwa watu wanaotumia teknolojia za kusaidia hutegemea alama za uandishi zilizoandaliwa vizuri na zenye maana ya kimaelezo ili kuelewa kazi yako vyema zaidi.

**Kumbuka:** Kipengele cha `section` pia ni kipya katika HTML5, na kina maana ya kimaelezo tofauti kidogo na `article`. `article` ni kwa maudhui huru, na `section` ni kwa kuunganisha maudhui yanayohusiana kwa mandhari. Zinaweza kutumika ndani ya kila moja, kulingana na mahitaji. Kwa mfano, ikiwa kitabu ni `article`, basi kila sura ni `section`. Wakati hakuna uhusiano kati ya makundi ya maudhui, basi tumia `div`.

`<div>` - huunganisha maudhui
`<section>` - huunganisha maudhui yanayohusiana
`<article>` - huunganisha maudhui huru, yenye kujitegemea

# --instructions--

Camper Cat alitumia lebo za `article` kufunika machapisho kwenye ukurasa wake wa blogu, lakini alisahau kuzitumia kwenye chapisho la juu kabisa. Badilisha lebo ya `div` ili itumie lebo ya `article` badala yake.

# --hints--

Msimbo wako unapaswa kuwa na lebo tatu za `article`.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
