---
id: 587d7790367417b2b2512ab0
title: Tumia tabindex kuongeza makini ya kibodi kwa kipengele
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

Sifa ya kipengele ya HTML `tabindex` ina kazi tatu tofauti zinazohusiana na makini ya kibodi kwa kipengele. Inapokuwa kwenye lebo, inaonyesha kwamba kipengele kinaweza kupokea makini. Thamani yake (nambari kamili chanya, hasi, au sifuri) inaamua tabia yake.

Vipengele fulani, kama viungo na udhibiti wa fomu, hupokea makini ya kibodi moja kwa moja wakati mtumizi anapobofya tab kupitia ukurasa. Hii ni kwa mpangilio ule ule kama vipengele vinavyoonekana katika alama za uandishi za chanzo za HTML. Utendaji huu huo unaweza pia kupewa vipengele vingine, kama `div`, `span`, na `p`, kwa kuweka sifa ya kipengele ya `tabindex="0"` juu yao. Hapa kuna mfano:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**Kumbuka:** Thamani hasi ya `tabindex` (kawaida -1) inaonyesha kwamba kipengele kinaweza kupokea makini, lakini hakiwezi kufikiwa kwa kibodi. Njia hii kwa kawaida hutumika kuleta makini kwa maudhui kwa njia ya programu (kama wakati `div` inayotumika kwa dirisha la pop-up inapoanzishwa), na ni zaidi ya wigo wa changamoto hizi.

# --instructions--

Camper Cat ameunda utafiti mpya kukusanya taarifa kuhusu watumizi wake. Anajua sehemu za ingizo hupokea makini ya kibodi moja kwa moja, lakini anataka kuhakikisha watumizi wake wa kibodi wanakaa makini kwenye maelekezo wakati wanapobofya tab kupitia vitu. Ongeza sifa ya kipengele ya `tabindex` kwenye lebo ya `p` na weka thamani yake kuwa `0`. Ziada - kutumia `tabindex` pia huruhusu darasa la bandia la CSS `:focus` kufanya kazi kwenye lebo ya `p`.

# --hints--

Msimbo wako unapaswa kuongeza sifa ya kipengele ya `tabindex` kwenye lebo ya `p` inayoshikilia maelekezo ya fomu.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

Msimbo wako unapaswa kuweka sifa ya kipengele ya `tabindex` kwenye lebo ya `p` kwa thamani ya 0.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
