---
id: 587d7788367417b2b2512aa2
title: Fanya urambazaji wa kisomaji cha skrini kuwa rahisi zaidi kwa lebo ya nav
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVwWSv'
forumTopicId: 301024
dashedName: make-screen-reader-navigation-easier-with-the-nav-landmark
---

# --description--

Kipengele cha `nav` ni kipengele kingine cha HTML5 chenye kipengele cha alama za uandishi kilichoingizwa kwa ajili ya urambazaji rahisi wa kisomaji cha skrini. Lebo hii inakusudiwa kuzunguka viungo vikuu vya urambazaji kwenye ukurasa wako.

Kama kuna viungo vya kurudia vya tovuti chini ya ukurasa, si lazima kuvitumia lebo ya `nav` pia. Kutumia `footer` (itayajadiliwa katika changamoto inayofuata) ni vya kutosha.

# --instructions--

Camper Cat aliingiza viungo vya urambazaji juu ya ukurasa wake wa mafunzo, lakini aliwazunguka kwa `div`. Badilisha `div` kuwa lebo ya `nav` ili kuboresha ufikikaji kwenye ukurasa wake.

# --hints--

Msimbo wako unapaswa kuwa na lebo moja ya `nav`.

```js
assert.lengthOf(document.querySelectorAll('nav') , 1);
```

Lebos zako za `nav` zinapaswa kuzunguka `ul` na vitu vyake vya orodha.

```js
const nav = document.querySelector('nav');
const children = nav?.querySelectorAll(`:scope ${'ul'}`);
assert.lengthOf(children,1);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Kipengele chako cha `nav` kinapaswa kuwa na lebo ya kufunga.

```js
assert.isTrue(code.match(/<\/nav>/g)?.length === code.match(/<nav>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <div>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </div>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
