---
id: 587d778b367417b2b2512aa8
title: Ongeza chaguo la tarehe linalopatikana kwa urahisi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

Fomu mara nyingi hujumuisha uwanja wa `input`, ambao unaweza kutumika kuunda udhibiti wa fomu mbalimbali. Sifa ya kipengele `type` kwenye kipengele hiki inaonyesha aina gani ya kipengele cha `input` kitaundwa.

Huenda umeona aina za ingizo za `text` na `submit` katika changamoto zilizopita, na HTML5 ilizindua chaguo la kubainisha uwanja wa `date`. Kulingana na msaada wa kivinjari, chaguo la tarehe linaonekana kwenye uwanja wa `input` linapokuwa na makini, jambo ambalo hufanya kujaza fomu kuwa rahisi kwa watumizi wote.

Kwa vivinjari vya zamani, aina itakuwa `text` kwa chaguo-msingi, hivyo ni msaada kuonyesha watumizi muundo wa tarehe unaotarajiwa katika maandishi ya `label` au `placeholder` kwa tahadhari.

Hapa kuna mfano:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat anaandaa mashindano ya Mortal Kombat na anataka kumuuliza mshindani wake kuona ni tarehe gani inafaa zaidi. Ongeza lebo ya `input` yenye sifa ya kipengele `type` yenye thamani ya `date`, sifa ya kipengele `id` yenye thamani ya `pickdate`, na sifa ya kipengele `name` yenye thamani ya `date`.

# --hints--

Msimbo wako unapaswa kuongeza lebo moja ya `input` kwa uwanja wa kuchagua tarehe.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

Leboa yako ya `input` inapaswa kuwa na sifa ya kipengele `type` yenye thamani ya `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

Leboa yako ya `input` inapaswa kuwa na sifa ya kipengele `id` yenye thamani ya `pickdate`.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

Leboa yako ya `input` inapaswa kuwa na sifa ya kipengele `name` yenye thamani ya `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
