---
id: 587d778a367417b2b2512aa5
title: Boresha ufikikaji wa chati kwa kutumia kipengele cha figure
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5 ilizindua kipengele cha `figure` na `figcaption` kinachohusiana. Vinapotumika pamoja, vitu hivi vinazunguka uwakilishi wa picha (kama picha, mchoro, au chati) pamoja na maelezo yake ya picha. Kuzunguka vipengele hivi pamoja kunatoa ongezeko la upatikanaji mara mbili kwa kuunganisha kwa maana maudhui yanayohusiana na kutoa mbadala wa maandishi unaoelezea `figure`.

Kwa uwakilishi wa data kama chati, maelezo ya picha yanaweza kutumika kuonyesha kwa ufupi mwelekeo au hitimisho kwa watumizi wenye udhaifu wa kuona. Changamoto nyingine inahusu jinsi ya kuhamisha toleo la jedwali la data ya chati nje ya skrini (kutumia CSS) kwa watumizi wa kisomaji cha skrini.

Hapa kuna mfano - kumbuka kwamba `figcaption` huingia ndani ya lebo za `figure` na inaweza kuunganishwa na vipengele vingine:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat anafanya kazi kwa bidii kuunda chati ya mstari iliyopangwa ikionyesha muda wa wiki kwa mafunzo ya kuingia kimya kimya, mapigano, na silaha. Msaidie kuunda muundo bora wa ukurasa wake kwa kubadilisha lebo ya `div` aliyotumia kuwa lebo ya `figure`, na lebo ya `p` inayozunguka maelezo ya picha kuwa lebo ya `figcaption`.

# --hints--

Msimbo wako unapaswa kuwa na lebo moja ya `figure`.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

Msimbo wako unapaswa kuwa na lebo moja ya `figcaption`.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `p`.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

`figcaption` inapaswa kuwa mtoto wa lebo ya `figure`.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

Kipengele chako cha `figure` kinapaswa kuwa na lebo ya kufunga.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
