---
id: 587d778a367417b2b2512aa6
title: Boresha ufikikaji wa sehemu za fomu kwa kipengele cha lebo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

Kuboresha ufikikaji kwa kutumia alama za uandishi za semantic HTML kunahusisha kutumia majina ya lebo na sifa zinazofaa. Changamoto kadhaa zinazofuata zinashughulikia baadhi ya hali muhimu za kutumia sifa katika fomu.

Lebo ya `label` inazunguka maandishi ya kipengele maalum cha udhibiti wa fomu, kawaida ni jina au lebo ya chaguo. Hii huunganisha maana na kipengele na kufanya fomu iwe rahisi kusomeka. Sifa ya `for` kwenye lebo ya `label` inaunganisha wazi `label` na udhibiti wa fomu na hutumika na visomaji vya skrini.

Umejifunza kuhusu vitufe vya radio na lebo zao katika funzo la sehemu ya HTML Msingi. Katika funzo hilo, tulizunguka kipengele cha ingizo cha kitufe cha radio ndani ya kipengele cha `label` pamoja na maandishi ya lebo ili kufanya maandishi hayo yaweze kubofywa. Njia nyingine ya kufanikisha hili ni kwa kutumia sifa ya `for`, kama ilivyoelezwa katika funzo hili.

Thamani ya sifa ya `for` lazima iwe sawa na thamani ya sifa ya `id` ya udhibiti wa fomu. Hapa kuna mfano:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Cat anatarajia kupata hamu kubwa kwa machapisho yake ya blogu yenye mawazo na anataka kujumuisha fomu ya usajili wa barua pepe. Ongeza sifa ya `for` kwenye `label` ya barua pepe inayolingana na `id` kwenye sehemu yake ya `input`.

# --hints--

Msimbo wako unapaswa kuwa na sifa ya `for` kwenye lebo ya `label` ambayo si tupu.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

Thamani ya sifa yako ya `for` inapaswa kulingana na thamani ya `id` kwenye `input` ya barua pepe.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
