---
id: 587d7788367417b2b2512aa3
title: Fanya urambazaji wa kisomaji cha skrini kuwa rahisi zaidi kwa lebo ya kijachini
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVrDh8'
forumTopicId: 301022
dashedName: make-screen-reader-navigation-easier-with-the-footer-landmark
---

# --description--

Kama ilivyo kwa `header` na `nav`, kipengele cha `footer` kina kipengele kilichojengwa ndani cha lebo ya mahali ambacho huruhusu vifaa vya kusaidia kurambaza haraka hadi hapo. Kimsingi hutumika kuhifadhi taarifa za hakimiliki au viungo vya nyaraka zinazohusiana ambazo kawaida huwa chini ya ukurasa.

# --instructions--

Ukurasa wa mafunzo wa Camper Cat unaendelea vizuri. Badilisha `div` aliyetumia kufunika taarifa zake za hakimiliki chini ya ukurasa kuwa kipengele cha `footer`.

# --hints--

Msimbo wako unapaswa kuwa na lebo moja ya `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer') ,1);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Msimbo wako unapaswa kuwa na lebo ya ufunguzi na kufunga ya `footer`.

```js
assert.match(code,/<footer>\s*&copy; 2018 Camper Cat\s*<\/footer>/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>


  <div>&copy; 2018 Camper Cat</div>


</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>


  <footer>&copy; 2018 Camper Cat</footer>


</body>
```
