---
id: 587d778f367417b2b2512aac
title: Epuka matatizo ya upofu wa rangi kwa kutumia uwiano wa utofauti wa rangi unaotosha
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

Rangi ni sehemu kubwa ya muundo wa kuona, lakini matumizi yake huleta matatizo mawili ya ufikikaji. Kwanza, rangi pekee haipaswi kutumika kama njia pekee ya kuwasilisha taarifa muhimu kwa sababu watumiaji wa kisomaji cha skrini hawawezi kuiona. Pili, rangi za mbele na nyuma zinahitaji uwiano wa utofauti wa rangi unaotosha ili watumiaji wenye upofu wa rangi waweze kuzitofautisha.

Changamoto zilizopita zilijumuisha kuwa na mbadala za maandishi kushughulikia tatizo la kwanza. Changamoto ya mwisho ilianzisha zana za kukagua uwiano wa utofauti kusaidia tatizo la pili. Uwiano wa utofauti wa rangi unaopendekezwa na WCAG wa 4.5:1 unatumika kwa matumizi ya rangi pamoja na mchanganyiko wa rangi za kijivu.

Watumiaji wenye upofu wa rangi huwa na shida kutofautisha baadhi ya rangi na nyingine - kawaida kwa mwelekeo wa rangi lakini wakati mwingine pia kwa uangavu. Unaweza kukumbuka uwiano wa utofauti unahesabiwa kwa kutumia thamani za mwangaza wa jamaa (au uangavu) wa rangi za mbele na nyuma.

Katika matumizi halisi, uwiano wa 4.5:1 unaweza kufikiwa kwa kupiga kivuli (kuongeza rangi nyeusi) kwenye rangi yenye giza zaidi na kupiga rangi nyepesi (kuongeza rangi nyeupe) kwenye rangi yenye uangavu zaidi. Vivuli vya giza kwenye gurudumu la rangi huchukuliwa kuwa vivuli vya buluu, zambarau, magenta, na nyekundu, wakati rangi nyepesi zilizopigwa rangi ni machungwa, manjano, kijani, na buluu-kijani.

# --instructions--

Camper Cat anajaribu kutumia rangi kwa maandishi na nyuma ya blogu yake, lakini mchanganyiko wake wa sasa wa kijani kidogo `background-color` na maandishi ya maroon `color` una uwiano wa utofauti wa 2.5:1. Unaweza kubadilisha kwa urahisi uangavu wa rangi kwa kuwa alitangaza rangi hizo kwa kutumia sifa ya CSS `hsl()` (inayomaanisha mwelekeo wa rangi, unyevunyevu, uangavu) kwa kubadilisha hoja ya tatu. Ongeza thamani ya uangavu ya `background-color` kutoka 35% hadi 55%, na punguza thamani ya uangavu ya `color` kutoka 20% hadi 15%. Hii inaboresha uwiano wa utofauti hadi 5.9:1.

# --hints--

Msimbo wako unapaswa kubadilisha thamani ya uangavu kwa sifa ya maandishi `color` tu hadi thamani ya 15%.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

Msimbo wako unapaswa kubadilisha thamani ya uangavu kwa sifa ya `background-color` tu hadi thamani ya 55%.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
