---
id: 587d778b367417b2b2512aa7
title: Wazungushe vitufe vya radio katika kipengele cha seti ya sehemu kwa ufikikaji bora
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVefw'
forumTopicId: 301030
dashedName: wrap-radio-buttons-in-a-fieldset-element-for-better-accessibility
---

# --description--

Mada inayofuata ya fomu inahusu ufikikaji wa vitufe vya radio. Kila chaguo hupewa `label` yenye sifa ya `for` inayounganisha na `id` ya kipengele kinacholingana kama ilivyojadiliwa katika changamoto iliyopita. Kwa kuwa vitufe vya radio mara nyingi huja katika kundi ambapo mtumizi lazima achague kimoja, kuna njia ya kielezo kuonyesha kuwa chaguzi ni sehemu ya seti.

Lebo ya `fieldset` huzunguka kundi lote la vitufe vya radio kufanikisha hili. Mara nyingi hutumia lebo ya `legend` kutoa maelezo ya kundi hilo, ambayo visomaji vya skrini husoma kwa kila chaguo katika kipengele cha `fieldset`.

Kifungashio cha `fieldset` na lebo ya `legend` si lazima wakati chaguzi zinajieleza zenyewe, kama vile uchaguzi wa jinsia. Kutumia `label` yenye sifa ya `for` kwa kila kitufe cha radio ni vya kutosha.

Hapa kuna mfano:

```html
<form>
  <fieldset>
    <legend>Choose one of these three items:</legend>
    <input id="one" type="radio" name="items" value="one">
    <label for="one">Choice One</label><br>
    <input id="two" type="radio" name="items" value="two">
    <label for="two">Choice Two</label><br>
    <input id="three" type="radio" name="items" value="three">
    <label for="three">Choice Three</label>
  </fieldset>
</form>
```

# --instructions--

Camper Cat anataka taarifa kuhusu kiwango cha ninja cha watumizi wake wanapojisajili kwenye orodha yake ya barua pepe. Ameongeza seti ya vitufe vya radio na amejifunza kutoka funzo letu la mwisho kutumia lebo za `label` zenye sifa za `for` kwa kila chaguo. Hongera Camper Cat! Hata hivyo, msimbo wake bado unahitaji msaada. Badilisha lebo ya `div` inayozunguka vitufe vya radio kuwa lebo ya `fieldset`, na badilisha lebo ya `p` ndani yake kuwa lebo ya `legend`.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya `fieldset` inayozunguka seti ya vitufe vya radio.

```js
assert.lengthOf(document.querySelectorAll('fieldset'),1);
```

Kipengele cha `fieldset` kinapaswa kuwa na lebo ya kufunga.

```js
assert.isTrue(code.match(/<\/fieldset>/g)?.length === code.match(/<fieldset>/g)?.length);
```

Msimbo wako unapaswa kuwa na lebo ya `legend` inayozunguka maandishi yanayouliza kiwango cha ninja cha mtumizi.

```js
assert.lengthOf(document.querySelectorAll('legend') ,1);
```

Msimbo wako haupaswi kuwa na lebo za `div`.

```js
assert.lengthOf(document.querySelectorAll('div') ,0);
```

Msimbo wako haupaswi tena kuwa na lebo ya `p` inayozunguka maandishi yanayouliza kiwango cha ninja cha mtumizi.

```js
assert.lengthOf(document.querySelectorAll('p') ,4);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <!-- Only change code below this line -->
      <div>
        <p>What level ninja are you?</p>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </div>
      <!-- Only change code above this line -->


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">

      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </fieldset>

      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
