---
id: 587d7787367417b2b2512aa1
title: Fanya urambazaji wa kisomaji cha skrini kuwa rahisi zaidi kwa kichwa landmark
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

Kipengele kingine cha HTML5 kinachoongeza maana ya kimaelezo na kuboresha Ufikikaji ni lebo ya `header`. Inatumika kufunika taarifa za utangulizi au viungo vya urambazaji kwa lebo yake mama na hufanya kazi vizuri kuzunguka maudhui yanayojirudia juu katika kurasa nyingi.

`header` inashiriki kipengele cha landmark kilichoingizwa ulichokiona na `main`, kuruhusu teknolojia za kusaidia kurambaza haraka kwa maudhui hayo.

**Kumbuka:** `header` inakusudiwa kutumika katika lebo ya `body` ya hati yako ya HTML. Ni tofauti na kipengele cha `head`, ambacho kina kichwa cha ukurasa, taarifa za meta, n.k.

# --instructions--

Camper Cat anaandika makala nzuri kuhusu mafunzo ya ninja, na anataka kuongeza ukurasa kwa ajili yao kwenye tovuti yake. Badilisha lebo ya `div` ya juu ambayo kwa sasa ina `h1` kuwa lebo ya `header` badala yake.

# --hints--

Msimbo wako unapaswa kuwa na lebo moja ya `header`.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

Lebos zako za `header` zinapaswa kufunika `h1`.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

Msimbo wako haupaswi kuwa na lebo yoyote ya `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Kipengele chako cha `header` kinapaswa kuwa na lebo ya kufunga.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
