---
id: 587d778f367417b2b2512aae
title: Toa maana kwa viungo kwa kutumia maandishi ya kiungo yenye maelezo
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

Watumiaji wa kisomaji cha skrini wana chaguzi mbalimbali kuhusu aina ya maudhui ambayo kifaa chao husoma. Chaguzi hizi ni pamoja na kuruka hadi (au kupita) vipengele vya nanga, kuruka hadi maudhui makuu, au kupata muhtasari wa ukurasa kutoka kwa vichwa vya sehemu. Chaguo jingine ni kusikia tu viungo vilivyopo kwenye ukurasa.

Visomaji vya skrini hufanya hivi kwa kusoma maandishi ya kiungo, au kile kilicho kati ya tagi za nanga (`a`). Kuwa na orodha ya viungo vya "bofya hapa" au "soma zaidi" hakuwezi kusaidia. Badala yake, tumia maandishi mafupi lakini yenye maelezo ndani ya tagi za `a` kutoa maana zaidi kwa watumiaji hawa.

# --instructions--

Maandishi ya kiungo ambayo Camper Cat anatumia hayana maelezo ya kutosha bila muktadha wa karibu. Hamisha tagi za nanga (`a`) ili zizunguke maandishi ya `information about batteries` badala ya `Click here`.

# --hints--

Msimbo wako unapaswa kuhamisha tagi za nanga `a` kutoka kuzunguka maneno ya `Click here` hadi kuzunguka maneno ya `information about batteries`.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

Kipengele cha `a` kinapaswa kuwa na sifa ya `href` yenye thamani ya mfuatano wa herufi tupu `""`.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

Kipengele cha `a` kinapaswa kuwa na lebo ya kufunga.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
