---
id: 587d7790367417b2b2512ab1
title: Tumia tabindex Kuweka Mfululizo wa Makini ya Kibodi kwa Vipengele Kadhaa
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

Sifa ya `tabindex` pia inaweka mfululizo halisi wa tab kwa vipengele. Hii hufanikishwa wakati thamani ya sifa hiyo imewekwa kuwa nambari chanya ya 1 au zaidi.

Kuweka `tabindex="1"` kutaleta makini ya kibodi kwa kipengele hicho kwanza. Kisha hupitia mfululizo wa thamani za `tabindex` zilizobainishwa (2, 3, n.k.), kabla ya kuhamia kwa vitu vya default na `tabindex="0"`.

Ni muhimu kutambua kuwa wakati mfululizo wa tab umewekwa kwa njia hii, unaghairi mfululizo wa default (ambao hutumia chanzo cha HTML). Hii inaweza kuwachanganya watumizi wanaotarajia kuanza urambazaji kutoka juu ya ukurasa. Mbinu hii inaweza kuwa muhimu katika hali fulani, lakini kwa upande wa Ufikikaji, kuwa makini kabla ya kuitumia.

Hapa kuna mfano:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat ana uwanja wa utafutaji kwenye ukurasa wake wa Nukuu za Kuhamasisha ambao anapanga kuuweka kona ya juu kulia kwa kutumia CSS. Anataka udhibiti wa fomu wa utafutaji `input` na wasilisha `input` kuwa vitu viwili vya kwanza katika mfululizo wa tab. Ongeza sifa ya `tabindex` iliyowekwa kuwa `1` kwa lebo ya `search` `input`, na sifa ya `tabindex` iliyowekwa kuwa `2` kwa lebo ya `submit` `input`.

Jambo lingine la kuzingatia ni kwamba baadhi ya vivinjari vinaweza kukuweka katikati ya mfululizo wako wa tab wakati kipengele kinapobofyanwa. Kipengele kimeongezwa kwenye ukurasa kinachohakikisha utaanza kila mara mwanzoni mwa mfululizo wako wa tab.

# --hints--

Msimbo wako unapaswa kuongeza sifa ya `tabindex` kwa lebo ya `search` `input`.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

Msimbo wako unapaswa kuongeza sifa ya `tabindex` kwa lebo ya `submit` `input`.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

Msimbo wako unapaswa kuweka sifa ya `tabindex` kwenye lebo ya `search` `input` kwa thamani ya 1.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

Msimbo wako unapaswa kuweka sifa ya `tabindex` kwenye lebo ya `submit` `input` kwa thamani ya 2.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
