---
id: 587d774c367417b2b2512a9c
title: Ongeza mbadala wa maandishi kwa picha kwa ajili ya ufikikaji wa watu wenye udhaifu wa kuona
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Huenda umewahi kuona sifa ya `alt` kwenye lebo ya `img` katika changamoto nyingine. Maandishi ya `alt` huelezea yaliyomo kwenye picha na hutoa mbadala wa maandishi kwa picha hiyo. Sifa ya `alt` husaidia katika hali ambapo picha haijapakia au mtumizi hawezi kuiona. Vifaa vya utafutaji pia hutumia sifa hii kuelewa picha ina nini ili kuijumuisha katika matokeo ya utafutaji. Hapa kuna mfano:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Watu wenye udhaifu wa kuona hutegemea visomaji vya skrini kubadilisha yaliyomo kwenye mtandao kuwa interface ya sauti. Hawatapokea taarifa ikiwa itaonyeshwa kwa njia ya kuona tu. Kwa picha, visomaji vya skrini vinaweza kupata sifa ya `alt` na kusoma yaliyomo ili kutoa taarifa muhimu.

Maandishi mazuri ya `alt` hutoa msomaji maelezo mafupi ya picha. Unapaswa kila mara kujumuisha sifa ya `alt` kwenye picha zako. Kulingana na maelezo ya HTML5, hii sasa inachukuliwa kuwa lazima.

# --instructions--

Camper Cat ni mtaalamu wa programu na pia ni ninja halisi, ambaye anajenga tovuti ya kushiriki maarifa yake. Picha ya wasifu anayotaka kutumia inaonyesha ujuzi wake na inapaswa kuthaminiwa na wageni wote wa tovuti. Ongeza sifa ya `alt` kwenye lebo ya `img` inayofafanua kuwa Camper Cat anafanya karate. (Picha ya `src` haina uunganisho wa faili halisi, hivyo unapaswa kuona maandishi ya `alt` yanavyoonyeshwa.)

# --hints--

Lebo yako ya `img` inapaswa kuwa na sifa ya `alt` na haipaswi kuwa tupu.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
