---
id: 587d7790367417b2b2512aaf
title: Fanya Viungo Viweze Kupatikana kwa Kutumia Funguo za Upatikanaji za HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML hutoa sifa ya `accesskey` kuweka funguo za njia za mkato ili kuamsha au kuleta makini kwa kipengele. Kuongeza sifa ya `accesskey` kunaweza kufanya urambazaji kuwa rahisi zaidi kwa watumiaji wanaotumia kibodi pekee.

HTML5 inaruhusu sifa hii kutumika kwa kipengele chochote, lakini ni muhimu hasa inapoitwa kwa vipengele vinavyoshirikiana. Hii ni pamoja na viungo, vifungo, na udhibiti wa fomu.

Hapa kuna mfano:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat anataka viungo vinavyozunguka vichwa viwili vya makala za blogu viwe na funguo za mkato za kibodi ili watumiaji wa tovuti yake waweze kuhamia haraka kwenye hadithi kamili. Ongeza sifa ya `accesskey` kwa viungo vyote viwili na weka ile ya kwanza kuwa `g` (kwa Garfield) na ile ya pili kuwa `c` (kwa Chuck Norris).

# --hints--

Msimbo wako unapaswa kuongeza sifa ya `accesskey` kwenye lebo ya `a` yenye thamani ya `id` ya `first`.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

Msimbo wako unapaswa kuongeza sifa ya `accesskey` kwenye lebo ya `a` yenye thamani ya `id` ya `second`.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

Msimbo wako unapaswa kuweka sifa ya `accesskey` kwenye lebo ya `a` yenye thamani ya `id` ya `first` kuwa `g`. Kumbuka kuwa herufi kubwa na ndogo ni muhimu.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

Msimbo wako unapaswa kuweka sifa ya `accesskey` kwenye lebo ya `a` yenye thamani ya `id` ya `second` kuwa `c`. Kumbuka kuwa herufi kubwa na ndogo ni muhimu.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
