---
id: 587d778e367417b2b2512aab
title: Boresha Usomeka kwa Matini yenye Utofauti Mkubwa wa Rangi
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

Utofauti mdogo kati ya rangi za mbele na nyuma unaweza kufanya maandishi kuwa magumu kusoma. Utofauti wa kutosha huboresha usomeka wa maudhui yako, lakini "kutosha" hasa inamaanisha nini?

Miongozo ya Ufikikaji wa Maudhui ya Mtandao (WCAG) inapendekeza angalau uwiano wa utofauti wa rangi wa 4.5 kwa 1 kwa maandishi ya kawaida. Uwiano huu huhesabiwa kwa kulinganisha thamani za mwangaza wa rangi mbili. Hii huanzia 1:1 kwa rangi ile ile, au kutokuwa na utofauti, hadi 21:1 kwa rangi nyeupe dhidi ya nyeusi, ambayo ni utofauti mkubwa zaidi. Kuna zana nyingi za kuangalia uwiano wa utofauti wa rangi mtandaoni zinazohesabu uwiano huu kwa niaba yako.

# --instructions--

Chaguo la Camper Cat la maandishi ya kijivu hafifu kwenye nyuma ya rangi nyeupe kwa chapisho lake la hivi karibuni la blogu lina uwiano wa utofauti wa 1.5:1, na kufanya kusoma kuwa vigumu. Badilisha `color` ya maandishi kutoka kijivu kilichopo sasa (`#D3D3D3`) kwenda kijivu kizito (`#636363`) ili kuboresha uwiano wa utofauti wa rangi hadi 6:1.

# --hints--

Msimbo wako unapaswa kubadilisha `color` ya maandishi kwa `body` kuwa kijivu kizito.

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

Msimbo wako haupaswi kubadilisha `background-color` ya `body`.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
