---
id: 60f030d388cb74067cf291c3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

제목 아래에 단락 요소 안에 다음 텍스트를 사용하여 사용자가 등록하도록 권장하세요:

```md
Please fill out this form with the required information
```

# --hints--

`p` 안에 `body` 요소를 추가하세요.

```js
assert.exists(document.querySelector('body > p'));
```

`p` 아래에 `h1` 요소를 추가하세요.

```js
assert.exists(document.querySelector('h1 + p'));
```

`p` 요소에 `Please fill out this form with the required information` 텍스트를 지정하세요.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
--fcc-editable-region--
<body>
  <h1>Registration Form</h1>
  
</body>
--fcc-editable-region--
</html>
```
