---
id: 60f02e7361b68405e27b62a5
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 회원가입 양식을 디자인하여 양식을 스타일링하는 방법을 배우세요.

모든 HTML 보일러플레이트(`DOCTYPE`, `html`, `head`, `body`)가 제공되어 있습니다.

`body` 안에 `h1` 요소를 사용해 `Registration Form` 텍스트가 들어간 제목을 추가하여 양식의 맥락을 제공하세요.

# --hints--

`h1`를 `body` 안에 추가하세요.

```js
assert.exists(document.querySelector('body > h1'));
```

`h1`에 `Registration Form` 텍스트를 넣으세요.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
