---
id: 60f1a5e2d2c23707a4f9a660
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

이제 일부 브라우저가 기본으로 추가하는 `body` 기본 `margin` 를 `0` 로 설정해서 가로 스크롤바를 없애세요.

# --hints--

`margin` 요소 선택자 안에 `body` 를 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

`margin` 에 `0` 값을 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
