---
id: 60f1a9cbd23023082e149fee
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

더 좋아졌습니다. 이제 `background-color`의 `body`를 `#1b1b32`로 바꿔 배경을 눈에 편안하게 만드세요. 그런 다음, 텍스트가 보이도록 `color`를 `#f5f6f7`로 바꾸세요.

# --hints--

`background-color` 요소 선택자 안에 `body`를 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

`background-color`에 `#1b1b32` 값을 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

`color` 요소 선택자 안에 `body`를 추가하세요.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

`color`에 `#f5f6f7` 값을 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
