---
id: 60f1922fcbd2410527b3bd89
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이전 수업에서 뷰포트 단위를 다루는 방법을 배웠습니다.

`vh` 단위는 뷰포트 높이를 나타내며, 뷰포트의 `height` 1%와 같습니다. 따라서 뷰포트 높이에 상대적입니다.

이제 CSS로 프로젝트를 꾸밀 차례입니다. `body`에 `width`를 `100%`로, 그리고 `height`를 `100vh`로 지정하는 것부터 시작하세요.

# --hints--

`body` 요소 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body`에 `width`를 `100%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

`body`에 `height`를 `100vh`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
