---
id: 60f5c3e399ff1a05629964e4
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

제목에서 알 수 있듯이, 폼을 만들고 있습니다. 그래서 `p` 요소 다음에 `form`를 대상으로 하는 `action` 속성을 가진 `https://register-demo.freecodecamp.org`을 삽입하세요.

# --hints--

`form` 요소 옆에 `p` 요소를 추가하세요.

```js
assert.exists(document.querySelector('p + form'));
```

`form`에 `action` 속성을 부여하세요.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

`action`에 `https://register-demo.freecodecamp.org` 값을 부여하세요.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

`form` 요소는 닫는 태그 `</form>`를 가져야 합니다.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
