---
id: 60f5d2776c854e069560fbe6
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

첫 번째 `fieldset`는 이름, 이메일, 비밀번호 필드를 포함합니다. 첫 번째 `label`에 네 개의 `fieldset` 요소를 추가하는 것부터 시작하세요.

# --hints--

총 세 개의 `fieldset` 요소가 있어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

`fieldset` 요소들은 모두 `form` 요소의 직계 자식이어야 합니다.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

네 개의 `label` 요소는 모두 첫 번째 `fieldset` 요소의 직계 자식이어야 합니다. `label` 요소를 닫는 것을 잊지 마세요.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

마지막 두 개의 `fieldset` 요소는 비어 있어야 합니다.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

`label` 요소를 닫는 것을 꼭 확인하세요.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
