---
id: 60f5cb8875ab6a0610f05071
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

폼에 세 개의 구분된 섹션이 있으므로 `fieldset` 요소 안에 세 개의 `form` 요소를 추가하세요.

# --hints--

세 개의 `fieldset` 요소를 추가해야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

세 개의 `fieldset` 요소는 `form` 안에 있어야 합니다.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

세 개의 `fieldset` 요소는 형제 요소여야 합니다.

```js
assert.exists(document.querySelector('fieldset + fieldset + fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
    
--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
