---
id: 60f5dc35c07ac1078f140916
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

다음 텍스트를 `label` 요소에 추가하세요:

- `Enter Your First Name:`
- `Enter Your Last Name:`
- `Enter Your Email:`
- `Create a New Password:`

# --hints--

첫 번째 `label`에는 텍스트 `Enter Your First Name:`가 있어야 합니다.

```js
assert.match(document.querySelector('label')?.innerHTML, /Enter Your First Name:/i);
```

두 번째 `label`에는 텍스트 `Enter Your Last Name:`가 있어야 합니다.

```js
assert.match(document.querySelector('fieldset > label:nth-child(2)')?.innerHTML, /Enter Your Last Name:/i);
```

세 번째 `label`에는 텍스트 `Enter Your Email:`가 있어야 합니다.

```js
assert.match(document.querySelector('fieldset > label:nth-child(3)')?.innerHTML, /Enter Your Email:/i);
```

네 번째 `label`에는 텍스트 `Create a New Password:`가 있어야 합니다.

```js
assert.match(document.querySelector('fieldset > label:nth-child(4)')?.innerHTML, /Create a New Password:/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label></label>
      <label></label>
      <label></label>
      <label></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
