---
id: 60f84ec41116b209c280ba91
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

`type="password"`를 사용하면 `pattern` 속성으로 비밀번호가 유효하다고 간주되려면 일치해야 하는 정규 표현식을 정의할 수 있습니다.

비밀번호 `pattern` 요소에 `input` 속성을 추가하여 입력이 다음과 일치하도록 요구하세요: `[a-z0-5]{8,}`

위는 소문자 8자 이상 또는 숫자 `0`부터 `5`까지와 일치하는 정규 표현식입니다. 그런 다음 `minlength` 속성을 제거하고 테스트해 보세요.

# --hints--

비밀번호 `input` 요소에 `pattern` 속성을 추가해야 합니다.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

`pattern` 속성에 `[a-z0-5]{8,}` 값을 지정해야 합니다.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

비밀번호 `minlength` 요소에서 `input` 속성을 제거해야 합니다.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
