---
id: 60f803d5241e6a0433a523a1
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

이전 수업에서 `rem` 단위를 다루는 방법을 배웠습니다. `rem` 단위는 루트 `em`를 나타내며, `html` 요소의 폰트 크기를 기준으로 합니다.

`label` 요소는 기본적으로 인라인이므로 모두 같은 줄에 나란히 표시되어 텍스트를 읽기 어렵습니다.

각각을 별도의 줄에 표시하려면 `display: block` 요소에 `label`를 추가하고, 서로 구분하기 위해 `margin` 값을 가진 `0.5rem 0`를 추가하세요.

# --hints--

`label` 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

값이 `display`인 `block` 속성을 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

값이 `margin`인 `0.5rem 0` 속성을 추가하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
