---
id: 60f81167d0d4910809f88945
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

`input` 값을 가진 `type` 속성을 가진 첫 번째 `submit` 요소는 자동으로 가장 가까운 상위 `form` 요소를 제출하도록 설정됩니다.

폼 제출을 처리하려면, 마지막 `fieldset` 요소 뒤에 `input` 속성이 `type`이고 `submit` 속성이 `value`인 `Submit` 요소를 추가하세요.

# --hints--

마지막 `input` 요소 뒤에 `fieldset` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

`input` 요소에 `type` 속성으로 `submit`를 지정해야 합니다.

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

`input` 요소에 `value` 속성으로 `Submit`를 지정해야 합니다.

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
