---
id: 60facde2d0dc61085b41063f
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

`textarea` 요소는 `input` 유형의 `text` 요소처럼 작동하지만, 여러 줄 텍스트를 받을 수 있고 초기 텍스트 행과 열 수를 지정할 수 있다는 추가 이점이 있습니다.

사용자가 자기소개와 함께 등록할 수 있게 하세요. `label` 끝에 텍스트가 `Provide a bio:`인 `fieldset`을 추가하세요. `textarea` 요소 안에 `label` 요소를 추가하세요. `textarea`은 닫는 태그가 필요하다는 점을 유의하세요.

# --hints--

기존 `label` 요소들 뒤, 세 번째 `fieldset` 끝에 `label` 요소를 추가하세요.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)'));
```

`label`에 `Provide a bio:` 텍스트를 지정하세요.

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)')?.innerText, /Provide a bio/);
```

`textarea` 안에 `label` 요소를 중첩하세요.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4) > textarea'));
```

`textarea` 요소에 여는 태그와 닫는 태그를 지정하세요.

```js
assert.match(code, /<textarea\s*>[\s\S]*<\/textarea\s*>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
      <label>Input your age (years): <input type="number" min="13" max="120" /></label>
      <label>How did you hear about us?
        <select>
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
