---
id: 60fab9f17fa294054b74228c
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

마지막 `fieldset`으로 넘어갑니다. 사용자가 프로필 사진을 업로드할 수 있게 하려면 어떻게 해야 할까요?

`input` 유형 `file`가 바로 그 기능을 제공합니다. 텍스트 `label`가 있는 `Upload a profile picture: `를 추가하고, 파일 업로드를 허용하는 `input`를 중첩하세요.

# --hints--

텍스트 `label`가 있는 `Upload a profile picture: `를 추가해야 합니다.

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label')?.innerText, /Upload a profile picture:/i);
```

`input` 요소 안에 `label` 요소를 중첩해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label > input'));
```

`input` 요소에 `type`의 `file`를 지정해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label > input')?.type, 'file');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
