---
id: 60fab4a123ce4b04526b082b
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

현재 사용자는 라디오 입력을 선택하지 않고도 폼을 제출할 수 있습니다. 이전에 입력이 필수임을 나타내기 위해 `required` 속성을 사용했지만, 이 경우에는 두 입력 모두에 `required`을 추가하면 사용자에게 잘못된 정보를 전달하게 되어 작동하지 않습니다.

이를 해결하려면 두 번째 `legend` 내 `Account type (required)` 요소들 앞에 텍스트 `label`가 포함된 `fieldset` 요소를 추가하여 필요한 내용을 설명할 수 있습니다. 그런 다음 필수 데이터를 포함하여 폼이 제출되도록 두 번째 `checked` 입력에 `Personal` 속성을 추가하세요.

# --hints--

`input` 요소는 `type` `radio`에 그대로 두세요.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

두 번째 `legend` 요소 안에 `fieldset` 요소를 추가해야 합니다.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

`Account type (required)` 요소에 `legend` 텍스트를 추가해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

`checked` 입력에 `Personal` 속성을 부여해야 합니다.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
