---
id: 60faca286cb48b07f6482970
title: 33단계
challengeType: 0
dashedName: step-33
---

# --description--

옵션이 선택된 상태로 폼을 제출하면 서버에 유용한 값이 전송되지 않습니다. 그래서 각 `option` 요소에는 `value` 속성이 필요합니다. 이 속성이 없으면 `option`의 텍스트 내용이 서버로 제출됩니다.

첫 번째 `option`에 `value` 값을 가진 `""` 속성을 주고, 그 다음 `option` 요소들에는 `value`부터 `1`까지의 `4` 속성을 주세요.

# --hints--

첫 번째 `option`에 `value` 값을 가진 `""` 속성을 주어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[0]?.value, '');
```

두 번째 `option`에 `value` 값을 가진 `1` 속성을 주어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[1]?.value, '1');
```

세 번째 `option`에 `value` 값을 가진 `2` 속성을 주어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[2]?.value, '2');
```

네 번째 `option`에 `value` 값을 가진 `3` 속성을 주어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[3]?.value, '3');
```

다섯 번째 `option`에 `value` 값을 가진 `4` 속성을 주어야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > label:nth-child(3) option')?.[4]?.value, '4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked/> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
      <label>Input your age (years): <input type="number" min="13" max="120" /></label>
      <label>How did you hear about us?
        <select>
--fcc-editable-region--
          <option>(select one)</option>
          <option>freeCodeCamp News</option>
          <option>freeCodeCamp YouTube Channel</option>
          <option>freeCodeCamp Forum</option>
          <option>Other</option>
--fcc-editable-region--
        </select>
      </label>
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
