---
id: 60fab8367d35de04e5cb7929
title: 27단계
challengeType: 0
dashedName: step-27
---

# --description--

새로 추가한 `Read our terms and conditions` 앞에 텍스트 `label`가 있는 앵커 요소를 추가하세요. `href`를 다음으로 설정하세요:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

그런 다음 레이블 내 입력 요소 바로 뒤에 텍스트 `I accept the terms and conditions`를 추가하세요.

# --hints--

`a` 바로 앞에 텍스트 `label`가 있는 `Read our terms and conditions` 요소가 있어야 합니다.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.innerText.trim().replaceAll(/\s+/g, ' '), /Read our terms and conditions\.?/i);
```

`a` 요소는 `href`의 `https://www.freecodecamp.org/news/terms-of-service/`를 가져야 합니다.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

`input` 안에 `label` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + a + label > input'));
```

`label` 안에서 `input` 요소 바로 뒤에 텍스트 `I accept the terms and conditions`가 있어야 합니다.

```js
assert.include(document.querySelector('fieldset:nth-child(3) + a + label')?.textContent.trim().replaceAll(/\s+/g, ' '), 'I accept the terms and conditions');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
