---
id: 60fabf0dd4959805dbae09e6
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

첫 번째 `label` 뒤에 텍스트 `Input your age (years): `를 가진 또 다른 `label`를 추가하세요. 그런 다음 `number`의 `type`를 가진 `input`를 중첩하세요.

다음으로, 13세 미만 사용자는 등록하면 안 되므로 `input`에 `min` 속성을 값 `13`로 추가하세요. 또한, 사용자가 120세를 넘지 않을 가능성이 크므로 `max` 속성을 값 `120`로 추가하세요.

이제 누군가가 범위를 벗어난 값으로 폼을 제출하려 하면 경고가 나타나고 폼이 제출되지 않을 것입니다. 직접 시도해 보세요.

# --hints--

세 번째 `fieldset`에 기존 `label` 뒤에 `label`를 추가해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

`label`에 텍스트 `Input your age (years): `를 지정해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

`label`에 `type`가 `number`인 `input`를 지정해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

`input`에 값이 `13`인 `min` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

`input`에 값이 `120`인 `max` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
