---
id: 62ff8b9dab5ac88e4d3d43a3
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

접근성 모범 사례를 따라 `input` 요소와 `label` 요소를 `for` 속성으로 연결하세요.

각 `first-name` 속성의 값으로 `last-name`, `email`, `new-password`, `id`를 사용하세요.

# --hints--

첫 번째 `input` 요소는 `id` 값을 가진 `first-name` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

두 번째 `input` 요소는 `id` 값을 가진 `last-name` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

세 번째 `input` 요소는 `id` 값을 가진 `email` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

네 번째 `input` 요소는 `id` 값을 가진 `new-password` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

첫 번째 `label` 요소는 `for` 값을 가진 `first-name` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

두 번째 `label` 요소는 `for` 값을 가진 `last-name` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

세 번째 `label` 요소는 `for` 값을 가진 `email` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

네 번째 `label` 요소는 `for` 값을 가진 `new-password` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name: <input /></label>
      <label>Enter Your Last Name: <input /></label>
      <label>Enter Your Email: <input /></label>
      <label>Create a New Password: <input /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
