---
id: 62ff8e998d3e7eae14d6ae3b
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

사용자가 약관을 읽었는지 확인해야 합니다.

`label` 요소를 추가하세요. 새로 생성한 `label` 요소 안에 `input` 요소를 추가하고 `type` 속성을 `checkbox`로 설정하세요. 이 `input` 요소를 `required`로 만들어 사용자가 약관에 동의하지 않으면 가입할 수 없게 하세요.

접근성을 위해 요소들에 `id`와 `for` 속성을 값 `terms-and-conditions`로 추가하세요.

# --hints--

세 번째 `label` 요소 뒤에 `fieldset`을 추가해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

`input` 요소에 `label`를 추가해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

`type` 요소에 값이 `checkbox`인 `input` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

`required` 요소에 `input` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 요소는 `id`인 `terms-and-conditions`을 가져야 합니다.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 요소는 값이 `for`인 `terms-and-conditions` 속성을 가져야 합니다.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
