---
id: 63541ef4f96cd82e8e6c788a
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`method` 속성은 `action` 속성에 지정된 URL로 폼 데이터를 보내는 방식을 지정합니다. 폼 데이터는 `GET` 요청을 통해 URL 매개변수로 (`method="get"`와 함께) 보내거나 `POST` 요청을 통해 요청 본문에 데이터로 (`method="post"`와 함께) 보낼 수 있습니다.

`method` 속성을 설정하여 폼 데이터를 `POST` 요청으로 보내세요.

**참고**: `GET`와 `POST` 메서드는 앞으로의 수업과 워크숍에서 더 자세히 배우게 됩니다.

# --hints--

새로운 `form` 요소를 추가하지 마세요.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

`form` 요소에 `method` 속성이 있어야 합니다.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

`method` 속성은 `post`로 설정해야 합니다.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
