---
id: 6880995521770716269832cf
title: 23단계
challengeType: 0
dashedName: step-23
---

# --description--

라디오 버튼에 아직 값이 할당되지 않았습니다. 따라서 폼을 통해 보내는 POST 요청은 어떤 옵션이 선택되었는지 등록하지 못합니다.

요청을 완전히 견고하게 만들려면, 두 라디오 버튼 모두에 `value` 속성을 각각 `personal`과 `business` 값으로 추가하세요.

# --hints--

첫 번째 라디오 입력에 `value` 속성을 `personal` 값으로 할당하세요.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.value, 'personal');
```

두 번째 라디오 입력에 `value` 속성을 `business` 값으로 할당하세요.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.value, 'business');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      <label><input type="radio" name="account-type" /> Personal</label>
      <label><input type="radio" name="account-type" /> Business</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
