---
id: 669e2f60e83c011754f711f9
title: 建置一個旅行社頁面
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `DOCTYPE` 宣告。
1. 你應該有一個 `html` 元素，其 `lang` 設定為 `en`。
1. 你應該有一個 `head` 元素，其中包含一個 `charset` 設定為 `utf-8` 的 `meta` 空元素，以及一個文字為 `Travel Agency Page` 的 `title`。
1. 你的 `head` 元素中應該有一個 `meta` 標籤，其中包含你網站的簡短描述，以利 SEO。
1. 你應該有一個 `h1` 元素來呈現你的旅遊目的地。它應該是你頁面上唯一的 `h1` 元素。
1. 你應該在 `h1` 元素下方有一個介紹旅遊機會的段落。
1. 你應該有一個文字為 `Packages` 的 `h2` 元素。
1. 你應該有一個簡要介紹各種套裝行程的 `p` 元素。
1. 你應該有一個包含兩個列表項目的無序列表元素。這兩個列表項目的文字應分別為 `Group Travels` 和 `Private Tours`。每個列表項目的文字都應該由一個錨點元素包覆。
1. 你應該有一個文字為 `Top Itineraries` 的 `h2` 元素。
1. 你應該至少有三個 `figure` 元素，每個都包含一個錨點元素，以及一個描述該圖片內容的 `figcaption` 元素。
1. 這三個錨點元素的內容應該是一個 `img` 元素，具有適當的 `alt` 屬性，且其 `src` 屬性設定為有效的圖片。如果你願意，可以使用 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` 和 `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`。
1. 你的五個錨點元素都應該有一個值為 `https://www.freecodecamp.org/learn` 的 `href` 屬性，以及一個值為 `_blank` 的 `target` 屬性。

# --hints--

你的旅行社頁面應該有一個 `!DOCTYPE html` 宣告。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 元素，其 `lang` 設定為 `en`。

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

你的 `html` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/html\s*>/i);
```

你應該在 `html` 元素內有一個 `head` 元素。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

你的 `head` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/head\s*>/i);
```

你應該在 `head` 元素中有一個 `title` 元素。

```js
assert.isNotNull(document.querySelector('head title'));
```

你的 `title` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/title\s*>/i);
```

你應該在你的 `meta` 元素中有兩個 `head` 元素。

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

一個 `meta` 元素應該有一個 `name` 屬性，值為 `description`，且有一個非空的 `content` 屬性。

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

其中一個 `meta` 元素的 `charset` 屬性應設為 `utf-8`。

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的 `title` 元素應該包含你的旅遊代理名稱。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `html` 元素中應該有一個 `body` 元素。

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

你的 `body` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/body\s*>/i);
```

你應該有一個 `h1` 元素來呈現你的旅遊目的地。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你應該只擁有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你應該在 `h1` 元素下方放置一個 `p` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

你的第一個 `p` 元素應該介紹旅遊機會。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一個 `h2` 元素應該有文字 `Packages`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

你應該在第一個 `h2` 元素下方放置一個 `p` 元素。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

你的第二個 `p` 元素應該簡要介紹各種套裝行程。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

你應該在第二個 `p` 元素下方放置一個無序列表元素。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

你的無序列表中應該有兩個條款。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

你的兩個列表條款都應該包含一個錨點元素。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

你第一個列表條款的錨點元素應該包覆文字 `Group Travels`。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

你第二個列表條款的錨點元素應該包覆文字 `Private Tours`。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

你應該在無序列表之後有一個 `h2` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

你的第二個 `h2` 元素應該有文字 `Top Itineraries`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

你應該至少有三個 `figure` 元素。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

每個 `figure` 元素應該包含一個錨點元素，作為其第一個巢狀元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

每個 `figure` 元素應該包含一個 `figcaption` 元素，作為其第二個巢狀元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

每個 `figcaption` 應該包含一些文字。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

你 `figure` 元素中巢狀的每個 `a` 元素都應該包含一個影像。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

每個 `img` 元素應該有一個有效的 `src` 屬性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

每個 `img` 元素應該有一個適當的 `alt` 屬性值。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

每個 `a` 元素應該有一個 `href` 屬性，值為 `https://www.freecodecamp.org/learn`。別忘了列表條款中的連結。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

每個 `a` 元素應該有一個 `target` 屬性，值為 `_blank`。別忘了列表條款中的連結。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
