---
id: 669e2f60e83c011754f711f9
title: 構建一個旅行社頁面
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `DOCTYPE` 聲明。
1. 你應該有一個 `lang` 設置爲 `en` 的 `html` 元素。
1. 你應該有一個 `head` 元素，其中包含一個 `meta` 空元素，其 `charset` 設置爲 `utf-8`，`title` 的文本爲 `Travel Agency Page`。
1. 在你的 `head` 元素中，你應該有一個 `meta` 標籤，包含你網站用於 SEO 的簡短描述。
1. 你應該有一個 `h1` 元素來顯示你的旅行目的地。
1. 在 `h1` 元素下面，你應該有一個介紹旅行機會的段落。
1. 你應該有一個 `h2` 元素，文本爲 `Packages` 。
1. 你應該有一個`p`元素，簡要介紹各種旅遊方式。
1. 你應該有一個包含兩個列表項的無序列表元素。 這兩個列表項應該分別具有文本 `Group Travels` 和 `Private Tours`。 每個列表項的文本都應該被錨元素包圍。
1. 你應該有一個 `h2` 元素，文本爲 `Top Itineraries` 。
1. 你應該至少有三個 `figure` 元素，每個元素包含一個錨元素和一個 `figcaption` 元素。
1. 三個錨元素應該分別包含一個帶有適當描述的 `alt` 屬性和設置爲有效的圖像的 `src` 屬性的 `img` 元素。 如果願意，可以使用 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` 和 `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`。
1. 所有五個錨元素都應該有一個值爲 `https://www.freecodecamp.org/learn` 的 `href` 屬性和一個值爲 `_blank` 的 `target` 屬性。

# --hints--

你的旅行社頁面該有一個 `!DOCTYPE html` 聲明。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `lang` 設置爲 `en` 的 `html` 元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

你的 `html` 元素中應該有一個 `head` 元素。

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你的 `head` 元素中應該有一兩個個 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

一個 `meta` 元素應該有一個值爲 `description` 的 `name` 屬性和一個非空的 `content` 屬性。

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

一個 `meta` 元素應該將它的 `charset` 屬性設置爲 `UTF-8`。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你的 `head` 元素中應該有一個 `title` 元素。

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

你的 `title` 元素應該包含你的旅行社名稱。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `html` 元素中應該有一個 `body` 元素。

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你應該有一個 `h1` 元素來顯示你的旅行目的地。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你應該只有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素下面應該有一個 `p` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

你的第一個 `p` 元素應該介紹旅行機會。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一個 `h2` 元素應該包含文本 `Packages`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Packages');
```

你的第一個 `h2` 元素下面應該有一個 `p` 元素。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

你的第二個 `p` 元素應該簡要介紹各種旅遊方式。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

你的第二個 `p` 元素下面應該有一個無序列表元素。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

你的無序列表中應該有兩個項目。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

你的兩個列表項都應該包含一個錨元素。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

第一個列表項的錨元素應該包裹文本 `Group Travels`。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText, 'Group Travels');
```

第二個列表項的錨元素應該包裹文本 `Private Tours`。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText, 'Private Tours');
```

你的無序列表後面應該有一個 `h2` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

你的第二個 `h2` 元素應該包含文本 `Top Itineraries`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Top Itineraries');
```

你應該至少有三個 `figure` 元素。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

每個 `figure` 元素應包含一個錨元素作爲其第一個子元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

每個 `figure` 元素應包含一個 `figcaption` 元素作爲其第一個子元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

每個 `figcaption` 應該包含一些文本。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

作爲 `figure` 元素子元素的每個 `a` 元素都應該包含一個圖像。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

每個 `img` 元素應該有一個有效的 `src` 屬性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

每個 `img` 元素都應該有一個具有適當值的 `alt` 屬性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Each `a` element should have an `href` attribute with the value of `https://www.freecodecamp.org/learn`. 不要忘記列表項中的鏈接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

每個 `a` 元素應該有一個 `target` 屬性，其值爲 `_blank`。 不要忘記列表項中的鏈接。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
