---
id: 669e2f60e83c011754f711f9
title: 建置一個旅行社頁面
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `DOCTYPE` 宣告。
1. 你應該有一個 `html` 元素，其 `lang` 設定為 `en`。
1. 你應該有一個 `head` 元素，包含一個 `meta` 空元素，其 `charset` 設為 `utf-8`，以及一個 `title`，其文字為 `Travel Agency Page`。
1. 你應該在你的 `head` 元素中有一個包含網站簡短描述以利 SEO 的 `meta` 頁籤。
1. 你應該有一個 `h1` 元素來呈現你的旅遊目的地。
1. 你應該在 `h1` 元素下方有一段介紹旅遊機會的段落。
1. 你應該有一個 `h2` 元素，其 `Packages` 文字。
1. 你應該有一個 `p` 元素簡要介紹各種軟體包。
1. 你應該有一個無序列表元素，裡面有兩個列表項目。這兩個列表項目的文字分別應該是 `Group Travels` 和 `Private Tours`。每個列表項目的文字都應該被一個錨點元素包住。
1. 你應該有一個 `h2` 元素，裡面有文字 `Top Itineraries`。
1. 你應該至少有三個 `figure` 元素，每個元素都包含一個錨點元素和一個 `figcaption` 元素。
1. 這三個錨點元素應該包含一個具有適當 `alt` 屬性和 `src` 屬性設為有效影像的 `img` 元素作為其內容。如果你願意，可以使用 `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`、`https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` 和 `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`。
1. 你所有的五個錨點元素都應該有一個 `href` 屬性，其值為 `https://www.freecodecamp.org/learn`，以及一個 `target` 屬性，其值為 `_blank`。

# --hints--

你的旅行社頁面應該有一個 `!DOCTYPE html` 宣告。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 元素，其 `lang` 設定為 `en`。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

你應該在 `html` 元素內有一個 `head` 元素。

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你應該在你的 `head` 元素中有兩個 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

一個 `meta` 元素應該有一個 `name` 屬性，值為 `description`，且有一個非空的 `content` 屬性。

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

一個 `meta` 元素應該將其 `charset` 屬性設為 `UTF-8`。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你應該在你的 `head` 元素中包含 `title` 元素。

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

你的 `title` 元素應該包含你的旅行社名稱。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `html` 元素中應該有一個 `body` 元素。

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

你應該有一個 `h1` 元素來呈現你的旅遊目的地。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

你應該只擁有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你應該在 `h1` 元素下方有一個 `p` 元素。

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

你的第一個 `p` 元素應該介紹旅遊機會。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一個 `h2` 元素應該有 `Packages` 的文字。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

你應該在你的第一個 `h2` 元素下方有一個 `p` 元素。

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

你的第二個 `p` 元素應該簡要介紹各種軟體包。

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

你應該在第二個 `p` 元素下方有一個無序列表元素。

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

你的無序列表中應該有兩個項目。

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

你的列表項目都應包含一個錨點元素。

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

你第一個列表項目的錨點元素應該包裹 `Group Travels` 文字。

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

你第二個列表項目的錨點元素應該包裹 `Private Tours` 文字。

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

你的無序列表後面應該有一個 `h2` 元素。

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

你的第二個 `h2` 元素應該有文字 `Top Itineraries`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

你應該至少有三個 `figure` 元素。

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

每個 `figure` 元素應該包含一個錨點元素作為其第一個子元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

每個 `figure` 元素應該包含一個作為第二個子元素的 `figcaption` 元素。

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

每個 `figcaption` 都應該包含一些文字。

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

你的 `figure` 元素的每個子元素 `a` 都應該包含一個影像。

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

每個 `img` 元素都應該有一個有效的 `src` 屬性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

每個 `img` 元素都應該有一個帶有適當值的 `alt` 屬性。

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

每個 `a` 元素都應該有一個 `href` 屬性，其值為 `https://www.freecodecamp.org/learn`。不要忘記列表項目中的聯結。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

每個 `a` 元素都應該有一個值為 `_blank` 的 `target` 屬性。不要忘記列表項目中的連結。

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
