---
id: 587d78a7367417b2b2512ae0
title: "usa l'animazione css per cambiare lo stato hover di un pulsante"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4vZAa'
forumTopicId: 301073
dashedName: use-css-animation-to-change-the-hover-state-of-a-button
---

# --description--

Puoi usare l'animazione CSS `@keyframes` per cambiare il colore di un pulsante nel suo stato hover.

Ecco un esempio di come cambiare la larghezza di un'immagine al passaggio del mouse:

```html
<style>
  img {
    width: 30px;
  }
  img:hover {
    animation-name: width;
    animation-duration: 500ms;
  }

  @keyframes width {
    100% {
      width: 40px;
    }
  }
</style>

<img src="https://cdn.freecodecamp.org/curriculum/applied-visual-design/google-logo.png" alt="Google's Logo" />
```

# --instructions--

Nota che `ms` indica i millisecondi, dove 1000ms corrispondono a 1s.

Usa l'animazione CSS `@keyframes` per cambiare il `background-color` dell'elemento `button` in modo che diventi `#4791d0` quando un utente ci passa sopra con il mouse. La regola `@keyframes` dovrebbe avere solo una voce per `100%`.

# --hints--

La regola @keyframes dovrebbe usare il colore di sfondo `animation-name`.

```js
assert.match(code, /@keyframes\s+?background-color\s*?{/g);
```

Dovrebbe esserci una regola sotto `@keyframes` che cambia il `background-color` in `#4791d0` al 100%.

```js
assert.match(code, /100%\s*?{\s*?background-color:\s*?#4791d0;\s*?}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }


</style>

<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }

  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
  }

  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
