---
id: 587d781e367417b2b2512acb
title: Bloccare un elemento al suo genitore con il posizionamento assoluto
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJ7c3'
forumTopicId: 301060
dashedName: lock-an-element-to-its-parent-with-absolute-positioning
---

# --description--

L'opzione successiva per la proprietĂ  CSS `position` Ă¨ `absolute`, che blocca l'elemento in posizione rispetto al suo contenitore genitore. A differenza della posizione `relative`, questa rimuove l'elemento dal flusso normale del documento, quindi gli elementi circostanti lo ignorano. Le proprietĂ  di offset CSS (top o bottom e left o right) si usano per regolare la posizione.

Una sfumatura del posizionamento assoluto Ă¨ che sarĂ  bloccato rispetto al suo antenato *posizionato* piĂą vicino. Se dimentichi di aggiungere una regola di posizione all'elemento genitore, (di solito si fa usando `position: relative;`), il browser continuerĂ  a cercare lungo la catena e alla fine userĂ  come predefinito il tag `body`.

# --instructions--

Blocca l'elemento `#searchbar` in alto a destra rispetto al suo genitore `section` dichiarando la sua `position` come `absolute`. Dagli offset `top` e `right` di 50 pixel ciascuno.

# --hints--

L'elemento `#searchbar` dovrebbe avere una `position` impostata su `absolute`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.position, 'absolute');
```

Il tuo codice dovrebbe usare l'offset CSS `top` di 50 pixel sull'elemento `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.top, '50px');
```

Il tuo codice dovrebbe usare l'offset CSS `right` di 50 pixel sull'elemento `#searchbar`.

```js
const searchbarElement = document.querySelector('#searchbar');
const searchbarStyle = window.getComputedStyle(searchbarElement);
assert.equal(searchbarStyle?.right, '50px');
```

# --seed--

## --seed-contents--

```html
<style>
  #searchbar {



  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```

# --solutions--

```html
<style>
  #searchbar {
    position: absolute;
    top: 50px;
    right: 50px;
  }
  section {
    position: relative;
  }
</style>
<body>
  <h1>Welcome!</h1>
  <section>
    <form id="searchbar">
      <label for="search">Search:</label>
      <input type="search" id="search" name="search">
      <input type="submit" name="submit" value="Go!">
    </form>
  </section>
</body>
```
