---
id: 587d78a7367417b2b2512ae2
title: creare una direzione visiva facendo sfumare un elemento da sinistra a destra
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJqqAE'
forumTopicId: 301054
dashedName: create-visual-direction-by-fading-an-element-from-left-to-right
---

# --description--

Per questa sfida, modificherai il `opacity` di un elemento animato in modo che sfumi gradualmente mentre raggiunge il lato destro dello schermo.

Nell'animazione mostrata, l'elemento rotondo con lo sfondo sfumato si sposta verso destra fino al 50% dell'animazione secondo la regola `@keyframes`.

# --instructions--

Seleziona l'elemento con l'id `ball` e aggiungi la proprietà `opacity` impostata a 0.1 a `50%`, così l'elemento sfuma mentre si sposta verso destra.

# --hints--

La regola `keyframes` per la sfumatura dovrebbe impostare la proprietà `opacity` a 0.1 al 50%.

```js
assert.match(code,/@keyframes fade\s*?{\s*?50%\s*?{\s*?(?:left:\s*?60%;\s*?opacity:\s*?0?\.1;|opacity:\s*?0?\.1;\s*?left:\s*?60%;)/gi
);
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;

    }
  }

</style>

<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 70px;
    height: 70px;
    margin: 50px auto;
    position: fixed;
    left: 20%;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: fade;
    animation-duration: 3s;
  }

  @keyframes fade {
    50% {
      left: 60%;
      opacity: 0.1;
    }
  }
</style>
<div id="ball"></div>
```
