---
id: 587d781e367417b2b2512aca
title: Spostare un elemento posizionato relativamente con offset CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bQEA4'
forumTopicId: 301065
dashedName: move-a-relatively-positioned-element-with-css-offsets
---

# --description--

Gli offset CSS di `top` o `bottom` e `left` o `right` indicano al browser di quanto spostare un elemento rispetto a dove si troverebbe nel flusso normale del documento. Stai spostando un elemento lontano da una posizione data, il che muove l'elemento lontano dal lato di riferimento (effettivamente, nella direzione opposta). Come hai visto nella sfida precedente, usare l'offset `top` ha spostato `h2` verso il basso. Allo stesso modo, usare un offset `left` sposta un elemento verso destra.

# --instructions--

Usa gli offset CSS per spostare `h2` di 15 pixel verso destra e 10 pixel verso l'alto.

# --hints--

Il tuo codice dovrebbe usare un offset CSS per posizionare relativamente `h2` 10px verso l'alto. In altre parole, spostalo di 10px lontano dal `bottom` rispetto a dove si trova normalmente.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.bottom, '10px');
```

Il tuo codice dovrebbe usare un offset CSS per posizionare relativamente `h2` 15px verso destra. In altre parole, spostalo di 15px lontano dal `left` rispetto a dove si trova normalmente.

```js
const h2Element =document.querySelector('h2')
const h2Style = window.getComputedStyle(h2Element);
assert.equal(h2Style?.left , '15px');
```

# --seed--

## --seed-contents--

```html
<head>
<style>
  h2 {
    position: relative;


  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```

# --solutions--

```html
<head>
<style>
  h2 {
    position: relative;
    left: 15px;
    bottom: 10px;
  }
</style>
</head>
<body>
  <h1>On Being Well-Positioned</h1>
  <h2>Move me!</h2>
  <p>I still think the h2 is where it normally sits.</p>
</body>
```
