---
id: 587d78a8367417b2b2512ae3
title: animare elementi continuamente usando un conteggio di animazione infinito
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDVfq'
forumTopicId: 301041
dashedName: animate-elements-continually-using-an-infinite-animation-count
---

# --description--

Le sfide precedenti hanno trattato come usare alcune delle proprietà di animazione e la regola `@keyframes`. Un'altra proprietà di animazione è `animation-iteration-count`, che ti permette di controllare quante volte vuoi iterare l'animazione. Ecco un esempio:

```css
animation-iteration-count: 3;
```

In questo caso l'animazione si fermerà dopo 3 esecuzioni, ma è possibile farla andare continuamente impostando quel valore su `infinite`.

# --instructions--

Per mantenere la palla che rimbalza a destra in un loop continuo, cambia la proprietà `animation-iteration-count` in `infinite`.

# --hints--

La proprietà `animation-iteration-count` dovrebbe avere un valore `infinite`.

```js
const ballElement = document.querySelector('#ball');
const ballStyle = window.getComputedStyle(ballElement); 
assert.equal(ballStyle?.animationIterationCount, 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>

  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: 3;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```

# --solutions--

```html
<style>
  #ball {
    width: 100px;
    height: 100px;
    margin: 50px auto;
    position: relative;
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    animation-name: bounce;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  @keyframes bounce{
    0% {
      top: 0px;
    }
    50% {
      top: 249px;
      width: 130px;
      height: 70px;
    }
    100% {
      top: 0px;
    }
  }
</style>
<div id="ball"></div>
```
