---
id: 587d78a5367417b2b2512ada
title: usa la proprietà css transform scale per ingrandire un elemento al passaggio del mouse
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLPJuM'
forumTopicId: 301077
dashedName: use-the-css-transform-scale-property-to-scale-an-element-on-hover
---

# --description--

La proprietà `transform` ha diverse funzioni che ti permettono di ingrandire, spostare, ruotare, inclinare ecc. i tuoi elementi. Quando la usi con pseudo-classi come `:hover` che indicano uno stato specifico di un elemento, la proprietà `transform` può facilmente aggiungere interattività ai tuoi elementi.

Ecco un esempio per ingrandire gli elementi paragrafo a 2,1 volte la loro dimensione originale quando un utente ci passa sopra con il mouse:

```css
p:hover {
  transform: scale(2.1);
}
```

**Nota:** Applicare una trasformazione a un elemento `div` influenzerà anche tutti gli elementi figli contenuti nel div.

# --instructions--

Aggiungi una regola CSS per lo stato `hover` del `div` e usa la proprietà `transform` per ingrandire l'elemento `div` a 1,1 volte la sua dimensione originale quando un utente ci passa sopra con il mouse.

# --hints--

La dimensione dell'elemento `div` dovrebbe ingrandirsi di 1,1 volte quando l'utente ci passa sopra con il mouse.

```js
assert.match(code, /div:hover\s*?{\s*?transform:\s*?scale\(1\.1\);/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }



</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
    background: linear-gradient(
      53deg,
      #ccfffc,
      #ffcccf
    );
  }
  div:hover {
    transform: scale(1.1);
  }
</style>
<div></div>
```
