---
id: 587d78a3367417b2b2512ad0
title: Centrare un elemento orizzontalmente usando la proprietà margin
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLJqU4'
forumTopicId: 301043
dashedName: center-an-element-horizontally-using-the-margin-property
---

# --description--

Un'altra tecnica di posizionamento è centrare orizzontalmente un elemento di blocco. Un modo per farlo è impostare la sua `margin` su un valore auto.

Questo metodo funziona anche per le immagini. Le immagini sono elementi in linea per impostazione predefinita, ma possono diventare elementi di blocco impostando la proprietà `display` su `block`.

# --instructions--

Centra il `div` sulla pagina aggiungendo una proprietà `margin` con valore `auto`.

# --hints--

La `div` dovrebbe avere una `margin` impostata su `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('div')?.margin ,'auto');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;

  }
</style>
<div></div>
```

# --solutions--

```html
<style>
  div {
    background-color: blue;
    height: 100px;
    width: 100px;
    margin: auto;
  }
</style>
<div></div>
```
