---
id: 587d78a4367417b2b2512ad4
title: Regola la tinta di un colore
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp38TZ'
forumTopicId: 301036
dashedName: adjust-the-hue-of-a-color
---

# --description--

I colori hanno diverse caratteristiche tra cui tonalità, saturazione e luminosità. CSS3 ha introdotto la funzione `hsl()` come modo alternativo per scegliere un colore dichiarando direttamente queste caratteristiche.

La **tonalità** è ciò che generalmente si intende per "colore". Se immagini uno spettro di colori che inizia con il rosso a sinistra, passa per il verde al centro e arriva al blu a destra, la tonalità indica la posizione di un colore lungo questa linea. In `hsl()`, la tonalità usa il concetto di ruota dei colori invece dello spettro, dove l'angolo del colore sul cerchio è dato come valore tra 0 e 360.

La **saturazione** è la quantità di grigio in un colore. Un colore completamente saturo non contiene grigio, mentre un colore poco saturo è quasi completamente grigio. Questo valore è espresso in percentuale, con 100% che indica saturazione completa.

La **luminosità** è la quantità di bianco o nero in un colore. Viene indicata una percentuale che va da 0% (nero) a 100% (bianco), dove 50% rappresenta il colore normale.

Ecco alcuni esempi di utilizzo di `hsl()` con colori completamente saturi e luminosità normale:

<table><thead><tr><th>Colore</th><th>HSL</th></tr></thead><tbody><tr><td>red</td><td>hsl(0, 100%, 50%)</td></tr><tr><td>yellow</td><td>hsl(60, 100%, 50%)</td></tr><tr><td>green</td><td>hsl(120, 100%, 50%)</td></tr><tr><td>cyan</td><td>hsl(180, 100%, 50%)</td></tr><tr><td>blue</td><td>hsl(240, 100%, 50%)</td></tr><tr><td>magenta</td><td>hsl(300, 100%, 50%)</td></tr></tbody></table>

# --instructions--

Cambia la `background-color` di ogni elemento `div` in base ai nomi delle classi (`green`, `cyan` o `blue`) usando `hsl()`. Tutti e tre devono avere saturazione completa e luminosità normale.

# --hints--

Il tuo codice dovrebbe usare la funzione `hsl()` per dichiarare il colore green.

```js
assert.match(code,/\.green\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Il tuo codice dovrebbe usare la funzione `hsl()` per dichiarare il colore cyan.

```js
assert.match(code,/\.cyan\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

Il tuo codice dovrebbe usare la funzione `hsl()` per dichiarare il colore blue.

```js
assert.match(code,/\.blue\s*?{\s*?background-color\s*:\s*?hsl/gi);
```

L'elemento `div` con classe `green` dovrebbe avere una `background-color` di green.

```js
const greenElement = document.querySelector(".green");
const greenStyle = window.getComputedStyle(greenElement); 
assert.equal(greenStyle?.backgroundColor, 'rgb(0, 255, 0)');
```

L'elemento `div` con classe `cyan` dovrebbe avere una `background-color` di cyan.

```js
const cyanElement = document.querySelector(".cyan");
const cyanStyle = window.getComputedStyle(cyanElement); 
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

L'elemento `div` con classe `blue` dovrebbe avere una `background-color` di blue.

```js
const blueElement = document.querySelector(".blue");
const blueStyle = window.getComputedStyle(blueElement); 
assert.equal(blueStyle?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .blue {
    background-color: #000000;
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>

<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .green {
    background-color: hsl(120, 100%, 50%);
  }

  .cyan {
    background-color:   hsl(180, 100%, 50%);
  }

  .blue {
    background-color: hsl(240, 100%, 50%);
  }

  div {
    display: inline-block;
    height: 100px;
    width: 100px;
  }
</style>
<div class="green"></div>
<div class="cyan"></div>
<div class="blue"></div>
```
