---
id: 587d78a6367417b2b2512adc
title: usa la proprietà css transform skewY per inclinare un elemento lungo l’asse y
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZ2uB'
forumTopicId: 301075
dashedName: use-the-css-transform-property-skewy-to-skew-an-element-along-the-y-axis
---

# --description--

Dato che la funzione `skewX()` inclina l’elemento selezionato lungo l’asse X di un certo grado, non sorprende che la proprietà `skewY()` inclini un elemento lungo l’asse Y (verticale).

# --instructions--

Inclina l’elemento con id `top` di -10 gradi lungo l’asse Y usando la proprietà `transform`.

# --hints--

L’elemento con id `top` dovrebbe essere inclinato di -10 gradi lungo il suo asse Y.

```js
assert.match(code, /#top\s*?{\s*?.*?\s*?transform:\s*?skewY\(-10deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;

  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin: 50px auto;
  }
  #top {
    background-color: red;
    transform: skewY(-10deg);
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
