---
id: 587d781d367417b2b2512ac8
title: Regola lo stato hover di un tag àncora
challengeType: 0
videoUrl: 'https://scrimba.com/c/cakRGcm'
forumTopicId: 301035
dashedName: adjust-the-hover-state-of-an-anchor-tag
---

# --description--

Questa sfida riguarda l’utilizzo delle pseudo-classi. Una pseudo-classe è una parola chiave che può essere aggiunta ai selettori per selezionare uno stato specifico dell’elemento.

Per esempio, lo stile di un tag àncora può essere modificato per il suo stato hover usando il selettore pseudo-classe `:hover`. Ecco il CSS per cambiare il `color` del tag àncora in rosso durante il suo stato hover:

```css
a:hover {
  color: red;
}
```

# --instructions--

L’editor del codice ha una regola CSS che definisce lo stile di tutti i tag `a` in nero. Aggiungi una regola in modo che quando l’utente passa il mouse sopra il tag `a`, il `color` diventi blu.

# --hints--

Il tag àncora `color` dovrebbe rimanere nero, aggiungi solo regole CSS per lo stato `:hover`.

```js
const anchorElement = document.querySelector("a"); 
const anchorStyle = window.getComputedStyle(anchorElement);
assert.equal(anchorStyle?.color, 'rgb(0, 0, 0)');
```

Il tag àncora dovrebbe avere un `color` blu al passaggio del mouse.

```js
assert.match(code,
    /a:hover\s*?{\s*?color:\s*?(blue|rgba\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?,\s*?1\s*?\)|#00F|rgb\(\s*?0\s*?,\s*?0\s*?,\s*?255\s*?\))\s*?;\s*?}/gi
  );
```

# --seed--

## --seed-contents--

```html
<style>
  a {
    color: #000;
  }



</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```

# --solutions--

```html
<style>
  a {
    color: #000;
  }
  a:hover {
    color: rgba(0,0,255,1);
  }
</style>
<a href="https://freecatphotoapp.com/" target="_blank">CatPhotoApp</a>
```
