---
id: 587d78a8367417b2b2512ae7
title: "Cambiare la durata dell'animazione con parole chiave"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cJKvwCM'
forumTopicId: 301045
dashedName: change-animation-timing-with-keywords
---

# --description--

Nelle animazioni CSS, la proprietà `animation-timing-function` controlla la velocità con cui un elemento animato cambia durante la durata dell'animazione. Se l'animazione è un'auto che si muove dal punto A al punto B in un dato tempo (il tuo `animation-duration`), la `animation-timing-function` indica come l'auto accelera e decelera durante il percorso.

Ci sono diverse parole chiave predefinite disponibili per opzioni comuni. Per esempio, il valore predefinito è `ease`, che parte lentamente, accelera a metà e poi rallenta di nuovo alla fine. Altre opzioni includono `ease-out`, che è veloce all'inizio e poi rallenta, `ease-in`, che è lenta all'inizio e poi accelera alla fine, o `linear`, che applica una velocità costante all'animazione per tutta la durata.

# --instructions--

Per gli elementi con id `ball1` e `ball2`, aggiungi una proprietà `animation-timing-function` a ciascuno, e imposta `#ball1` su `linear`, e `#ball2` su `ease-out`. Nota la differenza nel modo in cui gli elementi si muovono durante l'animazione ma finiscono insieme, dato che condividono la stessa `animation-duration` di 2 secondi.

# --hints--

Il valore della proprietà `animation-timing-function` per l'elemento con id `ball1` dovrebbe essere `linear`.

```js
const ballOne =document.querySelector('#ball1'); 
const ballOneStyle = window.getComputedStyle(ballOne); 

const ball1Animation = __helpers.removeWhiteSpace(ballOneStyle?.animationTimingFunction);
assert.isTrue(ball1Animation == 'linear' || ball1Animation == 'cubic-bezier(0,0,1,1)');
```

Il valore della proprietà `animation-timing-function` per l'elemento con id `ball2` dovrebbe essere `ease-out`.

```js
const ballTwo = document.querySelector('#ball2'); 
const ballTwoStyle = window.getComputedStyle(ballTwo); 

const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)');
```

# --seed--

## --seed-contents--

```html
<style>

  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;

  }
  #ball2 {
    left:56%;

  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left:27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left:56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
