---
id: 587d781e367417b2b2512acc
title: Bloccare un elemento alla finestra del browser con il posizionamento fisso
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDNUR'
forumTopicId: 301061
dashedName: lock-an-element-to-the-browser-window-with-fixed-positioning
---

# --description--

Lo schema di layout successivo che CSS offre è la posizione `fixed`, un tipo di posizionamento assoluto che blocca un elemento rispetto alla finestra del browser. Simile al posizionamento assoluto, si usa con le proprietà offset CSS e rimuove anche l'elemento dal flusso normale del documento. Gli altri elementi non "percepiscono" più dove si trova, il che potrebbe richiedere alcune regolazioni del layout altrove.

Una differenza chiave tra le posizioni `fixed` e `absolute` è che un elemento con posizione fissa non si sposta quando l'utente scorre la pagina.

# --instructions--

La barra di navigazione nel codice ha un id di `navbar`. Cambia la sua `position` in `fixed` e impostala a 0 pixel dal `top` e 0 pixel dal `left`. Dopo aver aggiunto il codice, scorri la finestra di anteprima per vedere come la navigazione rimane fissa.

# --hints--

L'elemento `#navbar` dovrebbe avere una `position` impostata su `fixed`.

```js
const navbarElement = document.querySelector('#navbar');
const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.position, 'fixed');
```

Il tuo codice dovrebbe usare l'offset CSS `top` di 0 pixel sull'elemento `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
assert.equal(navbarStyle?.top, '0px');
```

Il tuo codice dovrebbe usare l'offset CSS `left` di 0 pixel sull'elemento `#navbar`.

```js
const navbarElement = document.querySelector('#navbar');
 const navbarStyle = window.getComputedStyle(navbarElement);
 assert.equal(navbarStyle?.left, '0px');
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {



    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```

# --solutions--

```html
<style>
  body {
    min-height: 150vh;
  }
  #navbar {
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    background-color: #767676;
  }
  nav ul {
    margin: 0px;
    padding: 5px 0px 5px 30px;
  }
  nav li {
    display: inline;
    margin-right: 20px;
  }
  a {
    text-decoration: none;
  }
</style>
<body>
  <header>
    <h1>Welcome!</h1>
    <nav id="navbar">
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Contact</a></li>
      </ul>
    </nav>
  </header>
  <p>I shift up when the #navbar is fixed to the browser window.</p>
</body>
```
