---
id: 587d78a9367417b2b2512ae8
title: Impara come funzionano le curve di Bézier
challengeType: 0
videoUrl: 'https://scrimba.com/c/c9bDrs8'
forumTopicId: 301058
dashedName: learn-how-bezier-curves-work
---

# --description--

L'ultima sfida ha introdotto la proprietà `animation-timing-function` e alcune parole chiave che modificano la velocità di un'animazione durante la sua durata. Il CSS offre un'opzione alternativa alle parole chiave che consente un controllo ancora più preciso su come si svolge l'animazione, tramite l'uso delle curve di Bézier.

Nelle animazioni CSS, le curve di Bézier si usano con la funzione `cubic-bezier`. La forma della curva rappresenta come si svolge l'animazione. La curva si trova su un sistema di coordinate 1 per 1. L'asse X di questo sistema rappresenta la durata dell'animazione (pensalo come una scala temporale), mentre l'asse Y rappresenta la variazione nell'animazione.

La funzione `cubic-bezier` consiste di quattro punti principali che si trovano su questa griglia 1 per 1: `p0`, `p1`, `p2` e `p3`. `p0` e `p3` sono impostati per te: sono i punti iniziale e finale che si trovano sempre rispettivamente nell'origine (0, 0) e in (1, 1). Imposti i valori x e y per gli altri due punti, e la loro posizione nella griglia determina la forma della curva che l'animazione seguirà. Questo si fa in CSS dichiarando i valori x e y dei punti "àncora" `p1` e `p2` nella forma: `(x1, y1, x2, y2)`. Riassumendo, ecco un esempio di curva di Bézier nel codice CSS:

```css
animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
```

Nell'esempio sopra, i valori x e y sono equivalenti per ogni punto (x1 = 0.25 = y1 e x2 = 0.75 = y2), il che, se ricordi dalla geometria, dà come risultato una linea che si estende dall'origine al punto (1, 1). Questa animazione è una variazione lineare di un elemento durante la durata dell'animazione, ed è equivalente all'uso della parola chiave `linear`. In altre parole, cambia a velocità costante.

# --instructions--

Per l'elemento con id `ball1`, cambia il valore della proprietà `animation-timing-function` da `linear` al suo valore equivalente con la funzione `cubic-bezier`. Usa i valori dei punti dati nell'esempio sopra.

# --hints--

Il valore della proprietà `animation-timing-function` per l'elemento con id `ball1` dovrebbe essere la funzione `cubic-bezier` equivalente a lineare.

```js
const ballOne = document.querySelector('#ball1');
const ballOneStyle = window.getComputedStyle(ballOne); 
assert.equal(
  ballOneStyle?.animationTimingFunction,
    'cubic-bezier(0.25, 0.25, 0.75, 0.75)'
);
```

Il valore della proprietà `animation-timing-function` per l'elemento con id `ball2` non dovrebbe cambiare.

```js
const ballTwoElement = document.querySelector('#ball2');
const ballTwoStyle = window.getComputedStyle(ballTwoElement); 
const ball2Animation = __helpers.removeWhiteSpace(ballTwoStyle?.animationTimingFunction);
assert.isTrue(
  ball2Animation == 'ease-out' || ball2Animation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: linear;
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }

</style>

<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```

# --solutions--

```html
<style>

  .balls{
    border-radius: 50%;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #ball1 {
    left: 27%;
    animation-timing-function: cubic-bezier(0.25, 0.25, 0.75, 0.75);
  }
  #ball2 {
    left: 56%;
    animation-timing-function: ease-out;
  }

  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id="ball1"></div>
<div class="balls" id="ball2"></div>
```
