---
id: 587d78a3367417b2b2512acf
title: Cambia la posizione degli elementi sovrapposti con la proprietà z-index
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM94aHk'
forumTopicId: 301046
dashedName: change-the-position-of-overlapping-elements-with-the-z-index-property
---

# --description--

Quando gli elementi sono posizionati in modo da sovrapporsi (cioè usando `position: absolute | relative | fixed | sticky`), l'elemento che compare dopo nel markup HTML apparirà, per impostazione predefinita, sopra gli altri elementi. Tuttavia, la proprietà `z-index` può specificare l'ordine con cui gli elementi sono impilati uno sopra l'altro. Deve essere un numero intero (cioè un numero intero e non decimale), e valori più alti per la proprietà `z-index` di un elemento lo spostano più in alto nella pila rispetto a quelli con valori più bassi.

# --instructions--

Aggiungi una proprietà `z-index` all'elemento con il nome della classe `first` (il rettangolo rosso) e impostala su un valore di 2 in modo che copra l'altro elemento (rettangolo blu).

# --hints--

L'elemento con la classe `first` dovrebbe avere un valore `z-index` di 2.

```js
const firstElement = document.querySelector('.first');
const firstStyle = window.getComputedStyle(firstElement); 
assert.equal(firstStyle?.zIndex, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;

  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>

<div class="first"></div>
<div class="second"></div>
```

# --solutions--

```html
<style>
  div {
    width: 60%;
    height: 200px;
    margin-top: 20px;
  }

  .first {
    background-color: red;
    position: absolute;
    z-index: 2;
  }
  .second {
    background-color: blue;
    position: absolute;
    left: 40px;
    top: 50px;
    z-index: 1;
  }
</style>
<div class="first"></div>
<div class="second"></div>
```
