---
id: 587d78a4367417b2b2512ad2
title: Scopri i colori terziari
challengeType: 0
forumTopicId: 301057
dashedName: learn-about-tertiary-colors
---

# --description--

I monitor dei computer e gli schermi dei dispositivi creano colori diversi combinando quantità di luce rossa, verde e blu. Questo è noto come modello di colore additivo RGB nella teoria moderna del colore. Rosso (R), verde (G) e blu (B) sono chiamati colori primari. Mescolare due colori primari crea i colori secondari ciano (G + B), magenta (R + B) e giallo (R + G). Hai visto questi colori nella sfida Colori complementari. Questi colori secondari sono il complemento del colore primario non usato nella loro creazione e sono opposti a quel colore primario sulla ruota dei colori. Per esempio, il magenta è fatto con rosso e blu ed è il complemento del verde.

I colori terziari sono il risultato della combinazione di un colore primario con uno dei suoi colori secondari vicini. Per esempio, nel modello di colore RGB, rosso (primario) e giallo (secondario) creano l'arancione (terziario). Questo aggiunge sei colori in più a una semplice ruota dei colori per un totale di dodici.

Esistono vari metodi per selezionare colori diversi che risultano in una combinazione armoniosa nel design. Un esempio che può usare i colori terziari si chiama schema di colore complemento diviso. Questo schema parte da un colore base, poi lo abbina ai due colori adiacenti al suo complemento. I tre colori forniscono un forte contrasto visivo in un design, ma sono più sottili rispetto all'uso di due colori complementari.

Ecco tre colori creati usando lo schema complemento diviso:

<table><thead><tr><th>Colore</th><th>Codice esadecimale</th></tr></thead><thead></thead><tbody><tr><td>arancione</td><td>#FF7F00</td></tr><tr><td>ciano</td><td>#00FFFF</td></tr><tr><td>lampone</td><td>#FF007F</td></tr></tbody></table>

# --instructions--

Cambia la proprietà `background-color` delle classi `orange`, `cyan` e `raspberry` con i rispettivi colori. Assicurati di usare i codici esadecimali e non i nomi dei colori.

# --hints--

L'elemento `div` con classe `orange` dovrebbe avere una `background-color` arancione.

```js
const orangeElement = document.querySelector('.orange');
const orangeStyle = window.getComputedStyle(orangeElement);
assert.equal(orangeStyle?.backgroundColor, 'rgb(255, 127, 0)');
```

L'elemento `div` con classe `cyan` dovrebbe avere una `background-color` ciano.

```js
const cyanElement = document.querySelector('.cyan');
const cyanStyle = window.getComputedStyle(cyanElement);
assert.equal(cyanStyle?.backgroundColor, 'rgb(0, 255, 255)');
```

L'elemento `div` con classe `raspberry` dovrebbe avere una `background-color` lampone.

```js
const raspberryElement = document.querySelector('.raspberry');
const raspberryStyle = window.getComputedStyle(raspberryElement);
assert.equal(raspberryStyle?.backgroundColor, 'rgb(255, 0, 127)');
```

Tutti i valori `background-color` per le classi colore dovrebbero essere codici esadecimali e non nomi di colori.

```js
assert.notMatch(code, /background-color:\s(orange|cyan|raspberry)/);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #000000;
  }

  .cyan {
    background-color: #000000;
  }

  .raspberry {
    background-color: #000000;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>

<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```

# --solutions--

```html
<style>
  body {
    background-color: #FFFFFF;
  }

  .orange {
    background-color: #FF7F00;
  }

  .cyan {
    background-color: #00FFFF;
  }

  .raspberry {
    background-color: #FF007F;
  }

  div {
    height: 100px;
    width: 100px;
    margin-bottom: 5px;
  }
</style>
<div class="orange"></div>
<div class="cyan"></div>
<div class="raspberry"></div>
```
