---
id: 587d78a6367417b2b2512adb
title: "usa la proprietà css transform skewX per inclinare un elemento lungo l'asse x"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cyLP8Sr'
forumTopicId: 301074
dashedName: use-the-css-transform-property-skewx-to-skew-an-element-along-the-x-axis
---

# --description--

La funzione successiva della proprietà `transform` è `skewX()`, che inclina l'elemento selezionato lungo il suo asse X (orizzontale) di un dato angolo.

Il codice seguente inclina l'elemento paragrafo di -32 gradi lungo l'asse X.

```css
p {
  transform: skewX(-32deg);
}
```

# --instructions--

Inclina l'elemento con id `bottom` di 24 gradi lungo l'asse X usando la proprietà `transform`.

# --hints--

L'elemento con id `bottom` dovrebbe essere inclinato di 24 gradi lungo il suo asse X.

```js
assert.match(code, /#bottom\s*?{\s*?.*?\s*?transform:\s*?skewX\(24deg\);/g);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;

  }
</style>

<div id="top"></div>
<div id="bottom"></div>
```

# --solutions--

```html
<style>
  div {
    width: 70%;
    height: 100px;
    margin:  50px auto;
  }
  #top {
    background-color: red;
  }
  #bottom {
    background-color: blue;
    transform: skewX(24deg);
  }
</style>
<div id="top"></div>
<div id="bottom"></div>
```
