---
id: 587d78a6367417b2b2512add
title: Crea una grafica usando CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cEDWPs6'
forumTopicId: 301048
dashedName: create-a-graphic-using-css
---

# --description--

Manipolando diversi selettori e proprietĂ , puoi creare forme interessanti. Una delle piĂą semplici da provare Ă¨ la forma di una luna crescente. Per questa sfida devi lavorare con la proprietĂ  `box-shadow` che imposta l'ombra di un elemento, insieme alla proprietĂ  `border-radius` che controlla la rotonditĂ  degli angoli dell'elemento.

Creerai un oggetto rotondo e trasparente con un'ombra nitida leggermente spostata di lato: l'ombra sarĂ  in realtĂ  la forma della luna che vedi.

Per creare un oggetto rotondo, la proprietĂ  `border-radius` dovrebbe essere impostata a un valore del 50%.

Potresti ricordare da una sfida precedente che la proprietĂ  `box-shadow` accetta valori per `offset-x`, `offset-y`, `blur-radius`, `spread-radius` e un valore `color` in quest'ordine. I valori `blur-radius` e `spread-radius` sono opzionali.

# --instructions--

Manipola l'elemento quadrato nell'editor per creare la forma della luna. Prima, cambia `background-color` in `transparent`, poi imposta la proprietĂ  `border-radius` al 50% per ottenere la forma circolare. Infine, modifica la proprietĂ  `box-shadow` per impostare `offset-x` a 25px, `offset-y` a 10px, `blur-radius` a 0, `spread-radius` a 0 e `color` a `blue`.

# --hints--

Il valore della proprietĂ  `background-color` dovrebbe essere impostato su `transparent`.

```js
assert.match(code,/background-color:\s*?transparent;/gi);
```

Il valore della proprietĂ  `border-radius` dovrebbe essere impostato su `50%`.

```js
assert.match(code,/border-radius:\s*?50%;/gi);
```

Il valore della proprietĂ  `box-shadow` dovrebbe essere impostato a 25px per `offset-x`, 10px per `offset-y`, 0 per `blur-radius`, 0 per `spread-radius` e infine `blue` per `color`.

```js
assert.match(code,/box-shadow:\s*?25px\s+?10px\s+?0(px)?\s+?0(px)?\s+?blue\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 0px;
    box-shadow: 25px 10px 10px 10px green;
  }

</style>
<div class="center"></div>
```

# --solutions--

```html
<style>
  .center {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    width: 100px;
    height: 100px;
    background-color: transparent;
    border-radius: 50%;
    box-shadow: 25px 10px 0 0 blue;
  }
</style>
<div class="center"></div>
```
