---
id: 587d78a5367417b2b2512ad8
title: Crea una texture aggiungendo un pattern sottile come immagine di sfondo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQdwJC8'
forumTopicId: 301052
dashedName: create-texture-by-adding-a-subtle-pattern-as-a-background-image
---

# --description--

Un modo per aggiungere texture e interesse a uno sfondo e farlo risaltare di più è aggiungere un pattern sottile. La chiave è l'equilibrio, perché non vuoi che lo sfondo risalti troppo e distolga l'attenzione dal primo piano. La proprietà `background` supporta la funzione `url()` per collegare un'immagine della texture o del pattern scelto. L'indirizzo del collegamento è racchiuso tra virgolette all'interno delle parentesi tonde.

# --instructions--

Usando l'url di `https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png`, imposta la proprietà `background` dell'intera pagina con il selettore `body`.

# --hints--

Il tuo elemento `body` dovrebbe avere una proprietà `background` impostata su un `url()` con il collegamento fornito.

```js
assert.match(code,/background(-image)?:\s*?url\(\s*("|'|)https:\/\/cdn-media-1\.freecodecamp\.org\/imgr\/MJAkxbh\.png\2\s*\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {

  }
</style>
```

# --solutions--

```html
<style>
  body {
    background: url("https://cdn-media-1.freecodecamp.org/imgr/MJAkxbh.png");
  }
</style>
```
