---
id: 587d78a3367417b2b2512ace
title: Spostare gli elementi a sinistra o a destra con la proprietà float
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MDqu2'
forumTopicId: 301066
dashedName: push-elements-left-or-right-with-the-float-property
---

# --description--

Lo strumento di posizionamento successivo non usa realmente `position`, ma imposta la proprietà `float` di un elemento. Gli elementi flottanti vengono rimossi dal flusso normale di un documento e spinti verso `left` o `right` del loro elemento genitore contenitore. Viene comunemente usato insieme alla proprietà `width` per specificare quanto spazio orizzontale richiede l'elemento flottante.

# --instructions--

Il markup fornito funzionerebbe bene come layout a due colonne, con gli elementi `section` e `aside` affiancati. Dai all'elemento `#left` una `float` di `left` e all'elemento `#right` una `float` di `right`.

# --hints--

L'elemento con id `left` dovrebbe avere un valore `float` di `left`.

```js
const leftElement = document.querySelector('#left');
const leftStyle = window.getComputedStyle(leftElement);
assert.equal(leftStyle?.float, 'left');
```

L'elemento con id `right` dovrebbe avere un valore `float` di `right`.

```js
const rightElement = document.querySelector('#right');
const rightStyle = window.getComputedStyle(rightElement);
assert.equal(rightStyle?.float, 'right');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
    #left {

      width: 50%;
    }
    #right {

      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```

# --solutions--

```html
<head>
  <style>
    #left {
      float: left;
      width: 50%;
    }
    #right {
      float: right;
      width: 40%;
    }
    aside, section {
      padding: 2px;
      background-color: #ccc;
    }
  </style>
</head>
<body>
  <header>
    <h1>Welcome!</h1>
  </header>
  <section id="left">
    <h2>Content</h2>
    <p>Good stuff</p>
  </section>
  <aside id="right">
    <h2>Sidebar</h2>
    <p>Links</p>
  </aside>
</body>
```
