---
id: 587d78a5367417b2b2512ad6
title: Crea un graduale gradiente lineare CSS
challengeType: 0
videoUrl: 'https://scrimba.com/c/cg4dpt9'
forumTopicId: 301047
dashedName: create-a-gradual-css-linear-gradient
---

# --description--

Applicare un colore agli elementi HTML non si limita a una sola tonalità piatta. CSS offre la possibilità di usare transizioni di colore, chiamate anche gradienti, sugli elementi. Questo si consulta tramite la proprietà `background` con la funzione `linear-gradient()`. Ecco la sintassi generale:

```css
background: linear-gradient(gradient_direction, color 1, color 2, color 3, ...);
```

Il primo argomento specifica la direzione da cui inizia la transizione di colore - può essere indicata come un grado, dove `90deg` crea un gradiente orizzontale (da sinistra a destra) e `45deg` crea un gradiente diagonale (dal basso a sinistra in alto a destra). Gli argomenti successivi specificano l'ordine dei colori usati nel gradiente.

Esempio:

```css
background: linear-gradient(90deg, red, yellow, rgb(204, 204, 255));
```

# --instructions--

Usa un `linear-gradient()` per la `div` dell'elemento `background`, e impostalo con una direzione di 35 gradi per cambiare il colore da `#CCFFFF` a `#FFCCCC`.

# --hints--

L'elemento `div` dovrebbe avere una `linear-gradient` `background` con la direzione e i colori specificati.

```js
const divElement = document.querySelector('div');
const divStyle = window.getComputedStyle(divElement); 
assert.match(divStyle?.background, /linear-gradient\(35deg, rgb\(204, 255, 255\), rgb\(255, 204, 204\)\)/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;

  }

</style>

<div></div>
```

# --solutions--

```html
<style>
  div {
    border-radius: 20px;
    width: 70%;
    height: 400px;
    margin: 50px auto;
    background: linear-gradient(35deg, #CCFFFF, #FFCCCC);
  }
</style>
<div></div>
```
