---
id: 587d78a9367417b2b2512aea
title: rendi il movimento più naturale usando una curva di Bézier
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7akWUv'
forumTopicId: 301063
dashedName: make-motion-more-natural-using-a-bezier-curve
---

# --description--

Questa sfida anima un elemento per replicare il movimento di una palla che viene fatta rimbalzare. Le sfide precedenti hanno trattato le curve di Bézier cubiche `linear` e `ease-out`, tuttavia nessuna delle due rappresenta accuratamente il movimento del rimbalzo. Devi personalizzare una curva di Bézier per questo.

La `animation-timing-function` fa automaticamente il loop ad ogni keyframe quando la `animation-iteration-count` è impostata su infinito. Poiché c’è una regola di keyframe impostata a metà della durata dell’animazione (a `50%`), si ottengono due progressioni di animazione identiche nel movimento verso l’alto e verso il basso della palla.

La seguente curva di Bézier cubica simula un movimento di rimbalzo:

```css
cubic-bezier(0.3, 0.4, 0.5, 1.6);
```

Nota che il valore di y2 è maggiore di 1. Anche se la curva di Bézier cubica è mappata su un sistema di coordinate 1 per 1 e può accettare solo valori x da 0 a 1, il valore y può essere impostato su numeri maggiori di uno. Questo produce un movimento di rimbalzo ideale per simulare la palla che viene fatta rimbalzare.

# --instructions--

Cambia il valore della `animation-timing-function` dell’elemento con id `green` in una funzione `cubic-bezier` con i valori x1, y1, x2, y2 impostati rispettivamente a 0.311, 0.441, 0.444, 1.649.

# --hints--

Il valore della proprietà `animation-timing-function` per l’elemento con id `green` dovrebbe essere una funzione `cubic-bezier` con i valori x1, y1, x2, y2 come specificato.

```js
const greenElement = document.querySelector('#green');
 const greenStyle = window.getComputedStyle(greenElement);
assert.equal(greenStyle?.animationTimingFunction, 'cubic-bezier(0.311, 0.441, 0.444, 1.649)');
```

# --seed--

## --seed-contents--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.69, 0.1, 1, 0.1);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```

# --solutions--

```html
<style>
  .balls {
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    top: 60%;
    animation-name: jump;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 25%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 50%;
    animation-timing-function: ease-out;
  }
  #green {
    background: green;
    left: 75%;
    animation-timing-function: cubic-bezier(0.311, 0.441, 0.444, 1.649);
  }

  @keyframes jump {
    50% {
      top: 10%;
    }
  }
</style>
<div class="balls" id="red"></div>
<div class="balls" id="blue"></div>
<div class="balls" id="green"></div>
```
