---
id: 587d78a7367417b2b2512ae1
title: "creare movimento usando l'animazione CSS"
challengeType: 0
videoUrl: 'https://scrimba.com/c/c7amZfW'
forumTopicId: 301051
dashedName: create-movement-using-css-animation
---

# --description--

Quando gli elementi hanno un `position` specificato, come `fixed` o `relative`, le proprietà di offset CSS `right`, `left`, `top` e `bottom` possono essere usate nelle regole di animazione per creare movimento.

Come mostrato nell'esempio qui sotto, puoi spingere l'elemento verso il basso e poi verso l'alto impostando la proprietà `top` del keyframe `50%` a 50px, mantenendola a 0px per il primo (`0%`) e l'ultimo (`100%`) keyframe.

```css
@keyframes rainbow {
  0% {
    background-color: blue;
    top: 0px;
  }
  50% {
    background-color: green;
    top: 50px;
  }
  100% {
    background-color: yellow;
    top: 0px;
  }
}
```

# --instructions--

Aggiungi un movimento orizzontale all'animazione `div`. Usando la proprietà di offset `left`, aggiungi alla regola `@keyframes` in modo che rainbow inizi a 0 pixel a `0%`, si sposti a 25 pixel a `50%` e finisca a -25 pixel a `100%`. Non sostituire la proprietà `top` nell'editor: l'animazione deve avere sia movimento verticale sia orizzontale.

# --hints--

La regola `@keyframes` per `0%` dovrebbe usare l'offset `left` di 0px.

```js
assert.match(code,/[^50]0%\s*?{[\s\S]*?left:\s*?0px(;[\s\S]*?|\s*?)}/gi);
```

La regola `@keyframes` per `50%` dovrebbe usare l'offset `left` di 25px.

```js
assert.match(code,/50%\s*?{[\s\S]*?left:\s*?25px(;[\s\S]*?|\s*?)}/gi);
```

La regola `@keyframes` per `100%` dovrebbe usare l'offset `left` di -25px.

```js
assert.match(code,/100%\s*?{[\s\S]*?left:\s*?-25px(;[\s\S]*?|\s*?)}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;

    }
    50% {
      background-color: green;
      top: 50px;

    }
    100% {
      background-color: yellow;
      top: 0px;

    }
  }
</style>

<div id="rect"></div>
```

# --solutions--

```html
<style>
  div {
    height: 40px;
    width: 70%;
    background: black;
    margin: 50px auto;
    border-radius: 5px;
    position: relative;
  }

  #rect {
    animation-name: rainbow;
    animation-duration: 4s;
  }

  @keyframes rainbow {
    0% {
      background-color: blue;
      top: 0px;
      left: 0px;
    }
    50% {
      background-color: green;
      top: 50px;
      left: 25px;
    }
    100% {
      background-color: yellow;
      top: 0px;
      left: -25px;
    }
  }
</style>
<div id="rect"></div>
```
