---
id: 587d78a6367417b2b2512ade
title: Crea una forma più complessa usando CSS e HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpz4fr'
forumTopicId: 301050
dashedName: create-a-more-complex-shape-using-css-and-html
---

# --description--

Una delle forme più popolari al mondo è la forma a cuore, e in questa sfida ne creerai una usando puro CSS. Ma prima, devi capire i pseudo-elementi `::before` e `::after`. `::before` crea un pseudo-elemento che è il primo figlio dell'elemento selezionato; `::after` crea un pseudo-elemento che è l'ultimo figlio dell'elemento selezionato. Nell'esempio seguente, un pseudo-elemento `::before` viene usato per aggiungere un rettangolo a un elemento con la classe `heart`:

```css
.heart::before {
  content: "";
  background-color: yellow;
  border-radius: 25%;
  position: absolute;
  height: 50px;
  width: 70px;
  top: -50px;
  left: 5px;
}
```

Perché i pseudo-elementi `::before` e `::after` funzionino correttamente, devono avere una proprietà `content` definita. Questa proprietà è solitamente usata per aggiungere cose come una foto o testo all'elemento selezionato. Quando i pseudo-elementi `::before` e `::after` sono usati per creare forme, la proprietà `content` è comunque richiesta, ma viene impostata a una stringa vuota. Nell'esempio sopra, l'elemento con la classe `heart` ha un pseudo-elemento `::before` che produce un rettangolo giallo con altezza e larghezza rispettivamente di `50px` e `70px`. Questo rettangolo ha angoli arrotondati grazie al suo `border-radius` del 25% ed è posizionato in modo assoluto a `5px` da sinistra e `50px` sopra la parte superiore dell'elemento.

# --instructions--

Trasforma l'elemento sullo schermo in un cuore. Nel selettore `.heart::after`, cambia la proprietà `background-color` in `pink` e `border-radius` in 50%.

Poi, seleziona l'elemento con la classe `heart` (solo `heart`) e completa la proprietà `transform`. Usa la funzione `rotate()` con -45 gradi.

Infine, nel selettore `.heart::before`, imposta la sua proprietà `content` a una stringa vuota.

# --hints--

La proprietà `background-color` del selettore `.heart::after` dovrebbe essere `pink`.

```js
const heartAfter = code.match(/\.heart::after\s*{[\s\S]+?[^\}]}/g)?.[0];
assert.match(heartAfter, /({|;)\s*background-color\s*:\s*pink\s*(;|})/g);
```

La `border-radius` del selettore `.heart::after` dovrebbe essere 50%.

```js
assert.lengthOf(code.match(/border-radius\s*?:\s*?50%/gi),2);
```

La proprietà `transform` per la classe `heart` dovrebbe usare una funzione `rotate()` impostata a -45 gradi.

```js
assert.match(code,/transform\s*?:\s*?rotate\(\s*?-45deg\s*?\)/gi);
```

La `content` del selettore `.heart::before` dovrebbe essere una stringa vuota.

```js
assert.match(code,/\.heart::before\s*?{\s*?content\s*?:\s*?("|')\1\s*?;/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: ;
  }
  .heart::after {
    background-color: blue;
    content: "";
    border-radius: 25%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: ;
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
  }
  .heart::after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart::before {
    content: "";
    background-color: pink;
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }
</style>
<div class="heart"></div>
```
