---
id: 587d78a9367417b2b2512ae9
title: Usa una curva di Bézier per muovere un grafico
challengeType: 0
videoUrl: 'https://scrimba.com/c/c6bnRCK'
forumTopicId: 301071
dashedName: use-a-bezier-curve-to-move-a-graphic
---

# --description--

Una sfida precedente ha parlato della parola chiave `ease-out` che descrive un cambiamento di animazione che accelera all’inizio e poi rallenta alla fine dell’animazione. A destra, viene mostrata la differenza tra la parola chiave `ease-out` (per l’elemento blu) e la parola chiave `linear` (per l’elemento rosso). Progressioni di animazione simili alla parola chiave `ease-out` si possono ottenere usando una funzione personalizzata di curva di Bézier cubica.

In generale, modificare i punti di àncora `p1` e `p2` porta alla creazione di curve di Bézier diverse, che controllano come l’animazione procede nel tempo. Ecco un esempio di curva di Bézier che usa valori per imitare lo stile `ease-out`:

```css
animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
```

Ricorda che tutte le funzioni `cubic-bezier` iniziano con `p0` in (0, 0) e finiscono con `p3` in (1, 1). In questo esempio, la curva si muove più velocemente lungo l’asse Y (parte da 0, va al valore y `p1` di 0, poi va al valore y `p2` di 1) rispetto all’asse X (da 0 all’inizio, poi 0 per `p1`, fino a 0.58 per `p2`). Di conseguenza, il cambiamento nell’elemento animato procede più velocemente rispetto al tempo dell’animazione per quel segmento. Verso la fine della curva, la relazione tra il cambiamento dei valori x e y si inverte: il valore y si muove da 1 a 1 (nessun cambiamento), mentre i valori x si muovono da 0.58 a 1, facendo sì che i cambiamenti dell’animazione procedano più lentamente rispetto alla durata dell’animazione.

# --instructions--

Per vedere l’effetto di questa curva di Bézier in azione, cambia la proprietà `animation-timing-function` dell’elemento con id `red` in una funzione `cubic-bezier` con x1, y1, x2, y2 impostati rispettivamente ai valori `0, 0, 0.58, 1`. Questo farà sì che entrambi gli elementi procedano in modo simile nell’animazione.

# --hints--

Il valore della proprietà `animation-timing-function` dell’elemento con id `red` dovrebbe essere una funzione `cubic-bezier` con i valori x1, y1, x2, y2 impostati rispettivamente a 0, 0, 0.58, 1 .

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.equal(
  redStyle?.animationTimingFunction, 'cubic-bezier(0, 0, 0.58, 1)'
);
```

L’elemento con id `red` non dovrebbe più avere la proprietà `animation-timing-function` impostata su `linear`.

```js
const redElement = document.querySelector('#red');
const redStyle = window.getComputedStyle(redElement);
assert.notEqual(redStyle?.animationTimingFunction, 'linear');
```

Il valore della proprietà `animation-timing-function` per l’elemento con id `blue` non dovrebbe cambiare.

```js
const blueElement = document.querySelector('#blue');
const blueStyle = window.getComputedStyle( blueElement);
const blueBallAnimation = __helpers.removeWhiteSpace(
  blueStyle?.animationTimingFunction
);
assert.isTrue(
  blueBallAnimation == 'ease-out' ||
    blueBallAnimation == 'cubic-bezier(0,0,0.58,1)'
);
```

# --seed--

## --seed-contents--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: linear;
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```

# --solutions--

```html
<style>
  .balls{
    border-radius: 50%;
    position: fixed;
    width: 50px;
    height: 50px;
    margin-top: 50px;
    animation-name: bounce;
    animation-duration: 2s;
    animation-iteration-count: infinite;
  }
  #red {
    background: red;
    left: 27%;
    animation-timing-function: cubic-bezier(0, 0, 0.58, 1);
  }
  #blue {
    background: blue;
    left: 56%;
    animation-timing-function: ease-out;
  }
  @keyframes bounce {
    0% {
      top: 0px;
    }
    100% {
      top: 249px;
    }
  }
</style>
<div class="balls" id= "red"></div>
<div class="balls" id= "blue"></div>
```
