---
id: 587d78a5367417b2b2512ad9
title: usa la proprietà css transform scale per cambiare la dimensione di un elemento
challengeType: 0
videoUrl: 'https://scrimba.com/c/c2MZVSg'
forumTopicId: 301076
dashedName: use-the-css-transform-scale-property-to-change-the-size-of-an-element
---

# --description--

Per cambiare la scala di un elemento, CSS ha la proprietà `transform` insieme alla sua funzione `scale()`. L'esempio di codice seguente raddoppia la dimensione di tutti gli elementi paragrafo nella pagina:

```css
p {
  transform: scale(2);
}
```

# --instructions--

Aumenta la dimensione dell'elemento con l'id `ball2` a 1,5 volte la sua dimensione originale.

# --hints--

La proprietà `transform` per `#ball2` dovrebbe essere impostata su scale per ingrandirlo a 1,5 volte la sua dimensione.

```js
assert.match(code,
    /#ball2\s*?{\s*?left:\s*?65%;\s*?transform:\s*?scale\(1\.5\);\s*?}|#ball2\s*?{\s*?transform:\s*?scale\(1\.5\);\s*?left:\s*?65%;\s*?}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;

  }


</style>

<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>
```

# --solutions--

```html
<style>
  .ball {
    width: 40px;
    height: 40px;
    margin: 50 auto;
    position: fixed;
    background: linear-gradient(
      35deg,
      #ccffff,
      #ffcccc
    );
    border-radius: 50%;
  }
  #ball1 {
    left: 20%;
  }
  #ball2 {
    left: 65%;
    transform: scale(1.5);
  }
</style>
<div class="ball" id= "ball1"></div>
<div class="ball" id= "ball2"></div>
```
