---
id: 58a7a6ebf9a6318348e2d5aa
title: "Modifica la modalità di riempimento di un'animazione"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJDmcE'
forumTopicId: 301064
dashedName: modify-fill-mode-of-an-animation
---

# --description--

Ottimo, ma ancora non funziona correttamente. Nota come l'animazione si resetta dopo che `500ms` è passato, facendo tornare il pulsante al colore originale. Vuoi che il pulsante rimanga evidenziato.

Questo si può fare impostando la proprietà `animation-fill-mode` su `forwards`. `animation-fill-mode` specifica lo stile applicato a un elemento quando l'animazione è terminata. Puoi impostarlo così:

```css
animation-fill-mode: forwards;
```

# --instructions--

Imposta la proprietà `animation-fill-mode` di `button:hover` su `forwards` in modo che il pulsante rimanga evidenziato quando un utente ci passa sopra con il mouse.

# --hints--

`button:hover` dovrebbe avere una proprietà `animation-fill-mode` con valore `forwards`.

```js
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-fill-mode\s*?:\s*?forwards\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-name\s*?:\s*?background-color\s*?;[\s\S]*}/gi);
assert.match(code,/button\s*?:\s*?hover\s*?{[\s\S]*animation-duration\s*?:\s*?500ms\s*?;[\s\S]*}/gi );
```

# --seed--

## --seed-contents--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    /* Only change code below this line */

    /* Only change code above this line */
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```

# --solutions--

```html
<style>
  button {
    border-radius: 5px;
    color: white;
    background-color: #0F5897;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    animation-name: background-color;
    animation-duration: 500ms;
    animation-fill-mode: forwards;
  }
  @keyframes background-color {
    100% {
      background-color: #4791d0;
    }
  }
</style>
<button>Register</button>
```
