---
id: 5f33071498eb2472b87ddee4
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu.

Start by adding your `<!DOCTYPE html>` and `html` element with a `lang` attribute set to `en`.

# --hints--

Você deve ter a declaração `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter uma tag de abertura `<html lang="en">` com o idioma definido para inglês.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve acrescentar uma tag de fechamento `</html>`.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
