---
id: 5f33294a6af5e9188dbdb8f3
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

Você deve acrescentar uma tag de abertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Você deve acrescentar uma tag de fechamento `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Você deve manter o elemento preexistente `body`. Verifique se você não excluiu a tag de fechamento.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

A tag `main` deve estar dentro da tag `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
