---
id: 5f344f9c805cd193c33d829c
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

In previous lectures, you learned how to add `CSS` properties and values like this:

```css
element {
 property: value;
}
```

Centralize o conteúdo do elemento `h1` definindo sua propriedade `text-align` para o valor `center`.

# --hints--

Você deve adicionar um elemento `h1` ao elemento `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Sua propriedade `text-align` deve ter um valor de `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

O seletor `h1` deve definir a propriedade `text-align` para `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

