---
id: 61fd5a93fd62bb35968adeab
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin your project by giving your `head` element a `link` element for your stylesheet.

# --hints--

Dein Code sollte ein `link`-Element enthalten.

```js
assert.match(code, /<link/);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` enthalten.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
