---
id: 61fd6ab779390f49148773bb
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstellen Sie unter Ihrem `h1`-Element ein `div`-Element. Geben Sie ihm ein `id`-Attribut mit dem Wert `years`. Dieses bestimmte Element soll für Screenreader verborgen sein, daher geben Sie ihm das `aria-hidden`-Attribut mit dem Wert `true`.

# --hints--

Sie sollten ein neues `div`-Element nach Ihrem `h1`-Element erstellen.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

Ihr neues `div`-Element sollte ein `id`-Attribut mit dem Wert `years` haben.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

Ihr neues `div`-Element sollte das `aria-hidden`-Attribut mit dem Wert `true` haben.

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
