---
id: 61fd70336ebb3e4f62ee81ba
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML-Tabellen verwenden das `caption`-Element, um zu beschreiben, worum es in der Tabelle geht. Das `caption`-Element sollte immer das erste Kind eines `table` sein, kann aber mit der CSS-Eigenschaft `caption-side` positioniert werden.

Fügen Sie Ihrem ersten `caption`-Element ein `table`-Element hinzu und geben Sie ihm den Text `Assets`.

# --hints--

Ihr erstes `table`-Element sollte ein `caption`-Element enthalten.

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

Ihr `caption`-Element sollte den Text `Assets` enthalten.

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
