---
id: 61fd6343fb42013d99bcd7f3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Verschachteln Sie innerhalb Ihres `body`-Elements ein `section`-Element innerhalb eines `main`-Elements.

# --hints--

Das erste Element innerhalb des `body`-Elements sollte Ihr `main`-Element sein.

```js
assert.strictEqual(
  document.querySelector('body')?.children?.[0]?.localName,
  'main'
);
```

Sie sollten ein `section`-Element innerhalb Ihres `main`-Elements haben.

```js
assert.strictEqual(
  document.querySelector('main')?.children?.[0]?.localName,
  'section'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
