---
id: 61fd5a93fd62bb35968adeab
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie Ihr Projekt, indem Sie Ihrem `head`-Element ein `link`-Element für Ihr Stylesheet hinzufügen.

# --hints--

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.match(code, /<link/);
```

Ihr `link`-Element sollte sich innerhalb Ihres `head`-Elements befinden.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
