---
id: 61fd9b7285bde783ad5b8aac
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Innerhalb des `tbody` füge ein `tr` mit der `class` von `total` hinzu. In diesem, füge ein `th` mit dem Text `Total Net Worth` hinzu, und verpacke `Net Worth` in ein `span` mit der `class` von `sr-only`.

Füge dann drei `td`-Elemente hinzu und weise dem dritten eine `class` mit dem Wert `current` zu und weise jedem Element den folgenden Text zu: `$-171`, `$136`, `$334`.

# --hints--

Dein `tbody`-Element sollte ein `tr`-Element enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
assert.lengthOf(tbody?.querySelectorAll('tr'), 1);
```

Dein `tr`-Element sollte die `class` auf `total` eingestellt haben.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
assert.isTrue(tbody?.querySelector('tr')?.classList?.contains('total'));
```

Dein `tr` sollte ein `th`-Element enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isNotNull(tableRow?.querySelector('th'));
```

Dein `th`-Element sollte den Text `Total Net Worth` enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelector('th')?.innerText.trim(), 'Total Net Worth');
```

Du solltest den Text `Net Worth` in einem `span`-Element einbetten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Net Worth'
);
```

Dein `span`-Element sollte das `class`-Attribut auf `sr-only` eingestellt haben.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('sr-only')
);
```

Du solltest drei `td`-Elemente haben.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Dein erstes `td`-Element sollte den Text `$-171` enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$-171');
```

Dein zweites `td`-Element sollte den Text `$136` enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$136');
```

Dein drittes `td`-Element sollte den Text `$334` enthalten.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$334');
```

Dein drittes `td`-Element sollte die `class` auf `current` eingestellt haben.

```js
const table = document.querySelectorAll('table')?.[2];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[0];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
                <th>Expenses <span class="description">Annual anticipated expenses, such as payroll.</span></th>
                <td>$200</td>
                <td>$300</td>
                <td class="current">$400</td>
              </tr>
              <tr class="data">
                <th>Credit <span class="description">The outstanding balance on our credit card.</span></th>
                <td>$50</td>
                <td>$50</td>
                <td class="current">$75</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Liabilities</span></th>
                <td>$750</td>
                <td>$600</td>
                <td class="current">$475</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
            <caption>
              Net Worth
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
            </tbody>
          </table>
--fcc-editable-region--
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
