---
id: 61fd94056e0355785fbba4d3
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Fügen Sie innerhalb des zweiten `tr` ein `th`-Element mit dem Text `Expenses Annual anticipated expenses, such as payroll.` hinzu. Wickeln Sie diesen Text, außer `Expenses `, in ein `span`-Element mit dem `class`, das auf `description` gesetzt ist.

Fügen Sie darunter drei `td`-Elemente hinzu und geben Sie ihnen der Reihe nach die folgenden Texte: `$200`, `$300` und `$400`. Geben Sie dem dritten `td`-Element ein `class`, das auf `current` gesetzt ist.

# --hints--

Ihr zweites `tr` sollte ein `th`-Element enthalten.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.isNotNull(tableRow?.querySelector('th'));
```

Ihr `th`-Element sollte den Text `Expenses Annual anticipated expenses, such as payroll.` haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Expenses Annual anticipated expenses, such as payroll.'
);
```

Sie sollten den Text `Annual anticipated expenses, such as payroll.` in ein `span`-Element einwickeln.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Annual anticipated expenses, such as payroll.'
);
```

Ihr `span`-Element sollte das `class`-Attribut auf `description` gesetzt haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

Sie sollten drei `td`-Elemente haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Ihr erstes `td`-Element sollte den Text `$200` haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$200');
```

Ihr zweites `td`-Element sollte den Text `$300` haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$300');
```

Ihr drittes `td`-Element sollte den Text `$400` haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$400');
```

Ihr drittes `td`-Element sollte das `class` auf `current` gesetzt haben.

```js
const table = document.querySelectorAll('table')?.[1];
const tbody = table?.querySelector('tbody');
const tableRow = tbody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
          <table>
            <caption>
              Liabilities
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only">2019</span></th>
                <th><span class="sr-only">2020</span></th>
                <th><span class="sr-only">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Loans <span class="description">The outstanding balance on our startup loan.</span></th>
                <td>$500</td>
                <td>$250</td>
                <td class="current">$0</td>
              </tr>
              <tr class="data">
--fcc-editable-region--
                
--fcc-editable-region--
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
