---
id: 61fd78621573aa5e8b512f5e
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Fügen Sie in Ihrem dritten `tr`-Element ein `th`-Element mit dem Text `Savings Funds set aside for emergencies.` hinzu. Wickeln Sie diesen Text, außer `Savings `, in ein `span`-Element mit dem `class`-Attribut, das auf `description` gesetzt ist.

Fügen Sie anschließend drei `td`-Elemente mit folgendem Text (in dieser Reihenfolge) hinzu: `$500`, `$650`, `$728`. Geben Sie dem dritten `td`-Element ein `class`-Attribut, das auf `current` gesetzt ist.

# --hints--

Ihr drittes `tr` sollte ein `th`-Element enthalten.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isNotNull(tableRow?.querySelector('th'));
```

Ihr `th`-Element sollte den Text `Savings Funds set aside for emergencies.` haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Savings Funds set aside for emergencies.'
);
```

Sie sollten den Text `Funds set aside for emergencies.` in ein `span`-Element einwickeln.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Funds set aside for emergencies.'
);
```

Ihr `span`-Element sollte das `class`-Attribut haben, das auf `description` gesetzt ist.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

Sie sollten drei `td`-Elemente haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

Ihr erstes `td`-Element sollte den Text `$500` haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$500');
```

Ihr zweites `td`-Element sollte den Text `$650` haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$650');
```

Ihr drittes `td`-Element sollte den Text `$728` haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$728');
```

Ihr drittes `td`-Element sollte das `class` auf `current` gesetzt haben.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[2];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
