---
id: 68c33e502cd97f2fd52cb63b
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kutengeneza vipengele vya jedwali vinavyopatikana kwa wote kwa kuunda ratiba ya mkutano wa teknolojia.

Anza kwa kuongeza kipengele cha `h1` chenye maandishi ya `Tech Conference 2025 Schedule`.

# --hints--

Unapaswa kuwa na kipengele cha `h1`.

```js
const h1 = document.querySelector("h1");
assert.exists(h1);
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `Tech Conference 2025 Schedule`.

```js
const h1 = document.querySelector("h1");
assert.strictEqual(h1?.textContent.trim(), "Tech Conference 2025 Schedule");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
    --fcc-editable-region--
    
    --fcc-editable-region--
</body>
</html>
```
