---
id: 68c346d03709343bff132707
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Kwa hatua chache zijazo, utaanzisha mwili wa jedwali.

Anza kwa kuongeza kipengele cha `tbody` chini ya kipengele cha `thead`. Kisha ndani ya `tbody`, ongeza kipengele cha `tr`.

# --hints--

Unapaswa kuwa na kipengele cha `tbody` chini ya kipengele cha `thead`.

```js
assert.exists(document.querySelector("tbody"));
```

Unapaswa kuwa na kipengele cha `tr` ndani ya kipengele cha `tbody`.

```js
assert.exists(document.querySelector("tbody tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>
    
    --fcc-editable-region--
    
    --fcc-editable-region--
  </table>
</body>
</html>
```
