---
id: 68c34f4aeb2b1a4170506751
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Ifuatayo, ongeza safu ya tano kwenye jedwali. Anza kwa kuongeza kipengele kingine cha `tr`. Ndani ya kipengele hicho cha `tr`, ongeza kipengele cha `th` chenye sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"` na maandishi ya `12:30 PM`. Kisha chini ya kipengele hicho cha `th`, ongeza kipengele cha `td` chenye maandishi ya `Lunch Break`. Kipengele chako cha `td` pia kinapaswa kuwa na sifa ya kipengele ya `colspan` iliyowekwa kuwa `3` ili kiweze kuenea kwenye njia zote tatu.

# --hints--

Unapaswa kuwa na kipengele cha `tr` cha tano ndani ya kipengele cha `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 5);
```

Unapaswa kuwa na kipengele cha `th` ndani ya kipengele cha `tr` cha tano.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("th"));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi ya `12:30 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.textContent.trim(), "12:30 PM");
```

Kipengele chako cha `th` kinapaswa kuwa na sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("th")?.getAttribute("scope"), "row");
```

Unapaswa kuwa na kipengele cha `td` ndani ya kipengele cha `tr` cha tano.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[4].querySelector("td"));
```

Kipengele chako cha `td` kinapaswa kuwa na maandishi ya `Lunch Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.textContent.trim(), "Lunch Break");
``` 

Kipengele chako cha `td` kinapaswa kuwa na sifa ya kipengele ya `colspan` iliyowekwa kuwa `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[4].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```
