---
id: 68c3420cb0176439a95af1b8
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Hatua inayofuata ni kuongeza kichwa cha jedwali.

Anza kwa kuongeza kipengele cha `thead` chini ya kipengele cha `caption`. Ndani ya `thead`, ongeza kipengele cha `tr`.

# --hints--

Unapaswa kuwa na kipengele cha `thead`.

```js
assert.exists(document.querySelector("table thead"));
```

Unapaswa kuwa na kipengele cha `tr` ndani ya `thead`.

```js
assert.exists(document.querySelector("table thead tr"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    --fcc-editable-region--
        
    --fcc-editable-region--
  </table>
</body>
</html>
```
