---
id: 68c353c4b114aa428c01f260
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Hatua ya mwisho ni kuongeza safu moja zaidi kwenye jedwali.

Anza kwa kuongeza kipengele kingine cha `tr`. Ndani ya kipengele hicho cha `tr`, ongeza kipengele cha `th` chenye sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"` na maandishi ya `2:00 PM`.

Kisha, ongeza vipengele vitatu vya `td` vyenye maandishi yafuatayo:

- `Voice UI Workshop`
- `Git & GitHub Essentials`
- `Color & Contrast in UI`

Kwa mabadiliko hayo ya mwisho, jedwali lako sasa limekamilika!

# --hints--

Unapaswa kuwa na kipengele cha sita cha `tr` ndani ya kipengele cha `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 6);
```

Unapaswa kuwa na kipengele cha `th` ndani ya kipengele cha sita cha `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[5].querySelector("th"));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi ya `2:00 PM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.textContent.trim(), "2:00 PM");
```

Kipengele chako cha `th` kinapaswa kuwa na sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelector("th")?.getAttribute("scope"), "row");
```

Kipengele chako cha sita cha `tr` kinapaswa kuwa na vipengele vitatu vya `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[5].querySelectorAll("td"), 3);
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maandishi ya `Voice UI Workshop`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[0]?.textContent.trim(), "Voice UI Workshop");
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maandishi ya `Git & GitHub Essentials`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[1]?.textContent.trim(), "Git & GitHub Essentials");
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maandishi ya `Color & Contrast in UI`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[5].querySelectorAll("td")[2]?.textContent.trim(), "Color & Contrast in UI");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

    --fcc-editable-region--
      
    --fcc-editable-region--
    </tbody>
  </table>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      <tr>
        <th scope="row">11:30 AM</th>
        <td>AR/VR in Education</td>
        <td>JavaScript Fundamentals</td>
        <td>Design Systems at Scale</td>
      </tr>

      <tr>
        <th scope="row">12:30 PM</th>
        <td colspan="3">Lunch Break</td>
      </tr>

      <tr>
        <th scope="row">2:00 PM</th>
        <td>Voice UI Workshop</td>
        <td>Git & GitHub Essentials</td>
        <td>Color & Contrast in UI</td>
      </tr>
    </tbody>
  </table>
</body>
</html>
```
