---
id: 68c34e201b347c40acda8111
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Sasa ni wakati wa kuongeza safu nyingine kwenye jedwali.

Anza kwa kuongeza kipengele kingine cha `tr`. Ndani ya kipengele hicho cha `tr`, ongeza kipengele cha `th` chenye sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"` na maudhui ya maandishi ya `11:30 AM`.

Kisha, ongeza vipengele vitatu vya `td` vyenye maudhui ya maandishi yafuatayo:

- `AR/VR in Education`
- `JavaScript Fundamentals`
- `Design Systems at Scale`

# --hints--

Unapaswa kuwa na kipengele cha nne cha `tr` ndani ya kipengele cha `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 4);
```

Unapaswa kuwa na kipengele cha `th` ndani ya kipengele cha nne cha `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[3].querySelector("th"));
```

Kipengele chako cha `th` kinapaswa kuwa na maudhui ya maandishi ya `11:30 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.textContent.trim(), "11:30 AM");
```

Kipengele chako cha `th` kinapaswa kuwa na sifa ya kipengele ya `scope` iliyowekwa kuwa `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelector("th")?.getAttribute("scope"), "row");
```

Kipengele chako cha nne cha `tr` kinapaswa kuwa na vipengele vitatu vya `td`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr")[3].querySelectorAll("td"), 3);
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maudhui ya maandishi ya `AR/VR in Education`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[0]?.textContent.trim(), "AR/VR in Education");
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maudhui ya maandishi ya `JavaScript Fundamentals`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[1]?.textContent.trim(), "JavaScript Fundamentals");
```   

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maudhui ya maandishi ya `Design Systems at Scale`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[3].querySelectorAll("td")[2]?.textContent.trim(), "Design Systems at Scale");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        <td colspan="3">Break</td>
      </tr>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
