---
id: 68c34329d430583a6cd5f7f1
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Ndani ya kipengele chako cha `tr`, ongeza vipengele vinne vya `th`. Kipengele cha kwanza cha `th` kinapaswa kuwa na maandishi ya `Time`, cha pili kiwe na maandishi ya `Track A`, cha tatu kiwe na maandishi ya `Track B`, na cha nne kiwe na maandishi ya `Track C`.

# --hints--

Kipengele chako cha `tr` kinapaswa kuwa na vipengele vinne vya `th`.

```js
assert.lengthOf(document.querySelectorAll("tr > th"), 4);
```

Kipengele chako cha kwanza cha `th` kinapaswa kuwa na maandishi ya `Time`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(1)")?.textContent.trim(), "Time");
```

Kipengele chako cha pili cha `th` kinapaswa kuwa na maandishi ya `Track A`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(2)")?.textContent.trim(), "Track A");
```

Kipengele chako cha tatu cha `th` kinapaswa kuwa na maandishi ya `Track B`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(3)")?.textContent.trim(), "Track B");
```

Kipengele chako cha nne cha `th` kinapaswa kuwa na maandishi ya `Track C`.

```js
assert.strictEqual(document.querySelector("tr > th:nth-of-type(4)")?.textContent.trim(), "Track C");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
