---
id: 68c341332d98fe38fae733af
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Sasa ni wakati wa kuanza kujenga jedwali.

Anza kwa kuongeza kipengele cha `table` chini ya kipengele chako cha `h1`. Ndani ya kipengele cha `table`, ongeza kipengele cha `caption` chenye maandishi ya `Schedule by Track and Time`.

# --hints--

Unapaswa kuwa na kipengele cha `table`.

```js
assert.exists(document.querySelector("table"));
```

Kipengele chako cha `table` kinapaswa kuwa na kipengele cha `caption`.

```js
assert.exists(document.querySelector("table caption"));
```

Kipengele chako cha `caption` kinapaswa kuwa na maandishi ya `Schedule by Track and Time`.

```js
assert.equal(document.querySelector("table caption")?.textContent.trim(), "Schedule by Track and Time");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

--fcc-editable-region--
    
--fcc-editable-region--
</body>
</html>
```
