---
id: 68c343692285523aef0fc6bf
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sifa ya kipengele `scope` hutumika kubainisha kama seli ya kichwa ni kichwa cha safu, safu wima, au kundi la safu au safu wima. Hapa kuna mfano:

```html
<th scope="col">Example Header</th>
```

Hii husaidia visomaji vya skrini kuelewa uhusiano kati ya kichwa na seli za data.

Kwa vipengele vyote vya `th`, ongeza sifa ya kipengele `scope` yenye thamani ya `col`.

# --hints--

Kila kipengele cha `th` kinapaswa kuwa na sifa ya kipengele `scope` yenye thamani `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
