---
id: 68c34d57ed874e4004f7b5d4
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Kwa sasa, kipengele cha `td` chenye maudhui ya maandishi ya `Break` kinachukua safu moja tu. Lakini itakuwa nzuri ikiwa kingechukua safu zote tatu.

Kama unavyokumbuka kutoka kwenye warsha na mafundisho ya awali, unaweza kutumia sifa ya kipengele ya `colspan` kufanya seli ya jedwali ichukue safu nyingi.

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Ongeza sifa ya kipengele ya `colspan` kwenye kipengele cha `td` na uweke thamani yake kuwa `3`.

# --hints--

Kipengele chako cha `td` chenye maudhui ya maandishi ya `Break` kinapaswa kuwa na sifa ya kipengele ya `colspan` iliyowekwa kuwa `3`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.getAttribute("colspan"), "3");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

      <tr>
        <th scope="row">11:00 AM</th>
        --fcc-editable-region--
        <td>Break</td>
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
