---
id: 68c3482278a14e3cf225e897
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Ndani ya kipengele chako cha `tr`, ongeza kipengele cha `th` chenye maandishi ya `9:00 AM`. Kisha chini ya kipengele hicho cha `th`, ongeza vipengele vitatu vya `td` vyenye maandishi ya `Keynote: Tech Future`, `Intro to Web Dev`, na `UX for All`. 

# --hints--

Unapaswa kuwa na kipengele cha `th` ndani ya kipengele chako cha `tr`.

```js
assert.exists(document.querySelector("tbody tr th"));
```

Kipengele chako cha `th` kinapaswa kuwa na maudhui ya maandishi ya `9:00 AM`.

```js
assert.equal(document.querySelector("tbody tr th")?.textContent.trim(), "9:00 AM");
```

Unapaswa kuwa na jumla ya vipengele vitatu vya `td` chini ya kipengele chako cha `th`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.lengthOf(tdElements, 3);
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maudhui ya maandishi ya `Keynote: Tech Future`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[0]?.textContent.trim(), "Keynote: Tech Future");
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maudhui ya maandishi ya `Intro to Web Dev`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[1]?.textContent.trim(), "Intro to Web Dev");
``` 

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maudhui ya maandishi ya `UX for All`.

```js
const tdElements = document.querySelectorAll("tbody tr td");
assert.equal(tdElements[2]?.textContent.trim(), "UX for All");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </tr>
    </tbody>
    
  </table>
</body>
</html>
```
