---
id: 68c34c608cf9e53f5d36b1d5
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Ifuatayo, ongeza safu ya tatu kwenye jedwali. Anza kwa kuongeza kipengele kingine cha `tr`. Ndani ya kipengele hicho cha `tr`, ongeza kipengele cha `th` chenye sifa ya `scope` iliyowekwa kuwa `"row"` na maandishi ya `11:00 AM`. Kisha chini ya kipengele hicho cha `th`, ongeza kipengele cha `td` chenye maandishi ya `Break`. 

# --hints--

Unapaswa kuwa na kipengele cha tatu cha `tr` ndani ya kipengele cha `tbody`.

```js
assert.lengthOf(document.querySelectorAll("tbody tr"), 3);
```

Unapaswa kuwa na kipengele cha `th` ndani ya kipengele cha tatu cha `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("th"));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi ya `11:00 AM`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.textContent.trim(), "11:00 AM");
```

Kipengele chako cha `th` kinapaswa kuwa na sifa ya `scope` iliyowekwa kuwa `"row"`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("th")?.getAttribute("scope"), "row");
```

Unapaswa kuwa na kipengele cha `td` ndani ya kipengele cha tatu cha `tr`.

```js
assert.isNotNull(document.querySelectorAll("tbody tr")[2].querySelector("td"));
```

Kipengele chako cha `td` kinapaswa kuwa na maandishi ya `Break`.

```js
assert.strictEqual(document.querySelectorAll("tbody tr")[2].querySelector("td")?.textContent.trim(), "Break");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
        <th scope="col">Time</th>
        <th scope="col">Track A</th>
        <th scope="col">Track B</th>
        <th scope="col">Track C</th>
      </tr>
    </thead>

    <tbody>
      <tr>
        <th scope="row">9:00 AM</th>
        <td>Keynote: Tech Future</td>
        <td>Intro to Web Dev</td>
        <td>UX for All</td>
      </tr>

      <tr>
        <th scope="row">10:00 AM</th>
        <td>Accessibility Deep Dive</td>
        <td>CSS for Beginners</td>
        <td>Inclusive Design Principles</td>
      </tr>

--fcc-editable-region--
      
--fcc-editable-region--
    </tbody>
    
  </table>
</body>
</html>
```
