---
id: 662693f82c91a66be46c881b
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

Вчитель закінчив оцінювати тести своїх учнів та потребує вашої допомоги, щоб обчислити середній бал за клас.

Виконайте функцію `getAverage`, яка отримує масив результатів тесту та повертає середній бал.

Середнє значення обчислюється шляхом додавання всіх балів та ділення на загальну кількість балів.

```md
average = sum of all scores / total number of scores
```

Вам надано кілька викликів функцій, щоб ви могли перевірити свій код.

**Поради**

- Ви можете скористатися циклом для ітерування над масивом `scores` та підсумувати всі бали.
- Щоб отримати загальну кількість балів, можна скористатися властивістю `length`.

# --hints--

Функція `getAverage` має повернути число.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` має повернути `71.7`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` має повернути `85.4`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` має повернути `92.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

Функція `getAverage` має повернути середній бал результатів тесту.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
