---
id: 6626b4c58c027d86478ff5eb
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Вчитель реально задоволений тією програмою, яку ви створили до цього часу. Але тепер вона хоче мати простий спосіб перевірити, чи є у студента прохідний бал. Прохідний бал — це все, що не є `"F"`.

Виконайте функцію `hasPassingGrade`, яка приймає бали студента як параметр. Функція має повернути `true`, якщо студент має прохідний бал, та `false`, якщо його немає.

**Поради**

- Використовуйте функцію `getGrade`, щоб отримати оцінку студента. Потім перевірте, чи оцінка прохідна, чи ні.

# --hints--

Функція `hasPassingGrade` має повернути булеве значення.

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

Функція `hasPassingGrade` має повернути `true`, якщо оцінка є `"A"`.

```js
assert.isTrue(hasPassingGrade(100));
```

Функція `hasPassingGrade` має повернути `false`, якщо оцінка є `"F"`.

```js
assert.isFalse(hasPassingGrade(53));
```

Функція `hasPassingGrade` має повернути `true` для всіх прохідних оцінок.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {

}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
