---
id: 6626a060c4006f793e10cb33
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер вчителю потрібна ваша допомога, щоб перетворити оцінку студента на буквену оцінку.

Виконайте функцію `getGrade`, яка приймає число `score` як параметр. Функція має повернути рядок, що представляє буквену оцінку на основі балу.

Ось бали та відповідні їм буквені оцінки:

| Діапазон балів | Оцінка  |
| -------------- | ------- |
| `100`          | `"A++"` |
| `90 - 99`      | `"A"`   |
| `80 - 89`      | `"B"`   |
| `70 - 79`      | `"C"`   |
| `60 - 69`      | `"D"`   |
| `0 - 59`       | `"F"`   |


**Поради**

- Пам’ятайте, то що ви дізналися про умовні оператори (`if`, `else if` та `else`).
- Пам’ятайте, то що ви дізналися про оператори порівняння (`>`, `<`, `>=`, `<=`, `===`).
- Remember that you learned about the logical AND operator (`&&`).

# --hints--

Функція `getGrade` має повернути `"A++"`, якщо оцінка `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Функція `getGrade` має повернути `"A"`, якщо оцінка `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Функція `getGrade` має повернути `"B"`, якщо оцінка між `80` та `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Функція `getGrade` має повернути `"C"`, якщо оцінка `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Функція `getGrade` має повернути `"D"`, якщо оцінка між `60` та `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Функція `getGrade` має повернути `"F"`, якщо оцінка `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Функція `getGrade` має повернути `"F"`, якщо оцінка між `0` та `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
