---
id: a77dbc43c33f39daa4429b4f
title: 불리언 검사 함수 만들기
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

이 실습에서는 값이 불리언 원시형으로 분류되는지 검사하는 함수를 만드세요.

불리언 원시형은 `true`와 `false`입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 하나의 인수를 받는 `booWho`라는 함수가 있어야 합니다.
1. 받은 인수가 불리언 원시형이면 함수는 `true`를 반환해야 합니다.
1. 인수가 다른 값이면 함수는 `false`를 반환해야 합니다.

# --hints--

`booWho` 함수가 있어야 합니다.

```js
assert.isFunction(booWho);
```

`booWho(true)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(booWho(true));
```

`booWho(false)`는 `true`를 반환해야 합니다.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")`는 `false`를 반환해야 합니다.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
