---
id: a103376db3ba46b2d50db289
title: تنفيذ محول صيغة Spinal Case
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

صيغة Spinal Case هي تنسيق سلسلة نصية حيث تكون كل الكلمات بأحرف صغيرة ومفصولة بشرطات. `"this-is-spinal-tap"` مثال على سلسلة نصية بصيغة spinal case.

في هذا المختبر، ستنشئ دالة تحول سلسلة نصية معطاة إلى صيغة spinal case.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `spinalCase`.
2. يجب أن تأخذ دالة `spinalCase` معلمة واحدة، وهي سلسلة نصية.
3. يجب أن تُرجع دالة `spinalCase` السلسلة النصية بصيغة spinal case. على سبيل المثال، إذا كانت المعلمة هي `ProductLanding page`، يجب أن تُرجع الدالة `product-landing-page`.

# --hints--

يجب أن تنشئ دالة باسم `spinalCase`.

```js
assert.isFunction(spinalCase);
```

يجب أن تأخذ دالة `spinalCase` معلمة واحدة.

```js
assert.lengthOf(spinalCase, 1);
```

يجب أن تُرجع `spinalCase("This Is Spinal Tap")` السلسلة النصية `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

يجب أن تُرجع `spinalCase("thisIsSpinalTap")` السلسلة النصية `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

يجب أن تُرجع `spinalCase("The_Andy_Griffith_Show")` السلسلة النصية `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

يجب أن تُرجع `spinalCase("Teletubbies say Eh-oh")` السلسلة النصية `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

يجب أن تُرجع `spinalCase("AllThe-small Things")` السلسلة النصية `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
