---
id: 6a3e5ccb99fd6053af9032a1
title: 第 6 步
challengeType: 20
dashedName: step-6
---

# --description--

既然你已经知道厨房里有什么了，你就可以用这些食材来做饭。为此，创建一个接受参数的新函数。把参数看作是你调用函数时传入数据的占位符。

这是一个接受两个参数 `a` 和 `b` 的函数，并将它们的和打印到终端。

```py
def calculate_sum(a, b):
    print(a + b)

calculate_sum(3, 1) # Output: 4
```

在你的 `check_kitchen_stock` 函数下面，定义一个名为 `use_eggs` 的空函数，它接受两个参数，顺序分别是 `available_eggs` 和 `eggs_to_use`。

注意，如果你给参数起了和全局变量（比如 `available_eggs`）完全相同的名字，函数会使用它自己的局部参数 `available_eggs`，而不是全局变量。这被称为名称遮蔽，因为局部变量“遮蔽”了全局变量。

暂时在函数体内使用 `pass` 关键字。

# --hints--

你应该定义一个名为 `use_eggs` 的函数。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("use_eggs")
`)
})
```

你的 `use_eggs` 函数应该有 `available_eggs` 和 `eggs_to_use` 作为它的参数。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

--fcc-editable-region--

--fcc-editable-region--
```
