---
id: 6a3e71a0a3a44260e675b157
title: 第 12 步
challengeType: 20
dashedName: step-12
---

# --description--

为了解开为什么你的鸡蛋数量没有变化并且实际上做了早餐的谜团，创建一个新函数。把这个函数想象成你的食谱，它会调用其他函数来完成任务。

在你的 `print(available_eggs)` 语句下面，定义一个名为 `make_fried_egg` 的空函数，它以 `available_eggs` 作为唯一参数。

暂时在函数体内使用 `pass` 关键字。

# --hints--

你应该定义一个名为 `make_fried_egg` 的函数。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("make_fried_egg")
`)
})
```

你的 `make_fried_egg` 函数应该以 `available_eggs` 作为参数。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").has_args("available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```
