---
id: 6a3e7284e89ce3367e709331
title: 第 13 步
challengeType: 20
dashedName: step-13
---

# --description--

在煎蛋之前，你首先需要查看是否有必需的食材。你可以通过比较可用的鸡蛋数量和食谱所需的数量来完成这一步。

在你的 `make_fried_egg` 函数中，将 `pass` 关键字替换为名为 `has_enough_eggs` 的变量。为它赋值，结果是查看 `available_eggs` 是否大于或相等于 `1`。

# --hints--

你应该将 `has_enough_eggs` 赋值为查看 `available_eggs` 是否大于或相等于 `1` 的结果。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0]
assert node.is_equivalent("has_enough_eggs = available_eggs >= 1") or node.is_equivalent("has_enough_eggs = 1 <= available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
