---
id: 6a3ea8a55317ba7a08e21d5b
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

现在你终于可以回答之前的问题了：为什么全局的鸡蛋数量保持在 `1`？当一个函数返回一个值时，它不会自动覆盖你现有的变量。你必须显式地储存那个新值！

你在 `make_fried_egg` 函数内部使用了 `available_eggs` 变量，但它不是代码顶部定义的那个变量。函数内部的那个是一个局部变量，只在函数运行时存在。函数结束后，局部变量被丢弃，全局的 `available_eggs` 保持不变。

要更新你实际的厨房库存，就用 `make_fried_egg` 函数返回的新鸡蛋数量覆盖全局的 `available_eggs` 变量。

在代码的最底部，调用 `make_fried_egg` 函数，传入 `available_eggs` 作为参数，并将结果赋值回全局的 `available_eggs` 变量。

# --hints--

你应该将 `available_eggs` 变量重新赋值为 `make_fried_egg(available_eggs)` 的结果。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
