---
id: 6a3e524fc3394ef52ed5b797
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

函数有它们自己的局部作用域。这意味着在函数内部创建的变量就像存放在那个厨房储藏室里的食材：它们只能在那个特定函数内部访问。

然而，函数并不是完全孤立的。它们可以读取定义在其作用域之外的全局变量。它们可以使用这些变量来执行计算、做出决策，或者创建新的局部变量。

在你的 `check_kitchen_stock` 函数内部，将 `pass` 关键字替换为一个名为 `total_items` 的新变量，并将 `available_eggs`、`available_flour` 和 `available_sugar` 的和赋值给它。

# --hints--

你应该在你的 `check_kitchen_stock` 函数内部定义一个名为 `total_items` 的变量。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_variable("total_items")
`)
})
```

你的 `total_items` 变量应该是 `available_eggs`、`available_flour` 和 `available_sugar` 的总和。

```js
({
  test: () => runPython(`
variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
assert any(
  variable.is_equivalent(code)
  for code in [
    "total_items = available_eggs + available_flour + available_sugar",
    "total_items = available_flour + available_eggs + available_sugar",
    "total_items = available_sugar + available_flour + available_eggs",
    "total_items = available_eggs + available_sugar + available_flour",
    "total_items = available_flour + available_sugar + available_eggs",
    "total_items = available_sugar + available_eggs + available_flour"
  ]
)
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--
    pass
--fcc-editable-region--
```
