---
id: 6a3e5475c35dbc138d77c6f4
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

既然你已经计算出了总数，就应该显示库存。回想之前的课程，你可以使用 `print()` 函数和 **f-strings** 来在输出中包含变量。

在你的 `check_kitchen_stock` 函数中，位于 `total_items` 变量下面，添加这四条完全相同的打印语句：

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

你的函数中应该包含 `print(f'The kitchen has {total_items} total items:')` 语句。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'The kitchen has {total_items} total items:')")
`)
})
```

`print(f'The kitchen has {total_items} total items:')` 语句应该是你函数中第一个调用的 `print()`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").find_calls("print")[0].is_equivalent("print(f'The kitchen has {total_items} total items:')")
`)
})
```

你的函数中应该包含 `print(f'- {available_eggs} eggs')` 语句。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_eggs} eggs')")
`)
})
```

你的函数中应该包含 `print(f'- {available_flour} flour')` 语句。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_flour} flour')")
`)
})
```

你的函数中应该包含 `print(f'- {available_sugar} sugar')` 语句。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("check_kitchen_stock").has_call("print(f'- {available_sugar} sugar')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
--fcc-editable-region--
    
--fcc-editable-region--
```
