---
id: 6a3e2b9b7eeca084dd4fb857
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

欢迎来到厨房！就像厨师依赖食谱来准备一顿饭一样，你将依赖函数来构建软件。在本次工作坊中，你将学习如何编写可重用的 Python 函数并掌握作用域，准确理解在任何给定时间你的代码中有哪些原料（变量）可用。

要开始，你应该为你的厨房库存定义一些基本数据。创建三个变量：`available_eggs`、`available_flour` 和 `available_sugar`，并分别赋值为 `1`、`2` 和 `3`。

# --hints--

你应该定义一个名为 `available_eggs` 的变量，并赋值为 `1`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")
`)
})
```

你应该定义一个名为 `available_flour` 的变量，并赋值为 `2`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")
`)
})
```

你应该定义一个名为 `available_sugar` 的变量，并赋值为 `3`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")
`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
