---
id: 6a3eaa2caac22f67811fe40f
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

真相时刻到了！你已经成功用 `make_fried_egg` 函数返回的新值更新了全局的 `available_eggs` 变量。

最后再查看一次你的厨房库存，确认鸡蛋的数量确实减少了。

在代码的最后，再次调用 `check_kitchen_stock()` 函数。运行代码时，你应该会看到厨房里现在有 `0` 个鸡蛋！

恭喜你成功完成了厨房库存追踪工作坊！

# --hints--

你应该在代码末尾第二次调用 `check_kitchen_stock()` 函数。

```js
({
  test: () => runPython(`
node = _Node(_code)
assert len(node.find_calls("check_kitchen_stock")) == 2
assert node.is_ordered("available_eggs = make_fried_egg(available_eggs)", "check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

available_eggs = make_fried_egg(available_eggs)

check_kitchen_stock()
```
