---
id: 6a3e6ad6b9554e31dcc30cd6
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

如果你没有足够的鸡蛋，就不应该尝试做饭！你需要一种方法来停止函数继续运行。

你可以使用 `return` 关键字来实现这一点。当函数遇到 `return` 语句时，它会立即停止执行，并将返回值传回调用它的地方。

例如：

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

在你的 `if` 块中，打印语句下面，使用 `return` 关键字返回 `available_eggs` 的值。

# --hints--

你应该在 `if` 语句内返回 `available_eggs`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")
`)
})
```

你应该在 `if` 块中 `print` 语句之后返回 `available_eggs`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs").find_ifs()[0]
assert node.is_ordered("print('The kitchen does not have enough eggs.')", "return available_eggs")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
