---
id: 6a3e70b94f67d9ee3ecd8486
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

你刚刚调用了 `use_eggs(available_eggs, 1)` 来使用一个鸡蛋。这意味着你厨房里现在应该剩下 `0` 个鸡蛋，对吧？

你马上就会知道！在你的代码底部，使用 `print()` 函数打印 `available_eggs` 变量的值。

当你运行代码时，你会发现一个令人惊讶的现象：输出仍然是 `1`！为什么鸡蛋的数量没有变化？你将在接下来的几个步骤中探索这个谜团。

# --hints--

你应该在代码的末尾打印 `available_eggs` 变量。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("print(available_eggs)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
--fcc-editable-region--

--fcc-editable-region--
```
