---
id: 6a3e73eabddd614f91cc235a
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在你要解决之前的谜题了！当你的 `use_eggs` 函数返回剩余的鸡蛋数量时，你没有把这个新值储存到任何地方。它就这样消失在虚空中了！

要真正更新你的库存，你需要通过赋值将返回的值捕捉到一个变量中。

在你的 `make_fried_egg` 函数内，使用一个 `if/else` 语句来查看 `has_enough_eggs` 变量。

- 如果 `has_enough_eggs` 是 `True`，将你的 `available_eggs` 变量重新赋值为 `use_eggs(available_eggs, 1)` 的结果，并打印消息 `Made a fried egg. Yummy!`。

- 否则，在一个 `else` 块中打印消息 `Could not make a fried egg. Not enough eggs!`。

# --hints--

你应该有一个 `if` 语句，并使用 `has_enough_eggs` 作为它的条件。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")
`)
})
```

在你的 `if` 块内，你应该将 `available_eggs` 重新赋值为 `use_eggs(available_eggs, 1)` 的结果。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")
`)
})
```

当 `has_enough_eggs` 是 `True` 时，你应该打印消息 `Made a fried egg. Yummy!`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")
`)
})
```

你应该有一个 `else` 块，当 `has_enough_eggs` 是 `False` 时打印消息 `Could not make a fried egg. Not enough eggs!`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("make_fried_egg").find_if("has_enough_eggs")
conditions = node.find_conditions()
assert len(conditions) == 2 and conditions[1].is_empty()
assert node.find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)

def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1

--fcc-editable-region--
    
--fcc-editable-region--
```
