---
id: 6a3e6c0eb5d050ce07781709
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

如果 `if` 条件计算结果为 `False`（意味着你确实有足够的鸡蛋），函数将跳过 `if` 块并继续运行剩余的代码。

现在你需要处理你的数据并返回一个新的计算值。

在你的 `if` 块外（与 `if` 关键字处于相同的缩排级别），使用 `print()` 调用和 f-string 打印 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`。

然后，在下一行，通过从 `available_eggs` 中减去 `eggs_to_use` 来 `return` 剩余的鸡蛋数量。

# --hints--

你的 `use_eggs` 函数中应该有一个打印语句，打印 `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")
`)
})
```

你应该通过从 `available_eggs` 中减去 `eggs_to_use` 来返回剩余的鸡蛋数量。

```js
({
  test: () => runPython(`
assert _Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")
`)
})
```

你应该在 `print` 语句之后返回 `available_eggs - eggs_to_use`。

```js
({
  test: () => runPython(`
node = _Node(_code).find_function("use_eggs")
assert node.is_ordered("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')", "return available_eggs - eggs_to_use")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

--fcc-editable-region--
    
--fcc-editable-region--
```
