---
id: 6a3e58891e1ff13d07f4ab48
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

Python 会读取 `def` 块并记住该函数，但不会立即运行其中的代码。

要执行代码并查看输出，你需要**调用**（或调用）该函数。方法是输入函数名，后面跟上括号。

例如：

```py
my_function()
```

在代码底部、函数外部，调用你的 `check_kitchen_stock` 函数。

# --hints--

你应该在代码末尾调用 `check_kitchen_stock` 函数。确保它没有缩排，这样它才能在全局范围内运行。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("check_kitchen_stock()")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

--fcc-editable-region--

--fcc-editable-region--
```
