---
id: 6a3e6fd08125dfa615cb431b
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

现在你的 `use_eggs` 函数已经完成，你可以测试它了。当调用需要**参数**的函数时，你需要在括号内传入具体的值。这些值称为**参数**。

下面是一个如何带参数调用函数的示例：

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```

在这个例子中，前两行是函数定义，其中 `a` 和 `b` 是参数。最后一行 `calculate_sum(3, 1)` 是函数调用。值 `3` 和 `1` 是参数。它们按位置与参数匹配，所以 `3` 被赋值给 `a`，`1` 被赋值给 `b`。

在你的代码底部，调用你的 `use_eggs` 函数，并传入两个参数：全局的 `available_eggs` 变量和数字 `1`（表示你想使用的鸡蛋数量）。

# --hints--

你应该调用 `use_eggs` 函数，并传入 `available_eggs` 和 `1` 作为参数。

```js
({
  test: () => runPython(`
assert _Node(_code).has_call("use_eggs(available_eggs, 1)")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

--fcc-editable-region--

--fcc-editable-region--
```
