---
id: 6a3e4ae26f2b7e4b5b026cf8
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

现在是创建一个函数来查看你厨房里有什么的时候了。在 Python 中，你使用 `def` 关键字定义函数，后面跟着函数名、括号 `()` 和冒号 `:`。

下面是定义一个简单函数的示例：

```py
def my_function():
    pass
```

定义一个名为 `check_kitchen_stock()` 的空函数。暂时在函数体内使用 `pass` 关键字，这样 Python 就不会抛出语法误差。

记住，Python 依赖缩排来定义代码块。所以一定要缩排函数体。

# --hints--

你应该定义一个名为 `check_kitchen_stock` 的函数。

```js
({
  test: () => runPython(`
assert _Node(_code).has_function("check_kitchen_stock")
`)
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
