---
id: 66dadcf18df3a76104054d95
title: Побудуйте клас Linked List
challengeType: 14
dashedName: build-a-linked-list-class
---

# --description--

In this lab, you'll build a linked list, which is a linear collection of data elements, called <dfn>nodes</dfn>, where each node points to the next node in the linked list. Кожен вузол у зв’язаному списку містить два ключові фрагменти інформації: сам елемент і посилання на наступний вузол.

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. You should create a class named `LinkedList` to represent a linked list data structure.

1. Each time a new `LinkedList` object is instantiated, its `length` and `head` properties should be set to `0` and `null`, respectively.

1. The `LinkedList` class should have a property `Node`, which is a class instantiated with an `element` property set to the value passed at the instantiation, and a `next` property set to `null`.

1. The `LinkedList` class should have a method `isEmpty` that returns `true` if the linked list is empty, and `false` otherwise.

1. The `LinkedList` class should have a method `add` that creates a `Node` object from its argument and adds it to the linked list.

1. The `add` method should either set the head to the new node when the node to be added is the first, or set the `next` property of the last node in the linked list to the node to be added.

1. Each time a new node is added to the list the length of the list should be incremented.

1. The `LinkedList` class should have a method `remove` that takes a single argument and removes the corresponding node from the list.

1. Whenever a node is removed from the linked list, the `next` property of the previous node should be set to the following node, and the length of the list should be decremented.

# --hints--

Your `LinkedList` class should instantiate objects with the `length` property set to `0`.

```js
const test = new LinkedList();
assert.strictEqual(test.length, 0);
```

Your `LinkedList` class should instantiate objects with the `head` property set to `null`.

```js
const test = new LinkedList();
assert.strictEqual(test.head, null);
```

Your `LinkedList` class should have a property `Node`.

```js
const test = new LinkedList();
assert.property(test, "Node");
```

Your `Node` property should be a class that instantiates objects with an `element` property set to the value passed at the instantiation.

```js
const test = new LinkedList();
const node = new test.Node("pig");
assert.strictEqual(node.element, "pig");
```

Your `Node` property should be a class that instantiates objects with a `next` property set to `null`.

```js
const test = new LinkedList();
const node = new test.Node("pig");
assert.strictEqual(node.next, null);
```

Клас `LinkedList` повинен мати метод `add`.

```js
const test = new LinkedList();
assert.typeOf(test.add, 'function');
```

Your `add` method should assign `head` to the first node object added to the linked list.

```js
const test = new LinkedList();
test.add("pig");
const expected = { element: "pig", next: null }
assert.deepEqual(test.head, expected);
```

Your `add` method should set the `next` property of the last node in the linked list to the node to be added.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog")
const expected = { element: "pig", next: {element: "dog", next: null} }
assert.deepEqual(test.head, expected);
```

Each time a new node is added to the list the `length` of the list should be incremented.

```js
const test = new LinkedList();
test.add("pig");
assert.strictEqual(test.length, 1);
test.add("dog");
assert.strictEqual(test.length, 2);
test.add("sheep");
assert.strictEqual(test.length, 3);
```

Клас `LinkedList` повинен мати метод `isEmpty`.

```js
const test = new LinkedList();
assert.typeOf(test.isEmpty, 'function');
```

Метод `isEmpty` має повернути `false`, якщо у зв’язаному списку принаймні один елемент.

```js
const test = new LinkedList();
test.add('pig');
test.add('dog');
test.add('kitten');
assert.isFalse(test.isEmpty());
```

Метод `isEmpty` має повернути `true`, якщо у зв’язаному списку немає елементів.

```js
const test = new LinkedList();
assert.isTrue(test.isEmpty());
```

Попередній вузол у класі `LinkedList` має посилатися на найновіший створений вузол.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog")
test.add("sheep");
assert(test.head.next.element === 'dog' && test.head.next.next.element === "sheep");
```

Клас `LinkedList` повинен мати метод `remove`.

```js
const test = new LinkedList();
assert.typeOf(test.remove, 'function');
```

Метод `remove` має перепризначити `head` до другого вузла, якщо перший вузол видалено.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.remove("pig");
const expected = { element: "dog", next: {element: "sheep", next: null} }
assert.deepEqual(test.head, expected);
```

Метод `remove` має зменшувати `length`(довжину) зв’язаного списку на одиницю після кожного видалення вузла.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.remove("pig");
assert.strictEqual(test.length, 2);
test.remove("dog");
assert.strictEqual(test.length, 1);
test.remove("sheep");
assert.strictEqual(test.length, 0);
```

Метод `remove` має перепризначити посилання вузла, який був перед видаленим вузлом, до посилання `next` видаленого вузла.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.add("cat");
test.remove("sheep");
const expected = { element: "pig", next: {element: "dog", next: {element: "cat", next: null}} }
assert.deepEqual(test.head, expected);
```

Метод `remove` не повинен змінювати зв’язаний список, якщо заданого елемента не існує в цьому списку.

```js
const test = new LinkedList();
test.add("pig");
test.add("dog");
test.add("sheep");
test.add("cat");
test.remove("elephant");
const expected = { element: "pig", next: {element: "dog", next: {element: "sheep", next: {element: "cat", next: null}}} }
assert.deepEqual(test.head, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class LinkedList {
    constructor() {
        this.length = 0;
        this.head = null;
    }

    Node = class {
        constructor(element) {
            this.element = element;
            this.next = null;
        }
    }

    isEmpty() {
        return this.length === 0;
    }

    add(element) {
        const node = new this.Node(element)
        if (this.isEmpty()) {
            this.head = node;
        } else {
            let currentNode = this.head;
            while (currentNode.next !== null) {
                currentNode = currentNode.next;
            }
            currentNode.next = node;
        }
        this.length++;
    }

    remove(element) {
        if (this.isEmpty()) {
            return
        }
        let previousNode;
        let currentNode = this.head;
        while (currentNode.next !== null && currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }
        if (currentNode.next === null && currentNode.element !== element) {
            return;
        } else if (previousNode) {
            previousNode.next = currentNode.next
        } else {
            this.head = currentNode.next;
        }
        this.length--
    }
}
```
