---
id: 65c4eb814cc977c95cd7df0e
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 이진 탐색 트리(BST)를 만들 것입니다. BST는 각 노드가 최대 두 개의 자식을 가지며, 왼쪽 자식은 부모 노드보다 작은 값을, 오른쪽 자식은 부모 노드보다 큰 값을 포함하여 효율적인 검색과 정렬 작업이 가능한 자료 구조입니다.

이것이 이진 탐색 트리의 모습입니다:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

먼저 빈 `TreeNode` 클래스를 정의하세요. `TreeNode` 클래스는 이진 탐색 트리의 노드를 나타냅니다. `pass` 키워드를 사용해 클래스 본문을 채워 오류를 피하세요.

# --hints--

빈 `class`라는 이름의 클래스를 선언하려면 `TreeNode` 키워드를 사용하세요.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
