---
id: 65c4f0d32ca17ad4b1636b0e
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`__init__` 메서드 안에서 `pass`을 삭제하고 `root`를 값 `None`으로 초기화하세요.

`root` 속성은 이진 탐색 트리의 루트 노드를 나타냅니다. 새 `BinarySearchTree` 객체가 생성될 때 생성자이므로 빈 트리로 시작해서 `root` 속성은 `None`로 설정됩니다.

# --hints--

`pass` 메서드에서 `__init__` 문장을 제거해야 합니다.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

`root`를 사용해 `None` 속성을 `self.root`로 초기화해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
