---
id: 65c4f195de7e2ad5932be717
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

다음으로, 트리에 노드를 삽입하는 메커니즘을 정의해야 합니다. 이를 위해 `_insert` 메서드를 정의해야 하는데, 이 메서드는 도우미 함수이며 나중에 실제 `insert` 메서드에서 사용됩니다.

이 메서드는 재귀적이며, 적절한 새 노드 위치를 찾을 때까지 트리를 탐색하기 위해 자기 자신을 호출합니다.

`_insert`, `self`, `node` 매개변수를 가진 `key` 메서드를 정의하세요.

# --hints--

`_insert` 클래스 내에 `BinarySearchTree` 메서드를 정의해야 합니다. `pass`를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

`_insert` 메서드는 세 개의 매개변수 `self`, `node`, `key`를 받아야 합니다.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
