---
id: 65c4f013851cefd1a4fe4c96
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` 메서드는 두 개의 매개변수를 받습니다: `self` (생성 중인 클래스 인스턴스를 나타냄)와 `key` (노드에 저장할 값). 이 두 매개변수를 `__init__()` 메서드에 추가하세요.

# --hints--

메서드에 `self` 매개변수를 추가해야 합니다.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

메서드에 `key` 매개변수를 추가해야 합니다.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
