---
id: 65c4f22498d22ed775ef8efb
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

이제 `node` 매개변수가 `None`인지 확인해야 합니다. 만약 그렇다면, 이 메서드는 트리에서 새 노드를 삽입해야 하는 리프 노드 또는 빈 자리에 도달한 것입니다.

`_insert` 메서드 본문 안에서 `pass`을 `if`인지 확인하는 `node is None` 문으로 바꾸세요.

새로운 `if` 블록 안에서는 제공된 키로 새 `TreeNode(key)` 인스턴스를 생성하기 위해 `TreeNode`을 반환하세요. 이렇게 하면 새 리프 노드가 되어 키가 트리에 삽입됩니다.

# --hints--

`pass` 메서드 안에 `_insert`이 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

`if`인지 확인하는 `node is None` 문을 작성해야 합니다.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

`TreeNode(key)` 블록에서 `if`을 반환해야 합니다.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
