---
id: 65c4ef8463f869d0eaf87c67
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`TreeNode` 클래스 안에서 `pass`을 `__init__` 메서드로 바꿔서 객체의 속성을 초기화할 수 있게 하세요. 지금은 매개변수를 추가하지 마세요.

# --hints--

`pass` 클래스에서 `TreeNode` 키워드를 제거하고 `__init__` 메서드 안으로 옮기세요.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

`__init__` 클래스 안에 `TreeNode` 메서드를 정의하세요. 내부에 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
