---
id: 65c4f37ba6cbcfdc77e24165
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

사용자가 `insert` 메서드를 호출할 것입니다. `self` 매개변수 외에 `key` 매개변수도 필요합니다. 이 매개변수는 이진 탐색 트리에 삽입할 키 값이 될 것입니다.

함수 정의에 `key`를 두 번째 매개변수로 추가하세요.

# --hints--

`insert` 메서드는 `self` 매개변수를 포함해야 합니다.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

`insert` 메서드는 `key` 매개변수를 포함해야 합니다.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
