---
id: 65c4f33bc3c3d8db7f732787
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

이제 실제 삽입을 수행하려면 `insert` 클래스 내에 빈 `BinarySearchTree` 메서드를 정의하고 `self` 매개변수를 추가하세요.

# --hints--

`insert` 클래스 내에 `self` 매개변수를 가진 `BinarySearchTree` 메서드를 정의해야 합니다. `pass` 키워드를 기억하세요.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
