---
id: 65c4f07da204f4d2f325fbd3
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`__init__` 메서드 안에서 노드의 `left`와 `right` 속성을 `None`로 초기화하세요. 노드가 처음 생성될 때는 왼쪽이나 오른쪽 자식이 없기 때문입니다. `self` 키워드를 사용하는 것을 잊지 마세요.

# --hints--

`None`를 사용해 노드의 `left` 속성에 `self.left`를 할당해야 합니다.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

`None`를 사용해 노드의 `right` 속성에 `self.right`를 할당해야 합니다.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
