---
id: 65c4f09e074dd8d37830ea00
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이진 탐색 트리를 나타내는 `BinarySearchTree`라는 빈 클래스를 하나 더 만드세요.

# --hints--

`BinarySearchTree`라는 이름의 클래스가 있어야 합니다. `pass` 키워드를 사용하는 것을 기억하세요.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
