---
id: 65c4f2d9fd872fd99ac659dd
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`elif`이 `key`보다 큰지 확인하는 `node.key` 조건문을 추가하세요.

# --hints--

`elif`인지 확인하는 `key > node.key` 조건문을 추가해야 합니다. `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[1].is_equivalent("key > node.key")`))})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)

--fcc-editable-region--    

```
