---
id: 65c4f2851645e8d84d14f5d1
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 재귀적으로 트리를 순회하며 이진 트리 원칙에 따라 값을 삽입해야 합니다:

- 키보다 작은 값은 왼쪽 부분 트리에 배치합니다
- 키보다 큰 값은 오른쪽 부분 트리에 배치합니다

기존 if 문 다음에 `if` 문을 하나 더 작성하여 `key`이 `node.key`보다 작은지 확인하세요.

# --hints--

`key`이 `node.key`보다 작은지 확인하는 if 문을 하나 더 작성해야 합니다. `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```
