---
id: 65c4f300da28d8da361bfa93
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

`elif` 절 안에서, `pass`을 현재 노드의 오른쪽 자식을 첫 번째 인수로, `_insert`을 두 번째 인수로 하여 `key` 메서드를 호출하는 코드로 바꾸세요. 그 결과를 현재 노드의 `right` 속성에 할당하세요.

# --hints--

`pass` 블록에서 `elif` 키워드를 제거하세요.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

`self._insert()` 메서드를 호출할 때 `node.right`과 `key`를 인수로 전달하세요.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

`self._insert()` 호출 결과를 현재 노드의 올바른 속성(`node.right`)에 할당하세요.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
