---
id: 65c4f02dbd3557d23f12db62
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

`__init__` 메서드 안에서 `pass`을 삭제하고 `key` 매개변수의 값을 `key`를 사용해 노드의 `self.key` 속성에 할당하세요.

즉, `key` 인스턴스의 `TreeNode` 속성이 객체 생성 시 전달된 값으로 설정됩니다.

# --hints--

`pass` 메서드에서 `__init__` 문장을 삭제해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

`key`를 사용해 `key` 매개변수의 값을 노드의 `self.key` 속성에 할당해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
