---
id: 65c4f2b7178afed8e88f782f
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

만약 `key` < `node.key`가 `True`로 평가되면, 새 노드는 왼쪽 부분 트리에 배치되어야 합니다.

`pass`을 삭제하고 왼쪽 자식을 첫 번째 인수로, `_insert`를 두 번째 인수로 하여 `key` 메서드를 재귀적으로 호출하세요. 그 결과를 현재 노드의 `left` 속성에 할당하세요.

# --hints--

`pass` 블록에서 `if` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

`self._insert` 메서드를 호출할 때 `node.left`과 `key`을 인수로 전달해야 합니다.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

`self._insert()` 호출 결과를 현재 노드의 `left` 속성에 할당해야 합니다.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
