---
id: 65c4f0b82db314d3fcc0b8e1
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`BinarySearchTree` 클래스 내에서 `pass`을 `__init__` 메서드로 바꾸고 이 메서드에 `self` 매개변수를 추가하세요.

# --hints--

`pass` 클래스에서 `BinarySearchTree` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


`__init__` 메서드를 정의하고 이 메서드에 `self` 매개변수를 추가해야 합니다. 메서드 본문 안에서 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
