---
id: 65c4f3258d2e4cdacc919dfd
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`_insert` 메서드의 마지막에, 삽입 과정이 완료된 후 현재 노드를 반환하여 재귀 호출 스택의 상위 단계에서 트리 구조를 업데이트하세요.


# --hints--

조건문 블록 밖에서 현재 노드를 반환해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```
