---
id: 65c645b838e7deb080fc25e0
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

이제 노드가 삽입되었습니다. 올바르게 삽입되었는지 확인하려면 트리에서 노드를 검색할 수 있습니다.

for 루프 밖에서 `80` 인스턴스에서 노드 `bst`를 검색하고 `print` 호출에 추가하세요. 또한 `print` 함수의 첫 번째 인수는 `'Search for 80:'`여야 합니다.

# --hints--

`bst.search(80)`를 호출한 결과를 출력해야 하며 `print` 문은 첫 번째 인수로 `'Search for 80:'`를 가져야 합니다.

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
