---
id: 65c63fd9b387f0a6c5bf6a72
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

`search` 메서드 안에서 `pass`을 삭제하고 다음 인수들과 함께 헬퍼 메서드 `_search`를 호출하세요.

- `self.root`: 이진 탐색 트리의 루트입니다. 탐색은 루트부터 시작합니다.
- `key`: 사용자가 이진 탐색 트리에서 찾고자 하는 값입니다.

내부적으로 `search`는 실제 탐색 로직을 이진 탐색 트리 내에서 재귀 탐색을 수행하는 헬퍼 메서드 `_search`에 위임합니다.

# --hints--

`pass` 메서드에서 `search` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

`_search` 메서드 내에서 `search` 메서드를 호출해야 합니다.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

`_search`와 `self.root`을 인수로 전달하여 `key` 메서드를 호출해야 합니다.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
