---
id: 65c63e6962159fa372ecdbec
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

두 번째 `if` 문장이 `True`이 아니라면, 대상 키가 현재 노드 키보다 크거나 같다는 뜻입니다.

이진 탐색 트리에서 대상 키가 현재 노드 키보다 크면, 검색은 오른쪽 부분 트리에서 계속됩니다.

`if` 블록 뒤에 현재 노드의 오른쪽 자식과 키를 인수로 하여 `_search` 메서드를 호출한 결과를 반환하세요.

# --hints--

`_search`를 사용하여 `self._search()` 메서드를 재귀적으로 호출해야 합니다.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

`_search` 메서드를 호출할 때 두 개의 인수, `node.right`과 `key`을 전달해야 합니다.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

`_search` 메서드 호출 결과를 반환해야 합니다.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
