---
id: 65c646d4148ae3b2d1cbcac4
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

검색 결과가 `80: <__main__.TreeNode object at 0x108b3e0>`처럼 나오는 것을 확인하세요. 이것은 클래스 인스턴스를 출력할 때 기본 문자열 표현입니다.

유용한 값을 출력하도록 변경하려면 `__str__` 클래스에 `TreeNode`이라는 이름의 다른 메서드를 정의하세요. 이 메서드는 단일 인수 `self`를 받습니다.

`__str__`를 정의한 후에는 `__str__` 메서드가 아직 아무것도 반환하지 않기 때문에 콘솔에서 예외가 발생할 것입니다. 다음 단계에서 메서드 본문을 작업할 예정입니다.

# --hints--

단일 인수 `__str__`를 받는 `self` 메서드를 정의해야 합니다. `pass`를 사용하는 것을 잊지 마세요.

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
