---
id: 65c63eb01b9563a4b8a046b6
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

다음으로, `search` 클래스 안에 `BinarySearchTree`라는 메서드를 정의하세요. `search` 메서드에 `self`와 `key` 두 개의 매개변수를 주세요.

# --hints--

`search`라는 메서드를 정의해야 합니다. `pass` 키워드를 기억하세요.

```js
({ test: () => assert.match(code, /def\s+search/) });
```

`search` 메서드는 `self`와 `key` 두 개의 매개변수를 받아야 합니다.

```js
({ test: () => assert.match(code, /def\s+search\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
        
--fcc-editable-region--

--fcc-editable-region--
```
