---
id: 65c6409418806da8d0636ffc
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

`insert`와 `search` 기능이 완성되었으니, 이제 테스트할 시간입니다.

이 시점에서는 노드가 정렬되지 않고 단순히 삽입만 되었다는 점을 참고하세요. 나중에 중위 순회를 사용해 정렬 작업을 진행할 것입니다.

Python에서 클래스를 인스턴스화하는 방법은 다음과 같습니다:

```python
object_name = ClassName()
```

클래스 정의 외부에서 `BinarySearchTree` 클래스를 인스턴스화하고 변수 `bst`에 할당하세요.


# --hints--

`BinarySearchTree` 클래스의 인스턴스를 생성해야 합니다.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

새로운 `BinarySearchTree` 인스턴스를 변수 `bst`에 할당해야 합니다.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
