---
id: 65c63df529bd15a24c187c62
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

현재 노드의 키보다 대상 키가 더 작은지 확인하는 또 다른 `if` 문을 작성하세요.

`if` 블록 안에서, 현재 노드의 왼쪽 자식과 `_search`를 인수로 하여 `key` 메서드를 호출한 결과를 반환하세요.

# --hints--

현재 노드의 키 `if`보다 `key`가 작은지 확인하는 또 다른 `node.key` 문을 작성해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

새로운 `self._search(node.left, key)` 블록에서 `if`을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
