---
id: 65c63a4c4da62e9ae18e321a
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

검색 기능을 구현할 시간입니다. `_insert` 메서드에 대해 `insert`라는 도우미 메서드를 만든 것처럼, `_search` 메서드에 대해 `search`라는 도우미 메서드를 만들어야 합니다.

세 개의 매개변수 `_search`, `self`, `node`를 가진 `key` 메서드를 선언하세요.

# --hints--

`_search` 메서드를 선언해야 합니다. 앞에 밑줄이 붙은 것을 주의하고 `pass`를 사용하세요.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

`_search` 메서드는 세 개의 매개변수 `self`, `node`, `key`를 받아야 합니다.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
