---
id: 65c63da2ed6769a10e141341
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

이제 재귀 탐색의 기본 사례를 정의할 차례입니다. 현재 `pass`를 제거하고 두 가지 조건을 확인하는 `if` 문을 작성하세요:

- `node`가 `None`인 경우: 이는 탐색이 키를 찾지 못하고 가지의 끝에 도달했음을 나타냅니다.
- `node.key == key`인 경우: 이는 현재 노드에서 키를 찾았음을 의미합니다.

두 조건을 `or` 연산자로 결합하고 `if` 블록 안에서 현재 노드를 반환하세요.

# --hints--

`pass` 메서드에서 `_search` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

`if`이 `node`인지 또는 `None`가 `node.key`과 같은지 확인하는 `key` 문을 작성해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

`node` 블록에서 `if`를 반환해야 합니다.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
