---
id: 65c644829cfb63acf3479d09
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

리스트 `nodes`에 있는 노드들을 `bst` 인스턴스에 삽입하세요. 리스트를 반복하면서 각 노드에 `insert` 메서드를 호출하세요. 반복 변수로 `node`을 사용하세요.

# --hints--

`for node in nodes` 반복문이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_for_loops()[0].find_for_vars().is_equivalent("node") and _Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("nodes")`)) });
```

`bst.insert(node)` 반복문 안에 `for`가 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_for("node", "nodes").find_body().is_equivalent("bst.insert(node)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

--fcc-editable-region--
```
