---
id: 65c9ddd336596e30a4266a50
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

`__str__` 메서드 본문에서 `pass`을 삭제하고 `str()`을 인수로 하여 `self.key` 함수를 호출한 결과를 반환하세요. 이것은 노드에 연결된 값을 저장하는 현재 노드 객체의 속성입니다.

결과를 반환한 후에는 콘솔에서 예외가 사라지고 출력에 노드와 연결된 `key`의 값이 표시되어야 합니다.

# --hints--

`pass` 메서드에서 `__str__` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__str__").has_pass()`
      )
    );
  },
});

```

`self.key` 메서드에서 `__str__`의 문자열 값을 반환해야 합니다.

```js
({ test: () => assert.match(code, /^\s{8}return\s+str\(\s*self\.key\s*\)/m) })
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
    def __str__(self):
        pass

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
