---
id: 65c9de201959f73591b606e6
title: 32단계
challengeType: 20
dashedName: step-32
---

# --description--

보시다시피, 이제 콘솔에서 읽기 쉬운 출력 결과를 얻을 수 있습니다.

이제 `print` 호출을 주석 처리하세요.

# --hints--

`print('Search for 80:', bst.search(80))` 줄을 주석 처리해야 합니다.

```js
assert(code.match(/#\s*print\s*\(\s*'Search\s*for\s*80:\s*',\s*bst\.search\s*\(\s*80\s*\)\s*\)/));
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)


bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--

print('Search for 80:', bst.search(80))

--fcc-editable-region--

```
