---
id: 65ca0a1f27596a089b0363b9
title: 49단계
challengeType: 20
dashedName: step-49
---

# --description--

`delete` 메서드 안에서 `pass`을 삭제하고 이진 탐색 트리의 루트와 삭제할 키를 인수로 하여 헬퍼 메서드 `_delete`를 호출하세요.

또한 `_delete` 메서드에서 `self.root` 메서드의 결과를 `delete`에 할당하세요.

# --hints--

`pass` 메서드에서 `delete` 키워드를 제거해야 합니다.

```js
({ test: () => {
  assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("delete").has_pass()`))
}})
```

`delete` 메서드는 두 개의 인수 `_delete`과 `self.root`을 사용하여 `key` 메서드를 호출해야 합니다.

```js
({ test: () => assert.match(code, /self\._delete\(\s*self\.root\s*,\s*key\s*\)/) });
```

`_delete` 메서드에서 `self.root` 메서드의 결과를 `delete`에 할당해야 합니다.

```js
const after_split = code.split('def delete(self, key):')[1];
assert.match(after_split, /self\.root\s*=\s*self\._delete\(self\.root,\s*key\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node
    
    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key
        
--fcc-editable-region--
    def delete(self, key):
        pass        

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
