---
id: 65ca03bd8eb5faf24b250c56
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

다음 단계는 노드 삭제 작업입니다.
이를 위해 먼저 도우미 메서드를 정의한 다음 실제 메서드를 정의하는 동일한 방식을 따르세요.

`BinarySearchTree` 클래스 안에 `_delete`, `self`, `node` 세 개의 매개변수를 받는 `key`라는 새 도우미 메서드를 정의하세요.

# --hints--

`_delete`, `BinarySearchTree`, `self` 매개변수를 가진 `node` 클래스 안에 `key` 메서드를 정의해야 합니다. `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /def\s+_delete\(\s*self\s*,\s*node\s*,\s*key\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))

```
