---
id: 65ca07dd6fa8840491b7a5cd
title: 41단계
challengeType: 20
dashedName: step-41
---

# --description--

이전 조건 중 어느 것도 충족하지 않는다면, 노드가 왼쪽과 오른쪽 자식을 모두 가지고 있다는 뜻입니다.

후속 노드를 선택하려면 오른쪽 부분 트리에서 최솟값을 찾아야 합니다.
가장 작은 값이 현재 노드의 중위 순서 후속 노드가 됩니다.

가장 작은 값을 찾으려면 두 개의 매개변수 `_min_value`과 `self`를 받는 도우미 함수 `node`를 만드세요.

# --hints--

`_min_value`와 `self`를 매개변수로 하는 `node` 메서드를 정의하세요. `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
--fcc-editable-region--
    
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)
    
# print('Search for 80:', bst.search(80))


```
