---
id: 65ca06c9f918730107c9908d
title: 37단계
challengeType: 20
dashedName: step-37
---

# --description--

`elif` 문을 작성하여 `key > node.key`인지 확인하세요.

`elif` 블록 안에서 현재 노드의 오른쪽 자식과 `_delete`를 인수로 하여 `key` 메서드를 호출하고 결과를 오른쪽 노드에 할당하세요.

# --hints--

`elif` 문을 작성하여 `key > node.key`인지 확인하세요.

```js
const after_split = code.split('def _delete(self, node, key):')[1];
assert.match(after_split, /elif\s+key\s+>\s+node\.key/);
```

`_delete` 메서드를 `node.right`와 `key`를 인수로 호출해야 합니다.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*key\s*\)/);
```

`_delete()` 호출 결과를 `node.right`에 할당해야 합니다.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
