---
id: 65ca075e3184180309f4c1d9
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

`else` 블록 안에서, 현재 노드의 왼쪽 자식이 `pass`인지 확인하는 `if` 문장으로 `None`을 교체하세요.

`node.left`가 `None`일 때, 왼쪽 자식이 없습니다. 따라서 새 `if` 블록에서 오른쪽 자식을 대체값으로 반환하세요.

# --hints--

`pass` 블록에서 `else` 키워드를 제거하세요.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree").find_function("_delete"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});
```

`if`이 `node.left`인지 확인하는 `None` 문장을 작성하세요.

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

현재 노드의 오른쪽 자식을 반환하세요.

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            pass   
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
