---
id: 65ca05f7cba1e6fe70527534
title: 34단계
challengeType: 20
dashedName: step-34
---

# --description--

`_delete` 메서드 안에서, `pass`을 현재 `if`이 `node`인지 확인하는 `None` 문장으로 바꾸세요.

현재 노드가 `None`일 때, 삭제할 키를 찾지 못한 것입니다. 따라서 `node` 블록에서 `if`을 반환하세요.

# --hints--

`pass` 메서드에서 `_delete` 키워드를 삭제해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_delete").has_pass()`
      )
    );
  },
});

```

현재 `if`이 `node`인지 확인하는 `None` 문장을 작성해야 합니다.

```js
const after_split = code.split('def _delete(self, node, key):')[1];
assert.match(after_split, /if\s+node\s+is\s+None/);
```

`node` 블록 안에서 `if`을 반환해야 합니다.

```js
const after_split = code.split('def _delete(self, node, key):')[1];
assert.match(after_split, /return\s+node/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        pass

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))

```
