---
id: 65ca0d5adf39c410cd1177cc
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

이제 중위 순서 순회 방법을 기반으로 트리를 순회하는 작업을 하겠습니다. 중위 순서 순회는 깊이 우선 이진 트리 순회 알고리즘으로, 왼쪽 부분 트리, 현재 노드, 그리고 오른쪽 부분 트리를 방문합니다.

`_inorder_traversal` 클래스 내에 `BinarySearchTree` 메서드를 정의하고 세 개의 매개변수 `self`, `node`, `result`를 받도록 하세요. 여기서 `node`는 순회 중 현재 고려 중인 노드이고, `result`는 키들이 정렬된 순서로 추가되는 리스트입니다.

# --hints--

세 개의 매개변수 `_inorder_traversal`, `self`, `node`를 받는 `result` 메서드를 정의하세요. `pass`를 사용하는 것을 잊지 마세요.

```js
assert.match(code, /def\s+_inorder_traversal\(\s*self\s*,\s*node\s*,\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
