---
id: 65ca0e8eb2c9c215269d6a66
title: 54단계
challengeType: 20
dashedName: step-54
---

# --description--

`inorder_traversal` 메서드는 이진 탐색 트리를 중위 순서로 순회하는 역할을 합니다. 이 메서드는 노드의 키를 정렬된 순서로 반환합니다. 이를 위해 `_inorder_traversal` 보조 메서드를 사용할 것입니다.

`inorder_traversal`를 유일한 매개변수로 받는 `self` 메서드를 만드세요.

# --hints--

`inorder_traversal`를 매개변수로 받는 `self` 메서드를 정의하세요. `pass`를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert.match(code, /def\s+inorder_traversal\(\s*self\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

    def _inorder_traversal(self, node, result):
        if node:
            self._inorder_traversal(node.left, result)
            result.append(node.key)
            self._inorder_traversal(node.right, result)
--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
