---
id: 65cb45d130c97cb459439fac
title: 42단계
challengeType: 20
dashedName: step-42
---

# --description--

오른쪽 부분 트리에서 가장 작은 값을 찾으려면, 주어진 노드의 왼쪽 자식들을 순차적으로 반복하다가 부분 트리에서 가장 왼쪽(가장 작은) 노드에 도달해야 합니다.

이를 위해, `while` 루프를 작성하세요. 이 루프는 `node.left is not None`일 때 실행되며 `pass`를 `while` 블록 안으로 이동하세요. 이 조건은 왼쪽 자식이 있는지 확인합니다. 왼쪽 자식이 있는 한 루프는 계속되고 더 작은 값을 찾을 수 있습니다.

# --hints--

`pass` 메서드에서 `_min_value` 키워드를 제거해야 합니다.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_min_value").has_pass()`
      )
    );
  },
});
```

`node.left is not None` 루프에서 조건 `while`를 사용해야 합니다.

```js
({ test: () => assert.match(code, /while\s+node\.left\s+is\s+not\s+None\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left  

--fcc-editable-region--
    def _min_value(self, node):
        pass

--fcc-editable-region--   

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
