---
id: 65d8a6fcb15a3a239ba35dfd
title: 46단계
challengeType: 20
dashedName: step-46
---

# --description--

최솟값을 찾은 후에는 오른쪽 부분 트리에서 최솟값을 가진 노드를 재귀적으로 삭제해야 합니다.

이 단계는 이진 탐색 트리(BST) 속성을 유지하면서 최솟값을 가진 노드를 트리에서 제거하도록 합니다.

`_delete` 메서드를 `node.right`과 `node.key`를 인수로 하여 재귀적으로 호출하세요. `_delete()` 호출의 반환값을 현재 노드의 오른쪽 자식에 할당하세요.

# --hints--

`_delete` 메서드를 `node.right`와 `node.key`를 인수로 하여 재귀적으로 호출해야 합니다.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

`_delete()` 호출의 반환값을 현재 노드의 오른쪽 자식에 할당해야 합니다.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
