---
id: 680b3ef395479b0e449ecb6e
title: 選択ソートアルゴリズムを実装する
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

選択ソートは、多くのコンピューターサイエンスのコースで教えられる人気のあるソートアルゴリズムの一つです。

このアルゴリズムは、リストの未ソート部分から最小の要素を繰り返し見つけて、それを最初の未ソート要素と交換することで動作します。まずリスト全体の最小値を選択し、それを最初の要素と交換します。次に2番目の位置に移動し、残りの未ソート要素の中で最小の値を見つけて2番目の要素と交換します。このプロセスをリストの要素を1つずつ順に進めながら繰り返し、リスト全体がソートされるまで続けます。

選択ソートは、最良、平均、最悪のいずれのケースでも二次の時間計算量になります。空間計算量は一定です `O(1)`。これはソートがインプレースで行われ、リストのサイズに関わらず一定量のメモリしか使わないためです。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ユーザーストーリー:**

1. `selection_sort` という名前の関数を定義してください。
1. `selection_sort` 関数は、アイテムのリストを表す1つのパラメータを持つべきです。
1. `selection_sort` 関数はリストを受け取り、要素を最小から最大の順にインプレースでソートするべきです。
1. `selection_sort` 関数は入力リストをインプレースで変更し、ソート後にそれを返すべきです。
1. `selection_sort` 関数は選択ソートアルゴリズムに従い、リストの未ソート部分から最小の要素を最初の未ソート要素と交換するべきです。
1. `selection_sort` 関数は、最小の要素がすでに正しい位置にある場合は不要な交換を行わないようにしてください。
1. `selection_sort` 関数は組み込みの `sort()` メソッドや `sorted()` 関数を使用しないでください。

# --hints--

`selection_sort` という名前の関数を用意してください。

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

`selection_sort` 関数は1つのパラメータを持つべきです。

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

コード内でモジュールをインポートしたり、組み込みのソートメソッドを使ったりしないでください。

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

`selection_sort` は入力リストと同じリストを返すべきです。

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

`selection_sort` は入力リストをインプレースで変更するべきです。リストに要素を追加・削除するメソッドは使わないでください。

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

`selection_sort` 関数は選択ソートアルゴリズムに従い、未ソート部分の最小値を最初の未ソート要素と交換してください。最小値がすでに正しい位置にある場合は不要な交換を避けてください。

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` は `[1, 2, 33, 67, 89, 245]` を返すべきです。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` は `[3, 5, 12, 15, 16, 23, 72, 99, 567]` を返すべきです。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` は `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]` を返すべきです。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

`selection_sort` 関数は任意の数値リストを正しくソートできるべきです。

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
