---
id: 641d9a19bff38d34d5a5edb8
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Der gesamte HTML- und CSS-Code für dieses Projekt wurde für Sie bereitgestellt. Sie können sich die beiden Dateien ansehen, um sich mit ihnen vertraut zu machen.

Beginnen Sie damit, die `#author-container` und `#load-more-btn` Elemente mit der `.getElementById()` Methode zu holen. Weisen Sie sie den Variablen `authorContainer` beziehungsweise `loadMoreBtn` zu.

Die Variablen werden sich nicht ändern, also verwenden Sie `const`, um sie zu deklarieren.

# --hints--

Sie sollten `document.getElementById()` verwenden, um das `#author-container` Element zu holen.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)author\-container\1\s*\)/);
```

Sie sollten das `#author-container` Element der Variable `authorContainer` zuweisen. Vergessen Sie nicht, `const` zu verwenden, um die Variable zu deklarieren.

```js
assert.match(code, /const\s+authorContainer\s*\=\s*document\.getElementById\(\s*('|"|`)author\-container\1\s*\)/);
```

Sie sollten `document.getElementById()` verwenden, um das `#load-more-btn` Element zu holen.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)load\-more\-btn\1\s*\)/);
```

Sie sollten das `#load-more-btn` Element der Variable `loadMoreBtn` zuweisen. Vergessen Sie nicht, `const` zu verwenden, um die Variable zu deklarieren.

```js
assert.match(code, /const\s+loadMoreBtn\s*\=\s*document\.getElementById\(\s*('|"|`)load\-more\-btn\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
