---
id: 6995d0325fd7da40854fac40
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu utafutaji wa binary na jinsi unavyotofautiana na utafutaji wa mstari.

Katika warsha hii, utaweka hayo katika vitendo kwa kujenga algoriti ya utafutaji wa binary.

Kuanza, tengeneza kitendakazi cha `binarySearch` chenye vigezo vya `searchList` na `value`.

`searchList` ni safu ya data ya vipengele ambavyo kitendakazi kitaangalia ndani yake, na `value` ni lengo ambalo kitendakazi kitajaribu kulipata ndani ya `searchList`.

# --hints--

Unapaswa kutengeneza kitendakazi cha `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Kitendakazi chako cha `binarySearch` kinapaswa kuwa na vigezo vya `searchList` na `value`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
