---
id: 6996c60281feca6718cc15c1
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Ndani ya kitendakazi chako cha `binarySearch`, tengeneza kigezo cha `pathToTarget` na uanze kwa kukiweka kuwa safu ya data tupu.

# --hints--

Unapaswa kuwa na kigezo cha `pathToTarget` kilichowekwa kuwa safu ya data tupu `[]` ndani ya kitendakazi chako cha `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
