---
id: 6996cb5869acf58de751ae81
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Kumbuka kuwa kitendakazi kinapokea safu ya data inayoitwa `searchList` na `value` unayotafuta.

Sasa, fanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5], 3)` na `binarySearch([1, 2, 3, 4, 5, 9], 4)` na andika kauli zote mbili kwenye konsoli.

# --hints--

Unapaswa kufanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5], 3)` na kuandika kwenye konsoli.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*\]\s*,\s*3\s*\)\s*\)/
);
```

Unapaswa kufanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5, 9], 4)` na kuandika kwenye konsoli.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*9\s*\]\s*,\s*4\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--  

--fcc-editable-region--  
```
