---
id: 6996c60281feca6718cc15c1
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Ndani ya kitendakazi chako cha `binarySearch`, tengeneza kigezo cha `pathToTarget` na uanze kwa kukiweka kuwa safu ya data tupu.

# --hints--

Unapaswa kutangaza `pathToTarget` kama kigezo kilichowekwa kuwa safu ya data tupu ndani ya kitendakazi chako cha `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
