---
id: 6996c675b599baa7b103192d
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Katika utafutaji wa binary, kuna kielezo cha chini kabisa kinachowezekana na kielezo cha juu kabisa kinachowezekana.

Kielezo cha chini kabisa kinawakilisha mpaka wa kushoto kabisa katika sehemu ya sasa ya thamani zinazotafutwa, katika kesi hii, `searchList`. Kielezo cha juu kabisa, kwa upande mwingine, kinawakilisha mpaka wa kulia kabisa katika thamani zinazotafutwa.

Ili kuzingatia hayo mawili, lazima uzingatie safu yote ya data inayotafutwa. Hivyo, ndani ya kitendakazi cha `binarySearch`, fafanua tofauti ya `low` yenye thamani ya `0`, na tofauti ya `high` yenye thamani inayozingatia kielezo cha mwisho katika safu ya data inayotafutwa.

# --hints--

Unapaswa kuwa na tofauti iitwayo `low` iliyowekwa kuwa `0` katika kitendakazi chako cha `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Unapaswa kuwa na tofauti iitwayo `high` iliyowekwa kuwa `searchList.length - 1` katika kitendakazi chako cha `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
