---
id: 6996c8f94fe6b250f11a2a1e
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa, jumuisha `valueAtMiddle` katika safu ya data ya `pathToTarget` kufuatilia hatua zilizochukuliwa wakati wa utafutaji, bila kujali kama ni lengwa. Unaweza kufanya hivi kwa kutumia njia ya `push()`.

# --hints--

Unapaswa kutumia njia ya `push()` kuongeza `valueAtMiddle` katika safu ya data ya `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
