---
id: 6996ca21a74104721f2c7e5d
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Unahitaji kujaribu mambo ili uelewe mtiririko wa algoriti katika hatua hii ya mwanzo.

Ili kufanya hivyo, unahitaji kwanza kutoka kwenye mzunguko. Hii ni kwa sababu utekelezaji wa sasa utaruhusu mzunguko mmoja tu, hivyo kama sharti halitimizwi, kutakuwa na mzunguko usioisha.

Mara tu baada ya usemi wa `if`, tumia neno kuu la `break` kutoka kwenye mzunguko wa `while`. Kisha, baada ya mzunguko wa `while`, rudisha safu ya data tupu kuashiria kuwa thamani haikupatikana.

# --hints--

Unapaswa kutumia neno kuu la `break` kutoka kwenye mzunguko wa `while` baada ya tamko la `if`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Unapaswa kurudisha safu ya data tupu baada ya mzunguko wako wa `while`.

```js
assert.deepEqual(binarySearch([], NaN), []);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
