---
id: 6996c8375f4301004459ee85
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa baada ya kuhesabu kielezo cha `mid`, unahitaji kupata thamani kutoka `searchList` kwa kielezo hicho. Thamani hii ndiyo utakayolinganishwa na kigezo chako cha `value`, ambacho ni lengo unalotafuta.

Tengeneza kigezo cha `valueAtMiddle` na ukae thamani ya kipengele kutoka `searchList` kilichopo kwa kielezo cha `mid`.

# --hints--

Unapaswa kutengeneza kigezo cha `valueAtMiddle` kikiwa sawa na `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
