---
id: 6996d18b31aab5cbbffbc731
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Marekebisho madogo ya mwisho unayoweza kuongeza ni pia kuonyesha kielezo ambacho `value` lengwa linapatikana.

Kwa hivyo, badala ya kurudisha tu `pathToTarget` katika usemi wa `if`, ongeza usemi wa template wenye ujumbe `Value found at index ${mid}`.

Kwa hiyo, warsha yako ya algoriti ya utafutaji wa binary imekamilika!

# --hints--

Unapaswa kurudisha usemi wa `pathToTarget, \`Value found at index ${mid}\`` inside the `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--

    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];

}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }

  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
