---
id: 6996cfb6c01e9a802add38be
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Mwishowe, kama `value` haipatikani katikati na pia si kubwa kuliko `valueAtMiddle`, basi `value` lazima iwe ndogo kuliko `valueAtMiddle`, na lazima iwe upande wa kushoto. Kipande cha mwisho cha `else` kitashughulikia hili.

Hivyo, malizia mzunguko kwa kuondoa neno kuu la `break` na kuongeza kipande cha `else`. Ndani yake, sasisha kigezo cha `high` kwa kutoa `1` kutoka `mid`.

# --hints--

Unapaswa kubadilisha neno kuu la `break` na kipande cha `else` kinachosahihisha kigezo cha `high` kuwa `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--      
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
