---
id: 6996c95e5bede03bdf2937a2
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Jambo linalofuata ni kuunda hali itakayokagua kama `value` lengwa iko katikati. 

Ndani ya mzunguko wa `while`, unda usemi wa `if` unaokagua kama `value` lengwa ni sawa na `valueAtMiddle`. Ikiwa ni sawa, rudisha kigezo cha `pathToTarget`.

# --hints--

Unapaswa kuunda usemi wa `if` wenye hali ya `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Unapaswa kurudisha kigezo cha `pathToTarget` ndani ya usemi wa `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
