---
id: 6996cb5869acf58de751ae81
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Kumbuka kuwa kitendakazi kinapokea safu ya data inayoitwa `searchList` na `value` unayotafuta.

Sasa, fanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5], 3)` na `binarySearch([1, 2, 3, 4, 5, 9], 4)` na andika kauli zote mbili kwenye konsoli.

# --hints--

Unapaswa kufanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5], 3)` na kuandika kwenye konsoli.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

Unapaswa kufanya wito la kitendakazi kwa `binarySearch([1, 2, 3, 4, 5, 9], 4)` na kuandika kwenye konsoli.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
