---
id: 6996d1032feed88e2957dc8e
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Unaweza kuona kwamba wito la mwisho linachapisha safu ya data tupu. Ili kubainisha zaidi kwamba thamani haipo katika safu ya data inayotafutwa, sasisha usemi wa `return` kurudisha safu ya data yenye vipengele viwili: safu ya data tupu, na mfuatano wa herufi `Value not found`.

# --hints--

Unapaswa kusasisha usemi wa `return` ili kitendakazi kirudishe safu ya data yenye `[]` kama kipengele cha kwanza na `"Value not found"` kama kipengele cha pili.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const validExpressions = [
  "[[], 'Value not found']",
  "[[], \"Value not found\"]",
  "[[], `Value not found`]"
];
assert.isTrue(validExpressions.some(expr => binarySearch?.hasReturn(expr) || binarySearch?.hasReturn("(" + expr + ")")));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
