---
id: 6996c675b599baa7b103192d
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Katika utafutaji wa binary, kuna kielezo cha chini kabisa kinachowezekana na kielezo cha juu kabisa kinachowezekana.

Kielezo cha chini kabisa kinawakilisha mpaka wa kushoto kabisa katika sehemu ya sasa ya thamani zinazotafutwa, katika kesi hii, `searchList`. Kielezo cha juu kabisa, kwa upande mwingine, kinawakilisha mpaka wa kulia kabisa katika thamani zinazotafutwa.

Ili kuzingatia hayo mawili, lazima uzingatie safu ya data nzima inayotafutwa. Hivyo, ndani ya kitendakazi cha `binarySearch`, tangaza kigezo `low` chenye thamani ya `0`, na kigezo `high` chenye thamani inayozingatia kielezo cha mwisho katika safu ya data inayotafutwa.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `low` kilichowekwa kuwa `0` ndani ya kitendakazi chako cha `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Unapaswa kutangaza kigezo kinachoitwa `high` kilichowekwa kuwa `searchList.length - 1` ndani ya kitendakazi chako cha `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
