---
id: 6996c796385b6ea4e222d5e8
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Ili kuanza na mzunguko, unahitaji kupata kielezo cha katikati cha nafasi ya sasa ya utafutaji.

Unaweza kupata kielezo cha katikati kwa kuhesabu wastani wa vigezo vya `low` na `high`. Utahitaji kutumia mgawanyiko wa sakafu (floor division) kupata wastani baada ya kuongeza thamani hizo mbili, hivyo jibu litakuwa daima nambari kamili iliyozungushwa chini, kwani vigezo lazima viwe nambari kamili.

Ndani ya mzunguko wako wa `while`, tengeneza kigezo cha `mid` kilichowekwa kwa wastani wa `low` na `high`.

# --hints--

Ndani ya mzunguko wa while unapaswa kutengeneza kigezo cha `mid` kilichowekwa kwa `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
