---
id: 6996ccadf9048c79b857bb15
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Unaweza kuona kwamba wito la kwanza hurudisha `[3]`. Hii ni kwa sababu `3` ni thamani ya katikati katika `[1, 2, 3, 4, 5]`.

Kwa upande mwingine, wito la pili hurudisha safu ya data tupu. Hii hutokea kwa sababu, baada ya kukagua thamani ya katikati, mzunguko kwa sasa hauna mantiki ya kupunguza eneo la utafutaji kwa kusahihisha mabadiliko ya `low` au `high`. Inafanya tu ukaguzi mmoja wa katikati kisha inamaliza bila utafutaji zaidi.

Ili kuruhusu utafutaji wa binary kuendelea kupunguza utafutaji wake, ongeza tamko la `else if` linalokagua kama `value` ni kubwa kuliko `valueAtMiddle`.

# --hints--

Unapaswa kuongeza tamko la `else if` lenye sharti la `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
