---
id: 6996c7363a4a437839f1f15d
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa baada ya kufafanua mipaka ya utafutaji wako, ni wakati wa kuunda mzunguko utakaotekeleza algoriti ya utafutaji wa binary.

Utafutaji wa binary hufanya kazi kwa kupunguza mara kwa mara nafasi ya utafutaji. Mchakato huu unaendelea mradi tu kuna safu halali ya vipengele vya kukagua.

Unaweza kudhihirisha hili kwa mzunguko wa `while` ambao utaendelea mradi tu kiashiria chako cha `low` kiko chini au sawa na kiashiria cha `high`.

Ikiwa `low` itakuwa kubwa kuliko `high`, inamaanisha eneo la utafutaji limeisha, na `value` haipo katika safu ya data.

Ndani ya kitendakazi cha `binarySearch`, tengeneza mzunguko wa `while` wenye sharti linalokagua kama `low` ni ndogo au sawa na `high`.

# --hints--

Unapaswa kuunda mzunguko wa `while` unaoendelea mradi `low` iko chini au sawa na `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
