---
id: 5cc0bd7a49b71cb96132e54c
title: 数値の範囲を生成するジェネレーターを作成する
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

このラボでは、指定された範囲内の数値の配列を生成する`rangeOfNumbers`関数を作成します。

例:

- `rangeOfNumbers(3, 9)`は`[3, 4, 5, 6, 7, 8, 9]`を返します
- `rangeOfNumbers(5, 5)`は`[5]`を返します

要件:

- 再帰を使う（関数は自分自身を呼び出す必要があります）
- ループや配列メソッドは使わない（`for`、`while`、`Array.from`、`forEach`、`map`、`filter`、`reduce`）

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `rangeOfNumbers`という名前の関数を作成し、2つのパラメータ`startNum`と`endNum`を受け取る必要があります。
1. 関数は、`startNum`パラメータで表される数値から始まり、`endNum`パラメータで表される数値で終わる（含む）整数の配列を返す必要があります。
1. `startNum`は常に`endNum`以下です。
1. 関数は再帰を使って自分自身を呼び出す必要があります。ループ構文（`for`、`while`、`Array.from()`）や高階関数（`forEach`、`map`、`filter`、`reduce`）は使わないでください。
1. 基本ケースとして、`startNum`が`endNum`と等しい場合は、その単一の数値だけを含む配列を返す必要があります。
1. 再帰ケースでは、関数はパラメータを変更して自分自身を呼び出し、配列を作成した後、現在の数値を結果に追加する必要があります。
1. 配列をキャッシュまたは作成するためにグローバル変数を使わないでください。

# --hints--

`rangeOfNumbers`という名前の関数を用意してください。

```js
assert.isFunction(rangeOfNumbers);
```

`rangeOfNumbers`関数は2つのパラメータを受け取る必要があります。

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

関数は配列を返す必要があります。

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

コードはループ構文（`for`、`while`、`Array.from()`）や高階関数（`forEach`、`map`、`filter`、`reduce`）を使わないでください。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

`rangeOfNumbers`は再帰（自分自身を呼び出す）を使ってこのチャレンジを解決する必要があります。

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

`rangeOfNumbers(1, 5)`は`[1, 2, 3, 4, 5]`を返す必要があります。

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)`は`[6, 7, 8, 9]`を返す必要があります。

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)`は`[4]`を返す必要があります。

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)`は`[10, 11, 12, 13, 14, 15]`を返す必要があります。

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)`は`[2, 3, 4, 5, 6, 7, 8]`を返す必要があります。

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

配列をキャッシュするためにグローバル変数を使わないでください。

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
