---
id: 66eaddd04a9e533fba689001
title: 블로그 게시물 카드 디자인하기
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

이 실습에서는 블로그 게시물 카드를 만들어 배경과 테두리를 스타일링하는 방법을 연습합니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 모든 카드 요소를 담을 `div` 요소가 `blog-post-card` 클래스를 가져야 합니다.
2. `.blog-post-card` 요소 안에는 유효한 `alt` 속성과 텍스트가 있는 이미지 요소가 `post-img` 클래스를 가지고 있어야 합니다. 이미지의 `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` 속성에는 `src`를 사용할 수 있습니다.
3. `div` 요소 안에 `post-content`가 `.blog-post-card` 클래스를 가지고 있어야 하며, 다음을 포함해야 합니다:
    - 블로그 게시물 제목 텍스트가 있는 `h2` 요소가 `post-title` 클래스를 가져야 합니다.
    - 블로그 게시물을 요약하는 텍스트가 있는 `p` 요소가 `post-excerpt` 클래스를 가져야 합니다.
    - 텍스트 `a`가 있는 `read-more` 요소가 `Read More` 클래스를 가져야 합니다.
4. `body` 요소의 배경색을 흰색이 아닌 값으로 설정해야 합니다.
5. `.blog-post-card` 요소에 다음 스타일을 적용해야 합니다:
    - 흰색 배경
    - 둥근 모서리
    - 원하는 너비
    - 원하는 텍스트 정렬
6. `.post-img` 요소는 이미지가 카드 전체 너비를 채우고 아래쪽에 테두리가 있도록 스타일링해야 합니다.
7. `.post-content` 요소는 카드 내부에 패딩이 있도록 스타일링해야 합니다.
8. `.post-title`과 `.post-excerpt` 요소는 기본값이 아닌 텍스트 색상과 모든 면에 여백이 있어야 합니다.
9. `.read-more` 요소는 버튼처럼 스타일링하고 다음을 가져야 합니다:
    - 기본값이 아닌 텍스트 색상
    - 배경색
    - 모든 면에 여백
    - 모든 면에 패딩
    - 둥근 모서리
    - `display` 속성이 `inline-block`로 설정됨
    - 배경색이 변경되는 호버 효과

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용했는지 확인하세요.

# --hints--

`div`가 `blog-post-card` 클래스를 가져야 합니다.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

`img` 요소가 `class`가 `post-img`인 것을 가져야 합니다. 이미지에 텍스트가 있는 `alt` 속성과 값이 있는 `src` 속성이 있는지 확인하세요.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

`div`가 `post-content` 클래스를 가져야 합니다.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

`h2` 요소가 `post-title` 클래스를 가져야 합니다. 블로그 게시물 제목 텍스트가 포함되어 있는지 확인하세요.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

`p`가 `post-excerpt` 클래스를 가져야 합니다. 블로그 게시물을 요약하는 텍스트가 포함되어 있는지 확인하세요.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

`a` 요소가 `class`가 `read-more`인 것을 가져야 합니다.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

`.read-more` 요소에 텍스트 `Read More`가 있어야 합니다.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

`.blog-post-card` 요소에 `border-radius` 속성이 값과 함께 있어야 합니다 (`0` 또는 음수 값이 아니어야 합니다).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

`.blog-post-card` 요소에 흰색 배경이 있어야 합니다.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

`body` 요소에 흰색이 아닌 배경색이 있어야 합니다.

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

`.blog-post-card`를 대상으로 `width` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

`.blog-post-card`를 대상으로 `text-align` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

`.post-content`를 대상으로 `padding` 속성을 설정해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

`.read-more` 요소에 호버 효과가 있어야 합니다.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

`.read-more`를 대상으로 `color` 속성을 설정해야 합니다.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

`.read-more`를 대상으로 `background-color` 속성을 설정해야 합니다.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

`.read-more`를 대상으로 `margin` 속성을 설정해야 합니다. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

`.read-more`를 대상으로 `display` 속성을 설정해야 합니다. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

`.read-more`를 대상으로 `border-radius` 속성을 설정해야 합니다. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

`.read-more`를 대상으로 `padding` 속성을 설정해야 합니다. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

`.post-img` 요소가 카드의 `width`를 채워야 합니다.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

`.post-img` 요소에 `border-bottom` 값이 있어야 합니다.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

`.post-title`와 `.post-excerpt` 요소에 여백과 기본값이 아닌 텍스트 색상이 있어야 합니다.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
