---
id: 699f434e294a9e11cf4f65f0
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Ora è il momento di aggiungere un altro aiuto matematico al tuo progetto toolkit.

Inizia creando una funzione chiamata `getAverage` che prende un parametro. La tua funzione dovrebbe restituire la media di un array di numeri.

Nei prossimi passaggi aggiungerai le annotazioni di tipo per questa funzione.

# --hints--

Dovresti avere una funzione chiamata `getAverage`.

```js
assert.isFunction(getAverage);
```

La tua funzione dovrebbe prendere un parametro.

```js
assert.strictEqual(getAverage.length, 1);
```

La tua funzione dovrebbe restituire un numero.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

La tua funzione dovrebbe restituire il risultato corretto per qualsiasi array di numeri.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
