---
id: 699f4a3e7b4c6bfa538cc4f7
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Se apri la console, vedrai il seguente messaggio di errore:

```md
Type 'string' is not assignable to type 'number'.
```

Per risolvere questo errore, devi cambiare il valore di ritorno in `sum / numbers.length`. 

# --hints--

La tua funzione `getAverage` dovrebbe restituire un numero.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

La tua funzione `getAverage` dovrebbe restituire la media dei numeri nell'array.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
