---
id: 699f3738921fb79a651ded56
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora è il momento di chiamare la tua funzione.

Crea una variabile chiamata `result` e assegnale la chiamata della funzione `square` con un argomento `5`.

Sotto, stampa il valore di `result` nella console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile chiamata `result`.

```js
assert.isDefined(result);
```

La tua variabile `result` dovrebbe essere un numero.

```js
assert.isNumber(result);
```

La tua variabile `result` dovrebbe essere il quadrato di `5`.

```js
assert.strictEqual(result, 25);
```

Dovresti stampare la variabile `result` nella console.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
