---
id: 699f52220de88c8091032076
title: Passaggio 18
challengeType: 1
dashedName: step-18
---

# --description--

Se apri la console, il risultato è `NaN` perché il secondo argomento per `exponent` non è mai stato fornito. Ci sono un paio di aggiornamenti necessari per questa funzione `raiseTo` che risolveranno il problema.

La prima modifica è restituire `base ** 2`.

Nel passaggio successivo vedrai come gestire il parametro opzionale `exponent`.

# --hints--

La tua funzione `raiseTo` dovrebbe restituire un numero.

```js
assert.isNumber(raiseTo(4));
```

La tua funzione `raiseTo` dovrebbe restituire il valore di `base` elevato alla potenza di `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
