---
id: 699f1d72e906e4616155290d
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop continuerai a imparare sulle annotazioni di tipo e su come funzionano con le funzioni costruendo un toolkit matematico sicuro per i tipi.

Inizia creando una funzione chiamata `square` che accetta un parametro e restituisce il quadrato di un numero.

# --hints--

Dovresti avere una funzione chiamata `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

La tua funzione `square` dovrebbe avere un parametro.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

La tua funzione `square` dovrebbe restituire il quadrato di un numero.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
