---
id: 699f49669ea51a31d2ba89d8
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Come ricorderai dalle lezioni precedenti, puoi specificare esplicitamente il tipo di ritorno di una funzione in questo modo:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

In questo esempio, il tipo di ritorno della funzione `circleArea` è `number`. Se provi a restituire un tipo diverso, TypeScript mostrerà un errore.

Aggiorna la tua funzione `getAverage` per specificare esplicitamente il tipo di ritorno come `number`. 

# --hints--

La tua funzione `getAverage` dovrebbe avere un tipo di ritorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
