---
id: 699f4639ae0016caaf81b7c6
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ora è il momento di testare la tua funzione.

Inizia creando una variabile chiamata `avgResult` e assegnandole la chiamata della funzione `getAverage`. L'argomento per quella chiamata di funzione dovrebbe essere `[2, 14, 26, 8]`.

Poi stampa il valore di `avgResult` nella console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile chiamata `avgResult`.

```js
assert.isDefined(avgResult);
```

Dovresti assegnare `getAverage([2, 14, 26, 8])` alla tua variabile `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

La tua variabile `avgResult` dovrebbe essere un numero.

```js
assert.isNumber(avgResult);
```

Dovresti stampare `avgResult` nella console.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
