---
id: 699f3f636aa840b7201aa37c
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Se apri la console, dovresti vedere il seguente messaggio di errore di TypeScript:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Ora che il parametro è esplicitamente tipizzato come number, TypeScript riesce a intercettare l’errore di provare a passare un argomento stringa alla funzione `square`.

Cambia l’argomento passato alla funzione `square` nel numero `5` per correggere l’errore e vedere l’output corretto nella console.

# --hints--

La tua chiamata alla funzione `square` dovrebbe avere un argomento `5` invece di `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
