---
id: 699f4861dd5b5f164e1bda4a
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Un altro vantaggio di usare TypeScript è assicurarsi che i tipi di ritorno siano rispettati e corrispondano a ciò che ti aspetti. A volte TypeScript può dedurre il tipo di ritorno di una funzione, ma puoi anche specificarlo esplicitamente.

Per illustrare questo, cambia il tipo di ritorno della funzione `getAverage` nella stringa `"something"`.

# --hints--

La tua funzione `getAverage` dovrebbe restituire una stringa.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

La tua funzione `getAverage` dovrebbe restituire la stringa `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
