---
id: 699f3cb8d057f29954205d7d
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Come ricordi dalle lezioni precedenti, puoi aggiungere annotazioni di tipo ai parametri di una funzione in questo modo:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

In questo esempio, TypeScript mostrerà un errore se provi a chiamare la funzione `add` con argomenti che non sono numeri.

Aggiorna la tua funzione `square` per includere un'annotazione di tipo per il parametro `num`. Il tipo dovrebbe essere `number`.

# --hints--

Dovresti aggiungere un'annotazione di tipo al parametro `num` nella funzione `square`. Il tipo dovrebbe essere un `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
