---
id: 699f5063d2393c674b1f45dc
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Come ricorderai da un passaggio precedente, se chiami una funzione con un numero errato di argomenti, TypeScript lancerà un errore.

Esiste un modo per indicare i parametri come opzionali aggiungendo un `?` dopo il nome del parametro. Questo ti permette di chiamare la funzione con meno argomenti rispetto ai parametri.

Ecco un esempio:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

In questo esempio, se il parametro `greeting` non viene fornito quando si chiama la funzione `greet`, sarà `undefined` e la funzione restituirà un saluto predefinito.

Aggiorna la tua funzione `raiseTo` per rendere opzionale il parametro `exponent`. 

# --hints--

La tua funzione `raiseTo` dovrebbe avere il parametro `exponent` indicato come opzionale.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
