---
id: 699f52fac083214b6bf51de6
title: Passaggio 19
challengeType: 1
dashedName: step-19
---

# --description--

In questo ultimo passaggio, dovrai aggiungere un controllo condizionale per il parametro `exponent` nella funzione `raiseTo`.

Se viene fornito l'argomento `exponent`, la funzione dovrebbe restituire `base` elevato alla potenza di `exponent`. Se l'argomento `exponent` non viene fornito, dovrebbe restituire `base ** 2`.

Con questo ultimo passaggio, il tuo kit di strumenti matematici type-safe è completo!

# --hints--

La tua funzione `raiseTo` dovrebbe restituire `base ** exponent` quando viene fornito l'argomento `exponent`.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

La tua funzione `raiseTo` dovrebbe restituire `base ** 2` quando l'argomento `exponent` non viene fornito.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
