---
id: 699f4ec939f10da62fb56c3e
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Ora è il momento di testare la tua funzione.

Crea una variabile chiamata `raisedResult` e assegnale il risultato della chiamata a `raiseTo` con gli argomenti `2` e `3`.

Poi, stampa `raisedResult` nella console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Dovresti avere una variabile chiamata `raisedResult`.

```js
assert.isDefined(raisedResult);
```

La tua variabile `raisedResult` dovrebbe essere un numero.

```js
assert.isNumber(raisedResult);
```

Dovresti stampare `raisedResult` nella console.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
