---
id: 699f458e47fb88c756fdba81
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Al momento, il parametro `numbers` ha un tipo `any` implicito. Per risolvere questo, puoi specificare esplicitamente il tipo del parametro come un array di numeri.

Ricorda dalle lezioni precedenti che puoi definire un array di numeri in questo modo:

```ts
const numbers: number[] = [1, 2, 3];
```

Aggiorna la tua funzione `getAverage` per specificare esplicitamente il tipo del parametro `numbers` come un array di numeri.

# --hints--

La tua funzione `getAverage` dovrebbe avere un parametro chiamato `numbers` che è esplicitamente tipizzato come un array di numeri.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
