---
id: 699f412114c87221f16a99e3
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Un altro vantaggio di usare TypeScript è che rileva errori se passi troppi o troppo pochi argomenti a una chiamata di funzione.

Per illustrare questo, aggiorna la tua chiamata alla funzione `square` per prendere due argomenti `5` e `10`.

# --hints--

La tua chiamata alla funzione `square` dovrebbe avere due argomenti `5` e `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
