---
id: 699f41b8da20f816b5c7dece
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Se apri la console, dovresti vedere il seguente messaggio di errore:

```md
Expected 1 arguments, but got 2.
```

Per risolvere l'errore, aggiorna la tua chiamata alla funzione `square` in modo che prenda solo un argomento di tipo `5`.

# --hints--

La tua chiamata alla funzione `square` dovrebbe avere solo un argomento di tipo `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
