---
id: 699f4b903bf5427a599f791d
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Negli ultimi passaggi costruirai un'altra funzione per continuare a esercitarti con le annotazioni di tipo nelle funzioni.

Inizia creando una funzione chiamata `raiseTo` che prenda due parametri chiamati `base` e `exponent`. Entrambi i parametri devono essere di tipo `number`. La funzione deve restituire il risultato di elevare `base` alla potenza di `exponent`. Il tipo di ritorno della funzione deve essere annotato come `number`.

# --hints--

Dovresti avere una funzione `raiseTo`.

```js
assert.isFunction(raiseTo);
```

La tua funzione `raiseTo` dovrebbe avere due parametri, `base` e `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

I tuoi parametri `base` e `exponent` dovrebbero entrambi essere di tipo `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

La tua funzione `raiseTo` dovrebbe restituire un numero.

```js
assert.isNumber(raiseTo(2, 3));
```

La tua funzione `raiseTo` dovrebbe restituire il risultato di elevare `base` alla potenza di `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

La tua funzione `raiseTo` dovrebbe avere un'annotazione del tipo di ritorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
