---
id: 699f38937b053e4679878ff1
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Quando lavori con JavaScript puro, nulla ti impedisce di chiamare la funzione `square` con una stringa o un array, il che porterebbe a risultati inattesi.

Per illustrare questo, modifica la chiamata alla funzione per usare la stringa `"something"` invece del numero `5`.

Nella prossima lezione imparerai come usare TypeScript per aggiungere sicurezza di tipo al tuo codice e prevenire questo tipo di errori.

# --hints--

La tua chiamata alla funzione `square` dovrebbe avere un argomento di tipo `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
