---
id: 587d78ae367417b2b2512afe
title: Usa la proprietà abbreviata flex
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

Esiste una scorciatoia per impostare contemporaneamente diverse proprietà flex. Le proprietà `flex-grow`, `flex-shrink` e `flex-basis` possono essere tutte impostate insieme usando la proprietà `flex`.

Per esempio, `flex: 1 0 10px;` imposterà l'elemento su `flex-grow: 1;`, `flex-shrink: 0;` e `flex-basis: 10px;`.

Le impostazioni predefinite delle proprietà sono `flex: 0 1 auto;`.

# --instructions--

Aggiungi la proprietà CSS `flex` sia a `#box-1` che a `#box-2`. Dai a `#box-1` i valori in modo che la sua `flex-grow` sia `2`, la sua `flex-shrink` sia `2` e la sua `flex-basis` sia `150px`. Dai a `#box-2` i valori in modo che la sua `flex-grow` sia `1`, la sua `flex-shrink` sia `1` e la sua `flex-basis` sia `150px`.

Questi valori faranno sì che `#box-1` cresca per riempire lo spazio extra al doppio della velocità di `#box-2` quando il contenitore è più grande di 300px e si riduca al doppio della velocità di `#box-2` quando il contenitore è più piccolo di 300px. 300px è la dimensione combinata dei valori `flex-basis` dei due riquadri.

# --hints--

L'elemento `#box-1` dovrebbe avere la proprietà `flex` impostata su un valore `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

L'elemento `#box-2` dovrebbe avere la proprietà `flex` impostata su un valore `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Il tuo codice dovrebbe usare la proprietà `flex` per `#box-1` e `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
