---
id: 587d78ae367417b2b2512afd
title: usa la proprietà flex-basis per impostare la dimensione iniziale di un elemento
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

La proprietà `flex-basis` specifica la dimensione iniziale dell'elemento prima che il CSS applichi regolazioni con `flex-shrink` o `flex-grow`.

Le unità usate dalla proprietà `flex-basis` sono le stesse di altre proprietà di dimensione (`px`, `em`, `%`, ecc.). Il valore `auto` dimensiona gli elementi in base al contenuto.

# --instructions--

Imposta la dimensione iniziale dei riquadri usando `flex-basis`. Aggiungi la proprietà CSS `flex-basis` sia a `#box-1` che a `#box-2`. Dai a `#box-1` un valore `10em` e a `#box-2` un valore `20em`.

# --hints--

L'elemento `#box-1` dovrebbe avere la proprietà `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

L'elemento `#box-1` dovrebbe avere un valore `flex-basis` pari a `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

L'elemento `#box-2` dovrebbe avere la proprietà `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

L'elemento `#box-2` dovrebbe avere un valore `flex-basis` pari a `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
