---
id: 587d78ab367417b2b2512af0
title: 'Usa display: flex per posizionare due riquadri'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

Questa sezione usa stili di sfida alternati per mostrare come usare il CSS per posizionare gli elementi in modo flessibile. Prima, una sfida spiegherà la teoria, poi una sfida pratica con un semplice componente tweet applicherà il concetto di flexbox.

Impostare la proprietà CSS `display: flex;` su un elemento ti permette di usare altre proprietà flex per costruire una pagina responsiva.

# --instructions--

Aggiungi la proprietà CSS `display` a `#box-container` e imposta il suo valore su `flex`.

# --hints--

`#box-container` dovrebbe avere la proprietà `display` impostata con un valore di `flex`.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
