---
id: 587d78ad367417b2b2512afb
title: usa la proprietà flex-shrink per ridurre gli elementi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

Finora, tutte le proprietà nelle sfide si applicano al contenitore flex (il genitore degli elementi flex). Tuttavia, ci sono diverse proprietà utili per gli elementi flex.

La prima è la proprietà `flex-shrink`. Quando viene usata, permette a un elemento di ridursi se il contenitore flex è troppo piccolo. Gli elementi si riducono quando la larghezza del contenitore genitore è più piccola della somma delle larghezze di tutti gli elementi flex al suo interno.

La proprietà `flex-shrink` accetta numeri come valori. Più alto è il numero, più si ridurrà rispetto agli altri elementi nel contenitore. Per esempio, se un elemento ha un valore `flex-shrink` di `1` e l'altro ha un valore `flex-shrink` di `3`, quello con il valore `3` si ridurrà tre volte di più rispetto all'altro.

# --instructions--

Aggiungi la proprietà CSS `flex-shrink` sia a `#box-1` che a `#box-2`. Dai a `#box-1` un valore di `1` e a `#box-2` un valore di `2`.

# --hints--

L'elemento `#box-1` dovrebbe avere la proprietà `flex-shrink` impostata con un valore di `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

L'elemento `#box-2` dovrebbe avere la proprietà `flex-shrink` impostata con un valore di `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
