---
id: 587d78af367417b2b2512b00
title: usa la proprietà align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

L'ultima proprietà per gli elementi flex è `align-self`. Questa proprietà ti permette di regolare l'allineamento di ogni elemento singolarmente, invece di impostarli tutti insieme. Questo è utile perché altre tecniche di regolazione comuni usando le proprietà CSS `float`, `clear` e `vertical-align` non funzionano sugli elementi flex.

`align-self` accetta gli stessi valori di `align-items` e sovrascriverà qualsiasi valore impostato dalla proprietà `align-items`.

# --instructions--

Aggiungi la proprietà CSS `align-self` sia a `#box-1` che a `#box-2`. Imposta `#box-1` con il valore `center` e `#box-2` con il valore `flex-end`.

# --hints--

L'elemento `#box-1` dovrebbe avere la proprietà `align-self` impostata con il valore `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

L'elemento `#box-2` dovrebbe avere la proprietà `align-self` impostata con il valore `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
