---
id: 587d78ab367417b2b2512af2
title: usa la proprietà flex-direction per creare una riga
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

Aggiungere `display: flex` a un elemento lo trasforma in un contenitore flex. Questo permette di allineare i figli di quell’elemento in righe o colonne. Lo fai aggiungendo la proprietà `flex-direction` al genitore e impostandola su row o column. Creare una riga allineerà i figli orizzontalmente, mentre creare una colonna li allineerà verticalmente.

Altre opzioni per `flex-direction` sono `row-reverse` e `column-reverse`.

**Nota:** Il valore predefinito per la proprietà `flex-direction` è `row`.

# --instructions--

Aggiungi la proprietà CSS `flex-direction` all’elemento `#box-container` e impostala con il valore `row-reverse`.

# --hints--

L’elemento `#box-container` dovrebbe avere la proprietà `flex-direction` impostata su `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
