---
id: 587d78ae367417b2b2512aff
title: usa la proprietà order per riordinare gli elementi
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

La proprietà `order` serve a indicare a CSS l'ordine con cui gli elementi flex appaiono nel contenitore flex. Per impostazione predefinita, gli elementi appaiono nello stesso ordine in cui sono presenti nel codice HTML sorgente. La proprietà accetta numeri come valori, e si possono usare anche numeri negativi.

# --instructions--

Aggiungi la proprietà CSS `order` sia a `#box-1` sia a `#box-2`. Imposta `#box-1` con il valore `2` e `#box-2` con il valore `1`.

# --hints--

L'elemento `#box-1` dovrebbe avere la proprietà `order` impostata con il valore `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

L'elemento `#box-2` dovrebbe avere la proprietà `order` impostata con il valore `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
