---
id: 587d78ad367417b2b2512afa
title: usa la proprietà flex-wrap per andare a capo in una riga o colonna
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS flexbox ha una funzionalità per dividere un contenitore flex in più righe (o colonne). Per impostazione predefinita, un contenitore flex adatta tutti gli elementi flex insieme. Per esempio, una riga sarà tutta su una sola linea.

Tuttavia, usando la proprietà `flex-wrap` si dice a CSS di andare a capo sugli elementi. Questo significa che gli elementi in eccesso si spostano in una nuova riga o colonna. Il punto di interruzione in cui avviene il capo dipende dalla dimensione degli elementi e dalla dimensione del contenitore.

CSS ha anche opzioni per la direzione del capo:

<ul><li><code>nowrap</code>: questa è l'impostazione predefinita e non fa andare a capo gli elementi.</li><li><code>wrap</code>: fa andare a capo gli elementi su più righe dall'alto verso il basso se sono in righe e da sinistra a destra se sono in colonne.</li><li><code>wrap-reverse</code>: fa andare a capo gli elementi su più righe dal basso verso l'alto se sono in righe e da destra a sinistra se sono in colonne.</li></ul>

# --instructions--

L'attuale layout ha troppi riquadri per una sola riga. Aggiungi la proprietà CSS `flex-wrap` all'elemento `#box-container` e impostala con il valore `wrap`.

# --hints--

L'elemento `#box-container` dovrebbe avere la proprietà `flex-wrap` impostata con il valore `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
