---
id: 68a23004e96a105df3a21fc4
title: 构建一个卡牌计数助手
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

在 21 点游戏中，玩家可以通过计算牌桌上已经发放的卡牌的高低值来让自己在游戏中保持优势。 这被称为卡片计数。

牌桌上的大值的卡牌更多，对玩家有利。 如果卡牌的值大于 0，那么玩家应该追加赌注。 如果卡牌的值为 0 或负数，玩家应该追加少许赌注甚至不追加赌注。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该使用 `let` 声明一个名为 `count` 的全局的变量并将其设置为 `0`。
1. 你应该有一个名为 `cardCounter` 的函数。
1. `cardCounter` 函数应接收一个 `card` 参数，该参数可以是数字或字串。
    - 对于介于 `2` 到 `10` 之间的值，`card` 参数将是一个数字。
    - 对于所有其他值，`card` 参数将是一个字串。
1. `cardCounter` 函数应根据某些条件修改全局的 `count` 变量。
1. 全局的 `count` 变量应针对牌 `2`、`3`、`4`、`5` 或 `6` 增加 `1`
1. 全局的 `count` 变量对于卡片 `7`、`8`、`9` 应保持不变。
1. 全局的 `count` 变量应减少 `1`，针对牌 `10`、`"J"`、`"Q"`、`"K"`、`"A"`。
1. `cardCounter` 函数应返回包含当前计数和字串 `Bet` 的字串，如果计数为正。
1. `cardCounter` 函数应返回包含当前计数的字串，如果计数小于或相等于 `0`，则返回字串 `Hold`。
1. 在函数输出中，当前计数和玩家的决策（`Bet` 或 `Hold`）之间应以空格分隔。例如，`-3 Hold`。

# --hints--

你应该使用 `let` 声明一个名为 `count` 的全局的变量并将其设置为 `0`。 

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

你应该有一个名为 `cardCounter` 的函数。 

```js
assert.isFunction(cardCounter);
```

你的函数应返回 `count` 的值和文本（`Bet` 或 `Hold`），它们之间有一个空格字符。

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

在牌 `2`、`3`、`4`、`5` 之后，调用 `cardCounter(6)` 应该返回字串 `5 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

在牌 `7`、`8` 之后，调用 `cardCounter(9)` 应该返回字串 `0 Hold`。

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

在牌 `10`、`"J"`、`"Q"`、`"K"` 之后，调用 `cardCounter("A")` 应该返回字串 `-5 Hold`。

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

在牌 `3`、`7`、`"Q"`、`8` 之后，调用 `cardCounter("A")` 应该返回字串 `-1 Hold`。

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

在牌 `2`、`"J"`、`9`、`2` 之后，调用 `cardCounter(7)` 应该返回字串 `1 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

在牌 `2`、`2` 之后，调用 `cardCounter(10)` 应该返回字串 `1 Bet`。

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

在牌 `3`、`2`、`"A"`、`10` 之后，调用 `cardCounter("K")` 应该返回字串 `-1 Hold`。

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
