---
id: 67116d7584d0b469b14579bf
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Dentro do array `library`, crie um objeto com as seguintes propriedades e valores:

| Propriedade | Valor   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Você deve ter um objeto dentro do array `library`.

```js
assert.isObject(library[0])
```

Seu objeto deve ter uma propriedade `title`.

```js
assert.property(library[0], "title")
```

O valor da sua propriedade `title` deve ser `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Seu objeto deve ter uma propriedade `author`.

```js
assert.property(library[0], "author")
```

A sua propriedade `author` deve ter o valor `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Seu objeto deve ter uma propriedade `about`.

```js
assert.property(library[0], "about")
```

Sua propriedade `about` deve ter o valor `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Seu objeto deve ter uma propriedade `pages`.

```js
assert.property(library[0], "pages")
```

O valor da sua propriedade `pages` deve ser um número.

```js
assert.isNumber(library[0]?.pages)
```

A sua propriedade `pages` deve ser configurada para `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
