---
id: 67116d7584d0b469b14579c0
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Crie outro objeto dentro do array `library` com as seguintes propriedades e valores:

| Property | Value   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Você deve ter um objeto como o segundo item do seu array `library`.

```js
assert.isObject(library[1]);
```

Seu objeto deve ter uma propriedade `title`.

```js
assert.property(library[1], "title")
```

Sua propriedade `title` deve ter `"Atomic Habits"` como valor.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Seu objeto deve ter uma propriedade `author`.

```js
assert.property(library[1], "author")
```

A sua propriedade `author` deve ter `"James Clear"` como valor.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Seu objeto deve ter uma propriedade `about`.

```js
assert.property(library[1], "about")
```

Sua propriedade `about` deve ter o valor `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Seu objeto deve ter uma propriedade `pages`.

```js
assert.property(library[1], "pages")
```

O valor da sua propriedade `pages` deve ser um número.

```js
assert.isNumber(library[1]?.pages)
```

A sua propriedade `pages` deve ser configurada para `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
