---
id: 67116c175a77cb64c5c10d49
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nas lições anteriores, você aprendeu como trabalhar com funções de ordem superior como `map`, `filter` e `reduce`.

Neste workshop, você vai construir um aplicativo gerenciador de biblioteca que lhe dará a oportunidade de praticar o uso de diferentes funções de ordem superior.

Para começar, crie uma variável chamada `library` e atribua a ela um array vazio.

# --hints--

Você deve criar um array `library`.

```js
assert.isArray(library)
```

Seu array `library` deve estar vazio.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
