---
id: 68225ea2f1aae301fa4df0f4
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Enquanto o array de resultados está funcionando, o resultado final desejado deve ser uma string listando todos os livros com título e autor.

É aqui que o método `join()` entra em ação. Em lições anteriores, você aprendeu que o método `join()` é usado para retornar uma nova string com todos os elementos do array concatenados em uma única string, com um separador especificado entre cada elemento.

Aqui está uma revisão:

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

O separador neste caso é o `\n` que representa o caractere de nova linha. 

Encadeie o método `join()` com um `\n` como separador ao `map()`. Agora você deve ver uma string no console em vez do array de resultados. 

# --hints--

Sua função `getBookInformation` deve retornar uma string.

```js
assert.isString(getBookInformation(library));
```

Sua função `getBookInformation` deve retornar uma única string que representa o resultado final.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`)
}
--fcc-editable-region--

console.log(getBookInformation(library));
```
