---
id: 587d7dba367417b2b2512ba8
title: Auf Vorhandensein aller oder keines prüfen
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

Manchmal können die Muster, nach denen Sie suchen möchten, Teile enthalten, die möglicherweise vorhanden sind oder auch nicht. Dennoch kann es wichtig sein, diese zu überprüfen.

Sie können die mögliche Existenz eines Elements mit einem Fragezeichen `?` angeben. Dies prüft auf null oder eins des vorhergehenden Elements. Man kann sich dieses Symbol so vorstellen, dass das vorherige Element optional ist.

Zum Beispiel gibt es kleine Unterschiede zwischen amerikanischem und britischem Englisch, und Sie können das Fragezeichen verwenden, um beide Schreibweisen zu erfassen.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

Beide Verwendungen der `test`-Methode würden `true` zurückgeben.

# --instructions--

Ändern Sie das Regex `favRegex`, damit es sowohl die amerikanische englische (`favorite`) als auch die britische englische (`favourite`) Version des Wortes erfasst.

# --hints--

Ihr Regex sollte das optionale Symbol `?` verwenden.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

Ihr Regex sollte den String `favorite` erfassen.

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

Ihr Regex sollte den String `favourite` erfassen.

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

Ihr Regex sollte den String `fav` nicht erfassen.

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
