---
id: 587d7db7367417b2b2512b9e
title: Muster am Ende von Strings finden
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

In der letzten Aufgabe haben Sie gelernt, das Caret-Zeichen zu verwenden, um nach Mustern am Anfang von Strings zu suchen. Es gibt auch eine Möglichkeit, nach Mustern am Ende von Strings zu suchen.

Sie können das Ende von Strings mit dem Dollarzeichen `$` am Ende des Regex durchsuchen.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

Der erste `test` Aufruf würde `true` zurückgeben, während der zweite `false` zurückgeben würde.

# --instructions--

Verwenden Sie das Ankerzeichen (`$`), um den String `caboose` am Ende des Strings `caboose` zu finden.

# --hints--

Sie sollten nach `caboose` mit dem Dollarzeichen `$` Anker in Ihrem Regex suchen.

```js
assert(lastRegex.source == 'caboose$');
```

Ihr Regex sollte keine Flags verwenden.

```js
assert(lastRegex.flags == '');
```

Sie sollten `caboose` am Ende des Strings `The last car on a train is the caboose` finden.

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
