---
id: 587d7db3367417b2b2512b8e
title: Die Test-Methode verwenden
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

Reguläre Ausdrücke werden in Programmiersprachen verwendet, um Teile von Strings zu finden. Sie erstellen Muster, die Ihnen bei diesem Abgleich helfen.

Wenn Sie das Wort `the` im String `The dog chased the cat` finden möchten, könnten Sie den folgenden regulären Ausdruck verwenden: `/the/`. Beachten Sie, dass Anführungszeichen innerhalb des regulären Ausdrucks nicht erforderlich sind.

JavaScript bietet mehrere Möglichkeiten, Regex zu verwenden. Eine Möglichkeit, einen Regex zu testen, ist die Verwendung der `.test()`-Methode. Die `.test()`-Methode nimmt den Regex, wendet ihn auf einen String an (der in die runden Klammern gesetzt wird) und gibt `true` oder `false` zurück, je nachdem, ob Ihr Muster etwas findet oder nicht.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

Die `test`-Methode hier gibt `true` zurück.

# --instructions--

Wenden Sie den Regex `myRegex` auf den String `myString` mit der `.test()`-Methode an.

# --hints--

Sie sollten `.test()` verwenden, um den Regex zu testen.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Ihr Ergebnis sollte `true` zurückgeben.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
