---
id: 587d7db9367417b2b2512ba7
title: Exakte Anzahl von Übereinstimmungen angeben
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

Sie können die untere und obere Anzahl von Mustern mit Mengenangaben in geschweiften Klammern angeben. Manchmal möchten Sie nur eine bestimmte Anzahl von Übereinstimmungen.

Um eine bestimmte Anzahl von Mustern anzugeben, setzen Sie einfach diese eine Zahl zwischen die geschweiften Klammern.

Zum Beispiel, um nur das Wort `hah` mit dem Buchstaben `a` `3` Mal zu finden, wäre Ihr Regex `/ha{3}h/`.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

In der Reihenfolge würden die drei `test`-Aufrufe `false`, `true` und `false` zurückgeben.

# --instructions--

Ändern Sie den Regex `timRegex`, damit er das Wort `Timber` nur dann findet, wenn es genau vier Buchstaben `m` enthält.

# --hints--

Ihr Regex sollte geschweifte Klammern verwenden.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

Ihr Regex sollte den String `Timber` nicht finden.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

Ihr Regex sollte den String `Timmber` nicht finden.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

Ihr Regex sollte den String `Timmmber` nicht finden.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

Ihr Regex sollte den String `Timmmmber` finden.

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

Ihr Regex sollte den String `Timber` mit 30 `m`-Zeichen darin nicht finden.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
