---
id: 587d7db7367417b2b2512b9f
title: Alle Buchstaben und Zahlen finden
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

Mit Zeichenklassen konnten Sie mit `[a-z]` nach allen Buchstaben des Alphabets suchen. Diese Art von Zeichenklasse ist so gebräuchlich, dass es eine Abkürzung dafür gibt, die allerdings auch einige zusätzliche Zeichen einschließt.

Die nächstliegende Zeichenklasse in JavaScript, um das Alphabet zu erfassen, ist `\w`. Diese Abkürzung entspricht `[A-Za-z0-9_]`. Diese Zeichenklasse erfasst Groß- und Kleinbuchstaben sowie Zahlen. Beachten Sie, dass diese Zeichenklasse auch den Unterstrich (`_`) einschließt.

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

Alle vier dieser `test`-Aufrufe würden `true` zurückgeben.

Diese Abkürzungen für Zeichenklassen werden auch als <dfn>shorthand character classes</dfn> bezeichnet.

# --instructions--

Verwenden Sie die Abkürzung `\w` für Zeichenklassen, um die Anzahl der alphanumerischen Zeichen in verschiedenen Zitaten und Strings zu zählen.

# --hints--

Ihr Regex sollte das globale Flag verwenden.

```js
assert(alphabetRegexV2.global);
```

Ihr Regex sollte die Abkürzung `\w` verwenden, um alle alphanumerischen Zeichen zu erfassen.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

Ihr Regex sollte 31 alphanumerische Zeichen im String `The five boxing wizards jump quickly.` finden.

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

Ihr Regex sollte 32 alphanumerische Zeichen im String `Pack my box with five dozen liquor jugs.` finden.

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

Ihr Regex sollte 30 alphanumerische Zeichen im String `How vexingly quick daft zebras jump!` finden.

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

Ihr Regex sollte 36 alphanumerische Zeichen im String `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.` finden.

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
