---
id: 587d7db7367417b2b2512b9d
title: Übereinstimmung von Anfangsmustern in Strings
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

Frühere Aufgaben zeigten, dass reguläre Ausdrücke verwendet werden können, um mehrere Übereinstimmungen zu finden. Sie werden auch genutzt, um nach Mustern an bestimmten Positionen in Strings zu suchen.

In einer früheren Aufgabe haben Sie das Caret-Zeichen (`^`) innerhalb eines Zeichensatzes verwendet, um einen negierten Zeichensatz in der Form `[^thingsThatWillNotBeMatched]` zu erstellen. Außerhalb eines Zeichensatzes wird das Caret verwendet, um nach Mustern am Anfang von Strings zu suchen.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

Der erste `test`-Aufruf würde `true` zurückgeben, während der zweite `false` zurückgeben würde.

# --instructions--

Verwenden Sie das Caret-Zeichen in einem Regex, um `Cal` nur am Anfang des Strings `rickyAndCal` zu finden.

# --hints--

Ihr Regex sollte nach dem String `Cal` mit einem Großbuchstaben suchen.

```js
assert(calRegex.source == '^Cal');
```

Ihr Regex sollte keine Flags verwenden.

```js
assert(calRegex.flags == '');
```

Ihr Regex sollte den String `Cal` am Anfang des Strings finden.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

Ihr Regex sollte den String `Cal` nicht in der Mitte eines Strings finden.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
