---
id: 587d7db8367417b2b2512ba1
title: Alle Nicht-Zahlen finden
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

Die letzte Aufgabe zeigte, wie man mit der Abkürzung `\d` und einem Kleinbuchstaben `d` nach Ziffern sucht. Sie können auch nach Nicht-Ziffern suchen, indem Sie eine ähnliche Abkürzung verwenden, die stattdessen einen Großbuchstaben `D` nutzt.

Die Abkürzung, um nach Nicht-Ziffern zu suchen, ist `\D`. Diese entspricht der Zeichenklasse `[^0-9]`, die nach einem einzelnen Zeichen sucht, das keine Zahl zwischen null und neun ist.

# --instructions--

Verwenden Sie die Kurzzeichenklasse für Nicht-Ziffern `\D`, um zu zählen, wie viele Nicht-Ziffern in Filmtiteln vorkommen.

# --hints--

Ihr Regex sollte das Kurzzeichen verwenden, um Nicht-Ziffern zu finden

```js
assert(/\\D/.test(noNumRegex.source));
```

Ihr Regex sollte das globale Flag verwenden.

```js
assert(noNumRegex.global);
```

Ihr Regex sollte keine Nicht-Ziffern im String `9` finden.

```js
assert('9'.match(noNumRegex) == null);
```

Ihr Regex sollte 6 Nicht-Ziffern im String `Catch 22` finden.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

Ihr Regex sollte 11 Nicht-Ziffern im String `101 Dalmatians` finden.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

Ihr Regex sollte 15 Nicht-Ziffern im String `One, Two, Three` finden.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

Ihr Regex sollte 12 Nicht-Ziffern im String `21 Jump Street` finden.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

Ihr Regex sollte 17 Nicht-Ziffern im String `2001: A Space Odyssey` finden.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
