---
id: 587d7db3367417b2b2512b8f
title: Literalstrings finden
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

In der letzten Aufgabe haben Sie mit dem regulären Ausdruck `Hello` nach dem Wort `/Hello/` gesucht. Dieser reguläre Ausdruck suchte nach einer exakten Übereinstimmung des Strings `Hello`. Hier ist ein weiteres Beispiel, das nach einer exakten Übereinstimmung des Strings `Kevin` sucht:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

Dieser Aufruf `test` gibt `true` zurück.

Alle anderen Formen von `Kevin` werden nicht übereinstimmen. Zum Beispiel wird der reguläre Ausdruck `/Kevin/` nicht mit `kevin` oder `KEVIN` übereinstimmen.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

Dieser Aufruf `test` gibt `false` zurück.

Eine zukünftige Aufgabe zeigt, wie man auch diese anderen Formen findet.

# --instructions--

Vervollständigen Sie den regulären Ausdruck `waldoRegex`, um `"Waldo"` im String `waldoIsHiding` mit einer exakten Übereinstimmung zu finden.

# --hints--

Ihr regulärer Ausdruck `waldoRegex` sollte den String `Waldo` finden.

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

Ihr regulärer Ausdruck `waldoRegex` sollte nichts anderes suchen.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

Sie sollten mit Ihrem regulären Ausdruck eine exakte String-Übereinstimmung durchführen.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
