---
id: 587d7db6367417b2b2512b9a
title: "Zeichen finden, die null- oder mehrmals vorkommen"
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

Die letzte Aufgabe verwendete das Pluszeichen `+`, um nach Zeichen zu suchen, die einmal oder öfter vorkommen. Es gibt auch eine Möglichkeit, Zeichen zu finden, die null- oder mehrmals vorkommen.

Das Zeichen dafür ist der Stern oder Asterisk: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

In der Reihenfolge würden die drei `match`-Aufrufe die Werte `["goooooooo"]`, `["g"]` und `null` zurückgeben.

# --instructions--

Erstellen Sie einen Regex `chewieRegex`, der das `*`-Zeichen verwendet, um ein Großbuchstaben-`A`-Zeichen zu finden, dem null oder mehr Kleinbuchstaben-`a`-Zeichen in `chewieQuote` unmittelbar folgen. Ihr Regex benötigt keine Flags oder Zeichenklassen und sollte keine der anderen Zitate erfassen.

# --hints--

Ihr Regex `chewieRegex` sollte das `*`-Zeichen verwenden, um null oder mehr `a`-Zeichen zu erfassen.

```js
assert(/\*/.test(chewieRegex.source));
```

Ihr Regex sollte den String `A` in `chewieQuote` erfassen.

```js
assert(result[0][0] === 'A');
```

Ihr Regex sollte den String `Aaaaaaaaaaaaaaaa` in `chewieQuote` erfassen.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

Ihr Regex `chewieRegex` sollte 16 Zeichen in `chewieQuote` erfassen.

```js
assert(result[0].length === 16);
```

Ihr Regex sollte keine Zeichen im String `He made a fair move. Screaming about it can't help you.` erfassen.

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

Ihr Regex sollte keine Zeichen im String `Let him have it. It's not wise to upset a Wookiee.` erfassen.

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
