---
id: 587d7db4367417b2b2512b92
title: Übereinstimmungen extrahieren
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

Bisher haben Sie nur überprüft, ob ein Muster in einem String vorhanden ist oder nicht. Sie können auch die tatsächlichen Übereinstimmungen, die Sie gefunden haben, mit der `.match()`-Methode extrahieren.

Um die `.match()`-Methode zu verwenden, wenden Sie die Methode auf einen String an und übergeben das Regex innerhalb der runden Klammern.

Hier ist ein Beispiel:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

Hier würde die erste `match` `["Hello"]` zurückgeben und die zweite `["expressions"]`.

Beachten Sie, dass die `.match`-Syntax das „Gegenteil“ der `.test`-Methode ist, die Sie bisher verwendet haben:

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

Wenden Sie die `.match()`-Methode an, um den String `coding` zu extrahieren.

# --hints--

Die `result` sollte den String `coding` enthalten

```js
assert(result.join() === 'coding');
```

Ihr Regex `codingRegex` sollte nach dem String `coding` suchen

```js
assert(codingRegex.source === 'coding');
```

Sie sollten die `.match()`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
