---
id: 587d7db5367417b2b2512b95
title: Einzelnes Zeichen mit mehreren Möglichkeiten abgleichen
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

Sie haben gelernt, wie man wörtliche Muster (`/literal/`) und Platzhalterzeichen (`/./`) abgleicht. Das sind die Extreme bei regulären Ausdrücken: Das eine findet exakte Übereinstimmungen, das andere passt auf alles. Es gibt Optionen, die einen Mittelweg zwischen diesen Extremen darstellen.

Sie können nach einem wörtlichen Muster mit etwas Flexibilität suchen, indem Sie <dfn>Zeichenklassen</dfn> verwenden. Zeichenklassen erlauben es Ihnen, eine Gruppe von Zeichen zu definieren, die Sie abgleichen möchten, indem Sie sie in eckige Klammern (`[` und `]`) setzen.

Zum Beispiel möchten Sie `bag`, `big` und `bug` abgleichen, aber nicht `bog`. Sie können dafür den Regex `/b[aiu]g/` erstellen. Die `[aiu]` ist die Zeichenklasse, die nur die Zeichen `a`, `i` oder `u` abgleicht.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

In der Reihenfolge würden die vier `match`-Aufrufe die Werte `["big"]`, `["bag"]`, `["bug"]` und `null` zurückgeben.

# --instructions--

Verwenden Sie eine Zeichenklasse mit den Vokalen (`a`, `e`, `i`, `o`, `u`) in Ihrem Regex `vowelRegex`, um alle Vokale im String `quoteSample` zu finden.

**Hinweis:** Achten Sie darauf, sowohl Groß- als auch Kleinbuchstaben der Vokale abzugleichen.

# --hints--

Sie sollten alle 25 Vokale finden.

```js
assert(result.length == 25);
```

Ihr Regex `vowelRegex` sollte eine Zeichenklasse verwenden.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

Ihr Regex `vowelRegex` sollte das globale Flag verwenden.

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

Ihr Regex `vowelRegex` sollte das Flag für Groß-/Kleinschreibung ignorieren verwenden.

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

Ihr Regex sollte keine Konsonanten abgleichen.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
