---
id: 587d7db8367417b2b2512ba3
title: Leerzeichen erkennen
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

Die bisherigen Aufgaben haben das Erkennen von Buchstaben des Alphabets und Zahlen behandelt. Sie können auch Leerzeichen oder Abstände zwischen Buchstaben erkennen.

Sie können nach Leerzeichen suchen, indem Sie `\s` verwenden, was ein Kleinbuchstabe `s` ist. Dieses Muster erkennt nicht nur Leerzeichen, sondern auch Wagenrücklauf, Tabulator, Seitenvorschub und Zeilenumbruchzeichen. Man kann es sich ähnlich vorstellen wie die Zeichenklasse `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

Dieser `match`-Aufruf würde `[" ", " "]` zurückgeben.
# --instructions--

Ändern Sie den Regex `countWhiteSpace`, damit er mehrere Leerzeichen in einem String erkennt.

# --hints--

Ihr Regex sollte das globale Flag verwenden.

```js
assert(countWhiteSpace.global);
```

Ihr Regex sollte die Kurzschreibweise `\s` verwenden, um alle Leerzeichen zu erkennen.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

Ihr Regex sollte acht Leerzeichen im String `Men are from Mars and women are from Venus.` finden.

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

Ihr Regex sollte drei Leerzeichen im String `Space: the final frontier.` finden.

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

Ihr Regex sollte keine Leerzeichen im String `MindYourPersonalSpace` finden.

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
