---
id: 587d7db4367417b2b2512b93
title: Mehr als nur den ersten Treffer finden
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

Bisher konnten Sie ein Muster nur einmal extrahieren oder suchen.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

Hier würde `match` `["Repeat"]` zurückgeben.

Um ein Muster mehr als einmal zu suchen oder zu extrahieren, können Sie das globale Such-Flag verwenden: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

Und hier gibt `match` den Wert `["Repeat", "Repeat", "Repeat"]` zurück.

# --instructions--

Verwenden Sie den Regex `starRegex`, um beide `Twinkle` Wörter aus dem String `twinkleStar` zu finden und zu extrahieren.

**Hinweis**  
Sie können mehrere Flags in Ihrem Regex verwenden, wie `/search/gi`

# --hints--

Ihr Regex `starRegex` sollte das globale Flag `g` verwenden.

```js
assert(starRegex.flags.match(/g/).length == 1);
```

Ihr Regex `starRegex` sollte das Flag für Groß-/Kleinschreibung `i` verwenden.

```js
assert(starRegex.flags.match(/i/).length == 1);
```

Ihr Treffer sollte beide Vorkommen des Wortes `Twinkle` erfassen.

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

Ihr Treffer `result` sollte zwei Elemente enthalten.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
