---
id: 587d7db5367417b2b2512b96
title: Buchstaben des Alphabets abgleichen
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

Sie haben gesehen, wie Sie <dfn>Zeichensätze</dfn> verwenden können, um eine Gruppe von Zeichen zum Abgleichen anzugeben, aber das ist viel Tipparbeit, wenn Sie einen großen Bereich von Zeichen abgleichen müssen (zum Beispiel jeden Buchstaben im Alphabet). Glücklicherweise gibt es eine eingebaute Funktion, die das kurz und einfach macht.

Innerhalb eines Zeichensatzes können Sie einen Bereich von Zeichen zum Abgleichen mit einem Bindestrich definieren: `-`.

Zum Beispiel würden Sie, um Kleinbuchstaben `a` bis `e` abzugleichen, `[a-e]` verwenden.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

In der Reihenfolge würden die drei `match`-Aufrufe die Werte `["cat"]`, `["bat"]` und `null` zurückgeben.

# --instructions--

Gleichen Sie alle Buchstaben im String `quoteSample` ab.

**Hinweis**: Achten Sie darauf, sowohl Groß- als auch Kleinbuchstaben abzugleichen.

# --hints--

Ihr Regex `alphabetRegex` sollte 35 Treffer ergeben.

```js
assert(result.length == 35);
```

Ihr Regex `alphabetRegex` sollte das globale Flag verwenden.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

Ihr Regex `alphabetRegex` sollte das Flag für Groß-/Kleinschreibung ignorieren verwenden.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
