---
id: 587d7dbb367417b2b2512bab
title: Capture-Gruppen zum Suchen und Ersetzen verwenden
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

Suche ist nützlich. Sie können die Suche jedoch noch mächtiger machen, wenn sie auch den gefundenen Text ändert (oder ersetzt).

Sie können Text in einem String mit `.replace()` suchen und ersetzen. Die Eingaben für `.replace()` sind zuerst das Regex-Muster, nach dem Sie suchen möchten. Der zweite Parameter ist der String, der die Übereinstimmung ersetzt, oder eine Funktion, die etwas ausführt.

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

Der `replace`-Aufruf würde den String `The sky is blue.` zurückgeben.

Sie können auch auf Capture-Gruppen im Ersetzungsstring mit Dollarzeichen (`$`) zugreifen.

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

Der `replace`-Aufruf würde den String `Camp Code` zurückgeben.

# --instructions--

Schreiben Sie ein Regex `fixRegex` mit drei Capture-Gruppen, das nach jedem Wort im String `one two three` sucht. Aktualisieren Sie dann die Variable `replaceText`, um `one two three` durch den String `three two one` zu ersetzen, und weisen Sie das Ergebnis der Variable `result` zu. Stellen Sie sicher, dass Sie Capture-Gruppen im Ersetzungsstring mit der Dollarzeichen-Syntax (`$`) verwenden.

# --hints--

Sie sollten `.replace()` zum Suchen und Ersetzen verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

Ihr Regex sollte den String `one two three` in den String `three two one` ändern.

```js
assert(result === 'three two one');
```

Sie sollten die letzte Zeile nicht ändern.

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

`fixRegex` sollte mindestens drei Capture-Gruppen verwenden.

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

`replaceText` sollte parenthesierte Teilübereinstimmungsstrings verwenden (d. h. der n-te parenthesierte Teilübereinstimmungsstring, $n, entspricht der n-ten Capture-Gruppe).

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
