---
id: 587d7db9367417b2b2512ba6
title: Nur die untere Anzahl der Übereinstimmungen angeben
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

Sie können die untere und obere Anzahl von Mustern mit Mengenangaben in geschweiften Klammern angeben. Manchmal möchten Sie nur die untere Anzahl von Mustern ohne obere Grenze angeben.

Um nur die untere Anzahl von Mustern anzugeben, behalten Sie die erste Zahl gefolgt von einem Komma bei.

Zum Beispiel, um nur den String `hah` zu finden, bei dem der Buchstabe `a` mindestens `3` Mal vorkommt, wäre Ihr Regex `/ha{3,}h/`.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

In der Reihenfolge würden die drei `test`-Aufrufe `true`, `false` und `true` zurückgeben.

# --instructions--

Ändern Sie den Regex `haRegex`, damit er das Wort `Hazzah` nur dann findet, wenn es vier oder mehr Buchstaben `z` enthält.

# --hints--

Ihr Regex sollte geschweifte Klammern verwenden.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

Ihr Regex sollte den String `Hazzah` nicht finden.

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

Ihr Regex sollte den String `Hazzzah` nicht finden.

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

Ihr Regex sollte den String `Hazzzzah` finden.

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

Ihr Regex sollte den String `Hazzzzzah` finden.

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

Ihr Regex sollte den String `Hazzzzzzah` finden.

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

Ihr Regex sollte den String `Hazzah` mit 30 `z`-Zeichen darin finden.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
