---
id: 587d7db8367417b2b2512ba2
title: Mögliche Benutzernamen einschränken
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

Benutzernamen werden überall im Internet verwendet. Sie geben den Nutzenden auf ihren Lieblingsseiten eine einzigartige Identität.

Sie müssen alle Benutzernamen in einer Datenbank überprüfen. Hier sind einige einfache Regeln, die Nutzende beim Erstellen ihres Benutzernamens befolgen müssen.

1) Benutzernamen dürfen nur alphanumerische Zeichen enthalten.

2) Die einzigen Zahlen im Benutzernamen dürfen nur am Ende stehen. Dort können null oder mehr Zahlen stehen. Der Benutzername darf nicht mit einer Zahl beginnen.

3) Die Buchstaben im Benutzernamen können Klein- oder Großbuchstaben sein.

4) Benutzernamen müssen mindestens zwei Zeichen lang sein. Ein Benutzername mit zwei Zeichen darf nur Buchstaben enthalten.

# --instructions--

Ändern Sie den Regex `userCheck`, damit er den oben genannten Vorgaben entspricht.

# --hints--

Ihr Regex sollte den String `JACK` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

Ihr Regex sollte den String `J` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

Ihr Regex sollte den String `Jo` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

Ihr Regex sollte den String `Oceans11` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

Ihr Regex sollte den String `RegexGuru` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

Ihr Regex sollte den String `007` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

Ihr Regex sollte den String `9` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

Ihr Regex sollte den String `A1` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

Ihr Regex sollte den String `BadUs3rnam3` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

Ihr Regex sollte den String `Z97` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

Ihr Regex sollte den String `c57bT3` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

Ihr Regex sollte den String `AB1` erfassen

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

Ihr Regex sollte den String `J%4` nicht erfassen

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
