---
id: 587d7dbb367417b2b2512bac
title: Leerzeichen am Anfang und Ende entfernen
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

Manchmal sind Leerzeichen um Strings herum unerwünscht, aber vorhanden. Typische Verarbeitung von Strings besteht darin, die Leerzeichen am Anfang und Ende zu entfernen.

# --instructions--

Schreiben Sie einen Regex und verwenden Sie die passenden String-Methoden, um Leerzeichen am Anfang und Ende von Strings zu entfernen.

**Hinweis:** Die `String.prototype.trim()`-Methode würde hier funktionieren, aber Sie müssen diese Aufgabe mit regulären Ausdrücken lösen.

# --hints--

`result` sollte gleich dem String `Hello, World!` sein

```js
assert(result === 'Hello, World!');
```

Ihre Lösung sollte die `String.prototype.trim()`-Methode nicht verwenden.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

Die Variable `result` sollte nicht direkt auf einen String gesetzt werden

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

Der Wert der Variable `hello` sollte nicht verändert werden.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
