---
id: 587d7db8367417b2b2512ba0
title: Alles außer Buchstaben und Zahlen finden
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

Sie haben gelernt, dass Sie eine Abkürzung verwenden können, um alphanumerische Zeichen `[A-Za-z0-9_]` mit `\w` zu finden. Ein natürliches Muster, nach dem Sie suchen könnten, ist das Gegenteil von alphanumerischen Zeichen.

Sie können mit `\w` nach dem Gegenteil von `\W` suchen. Beachten Sie, dass das gegenteilige Muster einen Großbuchstaben verwendet. Diese Abkürzung entspricht `[^A-Za-z0-9_]`.

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

Der erste `match`-Aufruf würde den Wert `["%"]` zurückgeben und der zweite den Wert `["!"]`.

# --instructions--

Verwenden Sie die Kurzzeichenklasse `\W`, um die Anzahl der nicht alphanumerischen Zeichen in verschiedenen Zitaten und Strings zu zählen.

# --hints--

Ihr Regex sollte das globale Flag verwenden.

```js
assert(nonAlphabetRegex.global);
```

Ihr Regex sollte 6 nicht alphanumerische Zeichen im String `The five boxing wizards jump quickly.` finden.

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

Ihr Regex sollte das Kurzzeichen verwenden, um Zeichen zu finden, die nicht alphanumerisch sind.

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

Ihr Regex sollte 8 nicht alphanumerische Zeichen im String `Pack my box with five dozen liquor jugs.` finden.

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

Ihr Regex sollte 6 nicht alphanumerische Zeichen im String `How vexingly quick daft zebras jump!` finden.

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

Ihr Regex sollte 12 nicht alphanumerische Zeichen im String `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.` finden.

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
