---
id: 587d7db4367417b2b2512b90
title: Einen Literalstring mit verschiedenen MĂ¶glichkeiten abgleichen
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

Mit Regexes wie `/coding/` kĂ¶nnen Sie nach dem Muster `coding` in einem anderen String suchen.

Das ist sehr nĂĽtzlich, um einzelne Strings zu durchsuchen, aber es ist auf nur ein Muster beschrĂ¤nkt. Sie kĂ¶nnen nach mehreren Mustern suchen, indem Sie den `alternation`- oder `OR`-Operator verwenden: `|`.

Dieser Operator stimmt mit Mustern ĂĽberein, die entweder davor oder danach stehen. Wenn Sie zum Beispiel die Strings `yes` oder `no` abgleichen mĂ¶chten, ist der Regex, den Sie brauchen, `/yes|no/`.

Sie kĂ¶nnen auch nach mehr als nur zwei Mustern suchen. Das erreichen Sie, indem Sie weitere Muster mit weiteren `OR`-Operatoren dazwischen hinzufĂĽgen, wie `/yes|no|maybe/`.

# --instructions--

VervollstĂ¤ndigen Sie den Regex `petRegex`, um die Haustiere `dog`, `cat`, `bird` oder `fish` abzugleichen.

# --hints--

Ihr Regex `petRegex` sollte `true` fĂĽr den String `John has a pet dog.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

Ihr Regex `petRegex` sollte `false` fĂĽr den String `Emma has a pet rock.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

Ihr Regex `petRegex` sollte `true` fĂĽr den String `Emma has a pet bird.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

Ihr Regex `petRegex` sollte `true` fĂĽr den String `Liz has a pet cat.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

Ihr Regex `petRegex` sollte `false` fĂĽr den String `Kara has a pet dolphin.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

Ihr Regex `petRegex` sollte `true` fĂĽr den String `Alice has a pet fish.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

Ihr Regex `petRegex` sollte `false` fĂĽr den String `Jimmy has a pet computer.` zurĂĽckgeben

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
