---
id: 587d7db5367417b2b2512b94
title: Mit Platzhalterpunkt alles finden
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

Manchmal wissen Sie nicht (oder müssen nicht wissen), welche genauen Zeichen in Ihren Mustern vorkommen. Sich alle Wörter vorzustellen, die beispielsweise eine falsche Schreibweise enthalten, würde lange dauern. Glücklicherweise können Sie Zeit sparen, indem Sie das Platzhalterzeichen verwenden: `.`

Das Platzhalterzeichen `.` entspricht genau einem beliebigen Zeichen. Der Platzhalter wird auch `dot` und `period` genannt. Sie können das Platzhalterzeichen genauso verwenden wie jedes andere Zeichen im Regex. Wenn Sie zum Beispiel `hug`, `huh`, `hut` und `hum` finden möchten, können Sie den Regex `/hu./` verwenden, um alle vier Wörter zu erfassen.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Beide `test`-Aufrufe würden `true` zurückgeben.

# --instructions--

Vervollständigen Sie den Regex `unRegex`, sodass er die Strings `run`, `sun`, `fun`, `pun`, `nun` und `bun` findet. Ihr Regex soll das Platzhalterzeichen verwenden.

# --hints--

Sie sollten die `.test()`-Methode verwenden.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Sie sollten das Platzhalterzeichen in Ihrem Regex `unRegex` verwenden.

```js
assert(/\./.test(unRegex.source));
```

Ihr Regex `unRegex` sollte `run` im String `Let us go on a run.` finden.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

Ihr Regex `unRegex` sollte `sun` im String `The sun is out today.` finden.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

Ihr Regex `unRegex` sollte `fun` im String `Coding is a lot of fun.` finden.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

Ihr Regex `unRegex` sollte `pun` im String `Seven days without a pun makes one weak.` finden.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

Ihr Regex `unRegex` sollte `nun` im String `One takes a vow to be a nun.` finden.

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

Ihr Regex `unRegex` sollte `bun` im String `She got fired from the hot dog stand for putting her hair in a bun.` finden.

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

Ihr Regex `unRegex` sollte den String `There is a bug in my code.` nicht finden.

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

Ihr Regex `unRegex` sollte den String `Catch me if you can.` nicht finden.

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
