---
id: 587d7db9367417b2b2512ba4
title: Suche nach Nicht-Leerzeichen
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

Sie haben gelernt, wie man nach Leerzeichen mit `\s` sucht, mit einem Kleinbuchstaben `s`. Sie können auch nach allem außer Leerzeichen suchen.

Suchen Sie nach Nicht-Leerzeichen mit `\S`, das ein Großbuchstabe `s` ist. Dieses Muster findet keine Leerzeichen, Wagenrücklauf, Tabulator, Seitenvorschub und Zeilenumbruchzeichen. Man kann es sich ähnlich vorstellen wie die Zeichenklasse `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

Der Wert, der von der `.length`-Methode zurückgegeben wird, wäre `32`.

# --instructions--

Ändern Sie den Regex `countNonWhiteSpace`, um nach mehreren Nicht-Leerzeichen in einem String zu suchen.

# --hints--

Ihr Regex sollte das globale Flag verwenden.

```js
assert(countNonWhiteSpace.global);
```

Ihr Regex sollte die Kurzschreibweise `\S` verwenden, um alle Nicht-Leerzeichen zu finden.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

Ihr Regex sollte 35 Nicht-Leerzeichen im String `Men are from Mars and women are from Venus.` finden.

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

Ihr Regex sollte 23 Nicht-Leerzeichen im String `Space: the final frontier.` finden.

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

Ihr Regex sollte 21 Nicht-Leerzeichen im String `MindYourPersonalSpace` finden.

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
