---
id: 587d7db5367417b2b2512b97
title: Zahlen und Buchstaben des Alphabets abgleichen
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

Die Verwendung des Bindestrichs (`-`) zum Abgleichen eines Zeichenbereichs ist nicht auf Buchstaben beschränkt. Er funktioniert auch, um einen Bereich von Zahlen abzugleichen.

Zum Beispiel stimmt `/[0-5]/` mit jeder Zahl zwischen `0` und `5` überein, einschließlich `0` und `5`.

Außerdem ist es möglich, einen Bereich von Buchstaben und Zahlen in einem einzigen Zeichensatz zu kombinieren.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

Erstellen Sie einen einzelnen Regex, der einen Bereich von Buchstaben zwischen `h` und `s` sowie einen Bereich von Zahlen zwischen `2` und `6` abgleicht. Denken Sie daran, die passenden Flags im Regex einzuschließen.

# --hints--

Ihr Regex `myRegex` sollte 17 Elemente abgleichen.

```js
assert(result.length == 17);
```

Ihr Regex `myRegex` sollte das globale Flag verwenden.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Ihr Regex `myRegex` sollte das Flag für Groß- und Kleinschreibung ignorieren verwenden.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
