---
id: 587d7db9367417b2b2512ba5
title: Ober- und Untergrenze der Anzahl von Übereinstimmungen angeben
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

Erinnern Sie sich, dass Sie das Pluszeichen `+` verwenden, um nach einem oder mehreren Zeichen zu suchen, und den Stern `*`, um nach null oder mehr Zeichen zu suchen. Diese sind praktisch, aber manchmal möchten Sie einen bestimmten Bereich von Mustern abgleichen.

Sie können die untere und obere Anzahl von Mustern mit <dfn>Mengenangaben</dfn> festlegen. Mengenangaben werden mit geschweiften Klammern (`{` und `}`) verwendet. Sie setzen zwei Zahlen zwischen die geschweiften Klammern – für die untere und obere Anzahl von Mustern.

Zum Beispiel, um nur den Buchstaben `a` zu finden, der zwischen `3` und `5` Mal im String `ah` vorkommt, wäre Ihr Regex `/a{3,5}h/`.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

Der erste `test`-Aufruf würde `true` zurückgeben, während der zweite `false` zurückgeben würde.

# --instructions--

Ändern Sie den Regex `ohRegex`, sodass er den gesamten Ausdruck `Oh no` nur dann abgleicht, wenn er `3` bis `6` Buchstaben `h` enthält.

# --hints--

Ihr Regex sollte geschweifte Klammern verwenden.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

Ihr Regex sollte den String `Ohh no` nicht abgleichen.

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

Ihr Regex sollte den String `Ohhh no` abgleichen.

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

Ihr Regex sollte den String `Ohhhh no` abgleichen.

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

Ihr Regex sollte den String `Ohhhhh no` abgleichen.

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

Ihr Regex sollte den String `Ohhhhhh no` abgleichen.

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

Ihr Regex sollte den String `Ohhhhhhh no` nicht abgleichen.

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
