---
id: 587d7db7367417b2b2512b9c
title: Finde einen oder mehrere Kriminelle bei einer Suche
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

Zeit, eine Pause einzulegen und Ihre neuen Regex-Schreibfähigkeiten zu testen. Eine Gruppe von Kriminellen ist aus dem Gefängnis ausgebrochen und weggelaufen, aber Sie wissen nicht, wie viele es sind. Sie wissen jedoch, dass sie dicht beieinander bleiben, wenn sie sich in der Nähe anderer Menschen aufhalten. Sie sind dafür verantwortlich, alle Kriminellen auf einmal zu finden.

Hier ist ein Beispiel, um zu überprüfen, wie das funktioniert:

Der Regex `/z+/` findet den Buchstaben `z`, wenn er einmal oder mehrmals hintereinander vorkommt. Er würde in allen folgenden Strings Übereinstimmungen finden:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

Aber er findet keine Übereinstimmungen in den folgenden Strings, da keine Buchstaben `z` vorhanden sind:

```js
""
"ABC"
"abcabc"
```

# --instructions--

Schreiben Sie einen gierigen Regex, der einen oder mehrere Kriminelle innerhalb einer Gruppe anderer Personen findet. Ein Krimineller wird durch den Großbuchstaben `C` dargestellt.

# --hints--

Ihr Regex sollte einen Kriminellen (`C`) im String `C` finden

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

Ihr Regex sollte zwei Kriminelle (`CC`) im String `CC` finden

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

Ihr Regex sollte drei Kriminelle (`CCC`) im String `P1P5P4CCCcP2P6P3` finden.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

Ihr Regex sollte fünf Kriminelle (`CCCCC`) im String `P6P2P7P4P5CCCCCP3P1` finden

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

Ihr Regex sollte keine Kriminellen im leeren String `""` finden

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

Ihr Regex sollte keine Kriminellen im String `P1P2P3` finden

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

Ihr Regex sollte fünfzig Kriminelle (`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`) im String `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3` finden.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
