---
id: 587d7db6367417b2b2512b98
title: Einzelne Zeichen ausschließen
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

Bisher haben Sie eine Menge von Zeichen erstellt, die Sie abgleichen möchten, aber Sie können auch eine Menge von Zeichen erstellen, die Sie nicht abgleichen möchten. Diese Arten von Zeichensätzen nennt man <dfn>negierte Zeichensätze</dfn>.

Um einen negierten Zeichensatz zu erstellen, setzen Sie ein Caret-Zeichen (`^`) direkt nach der öffnenden eckigen Klammer und vor die Zeichen, die Sie nicht abgleichen möchten.

Zum Beispiel stimmt `/[^aeiou]/gi` mit allen Zeichen überein, die kein Vokal sind. Beachten Sie, dass Zeichen wie `.`, `!`, `[`, `@`, `/` und Leerzeichen übereinstimmen – der negierte Vokal-Zeichensatz schließt nur die Vokalzeichen aus.

# --instructions--

Erstellen Sie einen einzelnen Regex, der alle Zeichen abgleicht, die keine Zahl oder kein Vokal sind. Denken Sie daran, die passenden Flags im Regex zu verwenden.

# --hints--

Ihr Regex `myRegex` sollte 9 Elemente abgleichen.

```js
assert(result.length == 9);
```

Ihr Regex `myRegex` sollte das globale Flag verwenden.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Ihr Regex `myRegex` sollte das Flag für Groß-/Kleinschreibung ignorieren verwenden.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
