---
id: 5c3dda8b4d8df89bea71600f
title: Überprüfung gemischter Gruppierungen von Zeichen
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

Manchmal möchten wir mit einem regulären Ausdruck Gruppen von Zeichen überprüfen, und dafür verwenden wir runde Klammern `()`.

Wenn Sie entweder `Penguin` oder `Pumpkin` in einem String finden möchten, können Sie den folgenden regulären Ausdruck verwenden: `/P(engu|umpk)in/g`

Dann prüfen Sie mit der Methode `test()`, ob die gewünschten Zeichen-Gruppen im Test-String enthalten sind.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

Die Methode `test` würde hier `true` zurückgeben.

# --instructions--

Korrigieren Sie den Regex so, dass er die Namen `Franklin Roosevelt` oder `Eleanor Roosevelt` in Groß- und Kleinschreibung unterscheidet und dabei auch Mittelnamen berücksichtigt.

Dann korrigieren Sie den Code so, dass der von Ihnen erstellte Regex gegen `myString` geprüft wird und je nachdem, ob der Regex übereinstimmt, entweder `true` oder `false` zurückgegeben wird.

# --hints--

Ihr Regex `myRegex` sollte für den String `true` `Franklin D. Roosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

Ihr Regex `myRegex` sollte für den String `true` `Eleanor Roosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

Ihr Regex `myRegex` sollte für den String `false` `Franklin Rosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

Ihr Regex `myRegex` sollte für den String `false` `Frank Roosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

Ihr Regex `myRegex` sollte für den String `false` `FranklinRoosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

Ihr Regex `myRegex` sollte für den String `false` `EleanorRoosevelt` zurückgeben.

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

Sie sollten `.test()` verwenden, um den Regex zu testen.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Ihr Ergebnis sollte `true` zurückgeben.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
