---
id: 681cb1b0dab50c87ddb2e518
title: "Challenge 7: Targeted Sum"
challengeType: 29
dashedName: challenge-7
---

# --description--

Gegeben ist ein Array von Zahlen und ein ganzzahliges Ziel. Finden Sie zwei eindeutige Zahlen im Array, deren Summe dem Zielwert entspricht. Geben Sie ein Array mit den Indizes dieser beiden Zahlen zurück, oder `"Target not found"`, wenn keine zwei Zahlen die Summe des Ziels ergeben.

- Das zurückgegebene Array sollte die Indizes in aufsteigender Reihenfolge enthalten.

# --hints--

`find_target([2, 7, 11, 15], 9)` should return `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` should return `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` should return `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` should return `'Target not found'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```
