---
id: 681cb1b0dab50c87ddb2e519
title: "Challenge 8: Factorializer"
challengeType: 29
dashedName: challenge-8
---

# --description--

Gegeben ist eine ganze Zahl von null bis 20. Geben Sie die Fakultät dieser Zahl zurück. Die Fakultät einer Zahl ist das Produkt aller Zahlen zwischen 1 und der angegebenen Zahl.

- Die Fakultät von null ist 1.

# --hints--

`factorial(0)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` should return `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` should return `2432902008176640000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```
