---
id: 681cb1afdab50c87ddb2e517
title: "Aufgabe 6: Anagrammprüfer"
challengeType: 29
dashedName: challenge-6
---

# --description--

Gegeben sind zwei Strings. Bestimmen Sie, ob sie Anagramme voneinander sind (d. h. dieselben Zeichen in beliebiger Reihenfolge enthalten).

- Groß- und Kleinschreibung sowie Leerzeichen ignorieren.

# --hints--

`are_anagrams("listen", "silent")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("listen", "silent"), True)`)
}})
```

`are_anagrams("School master", "The classroom")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("School master", "The classroom"), True)`)
}})
```

`are_anagrams("A gentleman", "Elegant man")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("A gentleman", "Elegant man"), True)`)
}})
```

`are_anagrams("Hello", "World")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("Hello", "World"), False)`)
}})
```

`are_anagrams("apple", "banana")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("apple", "banana"), False)`)
}})
```

`are_anagrams("cat", "dog")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(are_anagrams("cat", "dog"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def are_anagrams(str1, str2):

    return str1
```

# --solutions--

```py
def are_anagrams(str1, str2):
    def clean(s):
        return sorted(s.replace(" ", "").lower())
    
    return clean(str1) == clean(str2)
```
