---
id: 681cb05adab50c87ddb2e513
title: "Aufgabe 2: Basisprüfung"
challengeType: 29
dashedName: challenge-2
---

# --description--

Gegeben ist ein String, der eine Zahl darstellt, und eine ganze Zahl als Basis von 2 bis 36. Bestimmen Sie, ob die Zahl in dieser Basis gültig ist.

- Der String kann Ziffern sowie Groß- oder Kleinbuchstaben enthalten.
- Die Prüfung sollte nicht zwischen Groß- und Kleinbuchstaben unterscheiden.
- Eine Zahl ist gültig, wenn jedes Zeichen eine gültige Ziffer in der angegebenen Basis ist.
- Beispiele für gültige Ziffern in verschiedenen Basen:
  - Basis 2: 0-1
  - Basis 8: 0-7
  - Basis 10: 0-9
  - Basis 16: 0-9 und A-F
  - Basis 36: 0-9 und A-Z

# --hints--

`is_valid_number("10101", 2)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10101", 2), True)`);
}})
```

`is_valid_number("10201", 2)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("10201", 2), False)`)
}})
```

`is_valid_number("76543210", 8)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("76543210", 8), True)`)
}})
```

`is_valid_number("9876543210", 8)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 8), False)`)
}})
```

`is_valid_number("9876543210", 10)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("9876543210", 10), True)`)
}})
```

`is_valid_number("ABC", 10)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 10), False)`)
}})
```

`is_valid_number("ABC", 16)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 16), True)`)
}})
```

`is_valid_number("Z", 36)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("Z", 36), True)`)
}})
```

`is_valid_number("ABC", 20)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("ABC", 20), True)`)
}})
```

`is_valid_number("4B4BA9", 16)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("4B4BA9", 16), True)`)
}})
```

`is_valid_number("5G3F8F", 16)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 16), False)`)
}})
```

`is_valid_number("5G3F8F", 17)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("5G3F8F", 17), True)`)
}})
```

`is_valid_number("abc", 10)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 10), False)`)
}})
```

`is_valid_number("abc", 16)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("abc", 16), True)`)
}})
```

`is_valid_number("AbC", 16)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("AbC", 16), True)`)
}})
```

`is_valid_number("z", 36)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("z", 36), True)`)
}})
```

`is_valid_number("", 10)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_number("", 10), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_number(n, base):

    return n
```

# --solutions--

```py
def is_valid_number(n, base):
    allChars = "0123456789abcdefghijklmnopqrstuvwxyz"
    newN = n.lower()

    availableChars = allChars[0:base]

    if not newN:
        return False

    for char in newN:
        if char not in availableChars:
            return False

    return True
```
