---
id: 6821ebce237de8297eaee790
title: "Challenge 14: Character Battle"
challengeType: 29
dashedName: challenge-14
---

# --description--

Gegeben sind zwei Strings, die Ihre Armee und eine gegnerische Armee darstellen. Jeder Charakter aus Ihrer Armee kämpft gegen den Charakter an derselben Position der gegnerischen Armee nach folgenden Regeln:

- Die Charaktere `a-z` haben jeweils eine Stärke von `1-26`.
- Die Charaktere `A-Z` haben jeweils eine Stärke von `27-52`.
- Ziffern `0-9` haben eine Stärke entsprechend ihrem Zahlenwert.
- Alle anderen Charaktere haben einen Wert von null.
- Jeder Charakter kann nur einen Kampf führen.

Für jeden Kampf gewinnt der stärkere Charakter. Die Armee mit den meisten Siegen gewinnt den Krieg. Geben Sie folgende Werte zurück:

- `"Opponent retreated"`, wenn Ihre Armee mehr Charaktere hat als die gegnerische Armee.
- `"We retreated"`, wenn die gegnerische Armee mehr Charaktere hat als Ihre.
- `"We won"`, wenn Ihre Armee mehr Kämpfe gewonnen hat.
- `"We lost"`, wenn die gegnerische Armee mehr Kämpfe gewonnen hat.
- `"It was a tie"`, wenn beide Armeen gleich viele Kämpfe gewonnen haben.

# --hints--

`battle("Hello", "World")` should return `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` should return `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` should return `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```
