---
id: 6821ebea237de8297eaee795
title: "Challenge 19: Candlelight"
challengeType: 29
dashedName: challenge-19
---

# --description--

Gegeben ist eine ganze Zahl, die die Anzahl der Kerzen darstellt, mit denen Sie starten, und eine ganze Zahl, die angibt, wie viele abgebrannte Kerzen benötigt werden, um eine neue zu erzeugen. Geben Sie die Anzahl der Kerzen zurück, die Sie insgesamt verwendet haben, nachdem Sie so viele wie möglich hergestellt und abgebrannt haben.

Wenn Sie zum Beispiel 7 Kerzen haben und es 2 abgebrannte Kerzen braucht, um eine neue zu machen:

1. Brennen Sie 7 Kerzen ab, um 7 Reste zu erhalten,
2. Recyceln Sie 6 Reste zu 3 neuen Kerzen (1 Rest bleibt übrig),
3. Brennen Sie 3 Kerzen ab, um 3 weitere Reste zu erhalten (insgesamt 4),
4. Recyceln Sie 4 Reste zu 2 neuen Kerzen,
5. Brennen Sie 2 Kerzen ab, um 2 Reste zu erhalten,
6. Recyceln Sie 2 Reste zu 1 neuer Kerze,
7. Brennen Sie 1 Kerze ab.

In diesem Beispiel haben Sie insgesamt 13 Kerzen abgebrannt.

# --hints--

`burn_candles(7, 2)` should return `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` should return `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` should return `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` should return `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` should return `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```
