---
id: 6821ebd4237de8297eaee791
title: "Challenge 15: camelCase"
challengeType: 29
dashedName: challenge-15
---

# --description--

Gegeben ist ein String. Geben Sie seine CamelCase-Version zurück, wobei folgende Regeln gelten:

- Wörter im String-Argument sind durch ein oder mehrere Zeichen aus der folgenden Menge getrennt: Leerzeichen (` `), Bindestrich (`-`) oder Unterstrich (`_`). Behandeln Sie jede Folge dieser Zeichen als Worttrennung.
- Das erste Wort soll komplett in Kleinbuchstaben sein.
- Jedes folgende Wort soll mit einem Großbuchstaben beginnen, der Rest davon in Kleinbuchstaben.
- Alle Leerzeichen und Trennzeichen sollen entfernt werden.

# --hints--

`to_camel_case("hello world")` should return `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` should return `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` should return `"secretAgentX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` should return `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```
