---
id: 6821ebda237de8297eaee792
title: "Challenge 16: Reverse Parenthesis"
challengeType: 29
dashedName: challenge-16
---

# --description--

Gegeben ist ein String, der korrekt verschachtelte runde Klammern enthält. Geben Sie die dekodierte Version des Strings zurück, indem Sie die folgenden Regeln anwenden:

- Alle Zeichen innerhalb jedes Klammernpaares sollen umgekehrt werden.
- Die runden Klammern sollen im Endergebnis entfernt werden.
- Wenn Klammern verschachtelt sind, soll zuerst das innerste Paar umgekehrt werden, und dessen Ergebnis dann in die Umkehrung des äußeren Paares einfließen.
- Gehen Sie davon aus, dass alle Klammern gleichmäßig ausbalanciert und korrekt verschachtelt sind.

# --hints--

`decode("(f(b(dc)e)a)")` should return `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` should return `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```
