---
id: 6821ebe4237de8297eaee794
title: "Aufgabe 18: Zweitbeste"
challengeType: 29
dashedName: challenge-18
---

# --description--

Gegeben ist ein Array ganzer Zahlen, die den Preis verschiedener Laptops darstellen, und eine ganze Zahl, die Ihr Budget darstellt. Geben Sie zurück:

1. Der zweitteuerste Laptop in der Liste, wenn er in Ihrem Budget liegt, oder
2. Der teuerste Laptop, den Sie sich sonst leisten können, oder
3. `0`, wenn kein Laptop in Ihrem Budget liegt.

- Doppelte Preise sollten ignoriert werden.

# --hints--

`get_laptop_cost([1500, 2000, 1800, 1400], 1900)` sollte `1800` zurückgeben

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900)` sollte `1800` zurückgeben

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1500, 2000, 2000, 1800, 1400], 1900), 1800)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 2200)` sollte `1899` zurückgeben

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 2200), 1899)`)
}})
```

`get_laptop_cost([2099, 1599, 1899, 1499], 1000)` sollte `0` zurückgeben

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([2099, 1599, 1899, 1499], 1000), 0)`)
}})
```

`get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` sollte `1400` zurückgeben

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_laptop_cost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_laptop_cost(laptops, budget):

    return laptops
```

# --solutions--

```py
def get_laptop_cost(laptops, budget):
    laptops = sorted(set(laptops), reverse=True)

    if len(laptops) >= 2 and budget >= laptops[1]:
        return laptops[1]
    if not laptops or budget < laptops[-1]:
        return 0

    for price in laptops[2:]:
        if budget >= price:
            return price

    return 0
```
