---
id: 68adce01c0e1144d0a902956
title: "Challenge 25: Vowel Repeater"
challengeType: 29
dashedName: challenge-25
---

# --description--

Gegeben ist ein String. Geben Sie eine neue Version des Strings zurück, bei der jeder Vokal jeweils um eine Wiederholung mehr erscheint als der vorherige Vokal, den Sie gefunden haben. Zum Beispiel bleibt der erste Vokal im Satz unverändert. Der zweite Vokal erscheint zweimal hintereinander. Der dritte Vokal erscheint dreimal hintereinander, und so weiter.

- Die Buchstaben `a`, `e`, `i`, `o` und `u`, sowohl in Groß- als auch in Kleinbuchstaben, gelten als Vokale.
- Der ursprüngliche Vokal behält seine Groß- oder Kleinschreibung.
- Wiederholte Vokale sollten in Kleinbuchstaben sein.
- Alle Nicht-Vokal-Zeichen behalten ihre ursprüngliche Groß- oder Kleinschreibung.

# --hints--

`repeat_vowels("hello world")` should return `"helloo wooorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("hello world"), "helloo wooorld")`)
}})
```

`repeat_vowels("freeCodeCamp")` should return `"freeeCooodeeeeCaaaaamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp")`)
}})
```

`repeat_vowels("AEIOU")` should return `"AEeIiiOoooUuuuu"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("AEIOU"), "AEeIiiOoooUuuuu")`)
}})
```

`repeat_vowels("I like eating ice cream in Iceland")` should return `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(repeat_vowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand")`)
}})
```

# --seed--

## --seed-contents--

```py
def repeat_vowels(s):

    return s
```

# --solutions--

```py
def repeat_vowels(s):
    vowels = "aeiouAEIOU"
    count = 0
    result = []

    for char in s:
        result.append(char)
        if char in vowels:
            result.append(char.lower() * count)
            count += 1

    return "".join(result)
```
