---
id: 68af0687ef34c76c28ffa54b
title: "Challenge 32: Reverse Sentence"
challengeType: 29
dashedName: challenge-32
---

# --description--

Gegeben ist ein String mit Wörtern. Geben Sie einen neuen String zurück, in dem die Wörter in umgekehrter Reihenfolge stehen. Zum Beispiel sollte das erste Wort am Ende des zurückgegebenen Strings stehen, und das letzte Wort sollte am Anfang des zurückgegebenen Strings stehen.

- Im gegebenen String können die Wörter durch ein oder mehrere Leerzeichen getrennt sein.
- Der zurückgegebene String sollte nur ein Leerzeichen zwischen den Wörtern enthalten.

# --hints--

`reverse_sentence("world hello")` should return `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` should return `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` should return `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` should return `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```
