---
id: 68b06e589bf2273243814775
title: "Challenge 38: Slug Generator"
challengeType: 29
dashedName: challenge-38
---

# --description--

Gegeben ist ein String. Geben Sie eine URL-freundliche Version des Strings zurück, die folgende Bedingungen erfüllt:

- Alle Buchstaben sollen Kleinbuchstaben sein.
- Alle Zeichen, die keine Buchstaben, Zahlen oder Leerzeichen sind, sollen entfernt werden.
- Alle Leerzeichen sollen durch den URL-kodierten Leerzeichen-Code `%20` ersetzt werden.
- Aufeinanderfolgende Leerzeichen sollen durch ein einzelnes `%20` ersetzt werden.
- Der zurückgegebene String darf keine führenden oder abschließenden `%20` enthalten.

# --hints--

`generate_slug("helloWorld")` should return `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("helloWorld"), "helloworld")`)
}})
```

`generate_slug("hello world!")` should return `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello world!"), "hello%20world")`)
}})
```

`generate_slug(" hello-world ")` should return `"helloworld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug(" hello-world "), "helloworld")`)
}})
```

`generate_slug("hello  world")` should return `"hello%20world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("hello  world"), "hello%20world")`)
}})
```

`generate_slug("  ?H^3-1*1]0! W[0%R#1]D  ")` should return `"h3110%20w0r1d"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_slug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_slug(str):

    return str
```

# --solutions--

```py
import re
def generate_slug(s):
    cleaned = re.sub(r'[^a-zA-Z0-9 ]+', '', s)
    cleaned = re.sub(r'\s+', ' ', cleaned).strip()
    return cleaned.lower().replace(' ', '%20')
```
