---
id: 68b06e589bf2273243814771
title: "Challenge 36: Thermostat Adjuster"
challengeType: 29
dashedName: challenge-36
---

# --description--

Gegeben sind die aktuelle Temperatur eines Raums und eine Zieltemperatur. Geben Sie einen String zurück, der angibt, wie die Raumtemperatur basierend auf diesen Vorgaben angepasst werden soll:

- Geben Sie `"heat"` zurück, wenn die aktuelle Temperatur unter der Zieltemperatur liegt.
- Geben Sie `"cool"` zurück, wenn die aktuelle Temperatur über der Zieltemperatur liegt.
- Geben Sie `"hold"` zurück, wenn die aktuelle Temperatur gleich der Zieltemperatur ist.

# --hints--

`adjust_thermostat(68, 72)` should return `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` should return `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` should return `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` should return `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` should return `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` should return `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```
