---
id: 68b06e589bf2273243814777
title: "Challenge 39: Fill The Tank"
challengeType: 29
dashedName: challenge-39
---

# --description--

Gegeben sind die Größe eines Kraftstofftanks, der aktuelle Kraftstoffstand und der Preis pro Gallone. Geben Sie die Kosten zurück, um den Tank vollständig zu füllen.

- `tankSize` ist die Gesamtkapazität des Tanks in Gallonen.
- `fuelLevel` ist die aktuelle Menge Kraftstoff im Tank in Gallonen.
- `pricePerGallon` ist der Preis für eine Gallone Kraftstoff.
- Der zurückgegebene Wert sollte auf zwei Dezimalstellen gerundet sein im Format: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` should return `"$80.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` should return `"$17.50"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` should return `"$29.25"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` should return `"$0.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` should return `"$21.89"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
