---
id: 68b06e589bf2273243814777
title: "Aufgabe 39: Den Tank füllen"
challengeType: 29
dashedName: challenge-39
---

# --description--

Gegeben sind die Größe eines Kraftstofftanks, der aktuelle Kraftstoffstand und der Preis pro Gallone. Geben Sie die Kosten zurück, um den Tank vollständig zu füllen.

- Das erste Argument ist die Gesamtkapazität des Tanks in Gallonen.
- Das zweite Argument ist die aktuelle Menge an Kraftstoff im Tank in Gallonen.
- Das dritte Argument ist der Preis für eine Gallone Kraftstoff.
- Der zurückgegebene Wert sollte auf zwei Dezimalstellen gerundet sein im Format: `"$d.dd"`.

# --hints--

`cost_to_fill(20, 0, 4.00)` sollte `"$80.00"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(20, 0, 4.00), "$80.00")`)
}})
```

`cost_to_fill(15, 10, 3.50)` sollte `"$17.50"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 10, 3.50), "$17.50")`)
}})
```

`cost_to_fill(18, 9, 3.25)` sollte `"$29.25"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(18, 9, 3.25), "$29.25")`)
}})
```

`cost_to_fill(12, 12, 4.99)` sollte `"$0.00"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(12, 12, 4.99), "$0.00")`)
}})
```

`cost_to_fill(15, 9.5, 3.98)` sollte `"$21.89"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cost_to_fill(15, 9.5, 3.98), "$21.89")`)
}})
```

# --seed--

## --seed-contents--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):

    return tank_size
```

# --solutions--

```py
def cost_to_fill(tank_size, fuel_level, price_per_gallon):
    gallons_needed = tank_size - fuel_level
    cost = gallons_needed * price_per_gallon
    return f"${cost:.2f}"
```
