---
id: 68b1f72371a5ac895ac70a06
title: "Challenge 42: Video Storage"
challengeType: 29
dashedName: challenge-42
---

# --description--

Gegeben sind eine Videogröße, eine Einheit für die Videogröße, eine Festplattenkapazität und eine Einheit für die Festplatte. Geben Sie die Anzahl der Videos zurück, die auf der Festplatte gespeichert werden können, unter Verwendung der folgenden Vorgaben:

- Die Einheit für die Videogröße kann Bytes (`"B"`), Kilobytes (`"KB"`), Megabytes (`"MB"`) oder Gigabytes (`"GB"`) sein.
- Wenn keine der oben genannten Videoeinheiten angegeben ist, geben Sie `"Invalid video unit"` zurück.
- Die Einheit der Festplattenkapazität kann Gigabytes (`"GB"`) oder Terabytes (`"TB"`) sein.
- Wenn keine der oben genannten Festplatteneinheiten angegeben ist, geben Sie `"Invalid drive unit"` zurück.
- Geben Sie die Anzahl der ganzen Videos zurück, die auf die Festplatte passen.
- Verwenden Sie die folgenden Umrechnungen:

| Einheit | Entspricht |
|:-------:|:---------:|
| 1 B     | 1 B       |
| 1 KB    | 1000 B    |
| 1 MB    | 1000 KB   |
| 1 GB    | 1000 MB   |
| 1 TB    | 1000 GB   |

Beispielsweise bestimmen Sie bei den Argumenten `500`, `"MB"`, `100` und `"GB"`, wie viele 500 MB Videos auf eine 100 GB Festplatte passen.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` should return `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` should return `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` should return `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` should return `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` should return `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```
