---
id: 68b7687dded630607aceccb1
title: "Challenge 48: Spam Detector"
challengeType: 29
dashedName: challenge-48
---

# --description--

Gegeben ist eine Telefonnummer im Format `"+A (BBB) CCC-DDDD"`, wobei jeder Buchstabe eine Ziffer wie folgt darstellt:

- `A` steht für die Landesvorwahl und kann eine beliebige Anzahl von Ziffern enthalten.
- `BBB` steht für die Vorwahl und besteht immer aus drei Ziffern.
- `CCC` und `DDDD` stehen für die lokale Nummer und haben jeweils immer drei bzw. vier Ziffern.

Bestimmen Sie, ob es sich um eine Spam-Nummer handelt, basierend auf den folgenden Kriterien:

- Die Landesvorwahl ist länger als zwei Ziffern oder beginnt nicht mit einer Null (`0`).
- Die Vorwahl ist größer als 900 oder kleiner als 200.
- Die Summe der ersten drei Ziffern der lokalen Nummer erscheint innerhalb der letzten vier Ziffern der lokalen Nummer.
- Die Nummer enthält dieselbe Ziffer vier- oder mehrmals hintereinander (ohne Berücksichtigung der Formatierungszeichen).

# --hints--

`is_spam("+0 (200) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (200) 234-0182"), False)`)
}})
```

`is_spam("+091 (555) 309-1922")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+091 (555) 309-1922"), True)`)
}})
```

`is_spam("+1 (555) 435-4792")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+1 (555) 435-4792"), True)`)
}})
```

`is_spam("+0 (955) 234-4364")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (955) 234-4364"), True)`)
}})
```

`is_spam("+0 (155) 131-6943")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (155) 131-6943"), True)`)
}})
```

`is_spam("+0 (555) 135-0192")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 135-0192"), True)`)
}})
```

`is_spam("+0 (555) 564-1987")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+0 (555) 564-1987"), True)`)
}})
```

`is_spam("+00 (555) 234-0182")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_spam("+00 (555) 234-0182"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_spam(number):

    return number
```

# --solutions--

```py
import re
def is_spam(number):
    digits = re.sub(r"\D", "", number)
    match = re.match(r"^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$", number)
    country_code, area_code, ccc, dddd = match.groups()

    if len(country_code) > 2 or not country_code.startswith("0"):
        return True

    area_num = int(area_code)
    if area_num > 900 or area_num < 200:
        return True

    sum_ccc = sum(int(d) for d in ccc)
    if str(sum_ccc) in dddd:
        return True

    if re.search(r"(\d)\1\1\1", digits):
        return True

    return False
```
