---
id: 68b7687dded630607aceccb3
title: "Challenge 49: CSV Header Parser"
challengeType: 29
dashedName: challenge-49
---

# --description--

Gegeben ist die erste Zeile einer durch Kommas getrennten Werte-Datei (CSV). Geben Sie ein Array zurück, das die Überschriften enthält.

- Die erste Zeile einer CSV-Datei enthält Überschriften, die durch Kommas getrennt sind.
- Entfernen Sie alle führenden oder nachfolgenden Leerzeichen von jeder Überschrift.

# --hints--

`get_headings("name,age,city")` should return `["name", "age", "city"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("name,age,city"), ["name", "age", "city"])`)
}})
```

`get_headings("first name,last name,phone")` should return `["first name", "last name", "phone"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("first name,last name,phone"), ["first name", "last name", "phone"])`)
}})
```

`get_headings("username , email , signup date ")` should return `["username", "email", "signup date"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_headings("username , email , signup date "), ["username", "email", "signup date"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_headings(csv):

    return csv
```

# --solutions--

```py
def get_headings(csv):
    return [h.strip() for h in csv.split(",")]
```
