---
id: 68b7cadffed0e75a517da673
title: "Challenge 52: Binary to Decimal"
challengeType: 29
dashedName: challenge-52
---

# --description--

Gegeben ist ein String, der eine Binärzahl darstellt. Geben Sie deren dezimale Entsprechung als Zahl zurück.

Eine Binärzahl verwendet nur die Ziffern `0` und `1`, um eine beliebige Zahl darzustellen. Um Binärzahlen in Dezimalzahlen umzuwandeln, multiplizieren Sie jede Ziffer mit einer Potenz von `2` und addieren diese zusammen. Beginnen Sie damit, die ganz rechts stehende Ziffer mit `2^0` zu multiplizieren, die nächste Ziffer links davon mit `2^1` und so weiter. Sobald alle Ziffern mit einer Potenz von `2` multipliziert wurden, addieren Sie das Ergebnis zusammen.

Zum Beispiel entspricht die Binärzahl `101` der Dezimalzahl `5`, weil:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` should return `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```
