---
id: 68c1a929005bf54d342aa8d3
title: "Challenge 56: Space Week Day 2: Exoplanet Search"
challengeType: 29
dashedName: challenge-56
---

# --description--

Für den zweiten Tag der Space Week erhalten Sie einen String, bei dem jedes Zeichen die Helligkeitsmessung eines Sterns darstellt. Bestimmen Sie, ob die Messungen mit der Transit-Methode einen Exoplaneten entdeckt haben. Die Transit-Methode ist, wenn ein Planet vor einem Stern vorbeizieht und dadurch dessen beobachtete Helligkeit verringert.

- Die Helligkeitsmessungen bestehen nur aus den Zeichen `0-9` und `A-Z`, wobei jede Messung den folgenden numerischen Werten entspricht:
- Die Zeichen `0-9` entsprechen Helligkeitsstufen `0-9`.
- Die Zeichen `A-Z` entsprechen Helligkeitsstufen `10-35`.

Ein Stern gilt als mit einem Exoplaneten versehen, wenn eine einzelne Messung kleiner oder gleich 80 % des Durchschnitts aller Messungen ist. Zum Beispiel gilt ein Stern mit einer durchschnittlichen Helligkeit von 10 als mit einem Exoplaneten versehen, wenn eine einzelne Messung 8 oder weniger beträgt.

# --hints--

`has_exoplanet("665544554")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```
