---
id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Landing Spot"
challengeType: 29
dashedName: challenge-58
---

# --description--

Am vierten Tag der Space Week erhalten Sie eine Matrix von Zahlen (ein Array von Arrays), die potenzielle Landeplätze für Ihren Rover darstellen. Finden Sie den sichersten Landeplatz basierend auf den folgenden Regeln:

- Jeder Platz in der Matrix enthält eine Zahl ab `0-9`, inklusive.
- Jede `0` stellt einen potenziellen Landeplatz dar.
- Jede Zahl außer `0` ist zu gefährlich zum Landen. Je höher die Zahl, desto gefährlicher.
- Der sicherste Platz ist definiert als die `0`-Zelle, deren umliegende Zellen (bis zu 4 Nachbarn, Diagonalen ignorieren) die geringste Gesamtrisiko haben.
- Ignorieren Sie Nachbarn außerhalb der Grenzen (Ecken und Kanten haben einfach weniger Nachbarn).
- Geben Sie die Indizes des sichersten Landeplatzes zurück. Es wird immer nur einen sichersten Platz geben.

Zum Beispiel, gegeben:

```js
[
  [1, 0],
  [2, 0]
]
```

Geben Sie `[0, 1]` zurück, die Indizes für die `0` im ersten Array.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` should return `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` should return `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` should return `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` should return `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```
