---
id: 68c1a929005bf54d342aa8d4
title: "Challenge 57: Space Week Day 3: Phone Home"
challengeType: 29
dashedName: challenge-57
---

# --description--

Für den dritten Tag der Space Week erhalten Sie ein Array von Zahlen, die Entfernungen (in Kilometern) zwischen Ihnen, Satelliten und Ihrem Heimatplaneten in einer Kommunikationsroute darstellen. Bestimmen Sie, wie lange eine Nachricht, die durch die Route gesendet wird, benötigt, um ihren Zielplaneten zu erreichen, unter den folgenden Bedingungen:

- Der erste Wert im Array ist die Entfernung von Ihrem Standort zum ersten Satelliten.
- Jeder folgende Wert, außer dem letzten, ist die Entfernung zum nächsten Satelliten.
- Der letzte Wert im Array ist die Entfernung vom vorherigen Satelliten zu Ihrem Heimatplaneten.
- Die Nachricht reist mit 300.000 km/s.
- Jeder Satellit, den die Nachricht **durchläuft**, fügt eine Übertragungsverzögerung von 0,5 Sekunden hinzu.
- Geben Sie eine Zahl zurück, die auf 4 Dezimalstellen gerundet ist, wobei nachfolgende Nullen entfernt werden.

# --hints--

`send_message([300000, 300000])` should return `2.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([300000, 300000]), 2.5)`)
}})
```

`send_message([384400, 384400])` should return `3.0627`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([384400, 384400]), 3.0627)`)
}})
```

`send_message([54600000, 54600000])` should return `364.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([54600000, 54600000]), 364.5)`)
}})
```

`send_message([1000000, 500000000, 1000000])` should return `1674.3333`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([1000000, 500000000, 1000000]), 1674.3333)`)
}})
```

`send_message([10000, 21339, 50000, 31243, 10000])` should return `2.4086`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([10000, 21339, 50000, 31243, 10000]), 2.4086)`)
}})
```

`send_message([802101, 725994, 112808, 3625770, 481239])` should return `21.1597`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(send_message([802101, 725994, 112808, 3625770, 481239]), 21.1597)`)
}})
```

# --seed--

## --seed-contents--

```py
def send_message(route):

    return route
```

# --solutions--

```py
def send_message(route):
    total_distance = sum(route)
    delay = (len(route) - 1) * 0.5
    time = total_distance / 300_000
    total = time + delay
    return round(total, 4)
```
