---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Goldilocks Zone"
challengeType: 29
dashedName: challenge-59
---

# --description--

Am fünften Tag der Space Week berechnen Sie die „Goldilocks-Zone“ eines Sterns – den Bereich um einen Stern, in dem die Bedingungen „genau richtig“ sind, damit flüssiges Wasser existieren kann.

Gegeben ist die Masse eines Sterns. Geben Sie ein Array mit den Start- und Endabständen seiner Goldilocks-Zone in Astronomischen Einheiten zurück.

Zur Berechnung der Goldilocks-Zone:

1. Bestimmen Sie die Leuchtkraft des Sterns, indem Sie seine Masse hoch 3,5 potenzieren.
2. Der Beginn der Zone ist 0,95 mal die Quadratwurzel seiner Leuchtkraft.
3. Das Ende der Zone ist 1,37 mal die Quadratwurzel seiner Leuchtkraft.

- Geben Sie die Abstände auf zwei Dezimalstellen gerundet zurück.

Zum Beispiel: Bei einer Masse von `1` geben Sie `[0.95, 1.37]` zurück.

# --hints--

`goldilocks_zone(1)` should return `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` should return `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` should return `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` should return `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` should return `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```
