---
id: 68c497f3aaefc9fd9f1b0e24
title: "Challenge 60: Space Week Day 6: Moon Phase"
challengeType: 29
dashedName: challenge-60
---

# --description--

Für den sechsten Tag der Space Week erhalten Sie ein Datum im Format `"YYYY-MM-DD"` und müssen die Mondphase für diesen Tag anhand der folgenden Regeln bestimmen:

Verwenden Sie einen vereinfachten Mondzyklus von 28 Tagen, der in vier gleiche Phasen unterteilt ist:

- `"New"`: Tage 1 - 7
- `"Waxing"`: Tage 8 - 14
- `"Full"`: Tage 15 - 21
- `"Waning"`: Tage 22 - 28

Nach Tag 28 wiederholt sich der Zyklus mit Tag 1, Neumond.

- Verwenden Sie `"2000-01-06"` als Referenz-Neumond (Tag 1 des Zyklus), um die Phase des angegebenen Tages zu bestimmen.
- Sie erhalten keine Daten vor dem Referenzdatum.
- Geben Sie die korrekte Phase als String zurück.

**Hinweis:** Tag 1 stellt den Tag des Neumonds dar, das heißt, es sind 0 Tage seit dem letzten Neumond vergangen.

# --hints--

`moon_phase("2000-01-12")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` should return `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` should return `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` should return `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```
