---
id: 68cae5b538ff798bbd4da004
title: "Aufgabe 66: HTML-Tag-Entferner"
challengeType: 29
dashedName: challenge-66
---

# --description--

Gegeben ist ein String mit HTML-Code. Entfernen Sie die Tags und geben Sie den reinen Textinhalt zurück.

- Der Eingabe-String enthält nur gültiges HTML.
- HTML-Tags können verschachtelt sein.
- Entfernen Sie die Tags und alle Attribute.

Zum Beispiel sollte `'<a href="#">Click here</a>'` `"Click here"` zurückgeben.

# --hints--

`strip_tags('<a href="#">Click here</a>')` sollte `"Click here"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` sollte `"Hello World!"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` sollte einen leeren String (`""`) zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` sollte `sectionsection` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```
