---
id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Missing Socks"
challengeType: 29
dashedName: challenge-69
---

# --description--

Gegeben ist eine ganze Zahl, die die Anzahl der Sockenpaare darstellt, mit denen Sie gestartet sind, und eine weitere ganze Zahl, die angibt, wie viele Waschzyklen Sie bereits durchlaufen haben. Geben Sie die Anzahl der vollständigen Sockenpaare zurück, die Sie aktuell haben, unter Verwendung der folgenden Bedingungen:

- Alle 2 Waschzyklen verlieren Sie eine einzelne Socke.
- Alle 3 Waschzyklen finden Sie eine einzelne fehlende Socke.
- Alle 5 Waschzyklen ist eine einzelne Socke abgenutzt und muss weggeworfen werden.
- Alle 10 Waschzyklen kaufen Sie ein Paar Socken.
- Sie können niemals weniger als null Socken insgesamt haben.
- Die Regeln können sich überschneiden. Zum Beispiel verlieren Sie im 10. Waschzyklus eine einzelne Socke, werfen eine einzelne Socke weg und kaufen ein neues Paar Socken.
- Geben Sie die Anzahl der vollständigen Sockenpaare zurück.

# --hints--

`sock_pairs(2, 5)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```
