---
id: 68cae5b538ff798bbd4da008
title: "Challenge 70: HTML Attribute Extractor"
challengeType: 29
dashedName: challenge-70
---

# --description--

Gegeben ist ein String mit einem gültigen HTML-Element. Geben Sie die Attribute des Elements nach den folgenden Kriterien zurück:

- Es wird nur ein Element übergeben.
- Attribute liegen im Format `attribute="value"` vor.
- Geben Sie ein Array von Strings zurück, wobei jede Eigenschaft und ihr Wert durch ein Komma getrennt sind, im Format `["attribute1, value1", "attribute2, value2"]`.
- Geben Sie die Attribute in der Reihenfolge zurück, in der sie angegeben sind.
- Wenn keine Attribute gefunden werden, geben Sie ein leeres Array zurück.

# --hints--

`extract_attributes('<span class="red"></span>')` should return `["class, red"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<span class="red"></span>'), ["class, red"])`)
}})
```

`extract_attributes('<meta charset="UTF-8" />')` should return `["charset, UTF-8"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<meta charset="UTF-8" />'), ["charset, UTF-8"])`)
}})
```

`extract_attributes("<p>Lorem ipsum dolor sit amet</p>")` should return `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes("<p>Lorem ipsum dolor sit amet</p>"), [])`)
}})
```

`extract_attributes('<input name="email" type="email" required="true" />')` should return `["name, email", "type, email", "required, true"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"])`)
}})
```

`extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>')` should return `["id, submit", "class, btn btn-primary"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(extract_attributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"])`)
}})
```

# --seed--

## --seed-contents--

```py
def extract_attributes(element):

    return element
```

# --solutions--

```py
import re
def extract_attributes(element):
    pattern = r'([\w-]+)="([^"]*)"'
    matches = re.findall(pattern, element)
    return [f"{attr}, {val}" for attr, val in matches]
```
