---
id: 68cae5b538ff798bbd4da001
title: "Challenge 63: Battle of Words"
challengeType: 29
dashedName: challenge-63
---

# --description--

Gegeben sind zwei Sätze, die Ihr Team und ein gegnerisches Team darstellen, wobei jedes Wort Ihres Teams gegen das entsprechende Wort des gegnerischen Teams antritt. Bestimmen Sie, welches Team gewinnt, anhand der folgenden Regeln:

- Die gegebenen Sätze enthalten immer die gleiche Anzahl an Wörtern.
- Wörter sind durch ein einzelnes Leerzeichen getrennt und bestehen nur aus Buchstaben.
- Der Wert jedes Wortes ist die Summe seiner Buchstaben.
- Die Buchstaben `a` bis `z` entsprechen den Werten `1` bis `26`. Zum Beispiel ist `a` `1`, und `z` ist `26`.
- Ein Großbuchstabe verdoppelt den Wert des Buchstabens. Zum Beispiel ist `A` `2`, und `Z` ist `52`.
- Die Wörter treten der Reihe nach gegeneinander an: das erste Wort Ihres Teams gegen das erste Wort des gegnerischen Teams, und so weiter.
- Ein Wort gewinnt, wenn sein Wert größer ist als der Wert des gegnerischen Wortes.
- Das Team mit mehr gewonnenen Wörtern ist der Sieger.

Geben Sie `"We win"` zurück, wenn Ihr Team gewinnt, `"We lose"`, wenn Ihr Team verliert, und `"Draw"`, wenn beide Teams gleich viele Siege haben.

# --hints--

`battle("hello world", "hello word")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "hello word"), "We win")`)
}})
```

`battle("Hello world", "hello world")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello world", "hello world"), "We win")`)
}})
```

`battle("lorem ipsum", "kitty ipsum")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("lorem ipsum", "kitty ipsum"), "We lose")`)
}})
```

`battle("hello world", "world hello")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("hello world", "world hello"), "Draw")`)
}})
```

`battle("git checkout", "git switch")` should return `"We win"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("git checkout", "git switch"), "We win")`)
}})
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` should return `"We lose"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose")`)
}})
```

`battle("We must never surrender", "Our team must win")` should return `"Draw"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("We must never surrender", "Our team must win"), "Draw")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(our_team, opponent):

    return our_team
```

# --solutions--

```py
def get_word_value(word):
    value = 0
    for char in word:
        val = ord(char.lower()) - ord('a') + 1
        if char.isupper():
            val *= 2
        value += val
    return value

def battle(our_team, opponent):
    my_wins = 0
    opp_wins = 0

    for my_word, opp_word in zip(our_team.split(), opponent.split()):
        my_val = get_word_value(my_word)
        opp_val = get_word_value(opp_word)

        if my_val > opp_val:
            my_wins += 1
        elif opp_val > my_val:
            opp_wins += 1

    if my_wins > opp_wins:
        return "We win"
    elif opp_wins > my_wins:
        return "We lose"
    else:
        return "Draw"
```
