---
id: 68cae5b538ff798bbd4da00a
title: "Challenge 72: Thermostat Adjuster 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Gegeben ist die aktuelle Raumtemperatur in Fahrenheit und eine Zieltemperatur in Celsius. Geben Sie einen String zurück, der angibt, wie die Raumtemperatur basierend auf diesen Vorgaben angepasst werden soll:

- Geben Sie `"Heat: X degrees Fahrenheit"` zurück, wenn die aktuelle Temperatur unter dem Zielwert liegt. Dabei ist `X` die Anzahl der Grad in Fahrenheit, um die der Raum geheizt werden muss, um das Ziel zu erreichen, auf eine Nachkommastelle gerundet.
- Geben Sie `"Cool: X degrees Fahrenheit"` zurück, wenn die aktuelle Temperatur über dem Zielwert liegt. Dabei ist `X` die Anzahl der Grad in Fahrenheit, um die der Raum gekühlt werden muss, um das Ziel zu erreichen, auf eine Nachkommastelle gerundet.
- Geben Sie `"Hold"` zurück, wenn die aktuelle Temperatur dem Zielwert entspricht.

Um Celsius in Fahrenheit umzurechnen, multiplizieren Sie die Celsius-Temperatur mit 1,8 und addieren Sie 32 zum Ergebnis (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` should return `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` should return `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` should return `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```
