---
id: 68cae5b538ff798bbd4da002
title: "Challenge 64: 24 to 12"
challengeType: 29
dashedName: challenge-64
---

# --description--

Gegeben ist ein String, der eine Uhrzeit im 24-Stunden-Format von `"HHMM"` darstellt. Geben Sie die Uhrzeit im entsprechenden 12-Stunden-Format von `"H:MM AM"` oder `"H:MM PM"` zurück.

- Die gegebene Eingabe ist immer ein vierstelliger String im 24-Stunden-Format, von `"0000"` bis `"2359"`.

# --hints--

`to_12("1124")` should return `"11:24 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1124"), "11:24 AM")`)
}})
```

`to_12("0900")` should return `"9:00 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0900"), "9:00 AM")`)
}})
```

`to_12("1455")` should return `"2:55 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("1455"), "2:55 PM")`)
}})
```

`to_12("2346")` should return `"11:46 PM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("2346"), "11:46 PM")`)
}})
```

`to_12("0030")` should return `"12:30 AM"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_12("0030"), "12:30 AM")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_12(time):

    return time
```

# --solutions--

```py
def convert_hours(hours):
    if hours == 0:
        return 12
    elif hours > 12:
        return hours - 12
    else:
        return hours

def to_12(time):
    hours = int(time[:2])
    minutes = time[2:]
    period = "AM" if hours < 12 else "PM"

    return f"{convert_hours(hours)}:{minutes} {period}"
```
