---
id: 68cae5b538ff798bbd4da006
title: "Challenge 68: Credit Card Masker"
challengeType: 29
dashedName: challenge-68
---

# --description--

Gegeben ist ein String mit Kreditkartennummern. Geben Sie eine maskierte Version davon zurück, die folgende Bedingungen erfüllt:

- Der String enthält vier Gruppen mit jeweils vier Ziffern (`0-9`), wobei alle Gruppen durch ein einzelnes Leerzeichen oder einen einzelnen Bindestrich (`-`) getrennt sind.
- Ersetzen Sie alle Zahlen außer den letzten vier durch ein Sternchen (`*`).
- Lassen Sie die übrigen Zeichen unverändert.

Beispielsweise soll bei Eingabe von `"4012-8888-8888-1881"` `"****-****-****-1881"` zurückgegeben werden.

# --hints--

`mask("4012-8888-8888-1881")` should return `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` should return `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` should return `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` should return `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```
