---
id: 68d2ba1468508398389487ce
title: "Challenge 73: Speak Wisely, You Must"
challengeType: 29
dashedName: challenge-73
---

# --description--

Gegeben ist ein Satz. Geben Sie eine Version davon zurück, die wie ein Rat eines weisen Lehrers klingt, und zwar nach den folgenden Regeln:

- Wörter sind durch ein einzelnes Leerzeichen getrennt.
- Finden Sie das erste Vorkommen eines der folgenden Wörter im Satz: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Verschieben Sie alle Wörter vor und einschließlich dieses Wortes ans Ende des Satzes und:
  - Bewahren Sie die Reihenfolge der Wörter beim Verschieben.
  - Schreiben Sie sie alle in Kleinbuchstaben.
  - Fügen Sie ein Komma und ein Leerzeichen davor ein.
- Großschreiben Sie den ersten Buchstaben des neuen ersten Wortes im Satz.
- Alle gegebenen Sätze enden mit einem einzelnen Satzzeichen. Behalten Sie die ursprüngliche Zeichensetzung bei und setzen Sie sie ans Ende des neuen Satzes.
- Geben Sie den neuen Satz zurück. Achten Sie darauf, dass zwischen den Wörtern jeweils nur ein Leerzeichen steht und keine Leerzeichen am Anfang oder Ende des Satzes sind.

Zum Beispiel, bei Eingabe `"You must speak wisely."` geben Sie `"Speak wisely, you must."` zurück.

# --hints--

`wise_speak("You must speak wisely.")` should return `"Speak wisely, you must."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You must speak wisely."), "Speak wisely, you must.")`)
}})
```

`wise_speak("You can do it!")` should return `"Do it, you can!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You can do it!"), "Do it, you can!")`)
}})
```

`wise_speak("Do you think you will complete this?")` should return `"Complete this, do you think you will?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("Do you think you will complete this?"), "Complete this, do you think you will?")`)
}})
```

`wise_speak("All your base are belong to us.")` should return `"Belong to us, all your base are."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("All your base are belong to us."), "Belong to us, all your base are.")`)
}})
```

`wise_speak("You have much to learn.")` should return `"Much to learn, you have."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(wise_speak("You have much to learn."), "Much to learn, you have.")`)
}})
```

# --seed--

## --seed-contents--

```py
def wise_speak(sentence):

    return sentence
```

# --solutions--

```py
def wise_speak(sentence):
    triggers = ["have", "must", "are", "will", "can"]
    punctuation = sentence[-1]
    words = sentence[:-1].split(" ")

    index = next(i for i, w in enumerate(words) if w in triggers)

    to_move = [w.lower() for w in words[:index + 1]]
    remaining = words[index + 1:]

    if remaining:
        remaining[0] = remaining[0][0].upper() + remaining[0][1:]

    return " ".join(remaining) + ", " + " ".join(to_move) + punctuation
```
