---
id: 68d2ba1468508398389487cf
title: "Challenge 74: Favorite Songs"
challengeType: 29
dashedName: challenge-74
---

# --description--

Erinnern Sie sich an iPods? Das erste Modell kam vor genau 24 Jahren heraus, am 23. Oktober 2001.

Gegeben ist ein Array von Song-Objekten, die Ihre iPod-Playlist darstellen. Geben Sie ein Array mit den Titeln der zwei meistgespielten Lieder zurück, wobei das meistgespielte Lied an erster Stelle steht.

- Jedes Objekt hat eine `"title"`-Eigenschaft (String) und eine `"plays"`-Eigenschaft (Ganzzahl).

# --hints--

`favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` should return `["Sync or Swim", "Earbud Blues"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"])`)
}})
```

`favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` should return `["Clickwheel Love", "99 Downloads"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"])`)
}})
```

`favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` should return `["Song B", "Song C"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(favorite_songs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"])`)
}})
```

# --seed--

## --seed-contents--

```py
def favorite_songs(playlist):

    return playlist
```

# --solutions--

```py
def favorite_songs(playlist):
    sorted_songs = sorted(playlist, key=lambda x: x["plays"], reverse=True)
    return [song["title"] for song in sorted_songs[:2]]
```
