---
id: 68d30845cc08266018fc46bd
title: "Challenge 77: Duration Formatter"
challengeType: 29
dashedName: challenge-77
---

# --description--

Gegeben ist eine ganze Zahl von Sekunden. Geben Sie einen String zurück, der dieselbe Dauer im Format `"H:MM:SS"` darstellt, wobei `"H"` die Anzahl der Stunden, `"MM"` die Anzahl der Minuten und `"SS"` die Anzahl der Sekunden ist. Geben Sie die Zeit nach folgenden Regeln zurück:

- Sekunden: Sollten immer zweistellig sein.
- Minuten: Führende Nullen sollten weggelassen werden, wenn sie nicht benötigt werden. Verwenden Sie `"0"`, wenn die Dauer weniger als eine Minute beträgt.
- Stunden: Sollten nur angegeben werden, wenn sie größer als null sind.

# --hints--

`format(500)` should return `"8:20"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(500), "8:20")`)
}})
```

`format(4000)` should return `"1:06:40"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(4000), "1:06:40")`)
}})
```

`format(1)` should return `"0:01"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(1), "0:01")`)
}})
```

`format(5555)` should return `"1:32:35"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(5555), "1:32:35")`)
}})
```

`format(99999)` should return `"27:46:39"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(format(99999), "27:46:39")`)
}})
```

# --seed--

## --seed-contents--

```py
def format(seconds):

    return seconds
```

# --solutions--

```py
def format(seconds):
    h = seconds // 3600
    m = (seconds % 3600) // 60
    s = seconds % 60

    seconds_str = f"{s:02d}"
    minutes_str = str(m)

    if h > 0:
        hours_str = str(h)
        return f"{hours_str}:{minutes_str.zfill(2)}:{seconds_str}"
    else:
        return f"{minutes_str}:{seconds_str}"
```
