---
id: 68d30fc57588d97fd3027b30
title: "Challenge 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

Am 28. Oktober 1994 wurde Netscape Navigator veröffentlicht, der Millionen half, das frühe Web zu erkunden.

Gegeben ist ein Array von Browserbefehlen, die Sie im Netscape Navigator ausgeführt haben. Geben Sie die aktuelle Seite zurück, auf der Sie sich nach Ausführung aller Befehle befinden, und zwar nach den folgenden Regeln:

- Sie starten immer auf der `"Home"`-Seite, die nicht im Befehls-Array enthalten ist.
- Gültige Befehle sind:
  - `"Visit Page"`: Dabei ist `"Page"` der Name der Seite, die Sie besuchen. Zum Beispiel bringt Sie `"Visit About"` zur `"About"`-Seite. Wenn Sie eine neue Seite besuchen, verwerfen Sie unbedingt jegliche Vorwärts-Historie.
  - `"Back"`: Bringt Sie zur vorherigen Seite in Ihrer Historie oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.
  - `"Forward"`: Bringt Sie vorwärts in der Historie zur Seite, von der Sie gekommen sind, oder bleibt auf der aktuellen Seite, falls keine vorhanden ist.

Zum Beispiel, gegeben `["Visit About Us", "Back", "Forward"]`, geben Sie `"About Us"` zurück.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` should return `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` should return `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` should return `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` should return `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```
