---
id: 68e39ed6106dac2f0a98fd65
title: "Challenge 83: Signature Validation"
challengeType: 29
dashedName: challenge-83
---

# --description--

Gegeben sind ein Nachrichten-String, ein geheimer Schlüssel-String und eine Signaturzahl. Bestimmen Sie, ob die Signatur mit dieser Codierungsmethode gültig ist:

- Buchstaben in der Nachricht und im geheimen Schlüssel haben folgende Werte:
  - `a` bis `z` haben jeweils die Werte `1` bis `26`.
  - `A` bis `Z` haben jeweils die Werte `27` bis `52`.
- Alle anderen Zeichen haben keinen Wert.
- Berechnen Sie die Signatur, indem Sie die Summe der Nachricht plus die Summe des geheimen Schlüssels nehmen.

Zum Beispiel, gegeben die Nachricht `"foo"` und den geheimen Schlüssel `"bar"`, wäre die Signatur `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Prüfen Sie, ob die berechnete Signatur mit der angegebenen Signatur übereinstimmt.

# --hints--

`verify("foo", "bar", 57)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 57), True)`)
}})
```

`verify("foo", "bar", 54)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("foo", "bar", 54), False)`)
}})
```

`verify("freeCodeCamp", "Rocks", 238)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("freeCodeCamp", "Rocks", 238), True)`)
}})
```

`verify("Is this valid?", "No", 210)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "No", 210), False)`)
}})
```

`verify("Is this valid?", "Yes", 233)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Is this valid?", "Yes", 233), True)`)
}})
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def verify(message, key, signature):

    return message
```

# --solutions--

```py
def verify(message, key, signature):
    def charValue(ch) -> int:
        if 'a' <= ch <= 'z':
            return ord(ch) - ord('a') + 1
        elif 'A' <= ch <= 'Z':
            return ord(ch) - ord('A') + 27
        else:
            return 0

    def compute_sum(s):
        return sum(charValue(ch) for ch in s)

    total = compute_sum(message) + compute_sum(key)
    return total == signature
```
