---
id: 68e39ed6106dac2f0a98fd66
title: "Challenge 84: Infected"
challengeType: 29
dashedName: challenge-84
---

# --description--

Am 2. November 1988 wurde der erste große Internetwurm freigesetzt, der nach nur einem Tag etwa 10 % der mit dem Internet verbundenen Computer infizierte.

In dieser Aufgabe erhalten Sie eine Anzahl von Tagen, die seit der Freisetzung eines Internetwurms vergangen sind, und Sie sollen bestimmen, wie viele Computer infiziert sind, basierend auf folgenden Regeln:

- Am Tag 0 ist der erste Computer infiziert.
- An jedem folgenden Tag verdoppelt sich die Anzahl der infizierten Computer.
- Alle 3 Tage wird nach der Ausbreitung des Virus ein Patch angewendet, der die Anzahl der infizierten Computer um 20 % reduziert. Die Anzahl der gepatchten Computer wird auf die nächste ganze Zahl aufgerundet.

Zum Beispiel gilt:

- Tag 0: 1 Computer ist insgesamt infiziert.
- Tag 1: 2 Computer sind insgesamt infiziert.
- Tag 2: 4 Computer sind insgesamt infiziert.
- Tag 3: 8 Computer sind insgesamt infiziert. Dann wird der Patch angewendet: 8 infizierte * 20 % = 1,6 gepatcht. 1,6 wird auf 2 aufgerundet. 8 infizierte Computer - 2 gepatchte = 6 Computer sind nach Tag 3 insgesamt infiziert.

Geben Sie die Gesamtzahl der infizierten Computer zurück, nachdem die angegebene Anzahl von Tagen vergangen ist.

# --hints--

`infected(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` should return `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` should return `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` should return `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```
