---
id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--

Ein Standard-Kartenspiel enthält 13 einzigartige Karten in jeder Farbe. Gegeben ist eine ganze Zahl, die die Anzahl der Karten angibt, die aus dem Deck gezogen werden sollen. Geben Sie die Anzahl der einzigartigen Kombinationen von Karten zurück, die Sie ziehen können.

- Die Reihenfolge spielt keine Rolle. Die Karte A zuerst zu ziehen und dann Karte B ist dasselbe wie Karte B zuerst zu ziehen und dann Karte A.

Zum Beispiel, gegeben `52`, geben Sie `1` zurück. Es gibt nur eine Kombination von 52 Karten, wenn Sie aus einem 52-Karten-Deck ziehen. Und gegeben `2`, geben Sie `1326` zurück. Es gibt 1326 Kartenkombinationen, wenn Sie 2 Karten aus dem Deck ziehen.

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```
