---
id: 68f6587287ad1f4ad39b0c80
title: "Challenge 94: Email Signature Generator"
challengeType: 29
dashedName: challenge-94
---

# --description--

Gegeben sind Strings für den Namen, den Titel und das Unternehmen einer Person. Geben Sie eine E-Mail-Signatur als einen einzigen String zurück, der die folgenden Regeln erfüllt:

- Der Name soll zuerst erscheinen, vorangestellt von einem Präfix, das von dem ersten Buchstaben des Namens abhängt. Für Namen, die (unabhängig von Groß- oder Kleinschreibung) mit folgendem Buchstaben beginnen:
  - `A-I`: Verwenden Sie `>>` als Präfix.
  - `J-R`: Verwenden Sie `--` als Präfix.
  - `S-Z`: Verwenden Sie `::` als Präfix.
- Nach dem Namen folgt ein Komma und ein Leerzeichen (`, `).
- Der Titel und das Unternehmen folgen nach dem Komma und Leerzeichen, getrennt durch `" at "` (mit Leerzeichen davor und danach).

Zum Beispiel, bei den Eingaben `"Quinn Waverly"`, `"Founder and CEO"` und `"TechCo"` geben Sie `"--Quinn Waverly, Founder and CEO at TechCo"` zurück.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` should return `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` should return `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` should return `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` should return `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` should return `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```
