---
id: 68f6587287ad1f4ad39b0c7c
title: "Challenge 90: Character Limit"
challengeType: 29
dashedName: challenge-90
---

# --description--

In dieser Aufgabe erhalten Sie einen String und müssen bestimmen, ob er in einen Social-Media-Beitrag passt. Geben Sie die folgenden Strings basierend auf den angegebenen Regeln zurück:

- `"short post"`, wenn er innerhalb des 40-Zeichen-Limits liegt.
- `"long post"`, wenn er länger als 40 Zeichen ist und innerhalb des 80-Zeichen-Limits liegt.
- `"invalid post"`, wenn er zu lang ist, um in eines der beiden Limits zu passen.

# --hints--

`can_post("Hello world")` should return `"short post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("Hello world"), "short post")`)
}})
```

`can_post("This is a longer message but still under eighty characters.")` should return `"long post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This is a longer message but still under eighty characters."), "long post")`)
}})
```

`can_post("This message is too long to fit into either of the character limits for a social media post.")` should return `"invalid post"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(can_post("This message is too long to fit into either of the character limits for a social media post."), "invalid post")`)
}})
```

# --seed--

## --seed-contents--

```py
def can_post(message):

    return message
```

# --solutions--

```py
def can_post(message):
    if len(message) <= 40:
        return "short post"
    elif len(message) <= 80:
        return "long post"
    else:
        return "invalid post"
```
