---
id: 68f6587287ad1f4ad39b0c7e
title: "Challenge 92: Extension Extractor"
challengeType: 29
dashedName: challenge-92
---

# --description--

Gegeben ist ein String, der einen Dateinamen darstellt. Geben Sie die Erweiterung der Datei zurück.

- Die Erweiterung ist der Teil des Dateinamens, der nach dem letzten Punkt (`.`) kommt.
- Wenn der Dateiname keinen Punkt enthält oder mit einem Punkt endet, geben Sie `"none"` zurück.
- Die Erweiterung soll unverändert zurückgegeben werden, einschließlich Groß- und Kleinschreibung.

# --hints--

`get_extension("document.txt")` should return `"txt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("document.txt"), "txt")`)
}})
```

`get_extension("README")` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("README"), "none")`)
}})
```

`get_extension("image.PNG")` should return `"PNG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("image.PNG"), "PNG")`)
}})
```

`get_extension(".gitignore")` should return `"gitignore"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension(".gitignore"), "gitignore")`)
}})
```

`get_extension("archive.tar.gz")` should return `"gz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("archive.tar.gz"), "gz")`)
}})
```

`get_extension("final.draft.")` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_extension("final.draft."), "none")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_extension(filename):

    return filename
```

# --solutions--

```py
def get_extension(filename):
    last_dot = filename.rfind('.')
    
    if last_dot == -1 or last_dot == len(filename) - 1:
        return "none"
    
    return filename[last_dot + 1:]
```
