---
id: 68ffb91507a5b645769328c6
title: "Challenge 103: LCM"
challengeType: 29
dashedName: challenge-103
---

# --description--

Gegeben sind zwei ganze Zahlen. Geben Sie das kleinste gemeinsame Vielfache (kgV) der beiden Zahlen zurück.

Das kgV von zwei Zahlen ist die kleinste positive ganze Zahl, die ein Vielfaches beider Zahlen ist. Zum Beispiel, gegeben `4` und `6`, geben Sie `12` zurück, weil:

- Vielfache von `4` sind `4`, `8`, `12` und so weiter.
- Vielfache von `6` sind `6`, `12`, `18` und so weiter.
- `12` ist die kleinste Zahl, die ein Vielfaches von beiden ist.

# --hints--

`lcm(4, 6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(4, 6), 12)`)
}})
```

`lcm(9, 6)` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(9, 6), 18)`)
}})
```

`lcm(10, 100)` should return `100`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(10, 100), 100)`)
}})
```

`lcm(13, 17)` should return `221`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(13, 17), 221)`)
}})
```

`lcm(45, 70)` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(lcm(45, 70), 630)`)
}})
```

# --seed--

## --seed-contents--

```py
def lcm(a, b):

    return a
```

# --solutions--

```py
def lcm(a, b):
    def gcd(x, y):
        return x if y == 0 else gcd(y, x % y)

    return abs(a * b) // gcd(a, b)
```
