---
id: 68ffb91507a5b645769328c4
title: "Challenge 101: Markdown Heading Converter"
challengeType: 29
dashedName: challenge-101
---

# --description--

Gegeben ist ein String, der eine Markdown-Überschrift darstellt. Geben Sie die entsprechende HTML-Überschrift zurück.

Eine gültige Markdown-Überschrift muss:

- Mit null oder mehr Leerzeichen beginnen, gefolgt von
- 1 bis 6 Rautezeichen (`#`) in Folge, dann
- Mindestens einem Leerzeichen. Und schließlich,
- Dem Überschriftentext.

Die Anzahl der Rautezeichen bestimmt die Überschriftenebene. Zum Beispiel entspricht ein Rautezeichen einem `h1`-Tag, und sechs Rautezeichen entsprechen einem `h6`-Tag.

Wenn der gegebene String nicht genau das oben beschriebene Format hat, geben Sie `"Invalid format"` zurück.

Beispielsweise, wenn `"# My level 1 heading"` gegeben ist, geben Sie `"<h1>My level 1 heading</h1>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`convert("# My level 1 heading")` should return `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` should return `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` should return `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` should return `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```
