---
id: 68ffb91507a5b645769328c5
title: "Challenge 102: Longest Word"
challengeType: 29
dashedName: challenge-102
---

# --description--

Gegeben ist ein Satz-String. Geben Sie das längste Wort im Satz zurück.

- Wörter sind durch ein einzelnes Leerzeichen getrennt.
- Nur Buchstaben (`a-z`, Groß- und Kleinschreibung wird nicht unterschieden) zählen zur Wortlänge.
- Wenn mehrere Wörter die gleiche Länge haben, geben Sie das zuerst vorkommende zurück.
- Geben Sie das Wort so zurück, wie es im gegebenen String erscheint, jedoch ohne Satzzeichen.

# --hints--

`longest_word("The quick red fox")` should return `"quick"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("The quick red fox"), "quick")`)
}})
```

`longest_word("Hello coding challenge.")` should return `"challenge"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Hello coding challenge."), "challenge")`)
}})
```

`longest_word("Do Try This At Home.")` should return `"This"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Do Try This At Home."), "This")`)
}})
```

`longest_word("This sentence... has commas, ellipses, and an exclamation point!")` should return `"exclamation"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation")`)
}})
```

`longest_word("A tie? No way!")` should return `"tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("A tie? No way!"), "tie")`)
}})
```

`longest_word("Wouldn't you like to know.")` should return `"Wouldnt"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(longest_word("Wouldn't you like to know."), "Wouldnt")`)
}})
```

# --seed--

## --seed-contents--

```py
def longest_word(sentence):

    return sentence
```

# --solutions--

```py
import re
def longest_word(sentence):
    words = sentence.split(" ")
    longest = ""

    for word in words:
        cleaned = re.sub(r'[^a-zA-Z]', '', word)
        if len(cleaned) > len(longest):
            longest = cleaned

    return longest
```
