---
id: 69162d64f96574d9bb629f01
title: "Challenge 116: Permutation Count"
challengeType: 29
dashedName: challenge-116
---

# --description--

Gegeben ist ein String. Geben Sie die Anzahl der unterschiedlichen Permutationen zurück, die aus seinen Zeichen gebildet werden können.

- Eine Permutation ist jede Umordnung der Zeichen im String.
- Zählen Sie keine doppelten Permutationen.
- Wenn der String wiederholte Zeichen enthält, sollen wiederholte Anordnungen nur einmal gezählt werden.
- Der String enthält nur Buchstaben (`A-Z`, `a-z`).

Zum Beispiel, für `"abb"` geben Sie `3` zurück, weil es drei einzigartige Möglichkeiten gibt, die Buchstaben anzuordnen: `"abb"`, `"bab"` und `"bba"`.

# --hints--

`count_permutations("abb")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abb"), 3)`)
}})
```

`count_permutations("abc")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("abc"), 6)`)
}})
```

`count_permutations("racecar")` should return `630`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("racecar"), 630)`)
}})
```

`count_permutations("freecodecamp")` should return `39916800`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count_permutations("freecodecamp"), 39916800)`)
}})
```

# --seed--

## --seed-contents--

```py
def count_permutations(s):

    return s
```

# --solutions--

```py
from math import factorial
from collections import Counter
def count_permutations(s):
    freq = Counter(s)
    n = len(s)

    result = factorial(n)
    for count in freq.values():
        result //= factorial(count)

    return result
```
