---
id: 69162d64f96574d9bb629efd
title: "Challenge 112: AI Detector"
challengeType: 29
dashedName: challenge-112
---

# --description--

Die heutige Aufgabe ist inspiriert von der Veröffentlichung von ChatGPT am 30. November 2022.

Gegeben ist ein String mit einem oder mehreren Sätzen. Bestimmen Sie, ob dieser wahrscheinlich von einer KI generiert wurde, anhand der folgenden Regeln:

- Er enthält zwei oder mehr Bindestriche (`-`).
- Er enthält zwei oder mehr Klammerpaare (`()`). Text kann innerhalb der Klammern stehen.
- Er enthält drei oder mehr Wörter mit 7 oder mehr Buchstaben.

- Wörter sind durch ein einzelnes Leerzeichen getrennt und bestehen nur aus Buchstaben (`A-Z`). Satzzeichen oder andere Nicht-Buchstaben zählen nicht als Teil eines Wortes.

Wenn der gegebene Satz eine der oben genannten Regeln erfüllt, geben Sie `"AI"` zurück, andernfalls `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```
