---
id: 69162d64f96574d9bb629efc
title: "Challenge 111: Ball Trajectory"
challengeType: 29
dashedName: challenge-111
---

# --description--

Die heutige Aufgabe ist inspiriert vom Videospiel Pong, das am 29. November 1972 veröffentlicht wurde.

Gegeben ist eine Matrix (Array von Arrays), die die Position des Balls (`2`) und die vorherige Position des Balls (`1`) enthält. Geben Sie die Matrixindizes für die nächste Position des Balls zurück.

- Der Ball bewegt sich immer in einer geraden Linie.
- Die Bewegungsrichtung wird bestimmt durch die Bewegung des Balls von `1` zu `2`.
- Die Ränder der Matrix gelten als Wände. Wenn der Ball eine:
  - obere oder untere Wand trifft, prallt er ab, indem er seine vertikale Richtung umkehrt.
  - linke oder rechte Wand trifft, prallt er ab, indem er seine horizontale Richtung umkehrt.
  - Ecke trifft, prallt er ab, indem er beide Richtungen umkehrt.

# --hints--

`get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` should return `[2, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` should return `[3, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0])`)
}})
```

`get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` should return `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` should return `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` should return `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_location(matrix):

    return matrix
```

# --solutions--

```py
def get_next_location(matrix):
    prev = None
    curr = None

    for r in range(len(matrix)):
        for c in range(len(matrix[0])):
            if matrix[r][c] == 1:
                prev = (r, c)
            if matrix[r][c] == 2:
                curr = (r, c)

    prev_row, prev_col = prev
    curr_row, curr_col = curr

    dir_x = curr_col - prev_col
    dir_y = curr_row - prev_row

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    max_row = len(matrix) - 1
    max_col = len(matrix[0]) - 1

    if next_col < 0 or next_col > max_col:
        dir_x *= -1
    if next_row < 0 or next_row > max_row:
        dir_y *= -1

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    return [next_row, next_col]
```
