---
id: 69162d64f96574d9bb629efb
title: "Challenge 110: Word Guesser"
challengeType: 29
dashedName: challenge-110
---

# --description--

Gegeben sind zwei Strings gleicher Länge, ein geheimes Wort und ein Rateversuch. Vergleichen Sie den Rateversuch mit dem geheimen Wort anhand der folgenden Regeln:

- Das geheime Wort und der Rateversuch bestehen nur aus Großbuchstaben (`"A"` bis `"Z"`);
- Für jeden Buchstaben im Rateversuch ersetzen Sie ihn durch eine Zahl, je nachdem, wie er mit dem geheimen Wort übereinstimmt:
  - `"2"`, wenn der Buchstabe im geheimen Wort an der richtigen Position steht.
  - `"1"`, wenn der Buchstabe im geheimen Wort vorkommt, aber an der falschen Position.
  - `"0"`, wenn der Buchstabe nicht im geheimen Wort vorkommt.
- Jeder Buchstabe im geheimen Wort kann höchstens einmal verwendet werden.
- Exakte Übereinstimmungen (`"2"`) werden zuerst zugewiesen, danach werden partielle Übereinstimmungen (`"1"`) von links nach rechts für die verbleibenden Buchstaben zugewiesen.
- Wenn ein Buchstabe im Rateversuch mehrfach vorkommt, kann er nur so oft übereinstimmen, wie er im geheimen Wort vorkommt.

Zum Beispiel, bei einem geheimen Wort `"APPLE"` und einem Rateversuch `"POPPA"`, geben Sie `"10201"` zurück:

Der erste `"P"` steht nicht an der richtigen Stelle (`"1"`), der `"O"` ist nicht im geheimen Wort (`"0"`), der zweite `"P"` steht an der richtigen Stelle (`"2"`), der dritte `"P"` ist eine Null (`"0"`), weil die zwei `"P"` im geheimen Wort bereits verwendet wurden, und der `"A"` steht nicht an der richtigen Stelle (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` should return `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` should return `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` should return `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` should return `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` should return `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` should return `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```
