---
id: 69162d64f96574d9bb629f00
title: "Challenge 115: Markdown Ordered List Item Converter"
challengeType: 29
dashedName: challenge-115
---

# --description--

Gegeben ist ein String, der ein Element einer geordneten Liste in Markdown darstellt. Geben Sie den entsprechenden HTML-String zurück.

Ein gültiges Element einer geordneten Liste in Markdown muss:

- Mit null oder mehr Leerzeichen beginnen, gefolgt von
- Einer Zahl (1 oder größer) und einem Punkt (`.`), gefolgt von
- Mindestens einem Leerzeichen, und dann
- Dem Text des Listenelements.

Wenn der String nicht genau dieses Format hat, geben Sie `"Invalid format"` zurück. Andernfalls umschließen Sie den Text des Listenelements mit `li`-Tags und geben den String zurück.

Zum Beispiel, bei `"1. My item"` geben Sie `"<li>My item</li>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`convert_list_item("1. My item")` should return `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` should return `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` should return `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```
