---
id: 69162d64f96574d9bb629eff
title: "Challenge 114: Camel to Snake"
challengeType: 29
dashedName: challenge-114
---

# --description--

Gegeben ist ein String im Camel Case. Geben Sie die Snake Case-Version des Strings zurück, wobei Sie die folgenden Regeln beachten:

- Der Eingabe-String enthält nur Buchstaben (`A-Z` und `a-z`) und beginnt immer mit einem Kleinbuchstaben.
- Jeder Großbuchstabe im Camel Case-String markiert den Beginn eines neuen Wortes.
- Wandeln Sie alle Buchstaben in Kleinbuchstaben um.
- Trennen Sie die Wörter mit einem Unterstrich (`_`).

# --hints--

`to_snake("helloWorld")` should return `"hello_world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("helloWorld"), "hello_world")`)
}})
```

`to_snake("myVariableName")` should return `"my_variable_name"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("myVariableName"), "my_variable_name")`)
}})
```

`to_snake("freecodecampDailyChallenges")` should return `"freecodecamp_daily_challenges"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_snake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_snake(camel):

    return camel
```

# --solutions--

```py
import re
def to_snake(camel):
    return re.sub(r'([A-Z])', r'_\1', camel).lower()
```
