---
id: 691b559495c5cb5a37b9b487
title: "Challenge 127: Speed Check"
challengeType: 29
dashedName: challenge-127
---

# --description--

Gegeben ist die Geschwindigkeit, mit der Sie in Meilen pro Stunde (MPH) fahren, und eine Geschwindigkeitsbegrenzung in Kilometern pro Stunde (KPH). Bestimmen Sie, ob Sie zu schnell fahren und ob Sie eine Verwarnung oder ein Bußgeld erhalten.

- 1 Meile entspricht 1,60934 Kilometern.
- Wenn Sie mit einer Geschwindigkeit fahren, die kleiner oder gleich der Geschwindigkeitsbegrenzung ist, geben Sie `"Not Speeding"` zurück.
- Wenn Sie bis zu 5 KPH über der Geschwindigkeitsbegrenzung fahren, geben Sie `"Warning"` zurück.
- Wenn Sie mehr als 5 KPH über der Geschwindigkeitsbegrenzung fahren, geben Sie `"Ticket"` zurück.

# --hints--

`speed_check(30, 70)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```
