---
id: 691b559495c5cb5a37b9b489
title: "Challenge 129: Markdown Blockquote Parser"
challengeType: 29
dashedName: challenge-129
---

# --description--

Gegeben ist ein String, der ein Blockzitat in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

Ein Blockzitat in Markdown ist jede Zeile, die:

- Mit null oder mehr Leerzeichen beginnt
- Gefolgt von einem Größer-als-Zeichen (`>`)
- Dann einem oder mehreren Leerzeichen
- Und schließlich dem Text des Blockzitats.

Geben Sie den Text des Blockzitats zurück, umgeben von öffnenden und schließenden HTML-`blockquote`-Tags.

Zum Beispiel, gegeben `"> This is a quote"`, geben Sie `<blockquote>This is a quote</blockquote>` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parse_blockquote("> This is a quote")` should return `"<blockquote>This is a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>")`)
}})
```

`parse_blockquote(" > This is also a quote")` should return `"<blockquote>This is also a quote</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>")`)
}})
```

`parse_blockquote("  >    So  Is  This")` should return `"<blockquote>So  Is  This</blockquote>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```
