---
id: 691b559495c5cb5a37b9b484
title: "Challenge 124: Inventory Update"
challengeType: 29
dashedName: challenge-124
---

# --description--

Gegeben ist ein 2D-Array, das das Inventar Ihres Geschäfts darstellt, und ein weiteres 2D-Array, das eine erhaltene Lieferung darstellt. Geben Sie Ihr aktualisiertes Inventar zurück.

- Jedes Element in den Arrays hat das Format: `[quantity, "item"]`, wobei `quantity` eine ganze Zahl und `"item"` ein String ist.
- Aktualisieren Sie die Artikel im Inventar, indem Sie die Menge aller passenden Artikel aus der Lieferung hinzufügen.
- Wenn ein empfangener Artikel im aktuellen Inventar nicht existiert, fügen Sie ihn als neuen Eintrag am Ende des Inventars hinzu.
- Geben Sie das Inventar in der Reihenfolge zurück, in der es angegeben wurde, wobei neue Artikel am Ende in der Reihenfolge erscheinen, in der sie in der Lieferung vorkommen.

Zum Beispiel, wenn das Inventar `[[2, "apples"], [5, "bananas"]]` und die Lieferung `[[1, "apples"], [3, "bananas"]]` ist, geben Sie `[[3, "apples"], [8, "bananas"]]` zurück.

# --hints--

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]])` should return `[[3, "apples"], [8, "bananas"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"]]), [[3, "apples"], [8, "bananas"]])`)
}})
```

`update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]])` should return `[[3, "apples"], [8, "bananas"], [4, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[2, "apples"], [5, "bananas"]], [[1, "apples"], [3, "bananas"], [4, "oranges"]]), [[3, "apples"], [8, "bananas"], [4, "oranges"]])`)
}})
```

`update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]])` should return `[[10, "apples"], [30, "bananas"], [20, "oranges"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([], [[10, "apples"], [30, "bananas"], [20, "oranges"]]), [[10, "apples"], [30, "bananas"], [20, "oranges"]])`)
}})
```

`update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]])` should return `[[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(update_inventory([[0, "Bowling Ball"], [0, "Dirty Socks"], [0, "Hair Pin"], [0, "Microphone"]], [[1, "Hair Pin"], [1, "Half-Eaten Apple"], [1, "Bowling Ball"], [1, "Toothpaste"]]), [[1, "Bowling Ball"], [0, "Dirty Socks"], [1, "Hair Pin"], [0, "Microphone"], [1, "Half-Eaten Apple"], [1, "Toothpaste"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def update_inventory(inventory, shipment):

    return inventory
```

# --solutions--

```py
def update_inventory(inventory, shipment):
    inventory_map = {item: i for i, (qty, item) in enumerate(inventory)}

    for qty, item in shipment:
        if item in inventory_map:
            index = inventory_map[item]
            inventory[index][0] += qty
        else:
            inventory.append([qty, item])
            inventory_map[item] = len(inventory) - 1

    return inventory
```
