---
id: 691b559495c5cb5a37b9b485
title: "Challenge 125: Game of Life"
challengeType: 29
dashedName: challenge-125
---

# --description--

Gegeben ist eine Matrix (Array von Arrays), die den aktuellen Zustand im Game of Life von Conway darstellt. Geben Sie den nächsten Zustand der Matrix zurück, indem Sie diese Regeln anwenden:

- Jede Zelle ist entweder `1` (lebendig) oder `0` (tot).
- Die Nachbarn einer Zelle sind die bis zu acht umliegenden Zellen (vertikal, horizontal und diagonal).
- Zellen am Rand haben weniger als acht Nachbarn.

Regeln zur Aktualisierung jeder Zelle:

- Jede lebendige Zelle mit weniger als zwei lebendigen Nachbarn stirbt (Unterbevölkerung).
- Jede lebendige Zelle mit zwei oder drei lebendigen Nachbarn bleibt am Leben.
- Jede lebendige Zelle mit mehr als drei lebendigen Nachbarn stirbt (Überbevölkerung).
- Jede tote Zelle mit genau drei lebendigen Nachbarn wird lebendig (Fortpflanzung).

Zum Beispiel, gegeben:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

return:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Jede Zelle wird entsprechend der Anzahl lebendiger Nachbarn aktualisiert. Zum Beispiel bleibt `[0][0]` tot (2 lebendige Nachbarn), `[0][1]` bleibt lebendig (2 lebendige Nachbarn), `[0][2]` stirbt (3 lebendige Nachbarn) und so weiter.

# --hints--

`game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` should return `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]])`)
}})
```

`game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` should return `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]])`)
}})
```

`game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` should return `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]])`)
}})
```

`game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` should return `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(game_of_life([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def game_of_life(grid):

    return grid
```

# --solutions--

```py
def game_of_life(grid):
    rows, cols = len(grid), len(grid[0])

    def count_live_neighbors(r, c):
        count = 0
        for i in range(r-1, r+2):
            for j in range(c-1, c+2):
                if 0 <= i < rows and 0 <= j < cols and (i != r or j != c):
                    count += grid[i][j]
        return count

    next_state = [row[:] for row in grid]

    for r in range(rows):
        for c in range(cols):
            live_neighbors = count_live_neighbors(r, c)

            if grid[r][c] == 1:
                if live_neighbors < 2 or live_neighbors > 3:
                    next_state[r][c] = 0
            else:
                if live_neighbors == 3:
                    next_state[r][c] = 1

    return next_state
```
