---
id: 691b5597f0f3e85a588a5c94
title: "Challenge 130: Checkerboard"
challengeType: 29
dashedName: challenge-130
---

# --description--

Gegeben ist ein Array mit zwei Zahlen, wobei die erste die Anzahl der Zeilen und die zweite die Anzahl der Spalten angibt. Geben Sie eine Matrix (ein Array von Arrays) zurück, die mit `"X"` und `"O"` Zeichen der angegebenen Größe gefüllt ist.

- Die Zeichen sollen sich wie bei einem Schachbrett abwechseln.
- Die Zelle oben links muss immer `"X"` sein.

Zum Beispiel, bei `[3, 3]`, geben Sie zurück:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`create_board([3, 3])` should return `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]])`)
}})
```

`create_board([6, 1])` should return `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]])`)
}})
```

`create_board([2, 10])` should return `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]])`)
}})
```

`create_board([5, 4])` should return `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(create_board([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]])`)
}})
```

# --seed--

## --seed-contents--

```py
def create_board(dimensions):

    return dimensions
```

# --solutions--

```py
def create_board(dimensions):
    rows, cols = dimensions
    board = []

    for r in range(rows):
        row = []
        for c in range(cols):
            row.append("X" if (r + c) % 2 == 0 else "O")
        board.append(row)

    return board
```
