---
id: 691b559495c5cb5a37b9b480
title: "Challenge 120: Pounds to Kilograms"
challengeType: 29
dashedName: challenge-120
---

# --description--

Gegeben ist ein Gewicht in Pfund als Zahl. Geben Sie den String `"(lbs) pounds equals (kgs) kilograms."` zurück.

- Ersetzen Sie `"(lbs)"` durch die Eingabezahl.
- Ersetzen Sie `"(kgs)"` durch die in Kilogramm umgerechnete Eingabe, auf zwei Dezimalstellen gerundet, und geben Sie immer zwei Dezimalstellen im Wert an.
- 1 Pfund entspricht 0,453592 Kilogramm.
- Wenn die Eingabe `1` ist, verwenden Sie `"pound"` anstelle von `"pounds"`.
- Wenn der umgerechnete Wert `1` ist, verwenden Sie `"kilogram"` anstelle von `"kilograms"`.

# --hints--

`convert_to_kgs(1)` should return `"1 pound equals 0.45 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(1), "1 pound equals 0.45 kilograms.")`)
}})
```

`convert_to_kgs(0)` should return `"0 pounds equals 0.00 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(0), "0 pounds equals 0.00 kilograms.")`)
}})
```

`convert_to_kgs(100)` should return `"100 pounds equals 45.36 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(100), "100 pounds equals 45.36 kilograms.")`)
}})
```

`convert_to_kgs(2.5)` should return `"2.5 pounds equals 1.13 kilograms."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.5), "2.5 pounds equals 1.13 kilograms.")`)
}})
```

`convert_to_kgs(2.20462)` should return `"2.20462 pounds equals 1.00 kilogram."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_to_kgs(2.20462), "2.20462 pounds equals 1.00 kilogram.")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_to_kgs(lbs):

    return lbs
```

# --solutions--

```py
def convert_to_kgs(lbs):
    KG_PER_POUND = 0.453592
    kgs = f"{lbs * KG_PER_POUND:.2f}"

    pound_word = "pound" if lbs == 1 else "pounds"
    kilogram_word = "kilogram" if kgs == "1.00" else "kilograms"

    return f"{lbs} {pound_word} equals {kgs} {kilogram_word}."
```
