---
id: 691b559495c5cb5a37b9b483
title: "Challenge 123: Roman Numeral Builder"
challengeType: 29
dashedName: challenge-123
---

# --description--

Gegeben ist eine ganze Zahl. Geben Sie ihren entsprechenden Wert in römischen Zahlen zurück.

Römische Zahlen verwenden diese Symbole:

| Symbol | Value |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

Römische Zahlen werden von der größten zur kleinsten von links nach rechts geschrieben, wobei folgende Regeln gelten:

- Addition wird verwendet, wenn ein Symbol von einem Symbol mit gleichem oder kleinerem Wert gefolgt wird. Zum Beispiel wird `18` als `XVIII` geschrieben (10 + 5 + 1 + 1 + 1 = 18).
- Subtraktion wird verwendet, wenn ein kleineres Symbol vor einem größeren steht, um 4 oder 9 an jeder Stellenwertposition darzustellen. Zum Beispiel wird 19 als `XIX` geschrieben (10 + (10 - 1)).
- Kein Symbol darf mehr als dreimal hintereinander wiederholt werden. Subtraktion wird verwendet, wenn Sie sonst ein Symbol mehr als dreimal hintereinander schreiben müssten. Die größte Zahl, die Sie schreiben können, ist also 3999.

Hier ist ein weiteres Beispiel: Gegeben `1464`, geben Sie `"MCDLXIV"` zurück (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` should return `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` should return `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` should return `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` should return `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` should return `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```
