---
id: 691b559495c5cb5a37b9b482
title: "Challenge 122: Markdown Bold Parser"
challengeType: 29
dashedName: challenge-122
---

# --description--

Gegeben ist ein String, der möglicherweise fettgedruckten Text in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

- Fettgedruckter Text in Markdown ist jeder Text, der mit zwei Sternchen (`**`) oder zwei Unterstrichen (`__`) beginnt und endet.
- Zwischen dem Text und den Sternchen oder Unterstrichen dürfen keine Leerzeichen sein, aber im Text selbst können Leerzeichen enthalten sein.
- Wandeln Sie alle fettgedruckten Vorkommen in HTML-`b`-Tags um und geben Sie den String zurück.

Zum Beispiel, bei Eingabe von `"**This is bold**"` geben Sie `"<b>This is bold</b>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parse_bold("**This is bold**")` should return `"<b>This is bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is bold**"), "<b>This is bold</b>")`)
}})
```

`parse_bold("__This is also bold__")` should return `"<b>This is also bold</b>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__This is also bold__"), "<b>This is also bold</b>")`)
}})
```

`parse_bold("**This is not bold **")` should return `"**This is not bold **"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("**This is not bold **"), "**This is not bold **")`)
}})
```

`parse_bold("__ This is also not bold__")` should return `"__ This is also not bold__"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("__ This is also not bold__"), "__ This is also not bold__")`)
}})
```

`parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog.")` should return `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_bold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_bold(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_bold(markdown):
    markdown = re.sub(r'\*\*(\S(?:.*?\S)?)\*\*', r'<b>\1</b>', markdown)
    markdown = re.sub(r'__(\S(?:.*?\S)?)__', r'<b>\1</b>', markdown)

    return markdown
```
