---
id: 6925e2068081f40f549ced1b
title: "Challenge 137: Snowflake Generator"
challengeType: 29
dashedName: challenge-137
---

# --description--

Gegeben ist ein mehrzeiliger String, der Zeilenumbrüche mit dem Zeichen `\n` darstellt. Geben Sie einen neuen String zurück, bei dem jede Zeile horizontal gespiegelt und an das Ende der ursprünglichen Zeile angehängt wird.

- Spiegeln Sie eine Zeile, indem Sie alle ihre Zeichen, einschließlich Leerzeichen, umkehren.

Zum Beispiel, gegeben `"* \n *\n* "`, das in der Konsole wie folgt ausgegeben wird:

```sh
* 
 *
* 
```

Geben Sie `"*  *\n ** \n*  *"` zurück, das in der Konsole wie folgt ausgegeben wird:

```sh
*  *
 ** 
*  *
```

Achten Sie genau auf die Leerzeichen in den gegebenen und zurückgegebenen Strings. Schneiden Sie keine davon ab.

# --hints--

`generate_snowflake("* \n *\n* ")` should return `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` should return `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` should return `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` should return `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` should return `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```
