---
id: 6925e2068081f40f549ced1c
title: "Challenge 138: Sum of Divisors"
challengeType: 29
dashedName: challenge-138
---

# --description--

Gegeben ist eine positive ganze Zahl. Geben Sie die Summe aller ihrer Teiler zurück.

- Ein Teiler ist jede ganze Zahl, die die Zahl ohne Rest teilt (der Rest ist `0`).
- Zählen Sie jeden Teiler nur einmal.

Zum Beispiel, gegeben `6`, geben Sie `12` zurück, weil die Teiler von `6` `1`, `2`, `3` und `6` sind, und die Summe davon `12` ergibt.

# --hints--

`sum_divisors(6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` should return `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` should return `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` should return `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` should return `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```
