---
id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--

Gegeben ist ein String mit einem Bild in Markdown. Geben Sie den entsprechenden HTML-String zurück.

Ein Markdown-Bild hat folgendes Format: `"![alt text](image_url)"`. Dabei gilt:

- `alt text` ist die Beschreibung des Bildes (der Wert des `alt`-Attributs).
- `image_url` ist die Quell-URL des Bildes (der Wert des `src`-Attributs).

Geben Sie einen String mit dem HTML-`img`-Tag zurück, bei dem das `src`-Attribut auf die Bild-URL gesetzt ist und das `alt`-Attribut den Alternativtext enthält.

Zum Beispiel, bei Eingabe von `"![Cute cat](cat.png)"` geben Sie `'<img src="cat.png" alt="Cute cat">'` zurück;

- Achten Sie darauf, dass Tag, Reihenfolge der Attribute, Leerzeichen und Anführungszeichen genau wie oben sind.

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an – prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```
