---
id: 69272dcf1c24b44fd79137c3
title: "Challenge 140: SCREAMING_SNAKE_CASE"
challengeType: 29
dashedName: challenge-140
---

# --description--

Gegeben ist ein String, der einen Variablennamen darstellt. Geben Sie den Variablennamen zurück, konvertiert in SCREAMING_SNAKE_CASE.

Die gegebenen Variablennamen sind in einem der folgenden Formate geschrieben:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

In den oben genannten Formaten sind Wörter durch einen Unterstrich (`_`), einen Bindestrich (`-`) getrennt oder ein neues Wort beginnt mit einem Großbuchstaben.

Um in SCREAMING_SNAKE_CASE zu konvertieren:

- Machen Sie alle Buchstaben groß
- Trennen Sie Wörter mit einem Unterstrich (`_`)

# --hints--

`to_screaming_snake_case("userEmail")` should return `"USER_EMAIL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("userEmail"), "USER_EMAIL")`)
}})
```

`to_screaming_snake_case("UserPassword")` should return `"USER_PASSWORD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("UserPassword"), "USER_PASSWORD")`)
}})
```

`to_screaming_snake_case("user_id")` should return `"USER_ID"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user_id"), "USER_ID")`)
}})
```

`to_screaming_snake_case("user-address")` should return `"USER_ADDRESS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("user-address"), "USER_ADDRESS")`)
}})
```

`to_screaming_snake_case("username")` should return `"USERNAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("username"), "USERNAME")`)
}})
```

`to_screaming_snake_case("my_variable_name")` should return `"MY_VARIABLE_NAME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_screaming_snake_case("my_variable_name"), "MY_VARIABLE_NAME")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_screaming_snake_case(variable_name):

    return variable_name
```

# --solutions--

```py
import re
def to_screaming_snake_case(variable_name):
    temp = re.sub(r'[-_]+', ' ', variable_name)
    temp = re.sub(r'([a-z0-9])([A-Z])', r'\1 \2', temp)
    words = temp.strip().split()
    return '_'.join(words).upper()
```
