---
id: 69272dcf1c24b44fd79137c6
title: "Challenge 143: Markdown Italic Parser"
challengeType: 29
dashedName: challenge-143
---

# --description--

Gegeben ist ein String, der möglicherweise kursiven Text in Markdown enthält. Geben Sie den entsprechenden HTML-String zurück.

- Kursiver Text in Markdown ist jeder Text, der mit einem einzelnen Sternchen (`*`) oder einem einzelnen Unterstrich (`_`) beginnt und endet.
- Zwischen dem Text und dem Sternchen oder Unterstrich dürfen keine Leerzeichen stehen, aber im Text selbst können Leerzeichen enthalten sein.
- Wandeln Sie alle kursiven Vorkommen in HTML-`i`-Tags um und geben Sie den String zurück.

Zum Beispiel, bei Eingabe von `"*This is italic*"` geben Sie `"<i>This is italic</i>"` zurück.

Hinweis: Die Konsole zeigt HTML-Tags in Strings möglicherweise nicht an, wenn Nachrichten protokolliert werden. Prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parse_italics("*This is italic*")` should return `"<i>This is italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is italic*"), "<i>This is italic</i>")`)
}})
```

`parse_italics("_This is also italic_")` should return `"<i>This is also italic</i>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_This is also italic_"), "<i>This is also italic</i>")`)
}})
```

`parse_italics("*This is not italic *")` should return `"*This is not italic *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("*This is not italic *"), "*This is not italic *")`)
}})
```

`parse_italics("_ This is also not italic_")` should return `"_ This is also not italic_"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("_ This is also not italic_"), "_ This is also not italic_")`)
}})
```

`parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog.")` should return `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_italics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_italics(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_italics(markdown):
    pattern = r'(\*|_)([^\s][^]*?[^\s])\1'
    pattern = r'(\*|_)([^\s].*?[^\s])\1'
    return re.sub(pattern, r'<i>\2</i>', markdown)
```
