---
id: 69306364df283fcaff2e1ad5
title: "Challenge 144: Resolution Streak"
challengeType: 29
dashedName: challenge-144
---

# --description--

Gegeben ist ein Array von Arrays, wobei jedes Teilarray einen Tag Ihrer Vorsatzaktivitäten darstellt und drei ganze Zahlen enthält: die an diesem Tag gegangenen Schritte, die Minuten der Bildschirmzeit an diesem Tag und die Anzahl der an diesem Tag gelesenen Seiten; bestimmen Sie, ob Sie Ihre Vorsätze einhalten.

- Das erste Teilarray ist Tag 1, das zweite Tag 2 und so weiter.

Ein Tag gilt als erfolgreich, wenn alle drei der folgenden Ziele erreicht werden:

- Sie sind mindestens 10.000 Schritte gegangen.
- Sie hatten nicht mehr als 120 Minuten Bildschirmzeit.
- Sie haben mindestens fünf Seiten gelesen.

Wenn alle angegebenen Tage erfolgreich sind, geben Sie `"Resolution on track: N day streak."` zurück, wobei `N` die Anzahl der erfolgreichen Tage ist.

Wenn ein oder mehrere Tage nicht erfolgreich sind, geben Sie `"Resolution failed on day X: N day streak."` zurück, wobei `X` die Tagesnummer des ersten nicht erfolgreichen Tages ist und `N` die Anzahl der erfolgreichen Tage vor dem ersten nicht erfolgreichen Tag ist.

# --hints--

`resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` should return `"Resolution on track: 3 day streak."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.")`)
}})
```

`resolution_streak([[10000, 120, 5], [10950, 121, 11]])` should return `"Resolution failed on day 2: 1 day streak."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.")`)
}})
```

`resolution_streak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` should return `"Resolution failed on day 3: 2 day streak."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.")`)
}})
```

`resolution_streak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` should return `"Resolution failed on day 4: 3 day streak."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.")`)
}})
```

`resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` should return `"Resolution on track: 10 day streak."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(resolution_streak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.")`)
}})
```

# --seed--

## --seed-contents--

```py
def resolution_streak(days):

    return days
```

# --solutions--

```py
def resolution_streak(days):
    streak = 0

    for i, day in enumerate(days):
        steps, screen, pages = day
        if steps >= 10000 and screen <= 120 and pages >= 5:
            streak += 1
        else:
            return f"Resolution failed on day {i+1}: {streak} day streak."

    return f"Resolution on track: {streak} day streak."
```
