---
id: 69306364df283fcaff2e1ad8
title: "Challenge 147: Leap Year Calculator"
challengeType: 29
dashedName: challenge-147
---

# --description--

Gegeben ist ein ganzzahliges Jahr. Bestimmen Sie, ob es ein Schaltjahr ist.

Ein Jahr ist ein Schaltjahr, wenn es die folgenden Regeln erfüllt:

- Das Jahr ist durch 4 ohne Rest teilbar, und
- Das Jahr ist nicht durch 100 ohne Rest teilbar, außer
- Das Jahr ist durch 400 ohne Rest teilbar.

# --hints--

`is_leap_year(2024)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2024), True)`)
}})
```

`is_leap_year(2023)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2023), False)`)
}})
```

`is_leap_year(2100)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2100), False)`)
}})
```

`is_leap_year(2000)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2000), True)`)
}})
```

`is_leap_year(1999)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(1999), False)`)
}})
```

`is_leap_year(2040)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2040), True)`)
}})
```

`is_leap_year(2026)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_leap_year(2026), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_leap_year(year):

    return year
```

# --solutions--

```py
def is_leap_year(year):
    if (year % 4 == 0 and year % 100 != 0) or (year % 400 == 0):
        return True
    else:
        return False
```
