---
id: 69306364df283fcaff2e1ad7
title: "Challenge 146: Left-Handed Seat at the Table"
challengeType: 29
dashedName: challenge-146
---

# --description--

Gegeben ist eine 4x2-Matrix (Array von Arrays), die die Sitzordnung für eine Mahlzeit darstellt. Bestimmen Sie, wie viele Plätze ein linkshändiger Mensch einnehmen kann.

- Ein linkshändiger Mensch darf nicht dort sitzen, wo sich unmittelbar links neben ihm ein rechtshändiger Mensch befindet.

In der gegebenen Matrix:

- Ein `"R"` ist ein Sitz, der von einem rechtshändigen Menschen besetzt ist.
- Ein `"L"` ist ein Sitz, der von einem linkshändigen Menschen besetzt ist.
- Ein `"U"` ist ein unbesetzter Sitz.
- Nur unbesetzte Sitze sind verfügbar.
- Die Sitze in der oberen Reihe sind nach „unten“ ausgerichtet, und die Sitze in der unteren Reihe sind nach „oben“ ausgerichtet (wie an einem Tisch), daher sind links und rechts relativ zur Sitzorientierung zu verstehen.
- Ecksitze haben nur einen angrenzenden Sitz.

Zum Beispiel in der gegebenen Matrix:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Der Sitz oben links kann nicht besetzt werden, weil sich links davon ein rechtshändiger Mensch befindet. Die anderen beiden freien Sitze können besetzt werden, da sich links davon kein rechtshändiger Mensch befindet.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```
