---
id: 69373793f5a867f769cde139
title: "Challenge 154: Par for the Hole"
challengeType: 29
dashedName: challenge-154
---

# --description--

Gegeben sind zwei ganze Zahlen: das Par für ein Golfloch und die Anzahl der Striche, die ein Golfer für dieses Loch benötigt hat. Geben Sie die Punktzahl des Golfers mit Golfterminen zurück.

Return:

- `"Hole in one!"`, wenn es einen Strich gebraucht hat.
- `"Eagle"`, wenn es zwei Striche weniger als Par gebraucht hat.
- `"Birdie"`, wenn es einen Strich weniger als Par gebraucht hat.
- `"Par"`, wenn es genauso viele Striche wie Par gebraucht hat.
- `"Bogey"`, wenn es einen Strich mehr als Par gebraucht hat.
- `"Double bogey"`, wenn es zwei Striche mehr als Par gebraucht hat.

# --hints--

`golf_score(3, 3)` should return `"Par"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 3), "Par")`)
}})
```

`golf_score(4, 3)` should return `"Birdie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 3), "Birdie")`)
}})
```

`golf_score(3, 1)` should return `"Hole in one!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(3, 1), "Hole in one!")`)
}})
```

`golf_score(5, 7)` should return `"Double bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 7), "Double bogey")`)
}})
```

`golf_score(4, 5)` should return `"Bogey"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(4, 5), "Bogey")`)
}})
```

`golf_score(5, 3)` should return `"Eagle"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(golf_score(5, 3), "Eagle")`)
}})
```

# --seed--

## --seed-contents--

```py
def golf_score(par, strokes):

    return par
```

# --solutions--

```py
def golf_score(par, strokes):
    if strokes == 1:
        return "Hole in one!"

    diff = strokes - par

    if diff == -2:
        return "Eagle"
    if diff == -1:
        return "Birdie"
    if diff == 0:
        return "Par"
    if diff == 1:
        return "Bogey"
    if diff == 2:
        return "Double bogey"
```
