---
id: 69373793f5a867f769cde13a
title: "Challenge 155: Plant the Crop"
challengeType: 29
dashedName: challenge-155
---

# --description--

Gegeben ist eine ganze Zahl, die die Größe Ihres Bauernfelds darstellt, sowie `"acres"` oder `"hectares"`, die die Einheit für die Größe Ihres Bauernfelds angeben, und eine Pflanzenart. Bestimmen Sie, wie viele Pflanzen dieser Art auf Ihr Feld passen.

- 1 Acre entspricht 4046,86 Quadratmetern.
- 1 Hektar entspricht 10.000 Quadratmetern.

Hier ist eine Liste von Pflanzen, die als Eingabe gegeben werden, und wie viel Platz eine einzelne Pflanze benötigt:

|Pflanze|Platz pro Pflanze|
|-|-|
|`"corn"`|1 Quadratmeter|
|`"wheat"`|0,1 Quadratmeter|
|`"soybeans"`|0,5 Quadratmeter|
|`"tomatoes"`|0,25 Quadratmeter|
|`"lettuce"`|0,2 Quadratmeter|

Geben Sie die Anzahl der Pflanzen zurück, die auf das Feld passen, abgerundet auf die nächstkleinere ganze Pflanze.

# --hints--

`get_number_of_plants(1, "acres", "corn")` should return `4046`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(1, "acres", "corn"), 4046)`)
}})
```

`get_number_of_plants(2, "hectares", "lettuce")` should return `100000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(2, "hectares", "lettuce"), 100000)`)
}})
```

`get_number_of_plants(20, "acres", "soybeans")` should return `161874`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(20, "acres", "soybeans"), 161874)`)
}})
```

`get_number_of_plants(3.75, "hectares", "tomatoes")` should return `150000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(3.75, "hectares", "tomatoes"), 150000)`)
}})
```

`get_number_of_plants(16.75, "acres", "tomatoes")` should return `271139`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_number_of_plants(16.75, "acres", "tomatoes"), 271139)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_number_of_plants(field_size, unit, crop):

    return field_size
```

# --solutions--

```py
import math
def get_number_of_plants(field_size, unit, crop):
    crop_space = {
        "corn": 1,
        "wheat": 0.1,
        "soybeans": 0.5,
        "tomatoes": 0.25,
        "lettuce": 0.2
    }

    field_in_sqm = field_size * 4046.86 if unit == "acres" else field_size * 10000

    return math.floor(field_in_sqm / crop_space[crop.lower()])
```
