---
id: 6939b873185d8e00d453563c
title: "Aufgabe 157: Markdown-Link-Parser"
challengeType: 29
dashedName: challenge-157
---

# --description--

Gegeben ist der String eines Links in Markdown. Geben Sie den entsprechenden HTML-String zurück.

Ein Markdown-Link hat folgendes Format: `"[link_text](link_url)"`. Geben Sie den String des HTML-`a`-Tags zurück, wobei das `href` auf die `link_url` gesetzt ist und der `link_text` als Inhalt des Tags dient.

Zum Beispiel, gegeben `"[freeCodeCamp](https://freecodecamp.org/)"`, geben Sie `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'` zurück;

Hinweis: Die Konsole zeigt HTML-Tags in Strings beim Protokollieren von Nachrichten möglicherweise nicht an – prüfen Sie die Browser-Konsole, um Protokolle mit enthaltenen Tags zu sehen.

# --hints--

`parse_link("[freeCodeCamp](https://freecodecamp.org/)")` sollte `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>')`)
}})
```

`parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` sollte `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>')`)
}})
```

`parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` sollte `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_link("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_link(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_link(markdown):
    match = re.match(r'\[(.*?)\]\((.*?)\)', markdown)
    text, url = match.groups()
    return f'<a href="{url}">{text}</a>'
```
