---
id: 6939b873185d8e00d453563f
title: "Challenge 160: Knight Moves"
challengeType: 29
dashedName: challenge-160
---

# --description--

Gegeben ist die Position eines Springers auf einem Schachbrett. Geben Sie die Anzahl der gültigen Felder zurück, auf die der Springer ziehen kann.

Ein Standard-Schachbrett ist 8x8, mit Spalten, die von `A` bis `H` beschriftet sind (von links nach rechts), und Reihen, die von `1` bis `8` beschriftet sind (von unten nach oben). Es sieht so aus:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Ein Springer zieht in einer „L“-Form: zwei Felder in eine Richtung (horizontal oder vertikal) und ein Feld in die senkrechte Richtung.

Das bedeutet, ein Springer kann bis zu acht mögliche Positionen erreichen, aber weniger, wenn er sich nahe am Rand des Brettes befindet. Zum Beispiel könnte ein Springer auf `A1` nur auf `B3` oder `C2` ziehen.

# --hints--

`knight_moves("A1")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("A1"), 2)`)
}})
```

`knight_moves("D4")` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("D4"), 8)`)
}})
```

`knight_moves("G6")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("G6"), 6)`)
}})
```

`knight_moves("B8")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("B8"), 3)`)
}})
```

`knight_moves("H3")` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("H3"), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def knight_moves(position):

    return position
```

# --solutions--

```py
def knight_moves(position):
    col = ord(position[0]) - ord('A')
    row = int(position[1]) - 1

    moves = [
        (2, 1), (2, -1),
        (-2, 1), (-2, -1),
        (1, 2), (1, -2),
        (-1, 2), (-1, -2)
    ]

    valid = 0

    for dx, dy in moves:
        new_col = col + dx
        new_row = row + dy

        if 0 <= new_col < 8 and 0 <= new_row < 8:
            valid += 1

    return valid
```
