---
id: 694596b0585c11170ac7c7f9
title: "Challenge 161: Free Shipping"
challengeType: 29
dashedName: challenge-161
---

# --description--

Gegeben ist ein Array von Strings, die Artikel in Ihrem Warenkorb darstellen, sowie eine Zahl für den Mindestbestellwert, um sich für kostenlosen Versand zu qualifizieren. Bestimmen Sie, ob die Artikel in Ihrem Warenkorb für kostenlosen Versand qualifizieren.

Das gegebene Array enthält Artikel aus der folgenden Liste:

| Item | Price |
| - | - |
| `"shirt"`	| 34.25 |
| `"jeans"`	| 48.50 |
| `"shoes"` |	75.00 |
| `"hat"`	| 19.95 |
| `"socks"` | 15.00 |
| `"jacket"` | 109.95 |

# --hints--

`gets_free_shipping(["shoes"], 50)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shoes"], 50), True)`)
}})
```

`gets_free_shipping(["hat", "socks"], 50)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks"], 50), False)`)
}})
```

`gets_free_shipping(["jeans", "shirt", "jacket"], 75)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["jeans", "shirt", "jacket"], 75), True)`)
}})
```

`gets_free_shipping(["socks", "socks", "hat"], 75)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["socks", "socks", "hat"], 75), False)`)
}})
```

`gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["shirt", "shirt", "jeans", "socks"], 100), True)`)
}})
```

`gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(gets_free_shipping(["hat", "socks", "hat", "jeans", "shoes", "hat"], 200), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def gets_free_shipping(cart, minimum):

    return cart
```

# --solutions--

```py
def gets_free_shipping(cart, minimum):
    prices = {
        "shirt": 34.25,
        "jeans": 48.50,
        "shoes": 75.00,
        "hat": 19.95,
        "socks": 15.00,
        "jacket": 109.95
    }

    total = 0

    for item in cart:
        total += prices[item]

    return total >= minimum
```
