---
id: 694596b0585c11170ac7c7fa
title: "Challenge 162: Energy Consumption"
challengeType: 29
dashedName: challenge-162
---

# --description--

Gegeben ist die Anzahl der während eines Trainings verbrannten Kalorien und die Anzahl der Wattstunden, die Ihre elektronischen Geräte während dieses Trainings verbraucht haben. Bestimmen Sie, welches von beiden mehr Energie verbraucht hat.

Um sie zu vergleichen, wandeln Sie beide Werte mit den folgenden Umrechnungen in Joule um:

- 1 Kalorie entspricht 4184 Joule.
- 1 Wattstunde entspricht 3600 Joule.

Return:

- `"Workout"`, wenn das Training mehr Energie verbraucht hat.
- `"Devices"`, wenn das Gerät mehr Energie verbraucht hat.
- `"Equal"`, wenn beide die gleiche Energiemenge verbraucht haben.

# --hints--

`compare_energy(250, 50)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```
