---
id: 694596b0585c11170ac7c7fb
title: "Challenge 163: Consonant Case"
challengeType: 29
dashedName: challenge-163
---

# --description--

Gegeben ist ein String, der einen Variablennamen darstellt. Wandeln Sie ihn in Konsonanten-Großschreibung um, indem Sie folgende Regeln anwenden:

- Alle Konsonanten sollen in Großbuchstaben umgewandelt werden.
- Alle Vokale (`a`, `e`, `i`, `o`, `u` in beliebiger Schreibweise) sollen in Kleinbuchstaben umgewandelt werden.
- Alle Bindestriche (`-`) sollen in Unterstriche (`_`) umgewandelt werden.

# --hints--

`to_consonant_case("helloworld")` should return `"HeLLoWoRLD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("helloworld"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("HELLOWORLD")` should return `"HeLLoWoRLD"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("HELLOWORLD"), "HeLLoWoRLD")`)
}})
```

`to_consonant_case("_hElLO-WOrlD-")` should return `"_HeLLo_WoRLD_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_hElLO-WOrlD-"), "_HeLLo_WoRLD_")`)
}})
```

`to_consonant_case("_~-generic_~-variable_~-name_~-here-~_")` should return `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.


```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_consonant_case("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_consonant_case(s):

    return s
```

# --solutions--

```py
def to_consonant_case(s):
    vowels = "aeiouAEIOU"
    result = ""

    for char in s:
        if char == "-":
            result += "_"
        elif char.isalpha():
            if char in vowels:
                result += char.lower()
            else:
                result += char.upper()
        else:
            result += char

    return result
```
