---
id: 696655d24b614176d4c9b78d
title: "Challenge 170: Odd or Even Day"
challengeType: 29
dashedName: challenge-170
---

# --description--

Gegeben ist ein Zeitstempel (Anzahl der Millisekunden seit dem Unix-Epoch). Geben Sie zurück:

- `"odd"`, wenn der Tag des Monats für diesen Zeitstempel ungerade ist.
- `"even"`, wenn der Tag des Monats für diesen Zeitstempel gerade ist.

Zum Beispiel, bei `1769472000000`, einem Zeitstempel für den 27. Januar 2026, geben Sie `"odd"` zurück, weil die Tageszahl (`27`) eine ungerade Zahl ist.

Hinweis: Der Zeitstempel ist in Millisekunden angegeben, und Sie sollten das Datum in der UTC-Zeitzone verwenden, nicht in Ihrer lokalen Zeit.

# --hints--

`odd_or_even_day(1769472000000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```
