---
id: 69738771fb5a7b8b24cca2a2
title: "Challenge 176: Groundhog Day"
challengeType: 29
dashedName: challenge-176
---

# --description--

Heute ist Murmeltiertag, an dem ein Murmeltier das Wetter vorhersagt, indem es beobachtet, ob es seinen Schatten sieht oder nicht.

Geben Sie basierend auf einem Wert, der das Erscheinen des Murmeltiers darstellt, die korrekte Vorhersage zurück:

- Wenn der gegebene Wert der Boolean `true` ist (das Murmeltier hat seinen Schatten gesehen), geben Sie `"Looks like we'll have six more weeks of winter."` zurück.
- Wenn der Wert der Boolean `false` ist (das Murmeltier hat seinen Schatten nicht gesehen), geben Sie `"It's going to be an early spring."` zurück.
- Wenn der Wert etwas anderes ist (das Murmeltier ist nicht aufgetaucht), geben Sie `"No prediction this year."` zurück.

# --hints--

`groundhog_day_prediction(True)` should return `"Looks like we'll have six more weeks of winter."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(True), "Looks like we'll have six more weeks of winter.")`)
}})
```

`groundhog_day_prediction(False)` should return `"It's going to be an early spring."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(False), "It's going to be an early spring.")`)
}})
```

`groundhog_day_prediction(None)` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(None), "No prediction this year.")`)
}})
```

`groundhog_day_prediction(" ")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction(" "), "No prediction this year.")`)
}})
```

`groundhog_day_prediction("True")` should return `"No prediction this year."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(groundhog_day_prediction("True"), "No prediction this year.")`)
}})
```

# --seed--

## --seed-contents--

```py
def groundhog_day_prediction(appearance):

    return appearance
```

# --solutions--

```py
def groundhog_day_prediction(appearance):
    if appearance is True:
        return "Looks like we'll have six more weeks of winter."
    if appearance is False:
        return "It's going to be an early spring."
    return "No prediction this year."
```
