---
id: 697a49e6ff50d756c9b69366
title: "Aufgabe 189: Winterspiele Tag 10: Alpiner Skisport"
challengeType: 29
dashedName: challenge-189
---

# --description--

Gegeben sind der Höhenunterschied, die horizontale Distanz und der Typ eines Skihangs. Bestimmen Sie die Schwierigkeitsbewertung des Hangs.

Zur Bestimmung der Bewertung:

- Berechnen Sie die Steilheit des Hangs, indem Sie den Höhenunterschied durch die Distanz teilen.
- Berechnen Sie dann die angepasste Steilheit basierend auf dem Hangtyp:
  - `"Downhill"`: multiplizieren Sie die Steilheit mit 1,2
  - `"Slalom"`: multiplizieren Sie die Steilheit mit 0,9
  - `"Giant Slalom"`: multiplizieren Sie die Steilheit mit 1,0

Geben Sie zurück:

- `"Green"`, wenn die angepasste Steilheit kleiner oder gleich 0,1 ist
- `"Blue"`, wenn die angepasste Steilheit größer als 0,1 und kleiner oder gleich 0,25 ist
- `"Black"`, wenn die angepasste Steilheit größer als 0,25 ist


# --hints--

`get_hill_rating(95, 900, "Slalom")` sollte `"Green"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(95, 900, "Slalom"), "Green")`)
}})
```

`get_hill_rating(620, 2800, "Downhill")` sollte `"Black"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(620, 2800, "Downhill"), "Black")`)
}})
```

`get_hill_rating(420, 1680, "Giant Slalom")` sollte `"Blue"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(420, 1680, "Giant Slalom"), "Blue")`)
}})
```

`get_hill_rating(250, 3000, "Downhill")` sollte `"Green"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(250, 3000, "Downhill"), "Green")`)
}})
```

`get_hill_rating(110, 900, "Slalom")` sollte `"Blue"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(110, 900, "Slalom"), "Blue")`)
}})
```

`get_hill_rating(380, 1500, "Giant Slalom")` sollte `"Black"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_hill_rating(380, 1500, "Giant Slalom"), "Black")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_hill_rating(drop, distance, hill_type):

    return drop
```

# --solutions--

```py
def get_hill_rating(drop, distance, hill_type):
    steepness = drop / distance

    if hill_type == "Downhill":
        steepness *= 1.2
    elif hill_type == "Slalom":
        steepness *= 0.9
    elif hill_type == "Giant Slalom":
        steepness *= 1.0

    if steepness <= 0.1:
        return "Green"
    elif steepness <= 0.25:
        return "Blue"
    else:
        return "Black"
```
