---
id: 697a49e6ff50d756c9b69362
title: "Challenge 185: 2026 Winter Games Day 6: Figure Skating"
challengeType: 29
dashedName: challenge-185
---

# --description--

Gegeben ist ein Array mit Wertungen von Juroren und optionalen Strafpunkten. Berechnen Sie die Endpunktzahl für eine Eiskunstlauf-Kür.

Das erste Argument ist ein Array mit 10 Wertungen von Juroren, jeweils eine Zahl von 0 bis 10. Entfernen Sie die höchste und die niedrigste Wertung und summieren Sie die verbleibenden 8 Wertungen, um die Grundpunktzahl zu erhalten.

Alle weiteren an die Funktion übergebenen Argumente sind Strafpunkte. Ziehen Sie alle Strafpunkte von der Grundpunktzahl ab, um die Endpunktzahl zu erhalten.

# --hints--

`compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` should return `64`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64)`)
}})
```

`compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` should return `80`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80)`)
}})
```

`compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` should return `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67)`)
}})
```

`compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` should return `67.5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5)`)
}})
```

`compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` should return `59`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compute_score([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59)`)
}})
```

# --seed--

## --seed-contents--

```py
def compute_score(judge_scores, *penalties):

    return judge_scores
```

# --solutions--

```py
def compute_score(judge_scores, *penalties):
    sorted_scores = sorted(judge_scores)
    base_score = sum(sorted_scores[1:9])
    total_penalty = sum(penalties)

    return base_score - total_penalty
```
