---
id: 697a49e6ff50d756c9b6935f
title: "Aufgabe 182: Winterspiele Tag 3: Biathlon"
challengeType: 29
dashedName: challenge-182
---

# --description--

Gegeben ist ein Array aus ganzen Zahlen, wobei jeder Wert die Anzahl der in einer Runde eines Biathlons getroffenen Ziele darstellt. Geben Sie die gesamte Strafdistanz zurück, die der Athlet laufen muss.

- Jede Runde besteht aus 5 Zielen.
- Jedes verfehlte Ziel führt zu einer Strafschleife von 150 Metern.

# --hints--

`calculate_penalty_distance([4, 4])` sollte `300` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 4]), 300)`)
}})
```

`calculate_penalty_distance([5, 5])` sollte `0` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 5]), 0)`)
}})
```

`calculate_penalty_distance([4, 5, 3, 5])` sollte `450` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 5, 3, 5]), 450)`)
}})
```

`calculate_penalty_distance([5, 4, 5, 5])` sollte `150` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([5, 4, 5, 5]), 150)`)
}})
```

`calculate_penalty_distance([4, 3, 0, 3])` sollte `1500` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_penalty_distance([4, 3, 0, 3]), 1500)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_penalty_distance(rounds):

    return rounds
```

# --solutions--

```py
def calculate_penalty_distance(rounds):
    total_penalty = 0

    for hits in rounds:
        misses = 5 - hits
        total_penalty += misses * 150

    return total_penalty
```
