---
id: 697a49e6ff50d756c9b6935d
title: "Challenge 180: 2026 Winter Games Day 1: Opening Day"
challengeType: 29
dashedName: challenge-180
---

# --description--

Heute beginnt die Winterspiele 2026. Die nĂ¤chsten 17 Tage erwarten Sie Programmieraufgaben, die von ihnen inspiriert sind.

FĂĽr die erste Aufgabe erhalten Sie einen zweibuchstabigen LĂ¤ndercode und sollen das Flaggen-Emoji fĂĽr dieses Land zurĂĽckgeben.

Verwenden Sie diese Liste:

| Land                       | Code       | Flagge   |
| -                         | -:         | :--:     |
| Albanien                  |            `"AL"` | `"đź‡¦đź‡±"` |
| Andorra                   |            `"AD"` | `"đź‡¦đź‡©"` |
| Argentinien               |            `"AR"` | `"đź‡¦đź‡·"` |
| Armenien                  |            `"AM"` | `"đź‡¦đź‡˛"` |
| Australien                |            `"AU"` | `"đź‡¦đź‡ş"` |
| Ă–sterreich                |            `"AT"` | `"đź‡¦đź‡ą"` |
| Aserbaidschan             |            `"AZ"` | `"đź‡¦đź‡ż"` |
| Belgien                   |            `"BE"` | `"đź‡§đź‡Ş"` |
| Benin                     |            `"BJ"` | `"đź‡§đź‡Ż"` |
| Bolivien                  |            `"BO"` | `"đź‡§đź‡´"` |
| Bosnien und Herzegowina   |            `"BA"` | `"đź‡§đź‡¦"` |
| Brasilien                 |            `"BR"` | `"đź‡§đź‡·"` |
| Bulgarien                 |            `"BG"` | `"đź‡§đź‡¬"` |
| Kanada                    |            `"CA"` | `"đź‡¨đź‡¦"` |
| Chile                     |            `"CL"` | `"đź‡¨đź‡±"` |
| China                     |            `"CN"` | `"đź‡¨đź‡ł"` |
| Kolumbien                 |            `"CO"` | `"đź‡¨đź‡´"` |
| Kroatien                  |            `"HR"` | `"đź‡­đź‡·"` |
| Zypern                    |            `"CY"` | `"đź‡¨đź‡ľ"` |
| Tschechien                |            `"CZ"` | `"đź‡¨đź‡ż"` |
| DĂ¤nemark                  |            `"DK"` | `"đź‡©đź‡°"` |
| Ecuador                   |            `"EC"` | `"đź‡Şđź‡¨"` |
| Eritrea                   |            `"ER"` | `"đź‡Şđź‡·"` |
| Estland                   |            `"EE"` | `"đź‡Şđź‡Ş"` |
| Finnland                  |            `"FI"` | `"đź‡«đź‡®"` |
| Frankreich                |            `"FR"` | `"đź‡«đź‡·"` |
| Georgien                  |            `"GE"` | `"đź‡¬đź‡Ş"` |
| Deutschland               |            `"DE"` | `"đź‡©đź‡Ş"` |
| GroĂźbritannien            |            `"GB"` | `"đź‡¬đź‡§"` |
| Griechenland              |            `"GR"` | `"đź‡¬đź‡·"` |
| Guinea-Bissau             |            `"GW"` | `"đź‡¬đź‡Ľ"` |
| Haiti                     |            `"HT"` | `"đź‡­đź‡ą"` |
| Hongkong                  |            `"HK"` | `"đź‡­đź‡°"` |
| Ungarn                    |            `"HU"` | `"đź‡­đź‡ş"` |
| Island                    |            `"IS"` | `"đź‡®đź‡¸"` |
| Indien                    |            `"IN"` | `"đź‡®đź‡ł"` |
| Iran                      |            `"IR"` | `"đź‡®đź‡·"` |
| Irland                    |            `"IE"` | `"đź‡®đź‡Ş"` |
| Israel                    |            `"IL"` | `"đź‡®đź‡±"` |
| Italien                   |            `"IT"` | `"đź‡®đź‡ą"` |
| Jamaika                   |            `"JM"` | `"đź‡Żđź‡˛"` |
| Japan                     |            `"JP"` | `"đź‡Żđź‡µ"` |
| Kasachstan                |            `"KZ"` | `"đź‡°đź‡ż"` |
| Kenia                     |            `"KE"` | `"đź‡°đź‡Ş"` |
| Kosovo                    |            `"XK"` | `"đź‡˝đź‡°"` |
| Kirgisistan               |            `"KG"` | `"đź‡°đź‡¬"` |
| Lettland                  |            `"LV"` | `"đź‡±đź‡»"` |
| Libanon                   |            `"LB"` | `"đź‡±đź‡§"` |
| Liechtenstein             |            `"LI"` | `"đź‡±đź‡®"` |
| Litauen                   |            `"LT"` | `"đź‡±đź‡ą"` |
| Luxemburg                 |            `"LU"` | `"đź‡±đź‡ş"` |
| Madagaskar                |            `"MG"` | `"đź‡˛đź‡¬"` |
| Malaysia                  |            `"MY"` | `"đź‡˛đź‡ľ"` |
| Malta                     |            `"MT"` | `"đź‡˛đź‡ą"` |
| Mexiko                    |            `"MX"` | `"đź‡˛đź‡˝"` |
| Moldawien                 |            `"MD"` | `"đź‡˛đź‡©"` |
| Monaco                    |            `"MC"` | `"đź‡˛đź‡¨"` |
| Mongolei                  |            `"MN"` | `"đź‡˛đź‡ł"` |
| Montenegro                |            `"ME"` | `"đź‡˛đź‡Ş"` |
| Marokko                   |            `"MA"` | `"đź‡˛đź‡¦"` |
| Niederlande               |            `"NL"` | `"đź‡łđź‡±"` |
| Neuseeland                |            `"NZ"` | `"đź‡łđź‡ż"` |
| Nigeria                   |            `"NG"` | `"đź‡łđź‡¬"` |
| Nordmazedonien            |            `"MK"` | `"đź‡˛đź‡°"` |
| Norwegen                  |            `"NO"` | `"đź‡łđź‡´"` |
| Pakistan                  |            `"PK"` | `"đź‡µđź‡°"` |
| Philippinen               |            `"PH"` | `"đź‡µđź‡­"` |
| Polen                     |            `"PL"` | `"đź‡µđź‡±"` |
| Portugal                  |            `"PT"` | `"đź‡µđź‡ą"` |
| Puerto Rico               |            `"PR"` | `"đź‡µđź‡·"` |
| RumĂ¤nien                  |            `"RO"` | `"đź‡·đź‡´"` |
| San Marino                |            `"SM"` | `"đź‡¸đź‡˛"` |
| Saudi-Arabien             |            `"SA"` | `"đź‡¸đź‡¦"` |
| Serbien                   |            `"RS"` | `"đź‡·đź‡¸"` |
| Singapur                  |            `"SG"` | `"đź‡¸đź‡¬"` |
| Slowakei                  |            `"SK"` | `"đź‡¸đź‡°"` |
| Slowenien                 |            `"SI"` | `"đź‡¸đź‡®"` |
| SĂĽdafrika                 |            `"ZA"` | `"đź‡żđź‡¦"` |
| SĂĽdkorea                  |            `"KR"` | `"đź‡°đź‡·"` |
| Spanien                   |            `"ES"` | `"đź‡Şđź‡¸"` |
| Schweden                  |            `"SE"` | `"đź‡¸đź‡Ş"` |
| Schweiz                   |            `"CH"` | `"đź‡¨đź‡­"` |
| Thailand                  |            `"TH"` | `"đź‡ąđź‡­"` |
| Trinidad und Tobago       |            `"TT"` | `"đź‡ąđź‡ą"` |
| TĂĽrkei                    |            `"TR"` | `"đź‡ąđź‡·"` |
| Ukraine                   |            `"UA"` | `"đź‡şđź‡¦"` |
| Vereinigte Arabische Emirate |        `"AE"` | `"đź‡¦đź‡Ş"` |
| Vereinigte Staaten        |            `"US"` | `"đź‡şđź‡¸"` |
| Uruguay                   |            `"UY"` | `"đź‡şđź‡ľ"` |
| Usbekistan                |            `"UZ"` | `"đź‡şđź‡ż"` |
| Venezuela                 |            `"VE"` | `"đź‡»đź‡Ş"` |

# --hints--

`get_flag("AL")` should return `"đź‡¦đź‡±"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "đź‡¦đź‡±")`)
}})
```

`get_flag("AD")` should return `"đź‡¦đź‡©"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "đź‡¦đź‡©")`)
}})
```

`get_flag("AR")` should return `"đź‡¦đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "đź‡¦đź‡·")`)
}})
```

`get_flag("AM")` should return `"đź‡¦đź‡˛"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "đź‡¦đź‡˛")`)
}})
```

`get_flag("AU")` should return `"đź‡¦đź‡ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "đź‡¦đź‡ş")`)
}})
```

`get_flag("AT")` should return `"đź‡¦đź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "đź‡¦đź‡ą")`)
}})
```

`get_flag("AZ")` should return `"đź‡¦đź‡ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "đź‡¦đź‡ż")`)
}})
```

`get_flag("BE")` should return `"đź‡§đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "đź‡§đź‡Ş")`)
}})
```

`get_flag("BJ")` should return `"đź‡§đź‡Ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "đź‡§đź‡Ż")`)
}})
```

`get_flag("BO")` should return `"đź‡§đź‡´"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "đź‡§đź‡´")`)
}})
```

`get_flag("BA")` should return `"đź‡§đź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "đź‡§đź‡¦")`)
}})
```

`get_flag("BR")` should return `"đź‡§đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "đź‡§đź‡·")`)
}})
```

`get_flag("BG")` should return `"đź‡§đź‡¬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "đź‡§đź‡¬")`)
}})
```

`get_flag("CA")` should return `"đź‡¨đź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "đź‡¨đź‡¦")`)
}})
```

`get_flag("CL")` should return `"đź‡¨đź‡±"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "đź‡¨đź‡±")`)
}})
```

`get_flag("CN")` should return `"đź‡¨đź‡ł"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "đź‡¨đź‡ł")`)
}})
```

`get_flag("CO")` should return `"đź‡¨đź‡´"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "đź‡¨đź‡´")`)
}})
```

`get_flag("HR")` should return `"đź‡­đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "đź‡­đź‡·")`)
}})
```

`get_flag("CY")` should return `"đź‡¨đź‡ľ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "đź‡¨đź‡ľ")`)
}})
```

`get_flag("CZ")` should return `"đź‡¨đź‡ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "đź‡¨đź‡ż")`)
}})
```

`get_flag("DK")` should return `"đź‡©đź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "đź‡©đź‡°")`)
}})
```

`get_flag("EC")` should return `"đź‡Şđź‡¨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "đź‡Şđź‡¨")`)
}})
```

`get_flag("ER")` should return `"đź‡Şđź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "đź‡Şđź‡·")`)
}})
```

`get_flag("EE")` should return `"đź‡Şđź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "đź‡Şđź‡Ş")`)
}})
```

`get_flag("FI")` should return `"đź‡«đź‡®"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "đź‡«đź‡®")`)
}})
```

`get_flag("FR")` should return `"đź‡«đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "đź‡«đź‡·")`)
}})
```

`get_flag("GE")` should return `"đź‡¬đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "đź‡¬đź‡Ş")`)
}})
```

`get_flag("DE")` should return `"đź‡©đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "đź‡©đź‡Ş")`)
}})
```

`get_flag("GB")` should return `"đź‡¬đź‡§"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "đź‡¬đź‡§")`)
}})
```

`get_flag("GR")` should return `"đź‡¬đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "đź‡¬đź‡·")`)
}})
```

`get_flag("GW")` should return `"đź‡¬đź‡Ľ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "đź‡¬đź‡Ľ")`)
}})
```

`get_flag("HT")` should return `"đź‡­đź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "đź‡­đź‡ą")`)
}})
```

`get_flag("HK")` should return `"đź‡­đź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "đź‡­đź‡°")`)
}})
```

`get_flag("HU")` should return `"đź‡­đź‡ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "đź‡­đź‡ş")`)
}})
```

`get_flag("IS")` should return `"đź‡®đź‡¸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "đź‡®đź‡¸")`)
}})
```

`get_flag("IN")` should return `"đź‡®đź‡ł"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "đź‡®đź‡ł")`)
}})
```

`get_flag("IR")` should return `"đź‡®đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "đź‡®đź‡·")`)
}})
```

`get_flag("IE")` should return `"đź‡®đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "đź‡®đź‡Ş")`)
}})
```

`get_flag("IL")` should return `"đź‡®đź‡±"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "đź‡®đź‡±")`)
}})
```

`get_flag("IT")` should return `"đź‡®đź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "đź‡®đź‡ą")`)
}})
```

`get_flag("JM")` should return `"đź‡Żđź‡˛"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "đź‡Żđź‡˛")`)
}})
```

`get_flag("JP")` should return `"đź‡Żđź‡µ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "đź‡Żđź‡µ")`)
}})
```

`get_flag("KZ")` should return `"đź‡°đź‡ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "đź‡°đź‡ż")`)
}})
```

`get_flag("KE")` should return `"đź‡°đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "đź‡°đź‡Ş")`)
}})
```

`get_flag("XK")` should return `"đź‡˝đź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "đź‡˝đź‡°")`)
}})
```

`get_flag("KG")` should return `"đź‡°đź‡¬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "đź‡°đź‡¬")`)
}})
```

`get_flag("LV")` should return `"đź‡±đź‡»"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "đź‡±đź‡»")`)
}})
```

`get_flag("LB")` should return `"đź‡±đź‡§"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "đź‡±đź‡§")`)
}})
```

`get_flag("LI")` should return `"đź‡±đź‡®"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "đź‡±đź‡®")`)
}})
```

`get_flag("LT")` should return `"đź‡±đź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "đź‡±đź‡ą")`)
}})
```

`get_flag("LU")` should return `"đź‡±đź‡ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "đź‡±đź‡ş")`)
}})
```

`get_flag("MG")` should return `"đź‡˛đź‡¬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "đź‡˛đź‡¬")`)
}})
```

`get_flag("MY")` should return `"đź‡˛đź‡ľ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "đź‡˛đź‡ľ")`)
}})
```

`get_flag("MT")` should return `"đź‡˛đź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "đź‡˛đź‡ą")`)
}})
```

`get_flag("MX")` should return `"đź‡˛đź‡˝"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "đź‡˛đź‡˝")`)
}})
```

`get_flag("MD")` should return `"đź‡˛đź‡©"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "đź‡˛đź‡©")`)
}})
```

`get_flag("MC")` should return `"đź‡˛đź‡¨"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "đź‡˛đź‡¨")`)
}})
```

`get_flag("MN")` should return `"đź‡˛đź‡ł"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "đź‡˛đź‡ł")`)
}})
```

`get_flag("ME")` should return `"đź‡˛đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "đź‡˛đź‡Ş")`)
}})
```

`get_flag("MA")` should return `"đź‡˛đź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "đź‡˛đź‡¦")`)
}})
```

`get_flag("NL")` should return `"đź‡łđź‡±"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "đź‡łđź‡±")`)
}})
```

`get_flag("NZ")` should return `"đź‡łđź‡ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "đź‡łđź‡ż")`)
}})
```

`get_flag("NG")` should return `"đź‡łđź‡¬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "đź‡łđź‡¬")`)
}})
```

`get_flag("MK")` should return `"đź‡˛đź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "đź‡˛đź‡°")`)
}})
```

`get_flag("NO")` should return `"đź‡łđź‡´"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "đź‡łđź‡´")`)
}})
```

`get_flag("PK")` should return `"đź‡µđź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "đź‡µđź‡°")`)
}})
```

`get_flag("PH")` should return `"đź‡µđź‡­"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "đź‡µđź‡­")`)
}})
```

`get_flag("PL")` should return `"đź‡µđź‡±"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "đź‡µđź‡±")`)
}})
```

`get_flag("PT")` should return `"đź‡µđź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "đź‡µđź‡ą")`)
}})
```

`get_flag("PR")` should return `"đź‡µđź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "đź‡µđź‡·")`)
}})
```

`get_flag("RO")` should return `"đź‡·đź‡´"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "đź‡·đź‡´")`)
}})
```

`get_flag("SM")` should return `"đź‡¸đź‡˛"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "đź‡¸đź‡˛")`)
}})
```

`get_flag("SA")` should return `"đź‡¸đź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "đź‡¸đź‡¦")`)
}})
```

`get_flag("RS")` should return `"đź‡·đź‡¸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "đź‡·đź‡¸")`)
}})
```

`get_flag("SG")` should return `"đź‡¸đź‡¬"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "đź‡¸đź‡¬")`)
}})
```

`get_flag("SK")` should return `"đź‡¸đź‡°"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "đź‡¸đź‡°")`)
}})
```

`get_flag("SI")` should return `"đź‡¸đź‡®"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "đź‡¸đź‡®")`)
}})
```

`get_flag("ZA")` should return `"đź‡żđź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "đź‡żđź‡¦")`)
}})
```

`get_flag("KR")` should return `"đź‡°đź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "đź‡°đź‡·")`)
}})
```

`get_flag("ES")` should return `"đź‡Şđź‡¸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "đź‡Şđź‡¸")`)
}})
```

`get_flag("SE")` should return `"đź‡¸đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "đź‡¸đź‡Ş")`)
}})
```

`get_flag("CH")` should return `"đź‡¨đź‡­"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "đź‡¨đź‡­")`)
}})
```

`get_flag("TH")` should return `"đź‡ąđź‡­"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "đź‡ąđź‡­")`)
}})
```

`get_flag("TT")` should return `"đź‡ąđź‡ą"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "đź‡ąđź‡ą")`)
}})
```

`get_flag("TR")` should return `"đź‡ąđź‡·"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "đź‡ąđź‡·")`)
}})
```

`get_flag("UA")` should return `"đź‡şđź‡¦"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "đź‡şđź‡¦")`)
}})
```

`get_flag("AE")` should return `"đź‡¦đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "đź‡¦đź‡Ş")`)
}})
```

`get_flag("US")` should return `"đź‡şđź‡¸"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "đź‡şđź‡¸")`)
}})
```

`get_flag("UY")` should return `"đź‡şđź‡ľ"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "đź‡şđź‡ľ")`)
}})
```

`get_flag("UZ")` should return `"đź‡şđź‡ż"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "đź‡şđź‡ż")`)
}})
```

`get_flag("VE")` should return `"đź‡»đź‡Ş"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "đź‡»đź‡Ş")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
