---
id: 697a49e9860d24853adef67e
title: "Aufgabe 193: Winterspiele Tag 14: Skibergsteigen"
challengeType: 29
dashedName: challenge-193
---

# --description--

Anhand der Schneetiefe und des Gefälles eines Bergs bestimmen Sie, ob eine Lawinengefahr besteht.

- Die Werte für die Schneetiefe sind `"Shallow"`, `"Moderate"` oder `"Deep"`.
- Die Werte für das Gefälle sind `"Gentle"`, `"Steep"` oder `"Very Steep"`.

Geben Sie `"Safe"` oder `"Risky"` basierend auf dieser Tabelle zurück:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalanche_risk("Shallow", "Gentle")` sollte `"Safe"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Steep")` sollte `"Safe"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Steep"), "Safe")`)
}})
```

`avalanche_risk("Shallow", "Very Steep")` sollte `"Safe"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Shallow", "Very Steep"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Gentle")` sollte `"Safe"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Moderate", "Steep")` sollte `"Risky"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Moderate", "Very Steep")` sollte `"Risky"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Moderate", "Very Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Gentle")` sollte `"Safe"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Gentle"), "Safe")`)
}})
```

`avalanche_risk("Deep", "Steep")` sollte `"Risky"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Steep"), "Risky")`)
}})
```

`avalanche_risk("Deep", "Very Steep")` sollte `"Risky"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(avalanche_risk("Deep", "Very Steep"), "Risky")`)
}})
```

# --seed--

## --seed-contents--

```py
def avalanche_risk(snow_depth, slope):

    return snow_depth
```

# --solutions--

```py
def avalanche_risk(snow_depth, slope):
    risk_table = {
        "Gentle": {
            "Shallow": "Safe",
            "Moderate": "Safe",
            "Deep": "Safe",
        },
        "Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
        "Very Steep": {
            "Shallow": "Safe",
            "Moderate": "Risky",
            "Deep": "Risky",
        },
    }

    return risk_table[slope][snow_depth]
```
