---
id: 697a49e9860d24853adef67b
title: "Aufgabe 190: Winterspiele Tag 11: Eishockey"
challengeType: 29
dashedName: challenge-190
---

# --description--

Gegeben ist ein Array mit 6 Eishockeyteams und deren Ergebnissen nach den Rundenspielen. Bestimmen Sie die Paarungen für die Halbfinalrunde.

- Jedes Array-Element enthält ein Team und dessen Ergebnis im Format `"TEAM: W-OTW-OTL-L"`. Dabei gilt:
  - `"W"` ist die Anzahl der regulären Siege, die jeweils 3 Punkte wert sind
  - `"OTW"` ist die Anzahl der Siege in der Verlängerung, die jeweils 2 Punkte wert sind
  - `"OTL"` ist die Anzahl der Niederlagen in der Verlängerung, die jeweils 1 Punkt wert sind
  - `"L"` ist die Anzahl der Niederlagen, die jeweils 0 Punkte wert sind
  
Zum Beispiel hätte `"FIN: 2-2-1-0"` nach Addition der Punkte 11 Punkte.
  
Berechnen Sie die Gesamtpunktzahl für jedes Team und geben Sie `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."` zurück. Zum Beispiel `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` sollte `"The semi-final games will be FIN vs SWE and CAN vs USA."` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.")`)
}})
```

`get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` sollte `"The semi-final games will be USA vs CZE and CAN vs FIN."` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.")`)
}})
```

`get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` sollte `"The semi-final games will be CAN vs ITA and USA vs CZE."` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.")`)
}})
```

`get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` sollte `"The semi-final games will be JPN vs ITA and AUT vs KOR."` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_semifinal_matchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_semifinal_matchups(teams):

    return teams
```

# --solutions--

```py
def get_semifinal_matchups(teams):
    team_points = []
    for team_str in teams:
        name, record = team_str.split(": ")
        W, OTW, OTL, L = map(int, record.split("-"))
        points = W * 3 + OTW * 2 + OTL * 1 + L * 0
        team_points.append((name, points))

    team_points.sort(key=lambda x: x[1], reverse=True)
    first, second, third, fourth = team_points[:4]

    return f"The semi-final games will be {first[0]} vs {fourth[0]} and {second[0]} vs {third[0]}."
```
