---
id: 697a49e6ff50d756c9b6935e
title: "Aufgabe 181: Winterspiele 2026 Tag 2: Snowboarden"
challengeType: 29
dashedName: challenge-181
---

# --description--

Gegeben sind die Ausgangsposition eines Snowboarders und eine Drehung in Grad. Bestimmen Sie die Landeposition.

- Die Position eines Snowboarders ist entweder `"Regular"` oder `"Goofy"`.
- Drehungen bei Tricks sind Vielfache von 90 Grad. Positive Werte bedeuten eine Drehung im Uhrzeigersinn, negative eine gegen den Uhrzeigersinn.
- Die Landeposition wechselt alle 180 Grad Drehung.

Zum Beispiel: Bei `"Regular"` und `90` geben Sie `"Regular"` zurück. Bei `"Regular"` und `180` Grad geben Sie `"Goofy"` zurück.

# --hints--

`get_landing_stance("Regular", 90)` sollte `"Regular"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 90), "Regular")`)
}})
```

`get_landing_stance("Regular", 180)` sollte `"Goofy"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 180), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -270)` sollte `"Regular"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -270), "Regular")`)
}})
```

`get_landing_stance("Regular", 2340)` sollte `"Goofy"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Regular", 2340), "Goofy")`)
}})
```

`get_landing_stance("Goofy", 2160)` sollte `"Goofy"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 2160), "Goofy")`)
}})
```

`get_landing_stance("Goofy", -540)` sollte `"Regular"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", -540), "Regular")`)
}})
```

`get_landing_stance("Goofy", 90)` sollte `"Goofy"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_landing_stance("Goofy", 90), "Goofy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_landing_stance(start_stance, rotation):

    return start_stance
```

# --solutions--

```py
def get_landing_stance(start_stance, rotation):
    flips = abs(rotation) // 180
    if flips % 2 == 0:
        return start_stance
    else:
        return "Goofy" if start_stance == "Regular" else "Regular"
```
