---
id: 697a49e9860d24853adef67f
title: "Aufgabe 194: Winterspiele Tag 15: Freestyle-Ski"
challengeType: 29
dashedName: challenge-194
---

# --description--

Gegeben ist ein Trickname, der aus zwei Wörtern besteht. Bestimmen Sie, ob es sich um einen gültigen Freestyle-Ski-Tricknamen handelt.

Ein Trick ist gültig, wenn das erste Wort in der Liste der gültigen ersten Wörter steht und das zweite Wort in der Liste der gültigen zweiten Wörter.

- Die beiden Wörter sind durch ein einzelnes Leerzeichen getrennt.

Gültige erste Wörter:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Gültige zweite Wörter:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`is_valid_trick("Polar Vortex")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Polar Vortex"), True)`)
}})
```

`is_valid_trick("Solar Icequake")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Icequake"), True)`)
}})
```

`is_valid_trick("Thunder Blizzard")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Thunder Blizzard"), True)`)
}})
```

`is_valid_trick("Phantom Frostbite")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Phantom Frostbite"), True)`)
}})
```

`is_valid_trick("Ghost Avalanche")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Ghost Avalanche"), True)`)
}})
```

`is_valid_trick("Snowstorm Shadow")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Snowstorm Shadow"), False)`)
}})
```

`is_valid_trick("Solar Sky")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_trick("Solar Sky"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_trick(trick_name):

    return trick_name
```

# --solutions--

```py
def is_valid_trick(trick_name):
    valid_first = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"]
    valid_second = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"]

    words = trick_name.split(" ")
    first, second = words
    return first in valid_first and second in valid_second
```
