---
id: 697a49e6ff50d756c9b6935d
title: "Aufgabe 180: Winterspiele Tag 1: Eröffnungstag"
challengeType: 29
dashedName: challenge-180
---

# --description--

Die nächsten 17 Tage erwarten Sie Programmieraufgaben, die von den Winterspielen inspiriert sind.

Für die erste Aufgabe erhalten Sie einen zweibuchstabigen Ländercode und sollen das Flaggen-Emoji für dieses Land zurückgeben.

Verwenden Sie diese Liste:

| Land                       | Code       | Flagge   |
| -                         | -:         | :--:     |
| Albanien                  |            `"AL"` | `"🇦🇱"` |
| Andorra                   |            `"AD"` | `"🇦🇩"` |
| Argentinien               |            `"AR"` | `"🇦🇷"` |
| Armenien                  |            `"AM"` | `"🇦🇲"` |
| Australien                |            `"AU"` | `"🇦🇺"` |
| Österreich                |            `"AT"` | `"🇦🇹"` |
| Aserbaidschan             |            `"AZ"` | `"🇦🇿"` |
| Belgien                   |            `"BE"` | `"🇧🇪"` |
| Benin                     |            `"BJ"` | `"🇧🇯"` |
| Bolivien                  |            `"BO"` | `"🇧🇴"` |
| Bosnien und Herzegowina   |            `"BA"` | `"🇧🇦"` |
| Brasilien                 |            `"BR"` | `"🇧🇷"` |
| Bulgarien                 |            `"BG"` | `"🇧🇬"` |
| Kanada                    |            `"CA"` | `"🇨🇦"` |
| Chile                     |            `"CL"` | `"🇨🇱"` |
| China                     |            `"CN"` | `"🇨🇳"` |
| Kolumbien                 |            `"CO"` | `"🇨🇴"` |
| Kroatien                  |            `"HR"` | `"🇭🇷"` |
| Zypern                    |            `"CY"` | `"🇨🇾"` |
| Tschechien                |            `"CZ"` | `"🇨🇿"` |
| Dänemark                  |            `"DK"` | `"🇩🇰"` |
| Ecuador                   |            `"EC"` | `"🇪🇨"` |
| Eritrea                   |            `"ER"` | `"🇪🇷"` |
| Estland                   |            `"EE"` | `"🇪🇪"` |
| Finnland                  |            `"FI"` | `"🇫🇮"` |
| Frankreich                |            `"FR"` | `"🇫🇷"` |
| Georgien                  |            `"GE"` | `"🇬🇪"` |
| Deutschland               |            `"DE"` | `"🇩🇪"` |
| Großbritannien            |            `"GB"` | `"🇬🇧"` |
| Griechenland              |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau             |            `"GW"` | `"🇬🇼"` |
| Haiti                     |            `"HT"` | `"🇭🇹"` |
| Hongkong                  |            `"HK"` | `"🇭🇰"` |
| Ungarn                    |            `"HU"` | `"🇭🇺"` |
| Island                    |            `"IS"` | `"🇮🇸"` |
| Indien                    |            `"IN"` | `"🇮🇳"` |
| Iran                      |            `"IR"` | `"🇮🇷"` |
| Irland                    |            `"IE"` | `"🇮🇪"` |
| Israel                    |            `"IL"` | `"🇮🇱"` |
| Italien                   |            `"IT"` | `"🇮🇹"` |
| Jamaika                   |            `"JM"` | `"🇯🇲"` |
| Japan                     |            `"JP"` | `"🇯🇵"` |
| Kasachstan                |            `"KZ"` | `"🇰🇿"` |
| Kenia                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                    |            `"XK"` | `"🇽🇰"` |
| Kirgisistan               |            `"KG"` | `"🇰🇬"` |
| Lettland                  |            `"LV"` | `"🇱🇻"` |
| Libanon                   |            `"LB"` | `"🇱🇧"` |
| Liechtenstein             |            `"LI"` | `"🇱🇮"` |
| Litauen                   |            `"LT"` | `"🇱🇹"` |
| Luxemburg                 |            `"LU"` | `"🇱🇺"` |
| Madagaskar                |            `"MG"` | `"🇲🇬"` |
| Malaysia                  |            `"MY"` | `"🇲🇾"` |
| Malta                     |            `"MT"` | `"🇲🇹"` |
| Mexiko                    |            `"MX"` | `"🇲🇽"` |
| Moldawien                 |            `"MD"` | `"🇲🇩"` |
| Monaco                    |            `"MC"` | `"🇲🇨"` |
| Mongolei                  |            `"MN"` | `"🇲🇳"` |
| Montenegro                |            `"ME"` | `"🇲🇪"` |
| Marokko                   |            `"MA"` | `"🇲🇦"` |
| Niederlande               |            `"NL"` | `"🇳🇱"` |
| Neuseeland                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                   |            `"NG"` | `"🇳🇬"` |
| Nordmazedonien            |            `"MK"` | `"🇲🇰"` |
| Norwegen                  |            `"NO"` | `"🇳🇴"` |
| Pakistan                  |            `"PK"` | `"🇵🇰"` |
| Philippinen               |            `"PH"` | `"🇵🇭"` |
| Polen                     |            `"PL"` | `"🇵🇱"` |
| Portugal                  |            `"PT"` | `"🇵🇹"` |
| Puerto Rico               |            `"PR"` | `"🇵🇷"` |
| Rumänien                  |            `"RO"` | `"🇷🇴"` |
| San Marino                |            `"SM"` | `"🇸🇲"` |
| Saudi-Arabien             |            `"SA"` | `"🇸🇦"` |
| Serbien                   |            `"RS"` | `"🇷🇸"` |
| Singapur                  |            `"SG"` | `"🇸🇬"` |
| Slowakei                  |            `"SK"` | `"🇸🇰"` |
| Slowenien                 |            `"SI"` | `"🇸🇮"` |
| Südafrika                 |            `"ZA"` | `"🇿🇦"` |
| Südkorea                  |            `"KR"` | `"🇰🇷"` |
| Spanien                   |            `"ES"` | `"🇪🇸"` |
| Schweden                  |            `"SE"` | `"🇸🇪"` |
| Schweiz                   |            `"CH"` | `"🇨🇭"` |
| Thailand                  |            `"TH"` | `"🇹🇭"` |
| Trinidad und Tobago       |            `"TT"` | `"🇹🇹"` |
| Türkei                    |            `"TR"` | `"🇹🇷"` |
| Ukraine                   |            `"UA"` | `"🇺🇦"` |
| Vereinigte Arabische Emirate |        `"AE"` | `"🇦🇪"` |
| Vereinigte Staaten        |            `"US"` | `"🇺🇸"` |
| Uruguay                   |            `"UY"` | `"🇺🇾"` |
| Usbekistan                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                 |            `"VE"` | `"🇻🇪"` |

# --hints--

`get_flag("AL")` sollte `"🇦🇱"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AL"), "🇦🇱")`)
}})
```

`get_flag("AD")` sollte `"🇦🇩"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AD"), "🇦🇩")`)
}})
```

`get_flag("AR")` sollte `"🇦🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AR"), "🇦🇷")`)
}})
```

`get_flag("AM")` sollte `"🇦🇲"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AM"), "🇦🇲")`)
}})
```

`get_flag("AU")` sollte `"🇦🇺"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AU"), "🇦🇺")`)
}})
```

`get_flag("AT")` sollte `"🇦🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AT"), "🇦🇹")`)
}})
```

`get_flag("AZ")` sollte `"🇦🇿"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AZ"), "🇦🇿")`)
}})
```

`get_flag("BE")` sollte `"🇧🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BE"), "🇧🇪")`)
}})
```

`get_flag("BJ")` sollte `"🇧🇯"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BJ"), "🇧🇯")`)
}})
```

`get_flag("BO")` sollte `"🇧🇴"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BO"), "🇧🇴")`)
}})
```

`get_flag("BA")` sollte `"🇧🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BA"), "🇧🇦")`)
}})
```

`get_flag("BR")` sollte `"🇧🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BR"), "🇧🇷")`)
}})
```

`get_flag("BG")` sollte `"🇧🇬"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("BG"), "🇧🇬")`)
}})
```

`get_flag("CA")` sollte `"🇨🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CA"), "🇨🇦")`)
}})
```

`get_flag("CL")` sollte `"🇨🇱"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CL"), "🇨🇱")`)
}})
```

`get_flag("CN")` sollte `"🇨🇳"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CN"), "🇨🇳")`)
}})
```

`get_flag("CO")` sollte `"🇨🇴"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CO"), "🇨🇴")`)
}})
```

`get_flag("HR")` sollte `"🇭🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HR"), "🇭🇷")`)
}})
```

`get_flag("CY")` sollte `"🇨🇾"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CY"), "🇨🇾")`)
}})
```

`get_flag("CZ")` sollte `"🇨🇿"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CZ"), "🇨🇿")`)
}})
```

`get_flag("DK")` sollte `"🇩🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DK"), "🇩🇰")`)
}})
```

`get_flag("EC")` sollte `"🇪🇨"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EC"), "🇪🇨")`)
}})
```

`get_flag("ER")` sollte `"🇪🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ER"), "🇪🇷")`)
}})
```

`get_flag("EE")` sollte `"🇪🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("EE"), "🇪🇪")`)
}})
```

`get_flag("FI")` sollte `"🇫🇮"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FI"), "🇫🇮")`)
}})
```

`get_flag("FR")` sollte `"🇫🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("FR"), "🇫🇷")`)
}})
```

`get_flag("GE")` sollte `"🇬🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GE"), "🇬🇪")`)
}})
```

`get_flag("DE")` sollte `"🇩🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("DE"), "🇩🇪")`)
}})
```

`get_flag("GB")` sollte `"🇬🇧"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GB"), "🇬🇧")`)
}})
```

`get_flag("GR")` sollte `"🇬🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GR"), "🇬🇷")`)
}})
```

`get_flag("GW")` sollte `"🇬🇼"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("GW"), "🇬🇼")`)
}})
```

`get_flag("HT")` sollte `"🇭🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HT"), "🇭🇹")`)
}})
```

`get_flag("HK")` sollte `"🇭🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HK"), "🇭🇰")`)
}})
```

`get_flag("HU")` sollte `"🇭🇺"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("HU"), "🇭🇺")`)
}})
```

`get_flag("IS")` sollte `"🇮🇸"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IS"), "🇮🇸")`)
}})
```

`get_flag("IN")` sollte `"🇮🇳"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IN"), "🇮🇳")`)
}})
```

`get_flag("IR")` sollte `"🇮🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IR"), "🇮🇷")`)
}})
```

`get_flag("IE")` sollte `"🇮🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IE"), "🇮🇪")`)
}})
```

`get_flag("IL")` sollte `"🇮🇱"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IL"), "🇮🇱")`)
}})
```

`get_flag("IT")` sollte `"🇮🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("IT"), "🇮🇹")`)
}})
```

`get_flag("JM")` sollte `"🇯🇲"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JM"), "🇯🇲")`)
}})
```

`get_flag("JP")` sollte `"🇯🇵"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("JP"), "🇯🇵")`)
}})
```

`get_flag("KZ")` sollte `"🇰🇿"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KZ"), "🇰🇿")`)
}})
```

`get_flag("KE")` sollte `"🇰🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KE"), "🇰🇪")`)
}})
```

`get_flag("XK")` sollte `"🇽🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("XK"), "🇽🇰")`)
}})
```

`get_flag("KG")` sollte `"🇰🇬"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KG"), "🇰🇬")`)
}})
```

`get_flag("LV")` sollte `"🇱🇻"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LV"), "🇱🇻")`)
}})
```

`get_flag("LB")` sollte `"🇱🇧"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LB"), "🇱🇧")`)
}})
```

`get_flag("LI")` sollte `"🇱🇮"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LI"), "🇱🇮")`)
}})
```

`get_flag("LT")` sollte `"🇱🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LT"), "🇱🇹")`)
}})
```

`get_flag("LU")` sollte `"🇱🇺"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("LU"), "🇱🇺")`)
}})
```

`get_flag("MG")` sollte `"🇲🇬"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MG"), "🇲🇬")`)
}})
```

`get_flag("MY")` sollte `"🇲🇾"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MY"), "🇲🇾")`)
}})
```

`get_flag("MT")` sollte `"🇲🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MT"), "🇲🇹")`)
}})
```

`get_flag("MX")` sollte `"🇲🇽"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MX"), "🇲🇽")`)
}})
```

`get_flag("MD")` sollte `"🇲🇩"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MD"), "🇲🇩")`)
}})
```

`get_flag("MC")` sollte `"🇲🇨"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MC"), "🇲🇨")`)
}})
```

`get_flag("MN")` sollte `"🇲🇳"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MN"), "🇲🇳")`)
}})
```

`get_flag("ME")` sollte `"🇲🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ME"), "🇲🇪")`)
}})
```

`get_flag("MA")` sollte `"🇲🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MA"), "🇲🇦")`)
}})
```

`get_flag("NL")` sollte `"🇳🇱"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NL"), "🇳🇱")`)
}})
```

`get_flag("NZ")` sollte `"🇳🇿"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NZ"), "🇳🇿")`)
}})
```

`get_flag("NG")` sollte `"🇳🇬"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NG"), "🇳🇬")`)
}})
```

`get_flag("MK")` sollte `"🇲🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("MK"), "🇲🇰")`)
}})
```

`get_flag("NO")` sollte `"🇳🇴"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("NO"), "🇳🇴")`)
}})
```

`get_flag("PK")` sollte `"🇵🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PK"), "🇵🇰")`)
}})
```

`get_flag("PH")` sollte `"🇵🇭"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PH"), "🇵🇭")`)
}})
```

`get_flag("PL")` sollte `"🇵🇱"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PL"), "🇵🇱")`)
}})
```

`get_flag("PT")` sollte `"🇵🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PT"), "🇵🇹")`)
}})
```

`get_flag("PR")` sollte `"🇵🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("PR"), "🇵🇷")`)
}})
```

`get_flag("RO")` sollte `"🇷🇴"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RO"), "🇷🇴")`)
}})
```

`get_flag("SM")` sollte `"🇸🇲"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SM"), "🇸🇲")`)
}})
```

`get_flag("SA")` sollte `"🇸🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SA"), "🇸🇦")`)
}})
```

`get_flag("RS")` sollte `"🇷🇸"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("RS"), "🇷🇸")`)
}})
```

`get_flag("SG")` sollte `"🇸🇬"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SG"), "🇸🇬")`)
}})
```

`get_flag("SK")` sollte `"🇸🇰"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SK"), "🇸🇰")`)
}})
```

`get_flag("SI")` sollte `"🇸🇮"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SI"), "🇸🇮")`)
}})
```

`get_flag("ZA")` sollte `"🇿🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ZA"), "🇿🇦")`)
}})
```

`get_flag("KR")` sollte `"🇰🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("KR"), "🇰🇷")`)
}})
```

`get_flag("ES")` sollte `"🇪🇸"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("ES"), "🇪🇸")`)
}})
```

`get_flag("SE")` sollte `"🇸🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("SE"), "🇸🇪")`)
}})
```

`get_flag("CH")` sollte `"🇨🇭"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("CH"), "🇨🇭")`)
}})
```

`get_flag("TH")` sollte `"🇹🇭"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TH"), "🇹🇭")`)
}})
```

`get_flag("TT")` sollte `"🇹🇹"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TT"), "🇹🇹")`)
}})
```

`get_flag("TR")` sollte `"🇹🇷"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("TR"), "🇹🇷")`)
}})
```

`get_flag("UA")` sollte `"🇺🇦"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UA"), "🇺🇦")`)
}})
```

`get_flag("AE")` sollte `"🇦🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("AE"), "🇦🇪")`)
}})
```

`get_flag("US")` sollte `"🇺🇸"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("US"), "🇺🇸")`)
}})
```

`get_flag("UY")` sollte `"🇺🇾"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UY"), "🇺🇾")`)
}})
```

`get_flag("UZ")` sollte `"🇺🇿"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("UZ"), "🇺🇿")`)
}})
```

`get_flag("VE")` sollte `"🇻🇪"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_flag("VE"), "🇻🇪")`)
}})
```



# --seed--

## --seed-contents--

```py
def get_flag(code):

    return code
```

# --solutions--

```py
def get_flag(code):
    return "".join(chr(127397 + ord(char)) for char in code)
```
