---
id: 697a49e6ff50d756c9b69365
title: "Challenge 188: 2026 Winter Games Day 9: Skeleton"
challengeType: 29
dashedName: challenge-188
---

# --description--

Gegeben ist ein String, der die Kurven auf einer Skeleton-Bahn darstellt. Bestimmen Sie die Schwierigkeit der Bahn.

- Der gegebene String besteht nur aus den Buchstaben:
  - `"L"` für eine Linkskurve
  - `"R"` für eine Rechtskurve
  - `"S"` für ein gerades Teilstück

- Jede Richtungsänderung fügt 15 Punkte hinzu (ein `"L"` gefolgt von einem `"R"` oder umgekehrt).
- Alle anderen Kurven fügen jeweils 5 Punkte hinzu (alle anderen `"L"` oder `"R"` Zeichen).
- Gerade Teilstücke fügen 0 Punkte hinzu.

Die Schwierigkeit der Bahn basiert auf der Gesamtpunktzahl. Geben Sie zurück:

- `"Easy"`, wenn die Gesamtpunktzahl 0 bis 100 beträgt
- `"Medium"`, wenn die Gesamtpunktzahl 101 bis 200 beträgt
- `"Hard"`, wenn die Gesamtpunktzahl über 200 liegt

# --hints--

`get_difficulty("SLSLLSRRLSRLRL")` should return `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLSLLSRRLSRLRL"), "Easy")`)
}})
```

`get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` should return `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard")`)
}})
```

`get_difficulty("SRRRRLSLLRLRSSRLSRL")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRRRRLSLLRLRSSRLSRL"), "Medium")`)
}})
```

`get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` should return `"Hard"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard")`)
}})
```

`get_difficulty("SLLSSLRLSLSLRSLSSLRL")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium")`)
}})
```

`get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` should return `"Easy"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_difficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_difficulty(track):

    return track
```

# --solutions--

```py
def get_difficulty(track):
  score = 0

  for i in range(len(track)):
    current = track[i]
    previous = track[i - 1] if i > 0 else None

    if current == 'S':
      score += 0
    elif current == 'L' or current == 'R':
      if previous and previous != 'S' and previous != current:
        score += 15
      else:
        score += 5

  if score <= 100:
    return 'Easy'
  elif score <= 200:
    return 'Medium'
  else:
    return 'Hard'
```
