---
id: 698a1a73ade5ac0e19180fa0
title: "Challenge 197: Blood Type Compatibility"
challengeType: 29
dashedName: challenge-197
---

# --description--

Gegeben sind eine Blutgruppe des Spenders und eine Blutgruppe des Empfängers. Bestimmen Sie, ob der Spender dem Empfänger Blut geben kann.

Jede Blutgruppe besteht aus:

- Einem Buchstaben: `"A"`, `"B"`, `"AB"` oder `"O"`
- Und einem Rhesusfaktor: `"+"` oder `"-"`

Blutgruppen bestehen aus einem der gültigen Buchstaben, gefolgt von einem Rhesusfaktor. Zum Beispiel sind `"AB+"` und `"O-"` gültige Blutgruppen.

Letter Rules:

- `"O"` kann an andere Buchstabentypen spenden.
- `"A"` kann an `"A"` und `"AB"` spenden.
- `"B"` kann an `"B"` und `"AB"` spenden.
- `"AB"` kann nur an `"AB"` spenden.

Rh Rules:

- Negativ (`"-"`) kann an sowohl `"-"` als auch `"+"` spenden.
- Positiv (`"+"`) kann nur an `"+"` spenden.

Sowohl die Buchstabenregel als auch die Rhesusregel müssen erfüllt sein, damit ein Spender dem Empfänger Blut geben kann.

# --hints--

`can_donate("B+", "B+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```
