---
id: 698a1a863194f1f4e63f645f
title: "Challenge 208: Trail Traversal"
challengeType: 29
dashedName: challenge-208
---

# --description--

Gegeben ist ein Array von Strings, die Ihre Pfadkarte darstellen. Geben Sie einen String mit den Bewegungen zurück, die nötig sind, um Ihr Ziel zu erreichen.

Die gegebenen Strings enthalten die Werte:

- `"C"`: Ihren aktuellen Standort
- `"G"`: Ihr Ziel
- `"T"`: Begehbare Teile des Pfads
- `"-"`: Unbegehbare Teile der Karte

Geben Sie einen String mit den Bewegungen zurück, die nötig sind, um dem Pfad von Ihrem Standort bis zu Ihrem Ziel zu folgen, wobei:

- `"R"` eine Bewegung nach rechts ist
- `"D"` eine Bewegung nach unten ist
- `"L"` eine Bewegung nach links ist
- `"U"` eine Bewegung nach oben ist

- Es gibt immer genau einen durchgehenden Pfad ohne Verzweigungen von Ihrem aktuellen Standort bis zu Ihrem Ziel.
- Jeder Pfadabschnitt hat maximal zwei begehbare Nachbarfelder.

Zum Beispiel, gegeben:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Return `"RDDRDD"`.

# --hints--

`navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` should return `"RDDRDD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD")`)
}})
```

`navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` should return `"RRUUURR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR")`)
}})
```

`navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` should return `"DLDDRRRRD"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD")`)
}})
```

`navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` should return `"RRRDDL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL")`)
}})
```

`navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` should return `"ULLLUUURRRRRRDDDR"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate_trail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate_trail(map):

    return map
```

# --solutions--

```py
def navigate_trail(map):
    rows = len(map)
    cols = len(map[0])

    directions = [
        (0, 1, "R"),
        (1, 0, "D"),
        (0, -1, "L"),
        (-1, 0, "U"),
    ]

    for r in range(rows):
        for c in range(cols):
            if map[r][c] == "C":
                row, col = r, c

    result = ""
    prev = None

    while map[row][col] != "G":
        for dr, dc, move in directions:
            new_row = row + dr
            new_col = col + dc

            if (
                0 <= new_row < rows and
                0 <= new_col < cols and
                (map[new_row][new_col] == "T" or
                 map[new_row][new_col] == "G") and
                (not prev or (new_row, new_col) != prev)
            ):
                result += move
                prev = (row, col)
                row, col = new_row, new_col
                break

    return result
```
