---
id: 698a1a863194f1f4e63f645e
title: "Challenge 207: Smallest Gap"
challengeType: 29
dashedName: challenge-207
---

# --description--

Gegeben ist ein String. Geben Sie den Teilstring zwischen den zwei identischen Zeichen zurück, die die kleinste Anzahl von Zeichen zwischen sich haben (kleinste Lücke).

- Es gibt immer mindestens ein Paar übereinstimmender Zeichen.
- Der zurückgegebene Teilstring soll die übereinstimmenden Zeichen ausschließen.
- Wenn zwei oder mehr Lücken gleich lang sind, geben Sie die Zeichen der ersten zurück.

Zum Beispiel, bei `"ABCDAC"`, geben Sie `"DA"` zurück, weil:

- Nur `"A"` und `"C"` wiederholen sich im String.
- Die Anzahl der Zeichen zwischen den beiden `"A"`-Zeichen ist 3, und zwischen den `"C"`-Zeichen ist 2.
- Also geben Sie den String zwischen den beiden `"C"`-Zeichen zurück.

# --hints--

`smallest_gap("ABCDAC")` should return `"DA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("ABCDAC"), "DA")`)
}})
```

`smallest_gap("racecar")` should return `"e"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("racecar"), "e")`)
}})
```

`smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` should return `"#q6e&rkf(p"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p")`)
}})
```

`smallest_gap("Hello World")` should return `""`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("Hello World"), "")`)
}})
```

`smallest_gap("The quick brown fox jumps over the lazy dog.")` should return `"fox"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(smallest_gap("The quick brown fox jumps over the lazy dog."), "fox")`)
}})
```

# --seed--

## --seed-contents--

```py
def smallest_gap(s):

    return s
```

# --solutions--

```py
def smallest_gap(s):
    min_gap = float("inf")
    result = ""

    for i in range(len(s)):
        for j in range(i + 1, len(s)):
            if s[i] == s[j]:
                gap = j - i - 1

                if gap < min_gap:
                    min_gap = gap
                    result = s[i + 1:j]

                break

    return result
```
