---
id: 698a1a73ade5ac0e19180fa4
title: "Challenge 201: Matrix Shift"
challengeType: 29
dashedName: challenge-201
---

# --description--

Gegeben ist eine Matrix (Array von Arrays) mit Zahlen und eine ganze Zahl. Verschieben Sie alle Werte in der Matrix um den angegebenen Betrag.

- Eine positive Verschiebung bewegt die Werte nach rechts.
- Eine negative Verschiebung bewegt die Werte nach links.

Die Werte sollen umschlagen:

- Behandeln Sie die Matrix als eine durchgehende Folge von Werten.
- Wenn ein Wert über das Ende einer Zeile hinaus verschoben wird, setzt er am Anfang der nächsten Zeile fort.
- Wenn ein Wert über die letzte Position der Matrix hinaus verschoben wird, springt er zur ersten Position zurück.
- Das Gleiche gilt umgekehrt bei einer Verschiebung nach links.

Zum Beispiel, gegeben:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

mit einer Verschiebung von `1`, verschieben Sie alle Zahlen um eins nach rechts:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shift_matrix([[1, 2, 3], [4, 5, 6]], 1)` should return `[[6, 1, 2], [3, 4, 5]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6]], -1)` should return `[[2, 3, 4], [5, 6, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` should return `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]])`)
}})
```

`shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` should return `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` should return `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]])`)
}})
```

`shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` should return `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(shift_matrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]])`)
}})
```

# --seed--

## --seed-contents--

```py
def shift_matrix(matrix, shift):

    return matrix
```

# --solutions--

```py
def shift_matrix(matrix, shift):
    rows = len(matrix)
    cols = len(matrix[0])
    flat = [num for row in matrix for num in row]
    n = len(flat)
    s = shift % n
    shifted = flat[-s:] + flat[:-s] if s != 0 else flat[:]
    result = []
    for i in range(rows):
        result.append(shifted[i*cols:(i+1)*cols])

    return result
```
