---
id: 6994cff2290543b3aec9f50f
title: "Challenge 210: HSL Validator"
challengeType: 29
dashedName: challenge-210
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um einen gültigen CSS `hsl()`-Farbwert handelt.

- Ein gültiger HSL-Wert muss im Format `"hsl(h, s%, l%)"` vorliegen, wobei:
  - `h` (Farbton) eine Zahl zwischen 0 und 360 (einschließlich) sein muss.
  - `s` (Sättigung) ein Prozentsatz zwischen 0 % und 100 % sein muss.
  - `l` (Helligkeit) ein Prozentsatz zwischen 0 % und 100 % sein muss.

- Leerzeichen sind nur erlaubt:
  - Nach der öffnenden Klammer
  - Vor und/oder nach den Kommas
  - Vor und/oder nach der schließenden Klammer

- Optional kann der Wert mit einem Semikolon (`";"`) enden.

Zum Beispiel ist `"hsl(240, 50%, 50%)"` ein gültiger HSL-Wert.

# --hints--

`is_valid_hsl("hsl(240, 50%, 50%)")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(240, 50%, 50%)"), True)`)
}})
```

`is_valid_hsl("hsl( 200 , 50% , 75% )")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl( 200 , 50% , 75% )"), True)`)
}})
```

`is_valid_hsl("hsl(99,60%,80%);")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(99,60%,80%);"), True)`)
}})
```

`is_valid_hsl("hsl(0, 0%, 0%) ;")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(0, 0%, 0%) ;"), True)`)
}})
```

`is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(  10  ,  20%   ,  30%   )    ;"), True)`)
}})
```

`is_valid_hsl("hsl(361, 50%, 80%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(361, 50%, 80%)"), False)`)
}})
```

`is_valid_hsl("hsl(300, 101%, 70%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(300, 101%, 70%)"), False)`)
}})
```

`is_valid_hsl("hsl(200, 55%, 75)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl(200, 55%, 75)"), False)`)
}})
```

`is_valid_hsl("hsl (80, 20%, 10%)")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_hsl("hsl (80, 20%, 10%)"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_hsl(hsl):

    return hsl
```

# --solutions--

```py
import re
def is_valid_hsl(hsl):
    pattern = r"^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$"
    
    match = re.match(pattern, hsl)
    if not match:
        return False

    hue = int(match.group(1))
    sat = int(match.group(2))
    light = int(match.group(3))

    if not (0 <= hue <= 360):
        return False
    if not (0 <= sat <= 100):
        return False
    if not (0 <= light <= 100):
        return False

    return True
```
