---
id: 699c8e045ee7cb94ed2322d5
title: "Challenge 214: Domino Chain Validator"
challengeType: 29
dashedName: challenge-214
---

# --description--

Gegeben ist ein 2D-Array, das eine Folge von Dominosteinen darstellt. Bestimmen Sie, ob es eine gültige Kette bildet.

- Jedes Element im Array stellt einen Dominostein dar und ist ein Array aus zwei Zahlen von 1 bis 6 (einschließlich).
- Damit die Kette gültig ist, muss die zweite Zahl jedes Dominosteins mit der ersten Zahl des nächsten Dominosteins übereinstimmen.
- Die erste Zahl des ersten Dominosteins und die letzte Zahl des letzten Dominosteins müssen mit nichts übereinstimmen.

# --hints--

`is_valid_domino_chain([[1, 3], [3, 6], [6, 5]])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[1, 3], [3, 6], [6, 5]]), True)`)
}})
```

`is_valid_domino_chain([[6, 2], [3, 4], [4, 1]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[6, 2], [3, 4], [4, 1]]), False)`)
}})
```

`is_valid_domino_chain([[2, 5], [5, 6], [5, 1]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 5], [5, 6], [5, 1]]), False)`)
}})
```

`is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]), True)`)
}})
```

`is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_domino_chain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_domino_chain(dominoes):

    return dominoes
```

# --solutions--

```py
def is_valid_domino_chain(dominoes):
    for i in range(len(dominoes) - 1):
        if dominoes[i][1] != dominoes[i + 1][0]:
            return False

    return True
```
