---
id: 699c8e045ee7cb94ed2322d6
title: "Challenge 215: Parking Fee Calculator"
challengeType: 29
dashedName: challenge-215
---

# --description--

Gegeben sind zwei Strings, die die Zeit darstellen, zu der Sie Ihr Auto geparkt haben, und die Zeit, zu der Sie es abgeholt haben. Berechnen Sie die Parkgebühr.

- Die gegebenen Strings haben das Format `"HH:MM"` und verwenden eine 24-Stunden-Uhr. `"14:00"` entspricht zum Beispiel 14 Uhr.
- Der erste String ist die Zeit, zu der Sie Ihr Auto geparkt haben, der zweite die Zeit, zu der Sie es abgeholt haben.
- Wenn die Abholzeit früher als die Parkzeit ist, bedeutet das, dass das Parken über Mitternacht in den nächsten Tag hinausging.

Fee rules:

- Jede geparkte Stunde kostet 3 $.
- Teilstunden werden auf die nächste volle Stunde aufgerundet.
- Wenn das Parken über Nacht (über Mitternacht) geht, wird eine zusätzliche Übernachtgebühr von 10 $ berechnet.
- Es gibt eine Mindestgebühr von 5 $ (wird nur angewendet, wenn die Gesamtsumme darunter liegen würde).

Geben Sie die Gesamtkosten im Format `"$cost"`, `"$5"` zurück.

# --hints--

`calculate_parking_fee("09:00", "11:00")` should return `"$6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("09:00", "11:00"), "$6")`)
}})
```

`calculate_parking_fee("10:00", "10:30")` should return `"$5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("10:00", "10:30"), "$5")`)
}})
```

`calculate_parking_fee("08:10", "10:45")` should return `"$9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("08:10", "10:45"), "$9")`)
}})
```

`calculate_parking_fee("14:40", "23:10")` should return `"$27"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("14:40", "23:10"), "$27")`)
}})
```

`calculate_parking_fee("18:15", "01:30")` should return `"$34"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("18:15", "01:30"), "$34")`)
}})
```

`calculate_parking_fee("11:11", "11:10")` should return `"$82"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_parking_fee("11:11", "11:10"), "$82")`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_parking_fee(park_time, pickup_time):

    return park_time
```

# --solutions--

```py
import math

def calculate_parking_fee(park_time, pickup_time):
    def to_minutes(time_str):
        hours, minutes = map(int, time_str.split(":"))
        return hours * 60 + minutes

    entry_minutes = to_minutes(park_time)
    exit_minutes = to_minutes(pickup_time)

    overnight = False

    if exit_minutes < entry_minutes:
        overnight = True
        total_minutes = (24 * 60 - entry_minutes) + exit_minutes
    else:
        total_minutes = exit_minutes - entry_minutes

    hours = math.ceil(total_minutes / 60)

    cost = hours * 3

    if overnight:
        cost += 10

    if cost < 5:
        cost = 5

    return f"${cost}"
```
