---
id: 699c8e045ee7cb94ed2322dd
title: "Challenge 222: Equinox Shadows"
challengeType: 29
dashedName: challenge-222
---

# --description--

Heute ist die Tagundnachtgleiche, wenn die Sonne direkt über dem Äquator steht und mittags genau im Zenit ist. Gegeben eine Uhrzeit, bestimmen Sie den Schatten, den ein 4 Fuß hoher senkrechter Pfosten wirft.

- Die Uhrzeit wird als String im `"HH:MM"` 24-Stunden-Format angegeben (zum Beispiel ist `"15:00"` 15 Uhr).
- Sie erhalten nur Uhrzeiten in 30-Minuten-Schritten.

Rules:

- Die Sonne geht um 6 Uhr morgens direkt `"east"` auf und um 18 Uhr direkt `"west"` unter.
- Ein Schatten zeigt immer entgegengesetzt zur Sonne.
- Die Länge des Schattens (in Fuß) ist die Anzahl der Stunden vom Mittag entfernt, hoch drei.
- Vor Sonnenaufgang (vor 6 Uhr), nach Sonnenuntergang (18 Uhr oder später) und mittags gibt es keinen Schatten.

Return:

- Wenn ein Schatten existiert, geben Sie `"(length)ft (direction)"` zurück. Zum Beispiel `"8ft west"`.
- Andernfalls geben Sie `"No shadow"` zurück.

Beispielsweise wird bei `"10:00"` `"8ft west"` zurückgegeben, weil 10 Uhr 2 Stunden vom Mittag entfernt ist, also 2<sup>3</sup> = 8 Fuß, und der Schatten nach Westen zeigt, weil die Sonne um 10 Uhr im Osten steht.

# --hints--

`get_shadow("10:00")` should return `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` should return `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` should return `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` should return `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` should return `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```
