---
id: 69a890af247de743333bd4d1
title: "Challenge 228: Movie Night"
challengeType: 29
dashedName: challenge-228
---

# --description--

Gegeben sind ein String für den Wochentag, ein weiterer String für die Vorstellungszeit und eine ganze Zahl für die Anzahl der Tickets. Geben Sie die Gesamtkosten der Kinokarten für diese Vorstellung zurück.

Der angegebene Tag wird einer der folgenden sein:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Die Vorstellungszeit wird im Format `"H:MMam"` oder `"H:MMpm"` angegeben. Zum Beispiel `"10:00am"` oder `"10:00pm"`.

Geben Sie die Gesamtkosten im Format `"$D.CC"` zurück, und zwar nach diesen Regeln:

- Wochenende (Freitag bis Sonntag): 12,00 $ pro Ticket.
- Wochentag (Montag bis Donnerstag): 10,00 $ pro Ticket.
- Matinee (vor 17:00 Uhr): 2,00 $ Rabatt pro Ticket (außer dienstags).
- Dienstags: alle Tickets kosten jeweils 5,00 $.

# --hints--

`get_movie_night_cost("Saturday", "10:00pm", 1)` should return `"$12.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Saturday", "10:00pm", 1), "$12.00")`)
}})
```

`get_movie_night_cost("Sunday", "10:00am", 1)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Sunday", "10:00am", 1), "$10.00")`)
}})
```

`get_movie_night_cost("Tuesday", "7:20pm", 2)` should return `"$10.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "7:20pm", 2), "$10.00")`)
}})
```

`get_movie_night_cost("Wednesday", "5:40pm", 3)` should return `"$30.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Wednesday", "5:40pm", 3), "$30.00")`)
}})
```

`get_movie_night_cost("Monday", "11:50am", 4)` should return `"$32.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Monday", "11:50am", 4), "$32.00")`)
}})
```

`get_movie_night_cost("Friday", "4:30pm", 5)` should return `"$50.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Friday", "4:30pm", 5), "$50.00")`)
}})
```

`get_movie_night_cost("Tuesday", "11:30am", 1)` should return `"$5.00"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_movie_night_cost("Tuesday", "11:30am", 1), "$5.00")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):

    return day
```

# --solutions--

```py
def get_movie_night_cost(day, showtime, number_of_tickets):
    if day == "Tuesday":
        price = 5
    else:
        if day in ["Friday", "Saturday", "Sunday"]:
            price = 12
        else:
            price = 10

        hour = int(showtime.split(":")[0])
        period = showtime[-2:]

        hour24 = hour

        if period == "pm" and hour != 12:
            hour24 += 12
        if period == "am" and hour == 12:
            hour24 = 0

        if hour24 < 17:
            price -= 2

    total = price * number_of_tickets

    return f"${total:.2f}"
```
