---
id: 69a890af247de743333bd4cd
title: "Challenge 224: Coffee Roast Detector"
challengeType: 29
dashedName: challenge-224
---

# --description--

Gegeben ist ein String, der die Bohnen darstellt, die zur Herstellung einer Tasse Kaffee verwendet werden. Bestimmen Sie die Röstung der Tasse.

- Der gegebene String enthält die folgenden Zeichen, die jeweils eine Bohnenart darstellen:

- Ein Apostroph (`'`) ist eine hell geröstete Bohne, die jeweils 1 Punkt wert ist.
  - Ein Bindestrich (`-`) ist eine mittel geröstete Bohne, die jeweils 2 Punkte wert ist.
  - Ein Punkt (`.`) ist eine dunkel geröstete Bohne, die jeweils 3 Punkte wert ist.

- Der Röstgrad wird durch den Durchschnitt aller Bohnen bestimmt.

Return:

- `"Light"`, wenn der Durchschnitt kleiner als 1,75 ist.
- `"Medium"`, wenn der Durchschnitt zwischen 1,75 und 2,5 liegt.
- `"Dark"`, wenn der Durchschnitt größer als 2,5 ist.

# --hints--

`detect_roast("''-''''''-'-''--''''")` should return `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` should return `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` should return `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` should return `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` should return `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```
