---
id: 69b5b2be76ec8135a7fbe973
title: "Challenge 244: Rook and Bishop Attack"
challengeType: 29
dashedName: challenge-244
---

# --description--

Gegeben ist ein String für die Position eines Turms auf einem Schachbrett und ein weiterer für die Position eines Läufers. Bestimmen Sie, ob eine Figur die andere angreifen kann.

Ein Standard-Schachbrett ist 8x8, mit Spalten, die von `A` bis `H` beschriftet sind (von links nach rechts), und Reihen, die von `1` bis `8` beschriftet sind (von unten nach oben). Es sieht so aus:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Türme können sich beliebig viele Felder horizontal oder vertikal bewegen.
- Läufer können sich beliebig viele Felder diagonal bewegen.
- Eine Figur kann eine andere angreifen, wenn sie sich auf die Position dieser Figur bewegen kann.

Return:

- `"rook"`, wenn der Turm den Läufer angreifen kann.
- `"bishop"`, wenn der Läufer den Turm angreifen kann.
- `"neither"`, wenn keine der beiden Figuren die andere angreifen kann.

# --hints--

`rook_bishop_attack("A1", "A5")` should return `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("A1", "A5"), "rook")`)
}})
```

`rook_bishop_attack("C3", "F6")` should return `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("C3", "F6"), "bishop")`)
}})
```

`rook_bishop_attack("D4", "D7")` should return `"rook"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("D4", "D7"), "rook")`)
}})
```

`rook_bishop_attack("B7", "H1")` should return `"bishop"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B7", "H1"), "bishop")`)
}})
```

`rook_bishop_attack("B3", "C5")` should return `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("B3", "C5"), "neither")`)
}})
```

`rook_bishop_attack("G3", "E8")` should return `"neither"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(rook_bishop_attack("G3", "E8"), "neither")`)
}})
```

# --seed--

## --seed-contents--

```py
def rook_bishop_attack(rook, bishop):

    return rook
```

# --solutions--

```py
def rook_bishop_attack(rook, bishop):
    col_diff = abs(ord(rook[0]) - ord(bishop[0]))
    row_diff = abs(int(rook[1]) - int(bishop[1]))

    if rook[0] == bishop[0] or rook[1] == bishop[1]:
        return "rook"
    if col_diff == row_diff:
        return "bishop"
    return "neither"
```
