---
id: 69b559d2903b9e4afe9075f8
title: "Challenge 238: Digit Rotation Escape"
challengeType: 29
dashedName: challenge-238
---

# --description--

Gegeben ist eine positive ganze Zahl. Bestimmen Sie, ob sie selbst oder eine ihrer Rotationen durch die Anzahl ihrer Ziffern ohne Rest teilbar ist.

Eine Rotation bedeutet, die erste Ziffer an das Ende zu verschieben. Zum Beispiel wird aus 123 nach einer Rotation 231.

- Überprüfen Sie zuerst die Rotation `0` (die gegebene Zahl).
- Gegebene Zahlen enthalten keine Nullen.
- Geben Sie die erste Rotationszahl zurück, die diese Bedingung erfüllt, oder `"none"`, wenn keine gefunden wird.

# --hints--

`get_rotation(123)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(123), 0)`)
}})
```

`get_rotation(13579)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(13579), 3)`)
}})
```

`get_rotation(24681)` should return `"none"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(24681), "none")`)
}})
```

`get_rotation(84138789345)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rotation(84138789345), 6)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rotation(n):

    return n
```

# --solutions--

```py
def get_rotation(n):
    s = str(n)
    digit_count = len(s)
    current = s

    for i in range(digit_count):
        if int(current) % digit_count == 0:
            return i
        current = current[1:] + current[0]

    return "none"
```
