---
id: 69bc6cb30c1d112a2e110a07
title: "Aufgabe 250: Versteckter Schlüssel"
challengeType: 29
dashedName: challenge-250
---

# --description--

Willkommen zur 250. täglichen Aufgabe!

Gegeben ist ein verschlüsselter String. Entschlüsseln Sie ihn mit einem Verschlüsselungsschlüssel und geben Sie das Ergebnis zurück.

So finden Sie den Schlüssel:

- Schauen Sie sich alle täglichen Aufgaben an, von der Aufgabe Nummer 1 bis heute, deren Aufgabennummer ein Vielfaches von 25 ist (einschließlich dieser).
- Nehmen Sie den ersten Buchstaben jedes dieser Aufgabentitel und fügen Sie sie zu einem String zusammen. Wenn der Titel mit einem Nicht-Buchstaben beginnt, finden Sie den ersten Buchstaben.

Um die Nachricht zu entschlüsseln, gehen Sie jeden Buchstaben der verschlüsselten Nachricht durch und:

  1. Schauen Sie sich den entsprechenden Buchstaben im Schlüssel an (wiederholen Sie den Schlüssel, falls die Nachricht länger ist als der Schlüssel).
  2. Wandeln Sie den Schlüsselbuchstaben in seine entsprechende Zahl um: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Verschieben Sie den verschlüsselten Buchstaben im Alphabet um diese Zahl rückwärts.
  4. Wenn die Verschiebung vor `"A"` liegt, springen Sie zum Ende des Alphabets zurück zu `"Z"`.

Zum Beispiel: Wenn die verschlüsselte Nachricht mit `"Y"` beginnt und der erste Schlüsselbuchstabe `"V"` (22) ist, verschieben Sie `"Y"` um 22 Stellen zurück, um `"C"` zu erhalten. Wiederholen Sie diesen Vorgang für jeden Buchstaben, um die gesamte Nachricht zu entschlüsseln.

- Nur Buchstaben werden verschoben, Leerzeichen bleiben unverändert.
- Alle gegebenen und zurückgegebenen Buchstaben sind Großbuchstaben.

# --hints--

`decode("YAVJYNXE")` sollte `"CONGRATS"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` sollte `"CODING LEGEND"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` sollte `"YOU ARE AWESOME"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` sollte `"KEEP GOING"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` sollte `"A WINNER IS A DREAMER WHO NEVER GIVES UP"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```
