---
id: 69bc6cb30c1d112a2e110a06
title: "Challenge 249: String Math"
challengeType: 29
dashedName: challenge-249
---

# --description--

Gegeben ist ein String mit Zahlen und anderen Zeichen. Führen Sie mathematische Operationen mit den Zahlen durch, basierend auf der Anzahl der Nicht-Ziffern-Zeichen zwischen den Zahlen.

- Wenn die Anzahl der Zeichen zwischen zwei Zahlen gerade ist, verwenden Sie Addition.
- Wenn sie ungerade ist, verwenden Sie Subtraktion.
- Aufeinanderfolgende Ziffern bilden eine einzelne Zahl.
- Die Operationen werden von links nach rechts angewendet.
- Ignorieren Sie führende und abschließende Zeichen, die keine Ziffern sind.

Zum Beispiel, gegeben `"3ab10c8"`, geben Sie `5` zurück. Addieren Sie 3 und 10, um 13 zu erhalten, weil eine gerade Anzahl von Zeichen zwischen ihnen liegt. Ziehen Sie dann 8 von 13 ab, weil eine ungerade Anzahl von Zeichen zwischen dem Ergebnis und 8 liegt.

# --hints--

`do_math("3ab10c8")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("3ab10c8"), 5)`)
}})
```

`do_math("6MINUS4")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("6MINUS4"), 2)`)
}})
```

`do_math("9plus3")` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("9plus3"), 12)`)
}})
```

`do_math("5fkwo#10i#%.<>15P=@20!#B/25")` should return `15`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("5fkwo#10i#%.<>15P=@20!#B/25"), 15)`)
}})
```

`do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` should return `67`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(do_math("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67)`)
}})
```

# --seed--

## --seed-contents--

```py
def do_math(s):

    return s
```

# --solutions--

```py
import re

def do_math(s):
    tokens = re.findall(r'\d+|\D+', s)
    result = None
    pending_op = None
    for token in tokens:
        if re.match(r'^\d+$', token):
            if result is None:
                result = int(token)
            elif pending_op == 'add':
                result += int(token)
            else:
                result -= int(token)
        else:
            pending_op = 'add' if len(token) % 2 == 0 else 'subtract'
    return result
```
