---
id: 69c5f3d787b1725d5f00c8bd
title: "Challenge 259: FizzBuzz Explosion"
challengeType: 29
dashedName: challenge-259
---

# --description--

Gegeben ist eine ganze Zahl. Geben Sie die Anzahl der Schritte zurück, die benötigt werden, um das Wort `"fizzbuzz"` in einen String mit mindestens der angegebenen Anzahl von `"z"`-Zeichen zu verwandeln, wobei die folgenden Regeln gelten:

- Beginnen Sie mit dem String `"fizzbuzz"`.
- Wenden Sie in jedem Schritt die Standard-FizzBuzz-Regeln auf die Buchstabenposition im String an (der erste `"f"` hat die Position 1).
  - Wenn die Buchstabenposition durch 3 teilbar ist, ersetzen Sie den Buchstaben durch `"fizz"`
  - Wenn sie durch 5 teilbar ist, ersetzen Sie den Buchstaben durch `"buzz"`
  - Wenn sie durch 3 und 5 teilbar ist, ersetzen Sie den Buchstaben durch `"fizzbuzz"`

Nach 1 Schritt verwandelt sich `"fizzbuzz"` in `"fifizzzbuzzfizzzz"`, das 9 `"z"`-Zeichen enthält.

# --hints--

`explode_fizzbuzz(9)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```
