---
id: 69cfca90e8a0a6d4d6871c51
title: "Challenge 267: Parsec Converter"
challengeType: 29
dashedName: challenge-267
---

# --description--

In einer fernen Galaxie werden Parsec sowohl zur Messung von Zeit als auch von Entfernung verwendet. Gegeben ist eine ganze Zahl von Parsec, geben Sie deren Äquivalent in Zeit oder Entfernung zurück.

- Wenn die gegebene ganze Zahl ungerade ist, stellt sie Zeit dar. Wenn sie gerade ist, stellt sie Entfernung dar.

Verwenden Sie diese Umrechnungswerte:

| Parsec | Zeit/Entfernung |
| - | - |
| 1 | 2 Stunden |
| 2 | 6 Lichtjahre |

Geben Sie den umgerechneten Wert als ganze Zahl zurück.

# --hints--

`convert_parsecs(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` should return `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` should return `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` should return `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` should return `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```
