---
id: 69cfca90e8a0a6d4d6871c52
title: "Challenge 268: Narcissistic Number"
challengeType: 29
dashedName: challenge-268
---

# --description--

Gegeben ist eine positive ganze Zahl. Bestimmen Sie, ob es sich um eine narzisstische Zahl handelt.

- Eine Zahl ist narzisstisch, wenn die Summe jeder ihrer Ziffern, potenziert mit der Gesamtanzahl der Ziffern, der Zahl selbst entspricht.

Zum Beispiel hat 153 drei Ziffern, und 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, daher ist sie narzisstisch.

# --hints--

`is_narcissistic(153)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

`is_narcissistic(154)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

`is_narcissistic(371)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

`is_narcissistic(512)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

`is_narcissistic(9)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

`is_narcissistic(11)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

`is_narcissistic(9474)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

`is_narcissistic(6549)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```
