---
id: 69e2383af7832c8032603b8f
title: "Challenge 273: ISBN-13 Validator"
challengeType: 29
dashedName: challenge-273
---

# --description--

Gegeben ist ein String. Bestimmen Sie, ob es sich um eine gültige ISBN-13-Nummer handelt.

A valid ISBN-13:

- Enthält nur Ziffern und Bindestriche
- Hat genau 13 Ziffern, nachdem die Bindestriche entfernt wurden
- Besteht die folgende Prüfung:
  1. Multiplizieren Sie jede Ziffer abwechselnd mit 1 oder 3 (die erste Ziffer mit 1, die zweite mit 3, die dritte mit 1 usw.).
  2. Die Summe der Ergebnisse muss durch 10 teilbar sein.

# --hints--

`is_valid_isbn_13("9780306406157")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
