---
id: 69f8c998d78ad3171a0713b9
title: "Aufgabe 286: Offene Issues"
challengeType: 29
dashedName: challenge-286
---

# --description--

Gegeben sind ein Array von Issue-Nummern und ein weiteres Array von Pull-Request-(PR-)Nummern. Geben Sie ein Array der Issues zurück, die nach dem Zusammenführen aller PRs noch offen sind.

- Ein PR schließt ein Issue, wenn deren Ziffern eine Rotation voneinander sind. Zum Beispiel würde Issue 123 durch PR 231 oder 312 geschlossen.
- Ein PR schließt kein Issue mit exakt derselben Nummer. Zum Beispiel schließt PR 123 nicht Issue 123. Ein Issue mit nur derselben Zahl kann also nicht geschlossen werden.
- Beide Nummern können führende Nullen verlieren. Zum Beispiel würde PR 201 Issue 12 (012, eine Rotation von 201) schließen. Ebenso würde Issue 201 durch PR 12 geschlossen.

Geben Sie die verbleibenden offenen Issues in der Reihenfolge zurück, in der sie angegeben wurden.

# --hints--

`get_open_issues([123, 234], [231])` sollte `[234]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` sollte `[345, 16]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` sollte `[456, 332, 15]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` sollte `[115, 296, 24]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` sollte `[95, 395, 754, 296, 709, 237, 1802]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```
