---
id: 69f8c998d78ad3171a0713bd
title: "Aufgabe 290: Pizza-Party"
challengeType: 29
dashedName: challenge-290
---

# --description--

Gegeben ist ein Array mit den heute pro Person geleisteten Arbeitsstunden. Geben Sie die Anzahl der Pizzen zurück, die für eine Pizza-Party bestellt werden müssen.

- Teilen Sie die Arbeitsstunden jeder Person durch 3, um die Anzahl der Pizzastücke zu erhalten.
- Sie können kein halbes Stück essen, also runden Sie die Anzahl der Stücke jeder Person auf die nächste ganze Zahl auf.
- Jede Person bekommt mindestens zwei Stücke.
- Jede Pizza hat 8 Stücke. Runden Sie die Gesamtzahl der Pizzen auf die nächste ganze Pizza auf.

# --hints--

`get_pizzas_to_order([8, 8, 8])` sollte `2` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8]), 2)`)
}})
```

`get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10])` sollte `3` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 9, 8, 2, 2, 6, 10]), 3)`)
}})
```

`get_pizzas_to_order([1, 2, 3, 4, 5])` sollte `2` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([1, 2, 3, 4, 5]), 2)`)
}})
```

`get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8])` sollte `3` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([8, 8, 8, 8, 8, 8, 8, 8]), 3)`)
}})
```

`get_pizzas_to_order([9, 9, 6])` sollte `1` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([9, 9, 6]), 1)`)
}})
```

`get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0])` sollte `5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_pizzas_to_order([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_pizzas_to_order(hours_worked):

    return hours_worked
```

# --solutions--

```py
import math

def get_pizzas_to_order(hours_worked):
    total_slices = sum(max(math.ceil(h / 3), 2) for h in hours_worked)
    return math.ceil(total_slices / 8)
```
