---
id: 69f8c998d78ad3171a0713be
title: "Aufgabe 291: FizzBuzz zählen"
challengeType: 29
dashedName: challenge-291
---

# --description--

Gegeben sind eine Start- und eine Endzahl. Zählen Sie, wie oft „fizz“ und „buzz“ im Bereich (einschließlich) vorkommen.

- Zahlen, die durch 3 teilbar sind, zählen als „fizz“.
- Zahlen, die durch 5 teilbar sind, zählen als „buzz“.
- Zahlen, die durch sowohl 3 als auch 5 teilbar sind, zählen sowohl als „fizz“ als auch als „buzz“.

Geben Sie ein Objekt oder Wörterbuch mit den Zählungen im Format `{ fizz, buzz }` zurück.

# --hints--

`fizz_buzz_count(1, 11)` sollte `{"fizz": 3, "buzz": 2}` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(1, 11), {"fizz": 3, "buzz": 2})`)
}})
```

`fizz_buzz_count(14, 41)` sollte `{"fizz": 9, "buzz": 6}` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(14, 41), {"fizz": 9, "buzz": 6})`)
}})
```

`fizz_buzz_count(24, 100)` sollte `{"fizz": 26, "buzz": 16}` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(24, 100), {"fizz": 26, "buzz": 16})`)
}})
```

`fizz_buzz_count(-635, -14)` sollte `{"fizz": 207, "buzz": 125}` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-635, -14), {"fizz": 207, "buzz": 125})`)
}})
```

`fizz_buzz_count(-5432, 6789)` sollte `{"fizz": 4074, "buzz": 2444}` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fizz_buzz_count(-5432, 6789), {"fizz": 4074, "buzz": 2444})`)
}})
```

# --seed--

## --seed-contents--

```py
def fizz_buzz_count(start, end):

    return start
```

# --solutions--

```py
def fizz_buzz_count(start, end):
    fizz = sum(1 for i in range(start, end + 1) if i % 3 == 0)
    buzz = sum(1 for i in range(start, end + 1) if i % 5 == 0)
    return {"fizz": fizz, "buzz": buzz}
```
