---
id: 69f8c998d78ad3171a0713bf
title: "Aufgabe 292: Breiteres Seitenverhältnis"
challengeType: 29
dashedName: challenge-292
---

# --description--

Gegeben sind zwei Strings für unterschiedliche Bildmaße. Geben Sie das Seitenverhältnis des Bildes mit dem größeren Breite-zu-Höhe-Verhältnis zurück.

- Die gegebenen Strings haben das Format `"WxH"`, zum Beispiel `"1920x1080"`.
- Das Seitenverhältnis ist das Verhältnis von Breite zu Höhe, gekürzt auf die kleinsten ganzen Zahlen. Zum Beispiel wird `"1920x1080"` zu `"16:9"` gekürzt.
- Geben Sie einen String im Format `"W:H"` zurück, zum Beispiel `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` sollte `"16:9"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` sollte `"4:3"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` sollte `"2:1"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` sollte `"9:16"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` sollte `"43:18"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` sollte `"2089:8539"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```
