---
id: 69f90c1329a94b37e2a2086d
title: "Challenge 294: Parentheses Combinations"
challengeType: 29
dashedName: challenge-294
---

# --description--

Gegeben ist eine ganze Zahl, `n`. Geben Sie die Anzahl der gültigen Kombinationen von `n` Paaren runder Klammern zurück.

- Eine gültige Kombination ist ein String, bei dem jede öffnende runde Klammer eine entsprechende schließende runde Klammer hat und keine schließende runde Klammer vor ihrer passenden öffnenden runden Klammer erscheint.

Zum Beispiel gibt es bei `2` 2 gültige Kombinationen:

```json
(())
()()
```

# --hints--

`get_combinations(2)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` should return `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` should return `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` should return `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```
