---
id: 69f90c1329a94b37e2a2086c
title: "Aufgabe 293: Beste Hand"
challengeType: 29
dashedName: challenge-293
---

# --description--

Gegeben ist ein Array aus fünf Strings, die Spielkarten darstellen. Geben Sie den Namen der besten Hand zurück.

- Jede Karte wird als zweistelliger String dargestellt: der Rang gefolgt von der Farbe, `"2h"` zum Beispiel.
  - Die Ränge, von niedrig bis hoch, sind: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` und `"A"`.
  - Die Farben sind: `"h"`, `"d"`, `"c"` und `"s"`.
- Asse (`"A"`) können in einer Straße als hoch oder niedrig verwendet werden.

Die Hände, sortiert von der schlechtesten bis zur besten, sind:

| Name | Beschreibung |
| - | - |
| `"High Card"` | Kein Paar oder besser |
| `"Pair"` | Zwei Karten gleichen Rangs |
| `"Two Pair"` | Zwei Karten gleichen Rangs und zwei Karten eines anderen Rangs |
| `"Three of a Kind"` | Drei Karten gleichen Rangs |
| `"Straight"` | Fünf aufeinanderfolgende Ränge |
| `"Flush"` | Fünf Karten derselben Farbe |
| `"Full House"` | Drei Karten gleichen Rangs und zwei Karten eines anderen Rangs |
| `"Four of a Kind"` | Vier Karten gleichen Rangs |
| `"Straight Flush"` | Fünf aufeinanderfolgende Ränge derselben Farbe |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` derselben Farbe |

Geben Sie den Namen der besten Hand zurück.

# --hints--

`get_best_hand(["7s", "7h", "7d", "2c", "5h"])` sollte `"Three of a Kind"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind")`)
}})
```

`get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"])` sollte `"Full House"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House")`)
}})
```

`get_best_hand(["2h", "5h", "7h", "9h", "Jh"])` sollte `"Flush"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["2h", "5h", "7h", "9h", "Jh"]), "Flush")`)
}})
```

`get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"])` sollte `"Four of a Kind"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind")`)
}})
```

`get_best_hand(["Ts", "Th", "9d", "9c", "8h"])` sollte `"Two Pair"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair")`)
}})
```

`get_best_hand(["9c", "8c", "7c", "6c", "5c"])` sollte `"Straight Flush"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush")`)
}})
```

`get_best_hand(["As", "Kh", "Jd", "8c", "5h"])` sollte `"High Card"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Kh", "Jd", "8c", "5h"]), "High Card")`)
}})
```

`get_best_hand(["As", "2h", "3d", "4c", "5h"])` sollte `"Straight"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "2h", "3d", "4c", "5h"]), "Straight")`)
}})
```

`get_best_hand(["Ts", "Th", "7c", "6d", "5h"])` sollte `"Pair"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["Ts", "Th", "7c", "6d", "5h"]), "Pair")`)
}})
```

`get_best_hand(["As", "Ks", "Qs", "Js", "Ts"])` sollte `"Royal Flush"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_best_hand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_best_hand(cards):

    return cards
```

# --solutions--

```py
def get_best_hand(cards):
    rank_order = "23456789TJQKA"

    ranks = [c[0] for c in cards]
    suits = [c[1] for c in cards]

    is_flush = len(set(suits)) == 1
    rank_indices = sorted(rank_order.index(r) for r in ranks)

    def is_straight(indices):
        is_regular = all(indices[i] == indices[i - 1] + 1 for i in range(1, len(indices)))
        is_low_ace = indices == [0, 1, 2, 3, 12]
        return is_regular or is_low_ace

    straight = is_straight(rank_indices)

    counts = {}
    for r in ranks:
        counts[r] = counts.get(r, 0) + 1
    freq = sorted(counts.values(), reverse=True)

    if is_flush and rank_indices == [8, 9, 10, 11, 12]:
        return "Royal Flush"
    if is_flush and straight:
        return "Straight Flush"
    if freq[0] == 4:
        return "Four of a Kind"
    if freq[0] == 3 and freq[1] == 2:
        return "Full House"
    if is_flush:
        return "Flush"
    if straight:
        return "Straight"
    if freq[0] == 3:
        return "Three of a Kind"
    if freq[0] == 2 and freq[1] == 2:
        return "Two Pair"
    if freq[0] == 2:
        return "Pair"
    return "High Card"
```
