---
id: 6a0dcd03ee4e68698080ef68
title: "Aufgabe 304: Reiseplananordnungen"
challengeType: 29
dashedName: challenge-304
---

# --description--

Gegeben ist ein Array mit mindestens zwei optionalen Stopps für einen Tagesausflug. Geben Sie die Anzahl der gültigen Reiseplananordnungen zurück.

Der Reiseplan enthält immer `"breakfast"`, `"lunch"` und `"dinner"`, diese werden nicht als Argumente übergeben. Die optionalen Stopps können an beliebiger Stelle im Reiseplan eingefügt werden, unter Beachtung der folgenden Regeln:

- `"breakfast"` steht immer an erster Stelle, mit mindestens einem Stopp vor `"lunch"`.
- `"lunch"` muss vor `"dinner"` erscheinen, mit mindestens einem Stopp dazwischen.
- Nach `"dinner"` darf höchstens ein optionaler Stopp folgen.

Geben Sie die Anzahl der gültigen Anordnungen zurück.

# --hints--

`get_itinerary_count(["library", "park"])` sollte `2` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` sollte `18` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` sollte `120` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` sollte `840` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` sollte `55440` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```
