---
id: 6a0dcd03ee4e68698080ef6d
title: "Aufgabe 309: Zahlen sortieren"
challengeType: 29
dashedName: challenge-309
---

# --description--

Gegeben ist ein String mit Zahlen, die durch Kommas getrennt sind. Geben Sie ein Array mit den Zahlen zurück, sortiert von der kleinsten bis zur größten Zahl.

# --hints--

`sort_numbers("3,1,2")` sollte `[1, 2, 3]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("3,1,2"), [1, 2, 3])`)
}})
```

`sort_numbers("5,3,8,1,9,2")` sollte `[1, 2, 3, 5, 8, 9]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9])`)
}})
```

`sort_numbers("12,61,49,80,19,50,77,38")` sollte `[12, 19, 38, 49, 50, 61, 77, 80]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80])`)
}})
```

`sort_numbers("0,6,-19,44,-2,7,0")` sollte `[-19, -2, 0, 0, 6, 7, 44]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sort_numbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44])`)
}})
```

# --seed--

## --seed-contents--

```py
def sort_numbers(s):

    return s
```

# --solutions--

```py
def sort_numbers(s):
    return sorted(int(x) for x in s.split(","))
```
