---
id: 6a0dcd03ee4e68698080ef6c
title: "Aufgabe 308: Kreditkarten-Validator"
challengeType: 29
dashedName: challenge-308
---

# --description--

Gegeben ist ein String aus Ziffern für eine Kreditkartennummer. Bestimmen Sie, ob es sich um eine gültige Kartennummer handelt, indem Sie die folgende Methode anwenden:

- Beginnen Sie mit der zweitletzten Ziffer und verdoppeln Sie jede zweite Ziffer nach links.
- Wenn das Verdoppeln einer Ziffer zu einer Zahl größer als 9 führt, ziehen Sie 9 ab.
- Addieren Sie alle Ziffern (verdoppelte und nicht verdoppelte).
- Wenn die Summe durch 10 teilbar ist, ist die Nummer gültig.

# --hints--

`is_valid_card("4532015112830366")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830366"), True)`)
}})
```

`is_valid_card("5425233430109903")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("5425233430109903"), True)`)
}})
```

`is_valid_card("371449635398431")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("371449635398431"), True)`)
}})
```

`is_valid_card("6011111111111117")` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("6011111111111117"), True)`)
}})
```

`is_valid_card("4532015112830367")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830367"), False)`)
}})
```

`is_valid_card("1234567890123456")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("1234567890123456"), False)`)
}})
```

`is_valid_card("4532015112830368")` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_card("4532015112830368"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_card(number):

    return number
```

# --solutions--

```py
def is_valid_card(number):
    total = 0
    for i, digit in enumerate(reversed(number)):
        d = int(digit)
        if i % 2 == 1:
            d *= 2
            if d > 9:
                d -= 9
        total += d
    return total % 10 == 0
```
