---
id: 6a0dcd03ee4e68698080ef66
title: "Aufgabe 302: Jetlag"
challengeType: 29
dashedName: challenge-302
---

# --description--

Gegeben sind eine Abfahrtsstadt, eine Ankunftsstadt, eine Flugdauer in Stunden und eine Reiserichtung. Geben Sie die Anzahl der Jetlag-Stunden zurück, die der Reisende erlebt.

Die angegebenen Städte stammen aus der folgenden Liste mit ihrem UTC-Versatz:

| Stadt | Versatz |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Um die Jetlag-Stunden zu berechnen:

1. Ermitteln Sie die Zeitzonendifferenz in Stunden zwischen den beiden Städten.
2. Bestimmen Sie den Richtungsfaktor. Wenn die Reise `"east"` ist, beträgt er 1,5, andernfalls 1,0.
3. Berechnen Sie die Jetlag-Stunden mit der Formel: Zeitzonendifferenz + (Flugdauer * 0,1) * Richtungsfaktor

Geben Sie die Jetlag-Stunden auf eine Nachkommastelle gerundet zurück.

# --hints--

`get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east")` sollte `6.5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Istanbul", "Hong Kong", 10, "east"), 6.5)`)
}})
```

`get_jet_lag_hours("London", "New York", 8, "west")` sollte `5.8` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("London", "New York", 8, "west"), 5.8)`)
}})
```

`get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east")` sollte `1.6` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Hong Kong", "Tokyo", 4, "east"), 1.6)`)
}})
```

`get_jet_lag_hours("Dubai", "London", 7, "west")` sollte `4.7` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Dubai", "London", 7, "west"), 4.7)`)
}})
```

`get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west")` sollte `17.5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Los Angeles", "Hong Kong", 15, "west"), 17.5)`)
}})
```

`get_jet_lag_hours("Tokyo", "Dubai", 9, "west")` sollte `5.9` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("Tokyo", "Dubai", 9, "west"), 5.9)`)
}})
```

`get_jet_lag_hours("New York", "Istanbul", 10, "east")` sollte `9.5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_jet_lag_hours("New York", "Istanbul", 10, "east"), 9.5)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):

    return departure_city
```

# --solutions--

```py
def get_jet_lag_hours(departure_city, arrival_city, flight_duration, direction):
    offsets = {
        "Los Angeles": -8,
        "New York": -5,
        "London": 0,
        "Istanbul": 3,
        "Dubai": 4,
        "Hong Kong": 8,
        "Tokyo": 9
    }
    timezone_diff = abs(offsets[arrival_city] - offsets[departure_city])
    multiplier = 1.5 if direction == "east" else 1.0
    score = timezone_diff + (flight_duration * 0.1) * multiplier
    return round(score, 1)
```
