---
id: 6a0dcc730cb92a616f86f0c5
title: "Aufgabe 301: Letzte Ladung"
challengeType: 29
dashedName: challenge-301
---

# --description--

Gegeben ist die Anzahl der verbleibenden Waschpulverportionen und ein Array, das angibt, wie viele Portionen Sie an den vorherigen Tagen verwendet haben. Geben Sie die Anzahl der vollen Tage zurück, für die das Waschpulver noch ausreicht.

Berechnen Sie Ihren durchschnittlichen Tagesverbrauch anhand der Nutzungshistorie und gehen Sie davon aus, dass Sie diese Menge an Waschpulver jeden Tag weiterhin verwenden.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` sollte `5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` sollte `7` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` sollte `11` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` sollte `12` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` sollte `1` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```
