---
id: 6a0dcc730cb92a616f86f0c0
title: "Aufgabe 296: Schema-Validator Teil 2"
challengeType: 29
dashedName: challenge-296
---

# --description--

Gegeben ist ein Objekt (JavaScript) oder ein Dictionary (Python). Bestimmen Sie, ob es dem folgenden Schema entspricht:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Zusätzliche Schlüssel sind erlaubt

# --hints--

`is_valid_schema({"username": "alice", "posts": 10, "verified": False})` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "alice", "posts": 10, "verified": False}), True)`)
}})
```

`is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25})` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25}), True)`)
}})
```

`is_valid_schema({"username": "frank", "posts": "21", "verified": True})` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "frank", "posts": "21", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "sam", "posts": 17, "verified": "false"})` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sam", "posts": 17, "verified": "false"}), False)`)
}})
```

`is_valid_schema({"username": "bill", "verified": True})` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bill", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fred", "verified": True})` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fred", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": 5, "posts": 10, "verified": True})` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 5, "posts": 10, "verified": True}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool)
    )
```
