---
id: 6a0dcd03ee4e68698080ef6b
title: "Aufgabe 307: Bebauungsvorschriften"
challengeType: 29
dashedName: challenge-307
---

# --description--

Gegeben ist ein 2D-Grid (Array von Arrays), das den Bebauungsplan einer Stadt darstellt. Geben Sie die Koordinaten aller Gebäude zurück, die gegen die Bebauungsvorschriften verstoßen.

Jede Zelle im Grid enthält eine der Bezeichnungen aus der folgenden Tabelle. Ein Gebäude verstößt gegen die Vorschriften, wenn einer seiner (bis zu) 4 Nachbarn, horizontal oder vertikal, ein Typ ist, an den es nicht angrenzen darf.

| Bezeichnung | Typ | Darf nicht angrenzen an |
| - | - | - |
| `"i"` | industriell | `"R"`, `"I"` |
| `"A"` | landwirtschaftlich | `"C"` |
| `"R"` | Wohngebiet | `"i"`, `"C"` |
| `"I"` | institutionell | `"i"` |
| `"C"` | gewerblich | `"R"`, `"A"` |
| `""` (leerer String) | unbebaut | keine Einschränkungen |

Geben Sie die Koordinaten aller Zellen mit Verstößen als Array von `[row, col]`-Paaren in beliebiger Reihenfolge zurück. Wenn keine Verstöße vorliegen, geben Sie ein leeres Array zurück.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` sollte `[[0, 0], [0, 1]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` sollte `[[0, 1], [1, 1]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` sollte `[]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` sollte `[[0, 1], [0, 2], [0, 3]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` sollte `[[2, 3], [2, 4], [3, 1], [3, 2]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```
