---
id: 6a15cadf5f240d05a2649558
title: "Aufgabe 314: Primfaktorzerlegung"
challengeType: 29
dashedName: challenge-314
---

# --description--

Gegeben ist eine ganze Zahl größer als 1. Geben Sie ihre Primfaktorzerlegung als Array von Zahlen in aufsteigender Reihenfolge zurück.

Eine Primfaktorzerlegung ist die Menge von Primzahlen, die miteinander multipliziert die gegebene Zahl ergeben. Jede Zahl hat genau eine solche Menge. Zum Beispiel ist die Primfaktorzerlegung von 20 `[2, 2, 5]`, weil 2 * 2 * 5 = 20.

Wenn die gegebene Zahl selbst eine Primzahl ist, geben Sie sie in einem Array mit einem Element zurück.

# --hints--

`prime_factorization(20)` sollte `[2, 2, 5]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(20), [2, 2, 5])`)
}})
```

`prime_factorization(17)` sollte `[17]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(17), [17])`)
}})
```

`prime_factorization(15)` sollte `[3, 5]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(15), [3, 5])`)
}})
```

`prime_factorization(35)` sollte `[5, 7]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(35), [5, 7])`)
}})
```

`prime_factorization(999)` sollte `[3, 3, 3, 37]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(999), [3, 3, 3, 37])`)
}})
```

`prime_factorization(360)` sollte `[2, 2, 2, 3, 3, 5]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(360), [2, 2, 2, 3, 3, 5])`)
}})
```

`prime_factorization(510510)` sollte `[2, 3, 5, 7, 11, 13, 17]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(prime_factorization(510510), [2, 3, 5, 7, 11, 13, 17])`)
}})
```

# --seed--

## --seed-contents--

```py
def prime_factorization(n):

    return n
```

# --solutions--

```py
def prime_factorization(n):
    factors = []
    divisor = 2
    while divisor * divisor <= n:
        while n % divisor == 0:
            factors.append(divisor)
            n //= divisor
        divisor += 1
    if n > 1:
        factors.append(n)
    return factors
```
