---
id: 6a15cadf5f240d05a264955f
title: "Aufgabe 321: Periodische Rechtschreibung"
challengeType: 29
dashedName: challenge-321
---

# --description--

Gegeben ist ein Wort. Bestimmen Sie, ob es mit Elementsymbolen aus dem Periodensystem buchstabiert werden kann.

- Ignorieren Sie die Groß- und Kleinschreibung beim Buchstabieren eines Wortes. `"neon"` kann mit den Symbolen `"Ne"`, `"O"` und `"N"` buchstabiert werden.

Hier ist eine vollständige Liste der Elementsymbole:

```json
["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"];
```

Geben Sie ein Array der Elemente zurück, die verwendet werden, um das Wort zu buchstabieren, in ihrer ursprünglichen Groß- und Kleinschreibung und in der Reihenfolge, um das Wort zu buchstabieren. Oder ein leeres Array, wenn es nicht buchstabiert werden kann.

# --hints--

`get_periodic_spelling("neon")` sollte `["Ne", "O", "N"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("neon"), ["Ne", "O", "N"])`)
}})
```

`get_periodic_spelling("rational")` sollte `["Ra", "Ti", "O", "N", "Al"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("rational"), ["Ra", "Ti", "O", "N", "Al"])`)
}})
```

`get_periodic_spelling("yarn")` sollte `["Y", "Ar", "N"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("yarn"), ["Y", "Ar", "N"])`)
}})
```

`get_periodic_spelling("carbon")` sollte `["C", "Ar", "B", "O", "N"]` oder `["Ca", "Rb", "O", "N"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("carbon"), [["C", "Ar", "B", "O", "N"], ["Ca", "Rb", "O", "N"]])`)
}})
```

`get_periodic_spelling("noisy")` sollte `["N", "O", "I", "S", "Y"]` oder `["No", "I", "S", "Y"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("noisy"), [["N", "O", "I", "S", "Y"], ["No", "I", "S", "Y"]])`)
}})
```

`get_periodic_spelling("bicycles")` sollte `["B", "I", "C", "Y", "Cl", "Es"]` oder `["Bi", "C", "Y", "Cl", "Es"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("bicycles"), [["B", "I", "C", "Y", "Cl", "Es"], ["Bi", "C", "Y", "Cl", "Es"]])`)
}})
```

`get_periodic_spelling("optics")` sollte `["O", "P", "Ti", "C", "S"]`, `["O", "P", "Ti", "Cs"]`, `["O", "Pt", "I", "C", "S"]` oder `["O", "Pt", "I", "Cs"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIn(get_periodic_spelling("optics"), [["O", "P", "Ti", "C", "S"], ["O", "P", "Ti", "Cs"], ["O", "Pt", "I", "C", "S"], ["O", "Pt", "I", "Cs"]])`)
}})
```

`get_periodic_spelling("value")` sollte `[]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_periodic_spelling("value"), [])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_periodic_spelling(word):

    return word
```

# --solutions--

```py
def get_periodic_spelling(word):
    elements = ["H","He","Li","Be","B","C","N","O","F","Ne","Na","Mg","Al","Si","P","S","Cl","Ar","K","Ca","Sc","Ti","V","Cr","Mn","Fe","Co","Ni","Cu","Zn","Ga","Ge","As","Se","Br","Kr","Rb","Sr","Y","Zr","Nb","Mo","Tc","Ru","Rh","Pd","Ag","Cd","In","Sn","Sb","Te","I","Xe","Cs","Ba","La","Ce","Pr","Nd","Pm","Sm","Eu","Gd","Tb","Dy","Ho","Er","Tm","Yb","Lu","Hf","Ta","W","Re","Os","Ir","Pt","Au","Hg","Tl","Pb","Bi","Po","At","Rn","Fr","Ra","Ac","Th","Pa","U","Np","Pu","Am","Cm","Bk","Cf","Es","Fm","Md","No","Lr","Rf","Db","Sg","Bh","Hs","Mt","Ds","Rg","Cn","Nh","Fl","Mc","Lv","Ts","Og"]
    lower = {e.lower(): e for e in elements}

    def spell(w, i):
        if i == len(w):
            return []
        two = w[i:i + 2]
        one = w[i:i + 1]
        if two in lower:
            rest = spell(w, i + 2)
            if rest is not None:
                return [lower[two]] + rest
        if one in lower:
            rest = spell(w, i + 1)
            if rest is not None:
                return [lower[one]] + rest
        return None

    result = spell(word.lower(), 0)
    return result if result is not None else []
```
