---
id: 6a151d32d772271dd2448c2d
title: "Aufgabe 311: Zauberwirker"
challengeType: 29
dashedName: challenge-311
---

# --description--

Gegeben ist ein String mit Zaubercodes, die Sie wirken. Berechnen Sie die Gesamtpunktzahl.

Jedes Zeichen im String steht für einen Zauber:

| Code | Zauber | Kategorie | Grundpunktzahl |
| - | - | - | - |
| `"f"` | Feuer | Zerstörung | 3 |
| `"l"` | Blitz | Zerstörung | 3 |
| `"i"` | Eis | Kontrolle | 2 |
| `"w"` | Wind | Kontrolle | 2 |
| `"h"` | Heilung | Wiederherstellung | 1 |
| `"s"` | Schild | Wiederherstellung | 1 |

Ein Kombinationsmultiplikator wird basierend darauf angewendet, wie viele Zauber hintereinander aus unterschiedlichen Kategorien gewirkt wurden:

- Der erste Zauber erzielt immer die Grundpunktzahl.
- Jeder aufeinanderfolgende Zauber aus einer anderen Kategorie als der vorherige erhöht den Multiplikator um 1.
- Wird ein Zauber aus derselben Kategorie wie der vorherige gewirkt, wird der Multiplikator wieder auf 1 zurückgesetzt.
- Die Punktzahl für jeden Zauber ist seine Grundpunktzahl multipliziert mit dem aktuellen Multiplikator.

Geben Sie die Gesamtpunktzahl aus der Folge der Zauber zurück.

# --hints--

`cast("fihwl")` sollte `33` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` sollte `45` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` sollte `37` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` sollte `50` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` sollte `101` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```
