---
id: 6a15cadf5f240d05a264955c
title: "Aufgabe 318: DNA-Mutationen"
challengeType: 29
dashedName: challenge-318
---

# --description--

Gegeben sind zwei DNA-Stränge gleicher Länge. Geben Sie ein Array mit den Indexen zurück, an denen sich die Stränge unterscheiden (Mutationen).

- DNA-Stränge sind Strings, die aus den Zeichen `"A"`, `"T"`, `"C"` und `"G"` bestehen
- Geben Sie die Indexe in aufsteigender Reihenfolge zurück
- Wenn keine Mutationen vorliegen, geben Sie ein leeres Array zurück

# --hints--

`detect_mutations("ATCG", "ATGG")` sollte `[2]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATCG", "ATGG"), [2])`)
}})
```

`detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` sollte `[4, 9, 11]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11])`)
}})
```

`detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` sollte `[]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), [])`)
}})
```

`detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` sollte `[15, 16, 17, 18]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18])`)
}})
```

`detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` sollte `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_mutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28])`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_mutations(strand1, strand2):

    return strand1
```

# --solutions--

```py
def detect_mutations(strand1, strand2):
    return [i for i, (a, b) in enumerate(zip(strand1, strand2)) if a != b]
```
