---
id: 6a15cadf5f240d05a264955d
title: "Aufgabe 319: Frontmatter-Parser"
challengeType: 29
dashedName: challenge-319
---

# --description--

Gegeben ist ein String, der einen Frontmatter-Block darstellt. Parsen Sie ihn und geben Sie ein Objekt (JavaScript) oder ein Dictionary (Python) mit den Schlüsseln und Werten zurück.

Frontmatter ist von `---`-Begrenzer umschlossen und enthält darin `key: value`-Paare, jeweils eines pro Zeile. Zum Beispiel:

```md
---
title: My Post
draft: false
views: 100
---
```

Soll zurückgeben:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Zahlen, Booleans und Strings sollten jeweils als ihren entsprechenden Typ zurückgegeben werden.
- Der gegebene String enthält Zeilenumbrüche, die durch das Newline-Zeichen (`"\n"`) getrennt sind. Das obige Beispiel würde so übergeben werden: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parse_frontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` sollte `{ title: "My Post", draft: False, views: 100 }` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\ntitle: My Post\\ndraft: false\\nviews: 100\\n---"), {"title": "My Post", "draft": False, "views": 100})`)
}})
```

`parse_frontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` sollte `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: False }` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nid: 6a174db57256a112f932195c\\ntitle: My Book\\nlocale: en\\nwordCount: 10000\\npublished: false\\n---"), {"id": "6a174db57256a112f932195c", "title": "My Book", "locale": "en", "wordCount": 10000, "published": False})`)
}})
```

`parse_frontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` sollte `{ version: "1.0.0", url: "https://example.com", private: True }` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nversion: 1.0.0\\nurl: https://example.com\\nprivate: true\\n---"), {"version": "1.0.0", "url": "https://example.com", "private": True})`)
}})
```

`parse_frontmatter("---\nrating: 4.5\nprice: 9.99\n---")` sollte `{ rating: 4.5, price: 9.99 }` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_frontmatter("---\\nrating: 4.5\\nprice: 9.99\\n---"), {"rating": 4.5, "price": 9.99})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_frontmatter(s):

    return s
```

# --solutions--

```py
def parse_frontmatter(s):
    lines = [line for line in s.split("\n") if line != "---" and line != ""]
    result = {}
    for line in lines:
        key, *rest = line.split(":")
        value = ":".join(rest).strip()
        if value == "true":
            result[key.strip()] = True
        elif value == "false":
            result[key.strip()] = False
        else:
            try:
                n = float(value)
                result[key.strip()] = int(n) if n == int(n) else n
            except ValueError:
                result[key.strip()] = value
    return result
```
