---
id: 6a15cadf5f240d05a264955b
title: "Aufgabe 317: BMI-Rechner"
challengeType: 29
dashedName: challenge-317
---

# --description--

Gegeben sind ein Gewicht in Pfund und eine Körpergröße in Zoll. Geben Sie den BMI (Body Mass Index) zurück, gerundet auf eine Nachkommastelle.

Um den BMI zu berechnen: Teilen Sie das Gewicht durch die Körpergröße zum Quadrat und multiplizieren Sie das Ergebnis dann mit 703.

# --hints--

`calculate_bmi(180, 70)` sollte `25.8` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(180, 70), 25.8)`)
}})
```

`calculate_bmi(140, 64)` sollte `24.0` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(140, 64), 24.0)`)
}})
```

`calculate_bmi(160, 76)` sollte `19.5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(160, 76), 19.5)`)
}})
```

`calculate_bmi(200, 60)` sollte `39.1` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(200, 60), 39.1)`)
}})
```

`calculate_bmi(150, 68)` sollte `22.8` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_bmi(150, 68), 22.8)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_bmi(weight, height):

    return weight
```

# --solutions--

```py
def calculate_bmi(weight, height):
    return round((weight / (height * height)) * 703, 1)
```
