---
id: 6a15cadf5f240d05a2649557
title: "Aufgabe 313: Mietkosten"
challengeType: 29
dashedName: challenge-313
---

# --description--

Gegeben sind ein Mietzeitstempel, ein Rückgabezeitstempel und eine Mietstufe. Geben Sie die Gesamtkosten der Miete einschließlich etwaiger Verspätungsgebühren zurück.

- Die angegebenen Zeitstempel sind UTC-ISO-Strings, zum Beispiel: `"2026-06-18T18:30:00Z"`.
- Die Mietstufe ist die Anzahl der Tage, bevor die Miete zurückgegeben werden muss: `1`, `3` oder `7`.
- Die Miete muss bis spätestens 12:00 Uhr UTC am letzten Tag der Mietdauer zurückgegeben werden. Zum Beispiel ist eine 1-Tages-Miete, die irgendwann am 15. März ausgeliehen wurde, bis 12:00 Uhr UTC am 16. März zurückzugeben.
- Für jeden Tag, der über das Fälligkeitsdatum und die Fälligkeitszeit hinausgeht, fällt eine Verspätungsgebühr an.

Die Preisgestaltung ist wie folgt:

| Stufe | Grundkosten | Verspätungsgebühr pro Tag |
| - | - | - |
| 1 Tag | 4,99 $ | 3,99 $ |
| 3 Tage | 3,99 $ | 2,99 $ |
| 7 Tage | 2,99 $ | 0,99 $ |

Geben Sie die Gesamtkosten gerundet auf zwei Dezimalstellen im Format `"$D.CC"` zurück.

# --hints--

`get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` sollte `"$4.99"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99")`)
}})
```

`get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` sollte `"$12.97"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97")`)
}})
```

`get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` sollte `"$3.99"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99")`)
}})
```

`get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` sollte `"$9.97"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97")`)
}})
```

`get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` sollte `"$2.99"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99")`)
}})
```

`get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` sollte `"$358.40"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_rental_cost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_rental_cost(rented, returned, tier):

    return rented
```

# --solutions--

```py
from datetime import datetime, timedelta
import math

def get_rental_cost(rented, returned, tier):
    pricing = {
        1: {"base": 4.99, "late": 3.99},
        3: {"base": 3.99, "late": 2.99},
        7: {"base": 2.99, "late": 0.99},
    }
    rented_dt = datetime.fromisoformat(rented.replace("Z", "+00:00"))
    returned_dt = datetime.fromisoformat(returned.replace("Z", "+00:00"))
    due = rented_dt + timedelta(days=tier)
    due = due.replace(hour=12, minute=0, second=0, microsecond=0)
    base = pricing[tier]["base"]
    late = pricing[tier]["late"]
    if returned_dt <= due:
        return f"${base:.2f}"
    days_late = math.ceil((returned_dt - due).total_seconds() / 86400)
    return f"${base + days_late * late:.2f}"
```
