---
id: 6a19b1062d1b153d8ac76d73
title: "Aufgabe 325: Glückszahl"
challengeType: 29
dashedName: challenge-325
---

# --description--

Gegeben ist ein String mit dem Vor- und Nachnamen einer Person. Berechnen Sie deren Glückszahl anhand der folgenden Regeln:

- Vor- und Nachnamen sind durch ein Leerzeichen getrennt
- Ermitteln Sie die Anzahl der Vokale und Konsonanten für jeden Namen
- Multiplizieren Sie die kleinere Anzahl der Vokale und Konsonanten miteinander und anschließend mit der Länge des kleineren Namens
- Machen Sie dasselbe für die beiden größeren Werte und den größeren Namen
- Ziehen Sie den kleineren Wert vom größeren ab, um die Glückszahl zu erhalten

Wenn der Endwert null (`0`) ist, geben Sie `13` zurück.

# --hints--

`get_lucky_number("John Doe")` sollte `21` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("John Doe"), 21)`)
}})
```

`get_lucky_number("Olivia Lewis")` sollte `52` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Olivia Lewis"), 52)`)
}})
```

`get_lucky_number("James Wilson")` sollte `18` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("James Wilson"), 18)`)
}})
```

`get_lucky_number("Elizabeth Hernandez")` sollte `81` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Elizabeth Hernandez"), 81)`)
}})
```

`get_lucky_number("Mike Walker")` sollte `32` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Mike Walker"), 32)`)
}})
```

`get_lucky_number("Chloe Perez")` sollte `13` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lucky_number("Chloe Perez"), 13)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lucky_number(name):

    return name
```

# --solutions--

```py
def get_lucky_number(name):
    vowels = "aeiouAEIOU"

    def name_stats(s):
        vowel_count = sum(1 for c in s if c in vowels)
        consonant_count = sum(1 for c in s if c.isalpha() and c not in vowels)
        return vowel_count, consonant_count, len(s)

    first, last = name.split(" ")
    first_vowels, first_consonants, first_len = name_stats(first)
    last_vowels, last_consonants, last_len = name_stats(last)

    larger_vowels = max(first_vowels, last_vowels)
    smaller_vowels = min(first_vowels, last_vowels)
    larger_consonants = max(first_consonants, last_consonants)
    smaller_consonants = min(first_consonants, last_consonants)
    larger_length = max(first_len, last_len)
    smaller_length = min(first_len, last_len)

    larger_value = larger_vowels * larger_consonants * larger_length
    smaller_value = smaller_vowels * smaller_consonants * smaller_length
    result = larger_value - smaller_value

    return 13 if result == 0 else result
```
