---
id: 6a1d9f98e819ed70a0e994df
title: "Aufgabe 333: Issue-Triage 2"
challengeType: 29
dashedName: challenge-333
---

# --description--

Angenommen, Sie haben einen Issue-Titel und ein Array mit aktuellen Bezeichnungen. Geben Sie ein aktualisiertes Array von Bezeichnungen zurück, das auf den folgenden Regeln basiert:

Wenn der Issue keine Bezeichnungen hat, fügen Sie hinzu:

- `"bug"` und `"needs triage"`, wenn der Titel `"error"` oder `"bug"` enthält
- `"enhancement"` und `"discussing"`, wenn der Titel `"feature"` oder `"add"` enthält

Andernfalls, wenn die gegebenen Bezeichnungen enthalten:

- `"needs triage"` und der Titel `"simple"` oder `"easy"` enthält, entfernen Sie `"needs triage"` und fügen Sie `"good first issue"` hinzu
- `"discussing"` und der Titel `"planned"` oder `"next"` enthält, entfernen Sie `"discussing"` und fügen Sie `"on the roadmap"` hinzu
- Andernfalls, wenn `"needs triage"` oder `"discussing"` vorhanden ist, entfernen Sie es und fügen Sie `"help wanted"` hinzu

Wenn der Titel enthält:

- `"security"`, fügen Sie eine Bezeichnung `"critical"` hinzu

# --hints--

`triage_issue("app crashes with error", [])` sollte `["bug", "needs triage"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", [])), sorted(["bug", "needs triage"]))`)
}})
```

`triage_issue("app crashes with error", ["bug", "needs triage"])` sollte `["bug", "help wanted"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("app crashes with error", ["bug", "needs triage"])), sorted(["bug", "help wanted"]))`)
}})
```

`triage_issue("add dark mode", [])` sollte `["enhancement", "discussing"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", [])), sorted(["enhancement", "discussing"]))`)
}})
```

`triage_issue("add dark mode", ["enhancement", "discussing"])` sollte `["enhancement", "help wanted"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("add dark mode", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted"]))`)
}})
```

`triage_issue("xss security bug", [])` sollte `["bug", "needs triage", "critical"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("xss security bug", [])), sorted(["bug", "needs triage", "critical"]))`)
}})
```

`triage_issue("security vulnerability in auth", [])` sollte `["critical"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("security vulnerability in auth", [])), sorted(["critical"]))`)
}})
```

`triage_issue("easy a11y fix", ["bug", "needs triage"])` sollte `["bug", "good first issue"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("easy a11y fix", ["bug", "needs triage"])), sorted(["bug", "good first issue"]))`)
}})
```

`triage_issue("planned api migration", ["enhancement", "discussing"])` sollte `["enhancement", "on the roadmap"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("planned api migration", ["enhancement", "discussing"])), sorted(["enhancement", "on the roadmap"]))`)
}})
```

`triage_issue("improve security", ["enhancement", "discussing"])` sollte `["enhancement", "help wanted", "critical"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(triage_issue("improve security", ["enhancement", "discussing"])), sorted(["enhancement", "help wanted", "critical"]))`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(title, labels):

    return title
```

# --solutions--

```py
def triage_issue(title, labels):
    result = list(labels)
    t = title.lower()

    def has(label):
        return label in result

    def add(label):
        if not has(label):
            result.append(label)

    def remove(label):
        if label in result:
            result.remove(label)

    if len(result) == 0:
        if "error" in t or "bug" in t:
            add("bug")
            add("needs triage")
        elif "feature" in t or "add" in t:
            add("enhancement")
            add("discussing")
    else:
        if has("needs triage") and ("simple" in t or "easy" in t):
            remove("needs triage")
            add("good first issue")
        elif has("discussing") and ("planned" in t or "next" in t):
            remove("discussing")
            add("on the roadmap")
        elif has("needs triage") or has("discussing"):
            remove("needs triage")
            remove("discussing")
            add("help wanted")

    if "security" in t:
        add("critical")

    return result
```
