---
id: 6a1d9f98e819ed70a0e994dc
title: "Aufgabe 330: Kleinbuchstaben-Wörter"
challengeType: 29
dashedName: challenge-330
---

# --description--

Gegeben ist ein String. Geben Sie nur die Wörter zurück, die vollständig aus Kleinbuchstaben bestehen, in ihrer ursprünglichen Reihenfolge und mit einem Leerzeichen zwischen den Wörtern.

# --hints--

`get_lowercase_words("hello GOOD world")` sollte `"hello world"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("hello GOOD world"), "hello world")`)
}})
```

`get_lowercase_words("these are all lowercase")` sollte `"these are all lowercase"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("these are all lowercase"), "these are all lowercase")`)
}})
```

`get_lowercase_words("less is NoT more")` sollte `"less is more"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("less is NoT more"), "less is more")`)
}})
```

`get_lowercase_words("DonT eat pizza every OTHER day")` sollte `"eat pizza every day"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("DonT eat pizza every OTHER day"), "eat pizza every day")`)
}})
```

`get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` sollte `"the quick brown fox jumped over the lazy dog"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_lowercase_words("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_lowercase_words(s):

    return s
```

# --solutions--

```py
def get_lowercase_words(s):
    return " ".join(word for word in s.split(" ") if word == word.lower())
```
