---
id: 6a1d9f98e819ed70a0e994de
title: "Aufgabe 332: Issue-Triage"
challengeType: 29
dashedName: challenge-332
---

# --description--

Angenommen, Sie haben eine Anzahl von Millisekunden seit dem letzten Beitrag zu einem Issue und die letzte Nachricht, die zu diesem Issue gepostet wurde. Bestimmen Sie, was Sie mit dem Issue gemäß diesen Regeln tun sollten:

- Wenn die letzte Nachricht weniger als 7 Tage alt ist, geben Sie `"leave it"` zurück
- Wenn die letzte Nachricht 7 oder mehr Tage alt ist und ihr Inhalt `"bump"` (Groß-/Kleinschreibung egal) enthält, geben Sie `"close it"` zurück
- Andernfalls geben Sie `"bump it"` zurück

# --hints--

`triage_issue(86400000, "Lets fix it")` sollte `"leave it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(86400000, "Lets fix it"), "leave it")`)
}})
```

`triage_issue(1209600000, "still waiting")` sollte `"bump it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(1209600000, "still waiting"), "bump it")`)
}})
```

`triage_issue(864000000, "bump")` sollte `"close it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(864000000, "bump"), "close it")`)
}})
```

`triage_issue(604800000, "Do we still want this?")` sollte `"bump it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Do we still want this?"), "bump it")`)
}})
```

`triage_issue(604800000, "Bumping this")` sollte `"close it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(604800000, "Bumping this"), "close it")`)
}})
```

`triage_issue(345600000, "I'll make a PR")` sollte `"leave it"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(triage_issue(345600000, "I'll make a PR"), "leave it")`)
}})
```

# --seed--

## --seed-contents--

```py
def triage_issue(ms, message):

    return ms
```

# --solutions--

```py
def triage_issue(ms, message):
    seven_days = 7 * 24 * 60 * 60 * 1000
    if ms < seven_days:
        return "leave it"
    if "bump" in message.lower():
        return "close it"
    return "bump it"
```
