---
id: 6a1d9f98e819ed70a0e994da
title: "Aufgabe 328: Kaprekars Routine"
challengeType: 29
dashedName: challenge-328
---

# --description--

Gegeben ist eine vierstellige Zahl. Geben Sie die Anzahl der Male zurück, die Sie Kaprekars Routine anwenden müssen, bis Sie 6174 erreichen.

Kaprekars Routine funktioniert wie folgt:

- Ordnen Sie die Ziffern in absteigender Reihenfolge, um die größte Zahl zu bilden
- Ordnen Sie die Ziffern in aufsteigender Reihenfolge, um die kleinste Zahl zu bilden (füllen Sie bei Bedarf mit führenden Nullen auf)
- Subtrahieren Sie die kleinere von der größeren Zahl
- Wiederholen Sie den Vorgang mit der neuen Zahl

# --hints--

`kaprekar(1234)` sollte `3` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(1234), 3)`)
}})
```

`kaprekar(2025)` sollte `6` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(2025), 6)`)
}})
```

`kaprekar(7173)` sollte `4` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(7173), 4)`)
}})
```

`kaprekar(3164)` sollte `7` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(3164), 7)`)
}})
```

`kaprekar(8082)` sollte `2` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(kaprekar(8082), 2)`)
}})
```

# --seed--

## --seed-contents--

```py
def kaprekar(n):

    return n
```

# --solutions--

```py
def kaprekar(n):
    steps = 0
    current = n
    while current != 6174:
        digits = list(str(current).zfill(4))
        desc = int("".join(sorted(digits, reverse=True)))
        asc = int("".join(sorted(digits)))
        current = desc - asc
        steps += 1
    return steps
```
