---
id: 6a2037a68a0bc2aef0005fff
title: "Aufgabe 337: Zählwerk"
challengeType: 29
dashedName: challenge-337
---

# --description--

Gegeben ist ein String aus Strichlisten. Geben Sie die insgesamt dargestellte Anzahl zurück.

- Jeder senkrechte Strich `"|"` stellt eine Zählung dar.
- Jeder fünfte Strich wird als Schrägstrich `"/"` dargestellt und vervollständigt eine Gruppe von fünf (`"||||/"`).
- Die Gruppen sind durch ein Leerzeichen getrennt.

# --hints--

`get_tally_count("||||")` sollte `4` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||"), 4)`)
}})
```

`get_tally_count("||||/")` sollte `5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/"), 5)`)
}})
```

`get_tally_count("||||/ |||")` sollte `8` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ |||"), 8)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||")` sollte `17` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||"), 17)`)
}})
```

`get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` sollte `41` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_tally_count("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_tally_count(s):

    return s
```

# --solutions--

```py
def get_tally_count(s):
    return len(s.replace(" ", ""))
```
