---
id: 6a2037a68a0bc2aef0006006
title: "Aufgabe 344: Goldener Schnitt"
challengeType: 29
dashedName: challenge-344
---

# --description--

Gegeben sind zwei Zahlen. Bestimmen Sie, ob ihr Verhältnis ungefähr dem goldenen Schnitt entspricht.

- Verwenden Sie einen goldenen Schnitt von `1.618`
- Erlauben Sie eine Toleranz von `0.01`

# --hints--

`is_golden_ratio(21, 34)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(21, 34), True)`)
}})
```

`is_golden_ratio(15, 20)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(15, 20), False)`)
}})
```

`is_golden_ratio(8, 13)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(8, 13), True)`)
}})
```

`is_golden_ratio(10, 16)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(10, 16), False)`)
}})
```

`is_golden_ratio(1618, 1000)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(1618, 1000), True)`)
}})
```

`is_golden_ratio(88, 55)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_golden_ratio(88, 55), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_golden_ratio(a, b):

    return a
```

# --solutions--

```py
def is_golden_ratio(a, b):
    golden_ratio = 1.618
    ratio = max(a, b) / min(a, b)
    return abs(ratio - golden_ratio) <= 0.01
```
