---
id: 6a2037a68a0bc2aef0006002
title: "Aufgabe 340: Pig Latin-Konverter"
challengeType: 29
dashedName: challenge-340
---

# --description--

Gegeben ist ein String, den Sie nach den folgenden Regeln in Pig Latin umwandeln sollen:

- Beginnt ein Wort mit einem Vokal (`"a"`, `"e"`, `"i"`, `"o"` oder `"u"`), fügen Sie am Ende `"way"` hinzu. Zum Beispiel wird aus `"universe"` `"universeway"`.
- Beginnt ein Wort mit einem oder mehreren Konsonanten, verschieben Sie diese an das Ende und fügen `"ay"` hinzu. Zum Beispiel wird aus `"hello"` `"ellohay"`.
- Bewahren Sie die Groß- und Kleinschreibung des ersten Buchstabens. Zum Beispiel wird aus `"Hello"` `"Ellohay"`.

# --hints--

`pig_latin("universe")` sollte `"universeway"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("universe"), "universeway")`)
}})
```

`pig_latin("hello")` sollte `"ellohay"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello"), "ellohay")`)
}})
```

`pig_latin("hello universe")` sollte `"ellohay universeway"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("hello universe"), "ellohay universeway")`)
}})
```

`pig_latin("Hello universe")` sollte `"Ellohay universeway"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Hello universe"), "Ellohay universeway")`)
}})
```

`pig_latin("Pig Latin is fun")` sollte `"Igpay Atinlay isway unfay"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("Pig Latin is fun"), "Igpay Atinlay isway unfay")`)
}})
```

`pig_latin("The quick brown fox jumped over the lazy dog")` sollte `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pig_latin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday")`)
}})
```

# --seed--

## --seed-contents--

```py
def pig_latin(s):

    return s
```

# --solutions--

```py
def pig_latin(s):
    vowels = "aeiou"

    def convert_word(word):
        lower_word = word.lower()
        if lower_word[0] in vowels:
            return word + "way"
        i = 0
        while i < len(word) and lower_word[i] not in vowels:
            i += 1
        consonants = word[:i]
        rest = word[i:]
        result = rest + consonants + "ay"
        if word[0].isupper():
            return result[0].upper() + result[1:].lower()
        return result

    return " ".join(convert_word(word) for word in s.split(" "))
```
