---
id: 6a2037a68a0bc2aef0006005
title: "Aufgabe 343: Aufzug-Stopps"
challengeType: 29
dashedName: challenge-343
---

# --description--

Gegeben ist eine Zahl für die aktuelle Etage eines Aufzugs und ein Array mit angefragten Etagen. Geben Sie ein Array mit der Reihenfolge zurück, in der der Aufzug diese Etagen besuchen soll, um die Anzahl der zurückgelegten Etagen zu minimieren.

- Bei Gleichstand zuerst nach oben fahren
- Etagen mit einer Anfrage müssen besucht werden, wenn der Aufzug sie zum ersten Mal passiert

# --hints--

`elevator_stops(5, [2, 8, 3, 9])` sollte `[3, 2, 8, 9]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(5, [2, 8, 3, 9]), [3, 2, 8, 9])`)
}})
```

`elevator_stops(6, [2, 10, 8, 3, 1, 9])` sollte `[8, 9, 10, 3, 2, 1]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(6, [2, 10, 8, 3, 1, 9]), [8, 9, 10, 3, 2, 1])`)
}})
```

`elevator_stops(1, [4, 8, 3, 6, 9])` sollte `[3, 4, 6, 8, 9]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(1, [4, 8, 3, 6, 9]), [3, 4, 6, 8, 9])`)
}})
```

`elevator_stops(12, [6, 10, 7, 3, 1, 4])` sollte `[10, 7, 6, 4, 3, 1]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(12, [6, 10, 7, 3, 1, 4]), [10, 7, 6, 4, 3, 1])`)
}})
```

`elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19])` sollte `[10, 9, 8, 6, 5, 2, 12, 19, 23]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(elevator_stops(11, [2, 8, 23, 5, 12, 10, 6, 9, 19]), [10, 9, 8, 6, 5, 2, 12, 19, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def elevator_stops(current_floor, stops):

    return current_floor
```

# --solutions--

```py
def elevator_stops(current_floor, stops):
    stops = sorted(stops)
    above = [s for s in stops if s > current_floor]
    below = [s for s in stops if s < current_floor][::-1]
    up_distance = above[-1] - current_floor if above else float('inf')
    down_distance = current_floor - below[-1] if below else float('inf')
    return below + above if down_distance < up_distance else above + below
```
