---
id: 6a2037a68a0bc2aef0006001
title: "Aufgabe 339: Array-Teilarrays"
challengeType: 29
dashedName: challenge-339
---

# --description--

Gegeben ist ein Array und eine Chunk-Größe. Geben Sie das Array aufgeteilt in Teilarrays dieser Größe zurück.

- Der letzte Teilarray kann kleiner sein, wenn das Array nicht gleichmäßig teilbar ist.

# --hints--

`chunk_array([1, 2, 3, 4, 5, 6], 3)` sollte `[[1, 2, 3], [4, 5, 6]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]])`)
}})
```

`chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` sollte `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]])`)
}})
```

`chunk_array([1, 2, 3, 4, 5], 3)` sollte `[[1, 2, 3], [4, 5]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]])`)
}})
```

`chunk_array(["a", "b", "c", "d", "e"], 1)` sollte `[["a"], ["b"], ["c"], ["d"], ["e"]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]])`)
}})
```

`chunk_array([1, 2, 3], 5)` sollte `[[1, 2, 3]]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(chunk_array([1, 2, 3], 5), [[1, 2, 3]])`)
}})
```

# --seed--

## --seed-contents--

```py
def chunk_array(arr, size):

    return arr
```

# --solutions--

```py
def chunk_array(arr, size):
    return [arr[i:i + size] for i in range(0, len(arr), size)]
```
