---
id: 6a2037a68a0bc2aef0006000
title: "Aufgabe 338: Haustier-Alter-Rechner"
challengeType: 29
dashedName: challenge-338
---

# --description--

Geben Sie einen Haustiertyp und das Alter in Menschenjahren an, und geben Sie das entsprechende Alter in Tierjahren anhand der folgenden Umrechnungstabelle zurück:

| Haustier | Multiplikator |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`pet_years("dog", 5)` sollte `35` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("dog", 5), 35)`)
}})
```

`pet_years("cat", 9)` sollte `54` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("cat", 9), 54)`)
}})
```

`pet_years("rabbit", 3)` sollte `24` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("rabbit", 3), 24)`)
}})
```

`pet_years("hamster", 4)` sollte `120` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("hamster", 4), 120)`)
}})
```

`pet_years("guinea pig", 5)` sollte `60` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("guinea pig", 5), 60)`)
}})
```

`pet_years("goldfish", 2)` sollte `12` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("goldfish", 2), 12)`)
}})
```

`pet_years("bird", 1)` sollte `5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(pet_years("bird", 1), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def pet_years(pet, age):

    return pet
```

# --solutions--

```py
def pet_years(pet, age):
    table = {
        "dog": 7,
        "cat": 6,
        "rabbit": 8,
        "hamster": 30,
        "guinea pig": 12,
        "goldfish": 6,
        "bird": 5,
    }
    return age * table[pet]
```
