---
id: 6a2037a68a0bc2aef0006004
title: "Aufgabe 342: Würfelwahrscheinlichkeiten"
challengeType: 29
dashedName: challenge-342
---

# --description--

Gegeben ist eine Anzahl von sechsseitigen Würfeln, die geworfen werden, und eine Zielsumme. Geben Sie die Wahrscheinlichkeit zurück, diese Summe zu würfeln, als String im Format `"1 in X"`.

- Die Anzahl der Würfel liegt zwischen 1 und 6.
- Die Zielsumme ist mit der gegebenen Anzahl von Würfeln immer erreichbar.
- Runden Sie `"X"` auf die nächste ganze Zahl.

# --hints--

`get_odds(1, 5)` sollte `"1 in 6"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(1, 5), "1 in 6")`)
}})
```

`get_odds(2, 4)` sollte `"1 in 12"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(2, 4), "1 in 12")`)
}})
```

`get_odds(3, 10)` sollte `"1 in 8"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(3, 10), "1 in 8")`)
}})
```

`get_odds(4, 7)` sollte `"1 in 65"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(4, 7), "1 in 65")`)
}})
```

`get_odds(5, 26)` sollte `"1 in 111"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(5, 26), "1 in 111")`)
}})
```

`get_odds(6, 35)` sollte `"1 in 7776"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_odds(6, 35), "1 in 7776")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_odds(dice, target):

    return dice
```

# --solutions--

```py
def get_odds(dice, target):
    valid_combos = 0
    total_combos = 6 ** dice

    def roll(dice_left, total):
        nonlocal valid_combos
        if dice_left == 0:
            if total == target:
                valid_combos += 1
            return
        for face in range(1, 7):
            roll(dice_left - 1, total + face)

    roll(dice, 0)
    return f"1 in {round(total_combos / valid_combos)}"
```
