---
id: 6a2037a68a0bc2aef0006003
title: "Aufgabe 341: Countdown zum Geburtstag"
challengeType: 29
dashedName: challenge-341
---

# --description--

Geben Sie das heutige Datum und einen Geburtstag vor, und geben Sie die Anzahl der Tage bis zum nächsten Geburtstag der Person zurück.

- Das heutige Datum wird als String im Format `"YYYY-MM-DD"` mit führenden Nullen angegeben, zum Beispiel: `"2026-07-16"`.
- Der Geburtstag wird als String im Format `"M/D"` ohne führende Nullen angegeben, zum Beispiel: `"9/7"`.
- Wenn heute ihr Geburtstag ist, geben Sie die Anzahl der Tage bis zum nächsten Geburtstag zurück (nicht `0`).
- Schaltjahre müssen berücksichtigt werden.

# --hints--

`days_until_birthday("2026-07-16", "9/7")` sollte `53` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "9/7"), 53)`)
}})
```

`days_until_birthday("2026-07-16", "3/22")` sollte `249` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "3/22"), 249)`)
}})
```

`days_until_birthday("2026-07-16", "7/16")` sollte `365` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2026-07-16", "7/16"), 365)`)
}})
```

`days_until_birthday("2024-02-28", "3/1")` sollte `2` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-02-28", "3/1"), 2)`)
}})
```

`days_until_birthday("2023-04-24", "12/30")` sollte `250` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2023-04-24", "12/30"), 250)`)
}})
```

`days_until_birthday("2024-03-01", "2/29")` sollte `1460` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2024-03-01", "2/29"), 1460)`)
}})
```

`days_until_birthday("2096-03-01", "2/29")` sollte `2920` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(days_until_birthday("2096-03-01", "2/29"), 2920)`)
}})
```

# --seed--

## --seed-contents--

```py
def days_until_birthday(today, birthday):

    return today
```

# --solutions--

```py
def days_until_birthday(today, birthday):
    from datetime import date

    year, month, day = [int(x) for x in today.split("-")]
    b_month, b_day = [int(x) for x in birthday.split("/")]

    today_date = date(year, month, day)

    def is_leap_year(y):
        return (y % 4 == 0 and y % 100 != 0) or y % 400 == 0

    def next_valid_birthday(from_year):
        for y in range(from_year, from_year + 9):
            if b_month == 2 and b_day == 29 and not is_leap_year(y):
                continue
            d = date(y, b_month, b_day)
            if d > today_date:
                return d

    return (next_valid_birthday(year) - today_date).days
```
