---
id: 6a22d77ddf034bc4e35b1d59
title: "Aufgabe 352: Kontrastbewertung 1"
challengeType: 29
dashedName: challenge-352
---

# --description--

Geben Sie ein Kontrastverhältnis und einen Boolean an, der angibt, ob der Text groß ist, und geben Sie die WCAG-Bewertung anhand der folgenden Tabelle zurück:

| Bewertung | Normaler Text | Großer Text |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | unter 4,5 | unter 3,0 |

# --hints--

`get_contrast_rating("7.5", False)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("7.5", False), "AAA")`)
}})
```

`get_contrast_rating("4.8", False)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.8", False), "AA")`)
}})
```

`get_contrast_rating("4.2", False)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.2", False), "Fail")`)
}})
```

`get_contrast_rating("4.5", True)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("4.5", True), "AAA")`)
}})
```

`get_contrast_rating("3.0", True)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("3.0", True), "AA")`)
}})
```

`get_contrast_rating("2.7", False)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating("2.7", False), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(ratio, is_large_text):

    return ratio
```

# --solutions--

```py
def get_contrast_rating(ratio, is_large_text):
    ratio = float(ratio)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
