---
id: 6a22d77ddf034bc4e35b1d58
title: "Aufgabe 351: Pronic-Zahl"
challengeType: 29
dashedName: challenge-351
---

# --description--

Gegeben ist eine Zahl. Bestimmen Sie, ob es sich um eine Pronic-Zahl handelt.

Eine Pronic-Zahl ist das Produkt von zwei aufeinanderfolgenden ganzen Zahlen. Zum Beispiel ist 6 eine Pronic-Zahl, weil 2 * 3 = 6 gilt.

# --hints--

`is_pronic(6)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(6), True)`)
}})
```

`is_pronic(15)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(15), False)`)
}})
```

`is_pronic(12)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(12), True)`)
}})
```

`is_pronic(132)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(132), True)`)
}})
```

`is_pronic(80)` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(80), False)`)
}})
```

`is_pronic(0)` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_pronic(0), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_pronic(n):

    return n
```

# --solutions--

```py
def is_pronic(n):
    k = int(n ** 0.5)
    return k * (k + 1) == n
```
