---
id: 6a22d77ddf034bc4e35b1d5d
title: "Aufgabe 356: Magic Square Solver"
challengeType: 29
dashedName: challenge-356
---

# --description--

Gegeben ist ein 3x3-Raster mit einer fehlenden Zahl (dargestellt als `0`). Geben Sie die fehlende Zahl zurück, die das magische Quadrat vervollständigt, oder `"impossible"`, wenn keine gültige Zahl existiert.

Ein magisches Quadrat ist ein Raster, bei dem jede Zeile, Spalte und Diagonale die gleiche Summe ergibt.

# --hints--

`solve_magic_square([[2, 7, 6], [9, 0, 1], [4, 3, 8]])` sollte `5` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[2, 7, 6], [9, 0, 1], [4, 3, 8]]), 5)`)
}})
```

`solve_magic_square([[0, 14, 12], [18, 10, 2], [8, 6, 16]])` sollte `4` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[0, 14, 12], [18, 10, 2], [8, 6, 16]]), 4)`)
}})
```

`solve_magic_square([[12, 17, 16], [19, 0, 10], [14, 13, 18]])` sollte `"impossible"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[12, 17, 16], [19, 0, 10], [14, 13, 18]]), "impossible")`)
}})
```

`solve_magic_square([[15, 35, 31], [43, 27, 11], [23, 19, 0]])` sollte `39` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[15, 35, 31], [43, 27, 11], [23, 19, 0]]), 39)`)
}})
```

`solve_magic_square([[26, 41, 14], [47, 35, 0], [32, 29, 44]])` sollte `"impossible"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(solve_magic_square([[26, 41, 14], [47, 35, 0], [32, 29, 44]]), "impossible")`)
}})
```

# --seed--

## --seed-contents--

```py
def solve_magic_square(grid):

    return grid
```

# --solutions--

```py
def solve_magic_square(grid):
    flat = [n for row in grid for n in row]
    zero_index = flat.index(0)
    row = zero_index // 3
    col = zero_index % 3

    target = next(sum(r) for r in grid if 0 not in r)

    missing = target - sum(grid[row])
    filled = [r[:] for r in grid]
    filled[row][col] = missing

    def check(arr):
        return sum(arr) == target

    for i in range(3):
        if not check(filled[i]):
            return "impossible"
        if not check([filled[r][i] for r in range(3)]):
            return "impossible"

    if not check([filled[0][0], filled[1][1], filled[2][2]]):
        return "impossible"
    if not check([filled[0][2], filled[1][1], filled[2][0]]):
        return "impossible"

    return missing
```
