---
id: 6a22d77ddf034bc4e35b1d54
title: "Aufgabe 347: Spieltheorie"
challengeType: 29
dashedName: challenge-347
---

# --description--

Gegeben sind zwei Strings gleicher Länge, die die Strategien von zwei Spielern für ein Spiel darstellen. Geben Sie die Punktzahlen als Array `[player1, player2]` zurück.

- Die gegebenen Strings enthalten nur einen von zwei Buchstaben: `"C"` (kooperieren) oder `"D"` (defekt).
- Jeder Buchstabe steht für eine Runde, die wie folgt bewertet wird:
  - Wenn beide Spieler kooperieren, erhält jeder 3 Punkte.
  - Wenn beide Spieler defekt spielen, erhält jeder 1 Punkt.
  - Wenn ein Spieler defekt spielt und der andere kooperiert, erhält der Defektor 5 Punkte und der Kooperierende 0 Punkte.

# --hints--

`play_game("CCCC", "CCCC")` sollte `[12, 12]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCC", "CCCC"), [12, 12])`)
}})
```

`play_game("DDDD", "DDDD")` sollte `[4, 4]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDDD", "DDDD"), [4, 4])`)
}})
```

`play_game("CCDD", "CDDD")` sollte `[5, 10]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCDD", "CDDD"), [5, 10])`)
}})
```

`play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD")` sollte `[24, 34]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34])`)
}})
```

`play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` sollte `[66, 21]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(play_game("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21])`)
}})
```

# --seed--

## --seed-contents--

```py
def play_game(p1, p2):

    return p1
```

# --solutions--

```py
def play_game(p1, p2):
    scores = [0, 0]
    for i in range(len(p1)):
        if p1[i] == 'C' and p2[i] == 'C':
            scores[0] += 3
            scores[1] += 3
        elif p1[i] == 'D' and p2[i] == 'D':
            scores[0] += 1
            scores[1] += 1
        elif p1[i] == 'D' and p2[i] == 'C':
            scores[0] += 5
        else:
            scores[1] += 5
    return scores
```
