---
id: 6a22d77ddf034bc4e35b1d5b
title: "Aufgabe 354: Kontrastbewertung 3"
challengeType: 29
dashedName: challenge-354
---

# --description--

Gegeben sind zwei Arrays, die RGB-Werte darstellen, und ein Boolean, der angibt, ob der Text groß ist. Geben Sie die WCAG-Kontrastbewertung mit der folgenden Methode zurück:

Konvertieren Sie zuerst jeden RGB-Wert in die relative Leuchtdichte:

- Teilen Sie jeden Kanal `[R, G, B]` durch 255, um einen Wert zwischen `0` und `1` zu erhalten
- Wenden Sie die Gamma-Korrekturformel auf jeden Kanal an:
  - Wenn der Kanalwert kleiner oder gleich `0.04045` ist: `channel / 12.92`
  - Andernfalls: `((channel + 0.055) / 1.055) ^ 2.4`
- Berechnen Sie die Leuchtdichte: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Berechnen Sie dann das Kontrastverhältnis, indem Sie zu jedem Leuchtdichtewert 0,05 addieren und anschließend den helleren Wert durch den dunkleren teilen. Der hellere Wert ist immer das erste Argument.

Geben Sie die Bewertung basierend auf dem Kontrastverhältnis mit der folgenden Tabelle zurück:

| Bewertung | Normaler Text | Großer Text |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | unter 4,5 | unter 3,0 |

# --hints--

`get_contrast_rating([255, 255, 255], [0, 0, 0], False)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([255, 255, 255], [0, 0, 0], False), "AAA")`)
}})
```

`get_contrast_rating([215, 188, 188], [55, 55, 55], False)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([215, 188, 188], [55, 55, 55], False), "AA")`)
}})
```

`get_contrast_rating([143, 144, 210], [46, 47, 61], False)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([143, 144, 210], [46, 47, 61], False), "Fail")`)
}})
```

`get_contrast_rating([167, 167, 210], [53, 10, 53], True)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([167, 167, 210], [53, 10, 53], True), "AAA")`)
}})
```

`get_contrast_rating([135, 147, 155], [60, 70, 90], True)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([135, 147, 155], [60, 70, 90], True), "AA")`)
}})
```

`get_contrast_rating([125, 210, 195], [105, 130, 90], True)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating([125, 210, 195], [105, 130, 90], True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):

    return rgb1
```

# --solutions--

```py
def get_contrast_rating(rgb1, rgb2, is_large_text):
    def to_luminance(rgb):
        channels = []
        for c in rgb:
            c = c / 255
            if c <= 0.04045:
                channels.append(c / 12.92)
            else:
                channels.append(((c + 0.055) / 1.055) ** 2.4)
        return 0.2126 * channels[0] + 0.7152 * channels[1] + 0.0722 * channels[2]

    l1 = to_luminance(rgb1)
    l2 = to_luminance(rgb2)
    ratio = (l1 + 0.05) / (l2 + 0.05)

    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
