---
id: 6a22d77ddf034bc4e35b1d5a
title: "Aufgabe 353: Kontrastbewertung 2"
challengeType: 29
dashedName: challenge-353
---

# --description--

Gegeben sind zwei relative Leuchtdichtewerte und ein Boolean, der angibt, ob der Text groß ist. Geben Sie die WCAG-Kontrastbewertung mit der folgenden Methode zurück:

Berechnen Sie das Kontrastverhältnis, indem Sie zu jedem Leuchtdichtewert 0,05 addieren und dann den helleren Wert durch den dunkleren teilen. Der hellere Wert ist immer das erste Argument.

Geben Sie die Bewertung basierend auf dem Kontrastverhältnis mit der folgenden Tabelle zurück:

| Bewertung | Normaler Text | Großer Text |
|--------|-------------|------------|
| `"AAA"` | 7,0+ | 4,5+ |
| `"AA"` | 4,5+ | 3,0+ |
| `"Fail"` | unter 4,5 | unter 3,0 |

# --hints--

`get_contrast_rating(1.0, 0.0, False)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(1.0, 0.0, False), "AAA")`)
}})
```

`get_contrast_rating(0.9015, 0.1364, False)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.9015, 0.1364, False), "AA")`)
}})
```

`get_contrast_rating(0.8965, 0.1628, False)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.8965, 0.1628, False), "Fail")`)
}})
```

`get_contrast_rating(0.7469, 0.0957, True)` sollte `"AAA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7469, 0.0957, True), "AAA")`)
}})
```

`get_contrast_rating(0.7489, 0.2018, True)` sollte `"AA"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.7489, 0.2018, True), "AA")`)
}})
```

`get_contrast_rating(0.6571, 0.1974, True)` sollte `"Fail"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_contrast_rating(0.6571, 0.1974, True), "Fail")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_contrast_rating(l1, l2, is_large_text):

    return l1
```

# --solutions--

```py
def get_contrast_rating(l1, l2, is_large_text):
    ratio = (l1 + 0.05) / (l2 + 0.05)
    if is_large_text:
        if ratio >= 4.5:
            return "AAA"
        if ratio >= 3.0:
            return "AA"
    else:
        if ratio >= 7.0:
            return "AAA"
        if ratio >= 4.5:
            return "AA"
    return "Fail"
```
