---
id: 6a26df95efa55a2524399743
title: "Aufgabe 362: Nonogramm-Validator"
challengeType: 29
dashedName: challenge-362
---

# --description--

Gegeben sind ein Array von Hinweiszahlen und ein Array von Zellen. Bestimmen Sie, ob die Zellen den Nonogramm-Hinweis erfüllen.

- Der Hinweis ist ein Array von Zahlen, die die Längen aufeinanderfolgender gefüllter Zellen in der Reihenfolge darstellen. Zum Beispiel bedeutet ein Hinweis von `[3, 2]`, dass es 3 aufeinanderfolgende gefüllte Zellen geben soll, gefolgt von 2 aufeinanderfolgenden gefüllten Zellen, getrennt durch mindestens eine leere Zelle.
- Die Zeile ist ein Array aus 1en (gefüllt) und 0en (leer).

# --hints--

`is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1])` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [1, 1, 1, 0, 1, 1]), True)`)
}})
```

`is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1])` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2], [0, 1, 1, 1, 1, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1])` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [1, 0, 1, 0, 1, 0, 1, 0, 1]), False)`)
}})
```

`is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0])` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([1, 1, 1, 1], [0, 1, 0, 1, 0, 0, 0, 1, 0, 1, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0])` sollte `True` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 1, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 0, 0]), True)`)
}})
```

`is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0])` sollte `False` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_nonogram([3, 2, 3], [0, 0, 0, 1, 0, 0, 1, 0, 0, 0]), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_nonogram(clue, cells):

    return clue
```

# --solutions--

```py
def is_valid_nonogram(clue, cells):
    runs = []
    count = 0

    for i in range(len(cells) + 1):
        if i < len(cells) and cells[i] == 1:
            count += 1
        elif count > 0:
            runs.append(count)
            count = 0

    if len(runs) != len(clue):
        return False
    return all(runs[i] == clue[i] for i in range(len(runs)))
```
