---
id: 6a26df3e2988bcdded204895
title: "Aufgabe 359: Golf-Handicap-Rechner"
challengeType: 29
dashedName: challenge-359
---

# --description--

Gegeben ist ein Array von Golf-Ergebnissen und ein entsprechendes Array von Kurs-Par-Werten. Geben Sie den Handicap-Index des Golfspielers mit der folgenden Methode zurück:

- Berechnen Sie für jede Runde die Differenz, indem Sie das Par vom Ergebnis subtrahieren, und geben Sie dann den Durchschnitt aller Differenzen gerundet auf eine Nachkommastelle zurück.

# --hints--

`calculate_handicap([72, 72, 72], [72, 72, 72])` sollte `0` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([72, 72, 72], [72, 72, 72]), 0)`)
}})
```

`calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72])` sollte `6` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([80, 76, 78, 78], [72, 72, 72, 72]), 6)`)
}})
```

`calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36])` sollte `8.3` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3)`)
}})
```

`calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` sollte `8.8` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8)`)
}})
```

`calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` sollte `11.7` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_handicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7)`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_handicap(scores, pars):

    return scores
```

# --solutions--

```py
def calculate_handicap(scores, pars):
    differentials = [s - p for s, p in zip(scores, pars)]
    average = sum(differentials) / len(differentials)
    return int(average * 10 + 0.5) / 10
```
