---
id: 6a26df3e2988bcdded204893
title: "Aufgabe 357: Nahrungskette"
challengeType: 29
dashedName: challenge-357
---

# --description--

Gegeben ist ein Array von `[predator, prey]`-Paaren. Geben Sie die Nahrungskette vom Spitzenprädator bis zum unteren Ende zurück.

- Der Spitzenprädator ist das Tier, das niemals Beute eines anderen Tieres ist.
- Geben Sie die Kette als Array von Strings zurück.

# --hints--

`get_food_chain([["cat", "mouse"]])` sollte `["cat", "mouse"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["cat", "mouse"]]), ["cat", "mouse"])`)
}})
```

`get_food_chain([["wolf", "deer"], ["deer", "grass"]])` sollte `["wolf", "deer", "grass"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["wolf", "deer"], ["deer", "grass"]]), ["wolf", "deer", "grass"])`)
}})
```

`get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]])` sollte `["hawk", "snake", "frog", "fly"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["hawk", "snake"], ["snake", "frog"], ["frog", "fly"]]), ["hawk", "snake", "frog", "fly"])`)
}})
```

`get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]])` sollte `["eagle", "fox", "rabbit", "grass"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["rabbit", "grass"], ["fox", "rabbit"], ["eagle", "fox"]]), ["eagle", "fox", "rabbit", "grass"])`)
}})
```

`get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]])` sollte `["orca", "seal", "salmon", "herring", "shrimp", "plankton"]` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_food_chain([["seal", "salmon"], ["herring", "shrimp"], ["orca", "seal"], ["shrimp", "plankton"], ["salmon", "herring"]]), ["orca", "seal", "salmon", "herring", "shrimp", "plankton"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_food_chain(pairs):

    return pairs
```

# --solutions--

```py
def get_food_chain(pairs):
    prey_set = {p for _, p in pairs}
    chain_map = dict(pairs)
    apex = next(predator for predator, _ in pairs if predator not in prey_set)
    chain = [apex]
    while chain[-1] in chain_map:
        chain.append(chain_map[chain[-1]])
    return chain
```
