---
id: 6a26df3e2988bcdded204896
title: "Aufgabe 360: Gesprochene Dauer"
challengeType: 29
dashedName: challenge-360
---

# --description--

Geben Sie eine Anzahl von Sekunden vor, und geben Sie die Dauer in gesprochenem Englisch zurück.

- Zerlegen Sie die Dauer in Stunden, Minuten und Sekunden.
- Überspringen Sie alle Nullwerte.
- Verwenden Sie Singular oder Plural, je nachdem, was passend ist (`"1 hour"`, `"2 hours"`).
- Falls vorhanden, verbinden Sie die letzten beiden Einheiten mit `"and"` und die zweit- und drittletzten Einheiten mit einem Komma (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`get_spoken_duration(3723)` sollte `"1 hour, 2 minutes and 3 seconds"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(3723), "1 hour, 2 minutes and 3 seconds")`)
}})
```

`get_spoken_duration(7295)` sollte `"2 hours, 1 minute and 35 seconds"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(7295), "2 hours, 1 minute and 35 seconds")`)
}})
```

`get_spoken_duration(8521)` sollte `"2 hours, 22 minutes and 1 second"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(8521), "2 hours, 22 minutes and 1 second")`)
}})
```

`get_spoken_duration(435)` sollte `"7 minutes and 15 seconds"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(435), "7 minutes and 15 seconds")`)
}})
```

`get_spoken_duration(14455)` sollte `"4 hours and 55 seconds"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(14455), "4 hours and 55 seconds")`)
}})
```

`get_spoken_duration(72000)` sollte `"20 hours"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(72000), "20 hours")`)
}})
```

`get_spoken_duration(1)` sollte `"1 second"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_duration(1), "1 second")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_duration(seconds):

    return seconds
```

# --solutions--

```py
def get_spoken_duration(seconds):
    hours = seconds // 3600
    minutes = (seconds % 3600) // 60
    secs = seconds % 60

    parts = []
    if hours:
        parts.append(f"{hours} {'hour' if hours == 1 else 'hours'}")
    if minutes:
        parts.append(f"{minutes} {'minute' if minutes == 1 else 'minutes'}")
    if secs:
        parts.append(f"{secs} {'second' if secs == 1 else 'seconds'}")

    if len(parts) == 1:
        return parts[0]
    if len(parts) == 2:
        return f"{parts[0]} and {parts[1]}"
    return f"{parts[0]}, {parts[1]} and {parts[2]}"
```
