---
id: 6a26df3e2988bcdded204897
title: "Aufgabe 361: Gesprochene Zeit"
challengeType: 29
dashedName: challenge-361
---

# --description--

Gegeben sind die Winkel für den Stunden- und Minutenzeiger einer analogen Uhr in Grad (im Uhrzeigersinn von 12 aus gemessen). Geben Sie die Zeit in gesprochenem Englisch zurück.

Wandeln Sie den Winkel des Minutenzeigers in Minuten um (360° = 60 Minuten) und verwenden Sie dann die folgenden Regeln:

| Minuten | Gesprochen |
|---------|------------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (außer 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (außer 45) | `"X minutes to Z"` (wobei `X` 60 minus Minuten ist) |

Dabei ist `Y` die aktuelle Stunde und `Z` die nächste Stunde, beide abgeleitet vom Winkel des Stundenzeigers (360° = 12 Stunden).

Hinweis: Die Zeigerwinkel landen möglicherweise nicht genau auf einer Zahl, berücksichtigen Sie daher eine geeignete Rundung.

# --hints--

`get_spoken_time(90, 0)` sollte `"3 o'clock"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(90, 0), "3 o'clock")`)
}})
```

`get_spoken_time(160, 120)` sollte `"20 minutes past 5"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(160, 120), "20 minutes past 5")`)
}})
```

`get_spoken_time(255, 180)` sollte `"half past 8"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(255, 180), "half past 8")`)
}})
```

`get_spoken_time(67.5, 92)` sollte `"quarter past 2"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(67.5, 92), "quarter past 2")`)
}})
```

`get_spoken_time(200, 240)` sollte `"20 minutes to 7"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(200, 240), "20 minutes to 7")`)
}})
```

`get_spoken_time(322.5, 273)` sollte `"quarter to 11"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(322.5, 273), "quarter to 11")`)
}})
```

`get_spoken_time(117.5, 335)` sollte `"5 minutes to 4"` zurückgeben.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_spoken_time(117.5, 335), "5 minutes to 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_spoken_time(hour_angle, minute_angle):

    return hour_angle
```

# --solutions--

```py
def get_spoken_time(hour_angle, minute_angle):
    minutes = int(minute_angle / 6)
    hour = int(hour_angle / 30) or 12
    next_hour = (hour % 12) + 1

    if minutes == 0:
        return f"{hour} o'clock"
    if minutes == 15:
        return f"quarter past {hour}"
    if minutes == 30:
        return f"half past {hour}"
    if minutes == 45:
        return f"quarter to {next_hour}"
    if minutes < 30:
        return f"{minutes} minutes past {hour}"
    return f"{60 - minutes} minutes to {next_hour}"
```
