---
id: a103376db3ba46b2d50db289
title: 实现一个 Spinal Case 转换器
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinal case 是一种字串形式，其中所有单词均为小写并用连字符分隔。`"this-is-spinal-tap"` 是 spinal case 字串的一个示例。

在本实验中，你将创建一个将给定字串转换为 spinal case 的函数。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `spinalCase` 的函数。
2. `spinalCase` 函数应接受一个单一参数，即一个字串。
3. `spinalCase` 函数应返回脊柱形式格式的字串。例如，如果参数是 `ProductLanding page`，函数应返回 `product-landing-page`。

# --hints--

你应该创建一个名为 `spinalCase` 的函数。

```js
assert.isFunction(spinalCase);
```

`spinalCase` 应该接受一个单一参数。

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` 应返回 `this-is-spinal-tap`。

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` 应返回 `this-is-spinal-tap`。

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` 应返回 `the-andy-griffith-show`。

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` 应返回 `teletubbies-say-eh-oh`。

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` 应返回 `all-the-small-things`。

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
