---
id: 67ef8f695db52583424bd118
title: تصميم صفحة مراجعة فيلم
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي على عنصر `main`.
2. داخل عنصر `main`، يجب أن يوجد عنصر `h1` لعنوان الفيلم.
3. أسفل عنصر `h1`، يجب أن يوجد عنصر `img` يعرض غلاف الفيلم. يجب أن يحتوي عنصر `img` على نص `alt` وصفي يصف الصورة. يمكنك استخدام الصورة التالية إذا رغبت: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. يجب أن يوجد عنصر `p` يحتوي على وصف مختصر للفيلم.
5. يجب أن يوجد عنصر `p` آخر لعرض تقييم الفيلم. بداخله، يجب أن تتوفر العناصر التالية بالترتيب:
   - عنصر `b` بالنص `Movie Rating:`.
   - عنصر `span` مع خاصية `aria-hidden` مضبوطة على `true` ويحتوي على تمثيل مرئي للتقييم باستخدام النجوم `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`.
   - قيمة رقمية تمثل التقييم بين أقواس عادية (مثال: `9.2/10`) بعد الـ span.
6. يجب أن يوجد عنصر `h2` بالنص `Cast Members`.
7. يجب أن يوجد عنصر `ul`.
8. داخل عنصر `ul`، يجب أن يوجد عدة عناصر `li`، كل منها يحتوي على عنصر `b` لاسم الممثل متبوعًا باسم الشخصية مسبوقًا بالنص `as`. (مثال: `James Holloway as Ethan Carter`).

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.isNotNull(document.querySelector("main"));
```

داخل عنصر `main`، يجب أن يكون لديك عنصر `h1` يحتوي على عنوان الفيلم.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

أسفل عنصر `h1`، يجب أن يكون لديك عنصر `img` مع خاصية `alt` وصفية.
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

يجب أن يكون لديك عنصر `p` بعد الصورة يحتوي على وصف الفيلم.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

يجب أن يكون لديك عنصر `p` آخر يحتوي على التقييم.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

داخل عنصر `p` الثاني، يجب أن يكون لديك عنصر `b` يحتوي على النص `Movie Rating: `.

```js
const boldEl = document.querySelector("main p:nth-of-type(2) b");
assert.match(boldEl?.innerText.trim(), /^Movie Rating:?$/);
```

يجب أن يكون هناك عنصر `span` داخل فقرة التقييم.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

يجب أن يحتوي عنصر `span` داخل فقرة التقييم على عشرة نجوم، إما ممتلئة (⭐) أو فارغة (☆)، يتبعها قيمة رقمية بين أقواس عادية بعد عنصر `span`.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

يجب أن يحتوي عنصر `span` داخل فقرة التقييم على خاصية `aria-hidden` مضبوطة على `true`.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

يجب أن يكون لديك عنصر `h2`.

```js
assert.exists(document.querySelector("h2"));
```

يجب أن يأتي عنصر `h2` بعد عنصر `p` الثاني.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

يجب أن يحتوي عنصر `h2` على النص `Cast Members`.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

يجب أن يكون لديك عنصر `ul` بعد عنصر `h2`.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

يجب أن يحتوي عنصر `ul` على عدة عناصر `li`.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

يجب أن يحتوي كل عنصر `li` على عنصر `b` لاسم الممثل متبوعًا باسم الشخصية مع النص `as` قبله. (مثلاً، `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const bold = li.querySelector("b");
  assert.isNotNull(bold);
  const fullText = li.textContent.trim();
  const actorName = bold.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <b>Movie Rating:</b>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><b>James Holloway</b> as Ethan Carter</li>
        <li><b>Olivia Sterling</b> as Lena Mitchell</li>
        <li><b>William Lancaster</b> as Professor Adams</li>
        <li>
          <b>Maria Collins</b> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
