---
id: 6814d8e1516e86b171929de4
title: "Desafío 1: Equilibrio de Vocales"
challengeType: 28
dashedName: challenge-1
---

# --description--

Dada una cadena, determina si el número de vocales en la primera mitad de la cadena es igual al número de vocales en la segunda mitad.

- La cadena puede contener cualquier carácter.
- Las letras `a`, `e`, `i`, `o` y `u`, ya sea en mayúsculas o minúsculas, se consideran vocales.
- Si hay un número impar de caracteres en la cadena, ignora el carácter central.

# --hints--

`isBalanced("racecar")` debería devolver `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` debería devolver `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` debería devolver `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` debería devolver `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` debería devolver `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` debería devolver `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` debería devolver `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
