---
id: 681cb1a2dab50c87ddb2e514
title: "Desafío 3: Secuencia de Fibonacci"
challengeType: 28
dashedName: challenge-3
---

# --description--

La secuencia de Fibonacci es una serie de números donde cada número es la suma de los dos anteriores. Al comenzar con `0` y `1`, los primeros 10 números en la secuencia son `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Dado un arreglo que contiene los dos primeros números de una secuencia de Fibonacci, y un entero que representa la longitud de la secuencia, devuelve un arreglo que contiene la secuencia de la longitud dada.

- Tu función debe manejar secuencias de cualquier longitud mayor o igual a cero.
- Si la longitud es cero, devuelve un arreglo vacío.
- Ten en cuenta que los números iniciales son parte de la secuencia.

# --hints--

`fibonacciSequence([0, 1], 20)` debería devolver `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` debería devolver `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` debería devolver `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` debería devolver `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` debería devolver `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
