---
id: 681cb1b0dab50c87ddb2e51a
title: "Desafío 9: Suma de Cuadrados"
challengeType: 28
dashedName: challenge-9
---

# --description--

Dado un entero positivo hasta 1,000, devuelve la suma de todos los enteros al cuadrado desde 1 hasta el número.

# --hints--

`sumOfSquares(5)` debería devolver `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` debería devolver `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` debería devolver `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` debería devolver `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` debería devolver `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
