---
id: 681cb1afdab50c87ddb2e515
title: "Desafío 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

Dada una cadena, elimina todos los espacios de la cadena, inserta dos espacios entre cada carácter, convierte todas las letras alfabéticas a mayúsculas y devuelve el resultado.

- Los caracteres no alfabéticos deben permanecer sin cambios (excepto los espacios).

# --hints--

`spaceJam("freeCodeCamp")` debería devolver `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` debería devolver `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` debería devolver `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` debería devolver `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` debería devolver `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
