---
id: 681cb1b0dab50c87ddb2e518
title: "Desafío 7: Suma Dirigida"
challengeType: 28
dashedName: challenge-7
---

# --description--

Dado un arreglo de números y un entero target, encuentra dos números únicos en el arreglo que sumen el valor target. Devuelve un arreglo con los índices de esos dos números, o `"Target not found"` si no hay dos números que sumen el target.

- El array retornado debe tener los índices en orden ascendente.

# --hints--

`findTarget([2, 7, 11, 15], 9)` debería devolver `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` debería devolver `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` debería devolver `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` debería devolver `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
