---
id: 681cb1afdab50c87ddb2e516
title: "Desafío 5: Jbelmud Text"
challengeType: 28
dashedName: challenge-5
---

# --description--

Dada una cadena, devuelve una versión mezclada de esa cadena donde cada palabra se transforma usando las siguientes restricciones:

- La primera y última letra de las palabras permanecen en su lugar
- Todas las letras entre la primera y la última letra están ordenadas alfabéticamente.
- las cadenas de entrada no contendrán puntuación y estarán completamente en minúsculas

# --hints--

`jbelmu("hello world")` debería devolver `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` debería devolver `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` debería devolver `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` debería devolver `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
