---
id: 681cb1b0dab50c87ddb2e519
title: "Desafío 8: Factorializer"
challengeType: 28
dashedName: challenge-8
---

# --description--

Dado un entero de cero a 20, devuelve el factorial de ese número. El factorial de un número es el producto de todos los números entre 1 y el número dado.

- El factorial de cero es 1.

# --hints--

`factorial(0)` debería devolver `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` debería devolver `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` debería devolver `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
