---
id: 681cb1b0dab50c87ddb2e51b
title: "Desafío 10: 3 Strikes"
challengeType: 28
dashedName: challenge-10
---

# --description--

Dado un entero entre 1 y 10,000, devuelve un conteo de cuántos números desde 1 hasta ese entero cuyo cuadrado contiene al menos un dígito 3.

# --hints--

`squaresWithThree(1)` debería devolver `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` debería devolver `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` debería devolver `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` debería devolver `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` debería devolver `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
