---
id: 681cb1afdab50c87ddb2e517
title: "Desafío 6: Verificador de Anagramas"
challengeType: 28
dashedName: challenge-6
---

# --description--

Dadas dos cadenas, determina si son anagramas entre sí (contienen los mismos caracteres en cualquier orden).

- Ignora mayúsculas, minúsculas y espacios en blanco.

# --hints--

`areAnagrams("listen", "silent")` debería devolver `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` debería devolver `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` debería devolver `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` debería devolver `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` debería devolver `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` debería devolver `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
