---
id: 6821ebfd237de8297eaee799
title: "Desafío 23: RGB a Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

Dada una cadena de color CSS `rgb(r, g, b)`, devuelve su equivalente hexadecimal.

Aquí tienes algunos ejemplos de salidas para una entrada dada:

EntradaResultado
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Convierte cualquier letra a minúsculas.
- Devuelve un `#` seguido de seis caracteres. No uses valores abreviados.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` debería devolver `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` debería devolver `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` debería devolver `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` debería devolver `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
