---
id: 6821ebd4237de8297eaee791
title: "Desafío 15: camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

Dada una cadena, devuelve su versión en camel case usando las siguientes reglas:

- Las palabras en el argumento de cadena están separadas por uno o más caracteres del siguiente conjunto: espacio (` `), guion (`-`) o guion bajo (`_`). Trata cualquier secuencia de estos como una separación de palabra.
- La primera palabra debe estar toda en minúsculas.
- Cada palabra subsecuente debe comenzar con una letra mayúscula, y el resto en minúscula.
- Todos los espacios y separadores deben ser eliminados.

# --hints--

`toCamelCase("hello world")` debería devolver `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` debería devolver `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` debería devolver `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` debería devolver `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` debería devolver `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
