---
id: 6821ebda237de8297eaee792
title: "Desafío 16: Paréntesis Invertidos"
challengeType: 28
dashedName: challenge-16
---

# --description--

Dada una cadena que contiene paréntesis correctamente anidados, devuelve la versión decodificada de la cadena usando las siguientes reglas:

- Todos los caracteres dentro de cada paréntesis deben invertirse.
- Los paréntesis deben eliminarse del resultado final.
- Si los paréntesis están anidados, el par más interno debe invertirse primero, y luego su resultado debe incluirse en la inversión del par externo.
- Asume que todos los paréntesis están equilibrados y correctamente anidados.

# --hints--

`decode("(f(b(dc)e)a)")` debería devolver `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` debería devolver `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` debería devolver `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
