---
id: 6821ebc9237de8297eaee78f
title: "Desafío 13: Primo No Natural"
challengeType: 28
dashedName: challenge-13
---

# --description--

Dado un entero, determina si ese número es un número primo o un número primo negativo.

- Un número primo es un entero positivo mayor que 1 que solo es divisible por 1 y por sí mismo.
- Un número primo negativo es la versión negativa de un número primo positivo.
- `1` y `0` no se consideran números primos.

# --hints--

`isUnnaturalPrime(1)` debería devolver `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` debería devolver `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` debería devolver `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` debería devolver `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` debería devolver `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` debería devolver `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` debería devolver `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` debería devolver `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` debería devolver `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
