---
id: 6821ebce237de8297eaee790
title: "Desafío 14: Batalla de Personajes"
challengeType: 28
dashedName: challenge-14
---

# --description--

Dadas dos cadenas que representan tu ejército y un ejército opuesto, cada carácter de tu ejército combate contra el carácter en la misma posición del ejército opuesto usando las siguientes reglas:

- Los caracteres `a-z` tienen una fuerza de `1-26`, respectivamente.
- Los caracteres `A-Z` tienen una fuerza de `27-52`, respectivamente.
- Los dígitos `0-9` tienen una fuerza igual a su valor facial.
- Todos los demás caracteres tienen un valor de cero.
- Cada carácter solo puede luchar en una batalla.

Para cada batalla, el personaje más fuerte gana. El ejército con más victorias gana la guerra. Devuelve los siguientes valores:

- `"Opponent retreated"` si tu ejército tiene más personajes que el ejército contrario.
- "Nos retiramos" si el ejército opuesto tiene más caracteres que tú.
- "Ganamos" si tu ejército ganó más batallas.
- "Perdimos" si el ejército contrario ganó más batallas.
- "Fue un empate" si ambos ejércitos ganaron el mismo número de batallas.

# --hints--

`battle("Hello", "World")` debería devolver `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` debería devolver `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` debería devolver `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` debería devolver `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` debería devolver `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` debería devolver `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` debería devolver `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
