---
id: 6821ebee237de8297eaee796
title: "Desafío 20: Duplicados en Arrays"
challengeType: 28
dashedName: challenge-20
---

# --description--

Dado un arreglo de enteros, devuelve un arreglo de enteros que aparecen más de una vez en el arreglo inicial, ordenados en orden ascendente. Si ningún valor aparece más de una vez, devuelve un arreglo vacío.

- Solo incluye una instancia de cada valor en el array devuelto.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` debería devolver `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` debería devolver `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` debería devolver `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
