---
id: 6821ebf8237de8297eaee798
title: "Desafío 22: Secuencia Tribonacci"
challengeType: 28
dashedName: challenge-22
---

# --description--

La secuencia Tribonacci es una serie de números donde cada número es la suma de los tres anteriores. Al comenzar con `0`, `0` y `1`, los primeros 10 números de la secuencia son `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Dado un arreglo que contiene los primeros tres números de una secuencia Tribonacci, y un entero que representa la longitud de la secuencia, devuelve un arreglo que contiene la secuencia de la longitud dada.

- Tu función debe manejar secuencias de cualquier longitud mayor o igual a cero.
- Si la longitud es cero, devuelve un arreglo vacío.
- Ten en cuenta que los números iniciales son parte de la secuencia.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` debería devolver `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` debería devolver `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` debería devolver `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` debería devolver `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` debería devolver `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` debería devolver `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
