---
id: 6821ebdf237de8297eaee793
title: "Desafío 17: Orden de las Operaciones"
challengeType: 28
dashedName: challenge-17
---

# --description--

Dado un arreglo de enteros y un arreglo de operadores de cadena, aplica las operaciones a los números secuencialmente de izquierda a derecha. Repite las operaciones según sea necesario hasta que se usen todos los números. Devuelve el resultado final.

Por ejemplo, dado `[1, 2, 3, 4, 5]` y `['+', '*']`, devuelve el resultado de evaluar `1 + 2 * 3 + 4 * 5` de izquierda a derecha ignorando el orden estándar de las operaciones.

- Los operadores válidos son `+`, `-`, `*`, `/` y `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` debería devolver `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` debería devolver `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` debería devolver `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` debería devolver `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` debería devolver `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```
