---
id: 68216ef80f957572e7c340c5
title: "Desafío 12: Decodificador de Mensajes"
challengeType: 28
dashedName: challenge-12
---

# --description--

Dada una cadena de mensaje secreto y un entero que representa el número de letras que se usaron para desplazar el mensaje y codificarlo, devuelve la cadena decodificada.

- Un número positivo significa que el mensaje fue desplazado hacia adelante en el alfabeto.
- Un número negativo significa que el mensaje fue desplazado hacia atrás en el alfabeto.
- Las mayúsculas y minúsculas importan, los caracteres decodificados deben conservar la estructura de mayúsculas y minúsculas de sus contrapartes codificadas.
- Los caracteres no alfabéticos no deben ser decodificados.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` debería devolver `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` debería devolver `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` debería devolver `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` debería devolver `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
