---
id: 6821ebe4237de8297eaee794
title: "Desafío 18: Segundo Mejor"
challengeType: 28
dashedName: challenge-18
---

# --description--

Dado un arreglo de enteros que representa el precio de diferentes laptops, y un entero que representa tu presupuesto, devuelve:

1. La segunda laptop más cara si está dentro de tu presupuesto, o
2. La laptop más cara que esté dentro de tu presupuesto, o
3. `0` si no hay laptops dentro de tu presupuesto.

- Los precios duplicados deben ser ignorados.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` debería devolver `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` debería devolver `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` debería devolver `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` debería devolver `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` debería devolver `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
