---
id: 68adce01c0e1144d0a902958
title: "Desafío 26: Validador de IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

Dada una cadena, determina si es una dirección IPv4 válida. Una dirección IPv4 válida consiste en cuatro números enteros separados por puntos (`.`). Cada número debe cumplir con las siguientes condiciones:

- Está entre 0 y 255 inclusive.
- No tiene ceros a la izquierda (por ejemplo, 0 está permitido, 01 no).
- Solo se permiten caracteres numéricos.

# --hints--

`isValidIPv4("192.168.1.1")` debería devolver `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` debería devolver `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` debería devolver `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` debería devolver `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` debería devolver `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` debería devolver `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` debería devolver `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
