---
id: 68adce01c0e1144d0a90295c
title: "Desafío 28: Analizador de Números Romanos"
challengeType: 28
dashedName: challenge-28
---

# --description--

Dada una cadena que representa un número romano, devuelve su valor entero.

Los números romanos consisten en los siguientes símbolos y valores:

SímboloValor
|--------|-------|
I1
V5
X10
L50
C100
D500
M1000

- Los números se leen de izquierda a derecha. Si un número menor aparece antes que uno mayor, se resta el valor. De lo contrario, los valores se suman.

# --hints--

`parseRomanNumeral("III")` debería devolver `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` debería devolver `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` debería devolver `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` debería devolver `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` debería devolver `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` debería devolver `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` debería devolver `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
