---
id: 68adce01c0e1144d0a90295e
title: "Desafío 29: Constructor de Acrónimos"
challengeType: 28
dashedName: challenge-29
---

# --description--

Dada una cadena que contiene una o más palabras, devuelve un acrónimo de las palabras usando las siguientes restricciones:

- El acrónimo debe consistir en la primera letra de cada palabra en mayúscula, a menos que se indique lo contrario.
- El acrónimo debe ignorar la primera letra de estas palabras a menos que sean la primera palabra de la cadena dada: `a`, `for`, `an`, `and`, `by`, y `of`.
- Las letras del acrónimo deben devolverse en el orden en que se proporcionan.
- El acrónimo no debe contener espacios.

# --hints--

`buildAcronym("Search Engine Optimization")` debería devolver `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` debería devolver `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` debería devolver `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` debería devolver `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` debería devolver `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` debería devolver `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` debería devolver `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
