---
id: 68adce01c0e1144d0a90295a
title: "Desafío 27: Rotar Matriz"
challengeType: 28
dashedName: challenge-27
---

# --description--

Dada una matriz (un arreglo de arreglos), rota la matriz 90 grados en el sentido de las agujas del reloj y devuélvela. Por ejemplo, dada `[[1, 2], [3, 4]]`, que se ve así:

12
|---|---|
34

Debes devolver `[[3, 1], [4, 2]]`, que se ve así:

31
|---|---|
42

# --hints--

`rotate([[1]])` debería devolver `[[1]]`.

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

`rotate([[1, 2], [3, 4]])` debería devolver `[[3, 1], [4, 2]]`.

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

`rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` debería devolver `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

`rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` debería devolver `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```
