---
id: 68adce01c0e1144d0a902956
title: "Desafío 25: Repetidor de Vocales"
challengeType: 28
dashedName: challenge-25
---

# --description--

Dada una cadena, devuelve una nueva versión de la cadena donde cada vocal se duplica una vez más que la vocal anterior que encontraste. Por ejemplo, la primera vocal en la oración debe permanecer sin cambios. La segunda vocal debe aparecer dos veces seguidas. La tercera vocal debe aparecer tres veces seguidas, y así sucesivamente.

- Las letras `a`, `e`, `i`, `o` y `u`, ya sea en mayúsculas o minúsculas, se consideran vocales.
- La vocal original debe mantener su mayúscula o minúscula.
- Las vocales repetidas deben estar en minúsculas.
- Todos los caracteres que no sean vocales deben mantener su caso original.

# --hints--

`repeatVowels("hello world")` debería devolver `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` debería devolver `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` debería devolver `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` debería devolver `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
