---
id: 68af0687ef34c76c28ffa54d
title: "Desafío 33: Tiempo frente a la pantalla"
challengeType: 28
dashedName: challenge-33
---

# --description--

Dado un arreglo de entrada de siete enteros, que representan el tiempo de una semana, donde cada entero es la cantidad de horas que pasaste en tu teléfono ese día, determina si es demasiado tiempo frente a la pantalla basado en estas restricciones:

- Si cualquier día tiene 10 horas o más, es demasiado.
- Si el promedio de cualquier tres días consecutivos es mayor o igual a 8 horas, es demasiado.
- Si el promedio de los siete días es mayor o igual a 6 horas, es demasiado.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` debería devolver `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` debería devolver `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` debería devolver `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` debería devolver `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` debería devolver `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` debería devolver `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` debería devolver `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```
