---
id: 68af0687ef34c76c28ffa547
title: "Desafío 30: Caracteres Únicos"
challengeType: 28
dashedName: challenge-30
---

# --description--

Dada una cadena, determina si todos los caracteres en la cadena son únicos.

- Las letras mayúsculas y minúsculas deben considerarse caracteres diferentes.

# --hints--

`allUnique("abc")` debería devolver `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` debería devolver `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` debería devolver `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` debería devolver `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` debería devolver `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` debería devolver `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` debería devolver `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
