---
id: 68af0687ef34c76c28ffa549
title: "Desafío 31: Array Diff"
challengeType: 28
dashedName: challenge-31
---

# --description--

Dado dos arreglos con valores de cadenas, devuelve un nuevo arreglo que contiene todos los valores que aparecen en solo uno de los arreglos.

- El array retornado debe estar ordenado en orden alfabético.

# --hints--

`arrayDiff(["apple", "banana"], ["apple", "banana", "cherry"])` debería devolver `["cherry"]`.

```js
assert.deepEqual(arrayDiff(["apple", "banana"], ["apple", "banana", "cherry"]), ["cherry"]);
```

`arrayDiff(["apple", "banana", "cherry"], ["apple", "banana"])` debería devolver `["cherry"]`.

```js
assert.deepEqual(arrayDiff(["apple", "banana", "cherry"], ["apple", "banana"]), ["cherry"]);
```

`arrayDiff(["one", "two", "three", "four", "six"], ["one", "three", "eight"])` debería devolver `["eight", "four", "six", "two"]`.

```js
assert.deepEqual(arrayDiff(["one", "two", "three", "four", "six"], ["one", "three", "eight"]), ["eight", "four", "six", "two"]);
```

`arrayDiff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"])` debería devolver `["five", "one", "seven", "three"]`.

```js
assert.deepEqual(arrayDiff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"]), ["five", "one", "seven", "three"]);
```

`arrayDiff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"])` debería devolver `["freeCodeCamp", "rocks"]`.

```js
assert.deepEqual(arrayDiff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"]), ["freeCodeCamp", "rocks"]);
```

# --seed--

## --seed-contents--

```js
function arrayDiff(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function arrayDiff(arr1, arr2) {
  let unique1 = arr1.filter((v, i) => arr1.indexOf(v) === i);
  let unique2 = arr2.filter((v, i) => arr2.indexOf(v) === i);

  let only1 = unique1.filter(v => !unique2.includes(v));
  let only2 = unique2.filter(v => !unique1.includes(v));

  return [...only1, ...only2].sort();
}
```
