---
id: 68af0687ef34c76c28ffa54b
title: "Desafío 32: Oración Invertida"
challengeType: 28
dashedName: challenge-32
---

# --description--

Dada una cadena de palabras, devuelve una nueva cadena con las palabras en orden inverso. Por ejemplo, la primera palabra debe estar al final de la cadena devuelta, y la última palabra debe estar al principio de la cadena devuelta.

- En la cadena dada, las palabras pueden estar separadas por uno o más espacios.
- La cadena devuelta debe tener solo un espacio entre palabras.

# --hints--

`reverseSentence("world hello")` debería devolver `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` debería devolver `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` debería devolver `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` debería devolver `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
