---
id: 68b06e589bf2273243814775
title: "Desafío 38: Generador de Slugs"
challengeType: 28
dashedName: challenge-38
---

# --description--

Dada una cadena, devuelve una versión amigable para URL de la cadena usando las siguientes restricciones:

- Todas las letras deben estar en minúsculas.
- Todos los caracteres que no sean letras, números o espacios deben ser eliminados.
- Todos los espacios deben ser reemplazados con el código de espacio codificado para URL `%20`.
- Los espacios consecutivos deben reemplazarse con un solo `%20`.
- La cadena devuelta no debe tener `%20` al principio ni al final.

# --hints--

`generateSlug("helloWorld")` debería devolver `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` debería devolver `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` debería devolver `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` debería devolver `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` debería devolver `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
