---
id: 68b06e589bf227324381476f
title: "Desafío 35: Frecuencia de Palabras"
challengeType: 28
dashedName: challenge-35
---

# --description--

Dado un párrafo, devuelve un arreglo con las tres palabras que ocurren con mayor frecuencia.

- Las palabras en el párrafo estarán separadas por espacios.
- Ignora mayúsculas y minúsculas en el párrafo dado. Por ejemplo, trata `Hello` y `hello` como la misma palabra.
- Ignora la puntuación en el párrafo dado. La puntuación consiste en comas (`,`), puntos (`.`) y signos de exclamación (`!`).
- El array devuelto debe tener todas las palabras en minúsculas.
- El array devuelto debe estar en orden descendente con la palabra que ocurre con más frecuencia primero.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` debería devolver `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` debería devolver `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` debería devolver `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
