---
id: 68b06e589bf2273243814771
title: "Desafío 36: Ajustador de Termostato"
challengeType: 28
dashedName: challenge-36
---

# --description--

Dada la temperatura actual de una habitación y una temperatura objetivo, devuelve una cadena que indique cómo ajustar la temperatura de la habitación según estas restricciones:

- Devuelve `"heat"` si la temperatura actual está por debajo del objetivo.
- Devuelve `"cool"` si la temperatura actual está por encima del objetivo.
- Devuelve `"hold"` si la temperatura actual es igual a la objetivo.

# --hints--

`adjustThermostat(68, 72)` debería devolver `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` debería devolver `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` debería devolver `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` debería devolver `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` debería devolver `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` debería devolver `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
