---
id: 68b06e589bf2273243814773
title: "Desafío 37: Capitalizador de Oraciones"
challengeType: 28
dashedName: challenge-37
---

# --description--

Dado un párrafo, devuelve un nuevo párrafo donde la primera letra de cada oración esté en mayúscula.

- Todos los demás caracteres deben conservarse.
- Las oraciones pueden terminar con un punto (`.`), uno o más signos de interrogación (`?`), o uno o más signos de exclamación (`!`).

# --hints--

`capitalize("this is a simple sentence.")` debería devolver `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` debería devolver `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` debería devolver `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` debería devolver `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` debería devolver `"There's a space before this period . Why is there a space before that period ?"` .

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
