---
id: 68b06e589bf2273243814777
title: "Desafío 39: Llena el tanque"
challengeType: 28
dashedName: challenge-39
---

# --description--

Dado el tamaño de un tanque de combustible, el nivel actual de combustible y el precio por galón, devuelve el costo para llenar el tanque completamente.

- `tankSize` es la capacidad total del tanque en galones.
- `fuelLevel` es la cantidad actual de combustible en el tanque en galones.
- `pricePerGallon` es el costo de un galón de combustible.
- El valor devuelto debe redondearse a dos decimales en el formato: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` debería devolver `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` debería devolver `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` debería devolver `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` debería devolver `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` debería devolver `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
