---
id: 68b1f72371a5ac895ac70a0a
title: "Desafío 44: Espejo de Cadenas"
challengeType: 28
dashedName: challenge-44
---

# --description--

Dadas dos cadenas, determina si la segunda cadena es un espejo de la primera.

- Una cadena se considera un espejo si contiene las mismas letras en orden inverso.
- Trata las letras mayúsculas y minúsculas como distintas.
- Ignora todos los caracteres que no sean alfabéticos.

# --hints--

`isMirror("helloworld", "helloworld")` debería devolver `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` debería devolver `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` debería devolver `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` debería devolver `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` debería devolver `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` debería devolver `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` debería devolver `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
