---
id: 68b1f72371a5ac895ac70a08
title: "Desafío 43: Dígitos vs Letras"
challengeType: 28
dashedName: challenge-43
---

# --description--

Dada una cadena, devuelve `"digits"` si la cadena tiene más dígitos que letras, `"letters"` si tiene más letras que dígitos, y `"tie"` si tiene la misma cantidad de dígitos y letras.

- Los dígitos consisten en `0-9`.
- Las letras consisten en `a-z` en mayúsculas o minúsculas.
- Ignora cualquier otro carácter.

# --hints--

`digitsOrLetters("abc123")` debería devolver `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` debería devolver `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` debería devolver `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` debería devolver `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` debería devolver `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` debería devolver `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
