---
id: 68b1f72371a5ac895ac70a04
title: "Desafío 41: Almacenamiento de Archivos"
challengeType: 28
dashedName: challenge-41
---

# --description--

Dado un tamaño de archivo, una unidad para el tamaño del archivo y la capacidad del disco duro en gigabytes (GB), devuelve el número de archivos que el disco duro puede almacenar usando las siguientes restricciones:

- La unidad para el tamaño del archivo puede ser bytes (`"B"`), kilobytes (`"KB"`), o megabytes (`"MB"`).
- Devuelve el número de archivos completos que el disco puede contener.
- Usa las siguientes conversiones:

| Unit | Equivalent |
|:----:|:----------:|
| 1 B  |   1 B      |
| 1 KB |   1000 B   |
| 1 MB |   1000 KB  |
| 1 GB |   1000 MB  |

Por ejemplo, dado `500`, `"KB"` y `1` como argumentos, determina cuántos archivos de 500 KB pueden caber en un disco duro de 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` debería devolver `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` debería devolver `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` debería devolver `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` debería devolver `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` debería devolver `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` debería devolver `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
