---
id: 68b1f72371a5ac895ac70a02
title: "Desafío 40: Almacenamiento de Fotos"
challengeType: 28
dashedName: challenge-40
---

# --description--

Dado un tamaño de foto en megabytes (MB) y la capacidad del disco duro en gigabytes (GB), devuelve el número de fotos que el disco duro puede almacenar usando las siguientes restricciones:

- 1 gigabyte equivale a 1000 megabytes.
- Devuelve el número de fotos completas que puede almacenar el disco.

# --hints--

`numberOfPhotos(1, 1)` debería devolver `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` debería devolver `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` debería devolver `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` debería devolver `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` debería devolver `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
