---
id: 68b7cadffed0e75a517da675
title: "Desafío 53: Decimal a Binario"
challengeType: 28
dashedName: challenge-53
---

# --description--

Dado un entero no negativo, devuelve su representación binaria como una cadena.

Un número binario usa solo los dígitos `0` y `1` para representar cualquier número. Para convertir un número decimal a binario, divide repetidamente el número entre `2` y registra el residuo. Repite hasta que el número sea cero. Lee los residuos desde el último registrado hasta el primero. Por ejemplo, para convertir `12` a binario:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` en binario es `1100`.

# --hints--

`toBinary(5)` debería devolver `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` debería devolver `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` debería devolver `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` debería devolver `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
