---
id: 68b7cadffed0e75a517da671
title: "Desafío 51: Formateador de Números Telefónicos"
challengeType: 28
dashedName: challenge-51
---

# --description--

Dada una cadena de once dígitos, devuelve la cadena como un número de teléfono en este formato: `"+D (DDD) DDD-DDDD"`.

# --hints--

`formatNumber("05552340182")` debería devolver `"+0 (555) 234-0182"`.

```js
assert.equal(formatNumber("05552340182"), "+0 (555) 234-0182");
```

`formatNumber("15554354792")` debería devolver `"+1 (555) 435-4792"`.

```js
assert.equal(formatNumber("15554354792"), "+1 (555) 435-4792");
```

# --seed--

## --seed-contents--

```js
function formatNumber(number) {

  return number;
}
```

# --solutions--

```js
function formatNumber(number) {
  const country = number[0];
  const area = number.slice(1, 4);
  const prefix = number.slice(4, 7);
  const line = number.slice(7);

  return `+${country} (${area}) ${prefix}-${line}`;
}
```
