---
id: 68b7cadffed0e75a517da677
title: "Desafío 54: ¡P@ssw0rd Str3ngth!"
challengeType: 28
dashedName: challenge-54
---

# --description--

Dada una cadena de contraseña, devuelve `"weak"`, `"medium"` o `"strong"` según la fortaleza de la contraseña.

Una contraseña se evalúa según las siguientes reglas:

- Tiene al menos 8 caracteres de longitud.
- Contiene tanto letras mayúsculas como minúsculas.
- Contiene al menos un número.
- Contiene al menos un carácter especial de este conjunto: `!`, `@`, `#`, `$`, `%`, `^`, `&` o `*`.

Devuelve `"weak"` si la contraseña cumple con menos de dos de las reglas.  
Devuelve `"medium"` si la contraseña cumple con 2 o 3 de las reglas.  
Devuelve `"strong"` si la contraseña cumple con las 4 reglas.

# --hints--

`checkStrength("123456")` debería devolver `"weak"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` debería devolver `"weak"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` debería devolver `"weak"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` debería devolver `"weak"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` debería devolver `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` debería devolver `"medium"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` debería devolver `"medium"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` debería devolver `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` debería devolver `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
