---
id: 68b7687dded630607aceccb3
title: "Desafío 49: Analizador de encabezados CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

Dada la primera línea de un archivo de valores separados por comas (CSV), devuelve un arreglo que contiene los encabezados.

- La primera línea de un archivo CSV contiene encabezados separados por comas.
- Elimina cualquier espacio en blanco al principio o al final de cada encabezado.

# --hints--

`getHeadings("name,age,city")` debería devolver `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` debería devolver `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` debería devolver `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
