---
id: 68b7687dded630607aceccb1
title: "Desafío 48: Detector de Spam"
challengeType: 28
dashedName: challenge-48
---

# --description--

Dado un número de teléfono en el formato `"+A (BBB) CCC-DDDD"`, donde cada letra representa un dígito de la siguiente manera:

- `A` representa el código del país y puede ser cualquier cantidad de dígitos.
- `BBB` representa el código de área y siempre tendrá tres dígitos.
- `CCC` y `DDDD` representan el número local y siempre tendrán tres y cuatro dígitos, respectivamente.

Determina si es un número de spam basado en los siguientes criterios:

- El código del país tiene más de 2 dígitos o no comienza con un cero (`0`).
- El código de área es mayor que 900 o menor que 200.
- La suma de los primeros tres dígitos del número local aparece dentro de los últimos cuatro dígitos del número local.
- El número tiene el mismo dígito cuatro o más veces seguidas (ignorando los caracteres de formato).

# --hints--

`isSpam("+0 (200) 234-0182")` debería devolver `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` debería devolver `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` debería devolver `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` debería devolver `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` debería devolver `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` debería devolver `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` debería devolver `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` debería devolver `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
