---
id: 68b7cadffed0e75a517da673
title: "Desafío 52: De binario a decimal"
challengeType: 28
dashedName: challenge-52
---

# --description--

Dada una cadena que representa un número binario, devuelve su equivalente decimal como un número.

Un número binario usa solo los dígitos `0` y `1` para representar cualquier número. Para convertir binario a decimal, multiplica cada dígito por una potencia de `2` y súmalos. Comienza multiplicando el dígito más a la derecha por `2^0`, el siguiente dígito a la izquierda por `2^1`, y así sucesivamente. Una vez que todos los dígitos hayan sido multiplicados por una potencia de `2`, suma el resultado.

Por ejemplo, el número binario `101` equivale a `5` en decimal porque:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` debería devolver `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` debería devolver `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` debería devolver `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` debería devolver `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
