---
id: 68b7687dded630607aceccab
title: "Desafío 45: Cuadrado Perfecto"
challengeType: 28
dashedName: challenge-45
---

# --description--

Dado un entero, determina si es un cuadrado perfecto.

- Un número es un cuadrado perfecto si puedes multiplicar un entero por sí mismo para obtener el número. Por ejemplo, 9 es un cuadrado perfecto porque puedes multiplicar 3 por sí mismo para obtenerlo.

# --hints--

`isPerfectSquare(9)` debería devolver `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` debería devolver `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` debería devolver `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` debería devolver `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` debería devolver `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` debería devolver `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` debería devolver `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` debería devolver `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
