---
id: 68c1a929005bf54d342aa8d4
title: "Desafío 57: Semana Espacial Día 3: Phone Home"
challengeType: 28
dashedName: challenge-57
---

# --description--

Para el tercer día de la Semana Espacial, se te da un arreglo de números que representan las distancias (en kilómetros) entre tú, los satélites y tu planeta de origen en una ruta de comunicación. Determina cuánto tiempo tomará un mensaje enviado a través de la ruta para llegar a su planeta destino usando las siguientes restricciones:

- El primer valor en el array es la distancia desde tu ubicación hasta el primer satélite.
- Cada valor subsecuente, excepto el último, es la distancia al siguiente satélite.
- El último valor en el array es la distancia desde el satélite anterior hasta tu planeta natal.
- El mensaje viaja a 300,000 km/s.
- Cada satélite por el que **pasa** el mensaje añade un retraso de transmisión de 0.5 segundos.
- Devuelve un número redondeado a 4 decimales, con ceros finales eliminados.

# --hints--

`sendMessage([300000, 300000])` debería devolver `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` debería devolver `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` debería devolver `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` debería devolver `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` debería devolver `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` debería devolver `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
