---
id: 68c1a929005bf54d342aa8d2
title: "Desafío 55: Semana Espacial Día 1: Clasificación Estelar"
challengeType: 28
dashedName: challenge-55
---

# --description--

El 4 de octubre marca el inicio de la Semana Mundial del Espacio. Los próximos siete días te traerán desafíos de programación con temática de astronomía.

Para el desafío de hoy, se te da la temperatura superficial de una estrella en Kelvin (K) y necesitas determinar su clasificación estelar basada en los siguientes rangos:

- "O": 30,000 K o más
- "B": 10,000 K - 29,999 K
- "A": 7,500 K - 9,999 K
- "F": 6,000 K - 7,499 K
- "G": 5,200 K - 5,999 K
- "K": 3,700 K - 5,199 K
- "M": 0 K - 3,699 K

- Devuelve la clasificación de la estrella dada.

# --hints--

`classification(5778)` debería devolver `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` debería devolver `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` debería devolver `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` debería devolver `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` debería devolver `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` debería devolver `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` debería devolver `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` debería devolver `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
