---
id: 68c1a929005bf54d342aa8d3
title: "Desafío 56: Semana del Espacio Día 2: Búsqueda de Exoplanetas"
challengeType: 28
dashedName: challenge-56
---

# --description--

Para el segundo día de la Semana del Espacio, se te da una cadena donde cada carácter representa la lectura de luminosidad de una estrella. Determina si las lecturas han detectado un exoplaneta usando el método de tránsito. El método de tránsito es cuando un planeta pasa frente a una estrella, reduciendo su luminosidad observada.

- Las lecturas de luminosidad solo comprenden los caracteres `0-9` y `A-Z`, donde cada lectura corresponde a los siguientes valores numéricos:
- Los caracteres `0-9` corresponden a niveles de luminosidad `0-9`.
- Los caracteres `A-Z` corresponden a niveles de luminosidad `10-35`.

Se considera que una estrella tiene un exoplaneta si alguna lectura individual es menor o igual al 80% del promedio de todas las lecturas. Por ejemplo, si la luminosidad promedio de una estrella es 10, se consideraría que tiene un exoplaneta si alguna lectura individual es 8 o menos.

# --hints--

`hasExoplanet("665544554")` debería devolver `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` debería devolver `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` debería devolver `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` debería devolver `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` debería devolver `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` debería devolver `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
