---
id: 68c497f3aaefc9fd9f1b0e24
title: "Desafío 60: Semana del Espacio Día 6: Fase Lunar"
challengeType: 28
dashedName: challenge-60
---

# --description--

Para el día seis de la Semana Espacial, se te dará una fecha en el formato `"YYYY-MM-DD"` y necesitarás determinar la fase de la luna para ese día usando las siguientes reglas:

Usa un ciclo lunar simplificado de 28 días, dividido en cuatro fases iguales:

- "Nuevo": días 1 - 7
- "Luna creciente": días 8 - 14
- "Full": días 15 - 21
- "Waning": días 22 - 28

Después del día 28, el ciclo se repite con el día 1, una luna nueva.

- Usa `"2000-01-06"` como referencia de luna nueva (día 1 del ciclo) para determinar la fase del día dado.
- No se te proporcionarán fechas anteriores a la fecha de referencia.
- Devuelve la fase correcta como una cadena.

**Nota:** El Día 1 representa el día de la luna nueva, lo que significa que han pasado 0 días desde la última luna nueva.

# --hints--

`moonPhase("2000-01-12")` debería devolver `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` debería devolver `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` debería devolver `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` debería devolver `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` debería devolver `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
