---
id: 68c497f3aaefc9fd9f1b0e26
title: "Desafío 62: Hexadecimal a Decimal"
challengeType: 28
dashedName: challenge-62
---

# --description--

Dada una cadena que representa un número hexadecimal (base 16), devuelve su valor decimal (base 10) como un entero.

Hexadecimal es un sistema numérico que usa 16 dígitos:

- `0-9` representan valores de `0` a `9`.
- `A-F` representan valores del `10` al `15`.

Aquí tienes una tabla de conversión parcial:

| Hexadecimal | Decimal |
|:-----------:|:-------:|
| 0           | 0       |
| 1           | 1       |
| ...         | ...     |
| 9           | 9       |
| A           | 10      |
| ...         | ...     |
| F           | 15      |
| 10          | 16      |
| ...         | ...     |
| 9F          | 159     |
| A0          | 160     |
| ...         | ...     |
| FF          | 255     |
| 100         | 256     |

- La cadena solo contendrá caracteres `0–9` y `A–F`.

# --hints--

`hexToDecimal("A")` debería devolver `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` debería devolver `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` debería devolver `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` debería devolver `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` debería devolver `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
