---
id: 68cae5b538ff798bbd4da00a
title: "Desafío 72: Ajustador de Termostato 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Dada la temperatura actual de una habitación en Fahrenheit y una temperatura objetivo en Celsius, devuelve una cadena que indique cómo ajustar la temperatura de la habitación según estas restricciones:

- Devuelve `"Heat: X degrees Fahrenheit"` si la temperatura actual está por debajo del objetivo. Con `X` siendo el número de grados en Fahrenheit para calentar la habitación y alcanzar el objetivo, redondeado a 1 decimal.
- Devuelve `"Cool: X degrees Fahrenheit"` si la temperatura actual está por encima del objetivo. Con `X` siendo el número de grados en Fahrenheit para enfriar la habitación y alcanzar el objetivo, redondeado a 1 decimal.
- Devuelve `"Hold"` si la temperatura actual es igual a la objetivo.

Para convertir Celsius a Fahrenheit, multiplica la temperatura en Celsius por 1.8 y suma 32 al resultado (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` debería devolver `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` debería devolver `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` debería devolver `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` debería devolver `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` debería devolver `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
