---
id: 68cae5b538ff798bbd4da007
title: "Desafío 69: Calcetines Perdidos"
challengeType: 28
dashedName: challenge-69
---

# --description--

Dado un entero que representa la cantidad de pares de calcetines con los que comenzaste, y otro entero que representa cuántos ciclos de lavado has realizado, devuelve el número de pares completos de calcetines que tienes actualmente usando las siguientes restricciones:

- Cada 2 ciclos de lavado, pierdes un solo calcetín.
- Cada 3 ciclos de lavado, encuentras un calcetín perdido.
- Cada 5 ciclos de lavado, un solo calcetín se desgasta y debe ser desechado.
- Cada 10 ciclos de lavado, compras un par de calcetines.
- Nunca puedes tener menos de cero calcetines en total.
- Las reglas pueden superponerse. Por ejemplo, en el ciclo de lavado 10, perderás un solo calcetín, tirarás un solo calcetín y comprarás un nuevo par de calcetines.
- Devuelve el número de pares completos de calcetines.

# --hints--

`sockPairs(2, 5)` debería devolver `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` debería devolver `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` debería devolver `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` debería devolver `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` debería devolver `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
