---
id: 68cae5b538ff798bbd4da009
title: "Desafío 71: Calculadora de propinas"
challengeType: 28
dashedName: challenge-71
---

# --description--

Dado el precio de tu comida y un porcentaje de propina personalizado, devuelve un arreglo con tres valores de propina; 15%, 20% y la cantidad personalizada.

- Los precios se darán en el formato: `"$N.NN"`.
- Los porcentajes personalizados de propina se darán en este formato: `"25%"`.
- Devuelve los montos en el mismo formato `"$N.NN"`, redondeados a dos decimales.

Por ejemplo, dado un precio de comida de `"$10.00"` y un valor de propina personalizada de `"25%"`, devuelve `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` debería devolver `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` debería devolver `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` debería devolver `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
