---
id: 68cae5b538ff798bbd4da002
title: "Desafío 64: de 24 a 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

Dada una cadena que representa una hora del día en el formato de 24 horas `"HHMM"`, devuelve la hora en su equivalente formato de 12 horas `"H:MM AM"` o `"H:MM PM"`.

- La entrada dada siempre será una cadena de cuatro dígitos en formato de hora de 24 horas, desde `"0000"` hasta `"2359"`.

# --hints--

`to12("1124")` debería devolver `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` debería devolver `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` debería devolver `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` debería devolver `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` debería devolver `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
