---
id: 68cae5b538ff798bbd4da005
title: "Desafío 67: Validador de Email"
challengeType: 28
dashedName: challenge-67
---

# --description--

Dada una cadena, determina si es una dirección de correo electrónico válida usando las siguientes restricciones:

- Debe contener exactamente un símbolo `@`.
- The local part (before the `@`):
  - Solo puede contener letras (`a-z`, `A-Z`), dígitos (`0-9`), puntos (`.`), guiones bajos (`_`) o guiones (`-`).
  - No puede comenzar ni terminar con un punto.
- The domain part (after the `@`):
  - Debe contener al menos un punto.
  - Debe terminar con un punto seguido de al menos dos letras.
- Ni la parte local ni la parte de dominio pueden tener dos puntos seguidos.

# --hints--

`validate("a@b.cd")` debería devolver `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` debería devolver `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` debería devolver `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` debería devolver `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` debería devolver `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` debería devolver `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` debería devolver `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` debería devolver `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` debería devolver `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` debería devolver `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
