---
id: 68cae5b538ff798bbd4da004
title: "Desafío 66: Removedor de Etiquetas HTML"
challengeType: 28
dashedName: challenge-66
---

# --description--

Dado un string de código HTML, elimina las etiquetas y devuelve el contenido de texto plano.

- La cadena de entrada contendrá solo HTML válido.
- Las etiquetas HTML pueden estar anidadas.
- Elimina las etiquetas y cualquier atributo.

Por ejemplo, `'<a href="#">Click here</a>'` debería devolver `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` debería devolver `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` debería devolver `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` debería devolver una cadena vacía (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` debería devolver `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
