---
id: 68cae5b538ff798bbd4da001
title: "Desafío 63: Batalla de Palabras"
challengeType: 28
dashedName: challenge-63
---

# --description--

Dadas dos oraciones que representan a tu equipo y a un equipo contrario, donde cada palabra de tu equipo batalla contra la palabra correspondiente del equipo contrario, determina qué equipo gana usando las siguientes reglas:

- Las oraciones dadas siempre contendrán el mismo número de palabras.
- Las palabras están separadas por un solo espacio y solo contendrán letras.
- El valor de cada palabra es la suma de sus letras.
- Las letras `a` a `z` corresponden a los valores del `1` al `26`. Por ejemplo, `a` es `1` y `z` es `26`.
- Una letra mayúscula duplica el valor de la letra. Por ejemplo, `A` es `2`, y `Z` es `52`.
- Las palabras luchan en orden: la primera palabra de tu equipo lucha contra la primera palabra del equipo contrario, y así sucesivamente.
- Una palabra gana si su valor es mayor que el valor de la palabra opuesta.
- El equipo con más palabras ganadoras es el ganador.

Devuelve `"We win"` si tu equipo es el ganador, `"We lose"` si tu equipo pierde, y `"Draw"` si ambos equipos tienen el mismo número de victorias.

# --hints--

`battle("hello world", "hello word")` debería devolver `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` debería devolver `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` debería devolver `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` debería devolver `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` debería devolver `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` debería devolver `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` debería devolver `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
