---
id: 68cae5b538ff798bbd4da008
title: "Desafío 70: Extractor de Atributos HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

Dada una cadena de un elemento HTML válido, devuelve los atributos del elemento usando los siguientes criterios:

- Solo se te dará un elemento.
- Los atributos estarán en el formato: `attribute="value"`.
- Devuelve un arreglo de cadenas con cada propiedad y valor del atributo, separados por una coma, en este formato: `["attribute1, value1", "attribute2, value2"]`.
- Devuelve los atributos en el orden en que se proporcionan.
- Si no se encuentran atributos, devuelve un arreglo vacío.

# --hints--

`extractAttributes('<span class="red"></span>')` debería devolver `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` debería devolver `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` debería devolver `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` debería devolver `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` debería devolver `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
