---
id: 68cae5b538ff798bbd4da003
title: "Desafío 65: Conteo de Cadenas"
challengeType: 28
dashedName: challenge-65
---

# --description--

Dadas dos cadenas, determina cuántas veces aparece la segunda cadena en la primera.

- La cadena patrón puede superponerse en la primera cadena. Por ejemplo, `"aaa"` contiene `"aa"` dos veces. Las primeras dos `a` y las segundas dos.

# --hints--

`count('abcdefg', 'def')` debería devolver `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` debería devolver `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` debería devolver `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` debería devolver `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` debería devolver `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
