---
id: 68d2ba1468508398389487ce
title: "Desafío 73: Habla sabiamente, debes"
challengeType: 28
dashedName: challenge-73
---

# --description--

Dada una oración, devuelve una versión de ella que suene como un consejo de un maestro sabio usando las siguientes reglas:

- Las palabras están separadas por un solo espacio.
- Encuentra la primera ocurrencia de una de las siguientes palabras en la oración: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Mueve todas las palabras antes y que incluyen esa palabra al final de la oración y:
  - Conserva el orden de las palabras cuando las muevas.
  - Ponlos todos en minúsculas.
  - Y añade una coma y un espacio antes de ellos.
- Pon en mayúscula la primera letra de la nueva primera palabra de la oración.
- Todas las oraciones dadas terminarán con un solo signo de puntuación. Mantén la puntuación original al final de la nueva oración.
- Devuelve la nueva oración, asegúrate de que haya un solo espacio entre cada palabra y no haya espacios al principio ni al final de la oración.

Por ejemplo, dado `"You must speak wisely."` devuelve `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` debería devolver `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` debería devolver `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` debería devolver `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` debería devolver `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` debería devolver `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
