---
id: 68d2ba1468508398389487d0
title: "Desafío 75: Tesoro Oculto"
challengeType: 28
dashedName: challenge-75
---

# --description--

Dado un arreglo 2D que representa un mapa del fondo oceánico que incluye un tesoro oculto, y un arreglo con las coordenadas (`[row, column]`) para la próxima inmersión de tu búsqueda del tesoro, devuelve `"Empty"`, `"Found"` o `"Recovered"` usando las siguientes reglas:

- El arreglo 2D dado contendrá exactamente un tesoro no recuperado, que ocupará múltiples celdas.
- Each cell in the 2D array will contain one of the following values:
  - `"-"`: Sin tesoro.
  - "O": Una parte del tesoro que no ha sido encontrada.
  - `"X"`: Una parte del tesoro que ya ha sido encontrada.

- Si la ubicación de la inmersión no tiene tesoro, devuelve `"Empty"`.
- Si la ubicación de la inmersión encuentra tesoro, pero al menos otra parte del tesoro permanece sin encontrar, devuelve `"Found"`.
- Si la ubicación de la inmersión encuentra la última parte del tesoro que no se había encontrado, devuelve `"Recovered"`.

Por ejemplo, dado:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

Y `[2, 1]` para las coordenadas del lugar de la inmersión, devuelve `"Recovered"` porque la inmersión encontró la última parte del tesoro que no se había encontrado.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` debería devolver `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` debería devolver `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` debería devolver `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` debería devolver `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` debería devolver `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` debería devolver `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
