---
id: 68d2ba1468508398389487cf
title: "Desafío 74: Canciones Favoritas"
challengeType: 28
dashedName: challenge-74
---

# --description--

¿Recuerdas los iPods? El primer modelo salió hace 24 años hoy, el 23 de octubre de 2001.

Dado un arreglo de objetos de canciones que representan la lista de reproducción de tu iPod, devuelve un arreglo con los títulos de las dos canciones más reproducidas, con la canción más reproducida primero.

- Cada objeto tendrá una propiedad `"title"` (cadena) y una propiedad `"plays"` (entero).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` debería devolver `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` debería devolver `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` debería devolver `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```
