---
id: 68d30845cc08266018fc46bd
title: "Desafío 77: Formateador de Duración"
challengeType: 28
dashedName: challenge-77
---

# --description--

Dado un número entero de segundos, devuelve una cadena que represente la misma duración en el formato `"H:MM:SS"`, donde `"H"` es el número de horas, `"MM"` es el número de minutos y `"SS"` es el número de segundos. Devuelve el tiempo usando las siguientes reglas:

- Segundos: Siempre deben ser dos dígitos.
- Minutos: Deben omitirse los ceros a la izquierda cuando no sean necesarios. Usa `"0"` si la duración es menor a un minuto.
- Horas: Deben incluirse solo si son mayores que cero.

# --hints--

`format(500)` debería devolver `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` debería devolver `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` debería devolver `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` debería devolver `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` debería devolver `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
