---
id: 68d30845cc08266018fc46bc
title: "Desafío 76: ADN complementario"
challengeType: 28
dashedName: challenge-76
---

# --description--

Dada una cadena que representa una secuencia de ADN, devuelve su hebra complementaria usando las siguientes reglas:

- El ADN consiste en las letras `"A"`, `"C"`, `"G"` y `"T"`.
- Las letras `"A"` y `"T"` se complementan entre sí.
- Las letras `"C"` y `"G"` se complementan entre sí.

Por ejemplo, dado `"ACGT"`, devuelve `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` debería devolver `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` debería devolver `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` debería devolver `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` debería devolver `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
