---
id: 68d30fc57588d97fd3027b30
title: "Desafío 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

El 28 de octubre de 1994, se lanzó Netscape Navigator, ayudando a millones a explorar la web temprana.

Dado un arreglo de comandos del navegador que ejecutaste en Netscape Navigator, devuelve la página actual en la que estás después de ejecutar todos los comandos usando las siguientes reglas:

- Siempre comienzas en la página `"Home"`, que no se incluirá en el arreglo de comandos.
- Los comandos válidos son:
  - "Visitar Page": Donde "Page" es el nombre de la página que estás visitando. Por ejemplo, "Visitar About" te lleva a la página "About". Cuando visites una nueva página, asegúrate de descartar cualquier historial hacia adelante que tengas.
  - "`Back`": Te lleva a la página anterior en tu historial o permanece en la página actual si no hay ninguna.
  - "Forward": Te lleva hacia adelante en el historial a la página de la que viniste o permanece en la página actual si no hay ninguna.

Por ejemplo, dado `["Visit About Us", "Back", "Forward"]`, devuelve `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` debería devolver `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` debería devolver `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` debería devolver `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` debería devolver `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` debería devolver `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` debería devolver `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` debería devolver `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
