---
id: 68d30845cc08266018fc46be
title: "Desafío 78: Secuencia de enteros"
challengeType: 28
dashedName: challenge-78
---

# --description--

Dado un entero positivo, devuelve una cadena con todos los enteros desde `1` hasta, e incluyendo, el número dado, en orden numérico.

Por ejemplo, dado `5`, devuelve `"12345"`.

# --hints--

`sequence(5)` debería devolver `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` debería devolver `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` debería devolver `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` debería devolver `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
