---
id: 68e39ed6106dac2f0a98fd65
title: "Desafío 83: Validación de Firma"
challengeType: 28
dashedName: challenge-83
---

# --description--

Dado un mensaje en forma de cadena, una clave secreta en forma de cadena y un número de firma, determina si la firma es válida usando este método de codificación:

- Las letras en el mensaje y la clave secreta tienen estos valores:
  - `a` a `z` tienen valores de `1` a `26` respectivamente.
  - `A` a `Z` tienen valores de `27` a `52` respectivamente.
- Todos los demás caracteres no tienen valor.
- Calcula la firma sumando el mensaje más la suma de la clave secreta.

Por ejemplo, dado el mensaje `"foo"` y la clave secreta `"bar"`, la firma sería `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Verifica si la firma calculada coincide con la firma proporcionada.

# --hints--

`verify("foo", "bar", 57)` debería devolver `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` debería devolver `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` debería devolver `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` debería devolver `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` debería devolver `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` debería devolver `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
