---
id: 68e39ed6106dac2f0a98fd64
title: "Desafío 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

Dada una cadena que representa un nombre de variable, conviértela a "spooky case" usando las siguientes restricciones:

- Reemplaza todos los guiones bajos (`_`) y guiones (`-`) con una virgulilla (`~`).
- Pon en mayúscula la primera letra de la cadena, y cada otra letra después de esa. Ignora el carácter tilde al contar. Pon todas las demás letras en minúscula.

Por ejemplo, dado `hello_world`, devuelve `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` debería devolver `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` debería devolver `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` debería devolver `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` debería devolver `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` debería devolver `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```
