---
id: 68e39ed6106dac2f0a98fd63
title: "Desafío 81: Nth Prime"
challengeType: 28
dashedName: challenge-81
---

# --description--

Un número primo es un entero positivo mayor que 1 que es divisible solo por 1 y por sí mismo. Los primeros cinco números primos son `2`, `3`, `5`, `7` y `11`.

Dado un entero positivo `n`, devuelve el número primo `n`-ésimo. Por ejemplo, dado `5` devuelve el 5º número primo: `11`.

# --hints--

`nthPrime(5)` debería devolver `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` debería devolver `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` debería devolver `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` debería devolver `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` debería devolver `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
