---
id: 68ee9e3066cfd4eb2328e8a6
title: "Desafío 87: Constructor de Matrices"
challengeType: 28
dashedName: challenge-87
---

# --description--

Dado dos enteros (un número de filas y un número de columnas), devuelve una matriz (un arreglo de arreglos) llena de ceros (`0`) del tamaño indicado.

Por ejemplo, dado `2` y `3`, devuelve:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` debería devolver `[[0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` debería devolver `[[0, 0], [0, 0], [0, 0]]`.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` debería devolver `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` debería devolver `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```
