---
id: 68ee9e3066cfd4eb2328e8a4
title: "Desafío 85: Contador de Palabras"
challengeType: 28
dashedName: challenge-85
---

# --description--

Dada una cadena de oración, devuelve el número de palabras que hay en la oración.

- Las palabras son cualquier secuencia de caracteres que no sean espacios y están separadas por un solo espacio.

# --hints--

`countWords("Hello world")` debería devolver `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` debería devolver `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` debería devolver `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` debería devolver `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` debería devolver `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
