---
id: 68f6587287ad1f4ad39b0c80
title: "Desafío 94: Generador de Firma de Correo Electrónico"
challengeType: 28
dashedName: challenge-94
---

# --description--

Dadas cadenas para el nombre de una persona, título y empresa, devuelve una firma de correo electrónico como una sola cadena usando las siguientes reglas:

- El nombre debe aparecer primero, precedido por un prefijo que depende de la primera letra del nombre. Para nombres que comienzan con (sin distinguir mayúsculas o minúsculas):
  - `A-I`: Usa `>>` como prefijo.
  - `J-R`: Usa `--` como prefijo.
  - `S-Z`: Usa `::` como prefijo.
- Una coma y un espacio (`, `) deben seguir al nombre.
- El título y la empresa deben seguir después de la coma y el espacio, separados por `" at "` (con espacios alrededor).

Por ejemplo, dado `"Quinn Waverly"`, `"Founder and CEO"` y `"TechCo"` devuelve `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` debería devolver `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` debería devolver `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` debería devolver `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` debería devolver `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` debería devolver `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
