---
id: 68f6587287ad1f4ad39b0c84
title: "Desafío 98: Conteo de Rectángulos"
challengeType: 28
dashedName: challenge-98
---

# --description--

Dado dos enteros positivos que representan el ancho y la altura de un rectángulo, determina cuántos rectángulos pueden caber dentro del dado.

- Cuenta solo los rectángulos con ancho y alto enteros.

Por ejemplo, dado `1` y `3`, devuelve `6`. Tres rectángulos de 1x1, dos rectángulos de 1x2 y un rectángulo de 1x3.

# --hints--

`countRectangles(1, 3)` debería devolver `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` debería devolver `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` debería devolver `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` debería devolver `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` debería devolver `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
