---
id: 68f6587287ad1f4ad39b0c82
title: "Desafío 96: ¿Es fin de semana?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Dada una fecha en el formato `"YYYY-MM-DD"`, devuelve el número de días que faltan hasta el fin de semana.

- El fin de semana comienza el sábado.
- Si la fecha dada es sábado o domingo, devuelve `"It's the weekend!"`.
- De lo contrario, devuelve `"X days until the weekend."`, donde `X` es el número de días hasta el sábado.
- Si `X` es `1`, usa `"day"` (singular) en lugar de `"days"` (plural).
- Asegúrate de que el cálculo ignore tu zona horaria local.

# --hints--

`daysUntilWeekend("2025-11-14")` debería devolver `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` debería devolver `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` debería devolver `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` debería devolver `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` debería devolver `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` debería devolver `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
