---
id: 68f6587287ad1f4ad39b0c7e
title: "Desafío 92: Extractor de Extensiones"
challengeType: 28
dashedName: challenge-92
---

# --description--

Dada una cadena que representa un nombre de archivo, devuelve la extensión del archivo.

- La extensión es la parte del nombre del archivo que viene después del último punto (`.`).
- Si el nombre del archivo no contiene un punto o termina con un punto, devuelve `"none"`.
- La extensión debe devolverse tal cual, preservando mayúsculas y minúsculas.

# --hints--

`getExtension("document.txt")` debería devolver `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` debería devolver `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` debería devolver `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` debería devolver `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` debería devolver `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` debería devolver `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
