---
id: 68f6587287ad1f4ad39b0c7f
title: "Desafío 93: Vocales y Consonantes"
challengeType: 28
dashedName: challenge-93
---

# --description--

Dada una cadena, devuelve un arreglo con el número de vocales y el número de consonantes en la cadena.

- Las vocales consisten en `a`, `e`, `i`, `o`, `u` en cualquier caso.
- Las consonantes consisten en todas las demás letras en cualquier caso.
- Ignora cualquier carácter que no sea una letra.

Por ejemplo, dado `"Hello World"`, devuelve `[3, 7]`.

# --hints--

`count("Hello World")` debería devolver `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` debería devolver `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` debería devolver `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` debería devolver `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` debería devolver `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` debería devolver `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
