---
id: 68f6587287ad1f4ad39b0c81
title: "Desafío 95: Desplazamiento de Array"
challengeType: 28
dashedName: challenge-95
---

# --description--

Dado un arreglo y un entero que representa cuántas posiciones desplazar el arreglo, devuelve el arreglo desplazado.

- Un entero positivo desplaza el arreglo hacia la izquierda.
- Un entero negativo desplaza el arreglo hacia la derecha.
- El desplazamiento se envuelve alrededor del array.

Por ejemplo, dado `[1, 2, 3]` y `1`, desplaza el arreglo 1 posición a la izquierda, devolviendo `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` debería devolver `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` debería devolver `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` debería devolver `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` debería devolver `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` debería devolver `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
