---
id: 68ffb91507a5b645769328c4
title: "Desafío 101: Convertidor de Encabezados Markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

Dada una cadena que representa un encabezado Markdown, devuelve el encabezado HTML equivalente.

Un encabezado válido de Markdown debe:

- Comienza con cero o más espacios, seguido de
- De 1 a 6 caracteres hash (`#`) seguidos, luego
- Al menos un espacio. Y finalmente,
- El texto del encabezado.

El número de símbolos de almohadilla determina el nivel del encabezado. Por ejemplo, un símbolo de almohadilla corresponde a una etiqueta `h1`, y seis símbolos de almohadilla corresponden a una etiqueta `h6`.

Si la cadena dada no tiene el formato exacto anterior, devuelve `"Invalid format"`.

Por ejemplo, dado `"# My level 1 heading"`, devuelve `"<h1>My level 1 heading</h1>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`convert("# My level 1 heading")` debería devolver `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` debería devolver `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` debería devolver `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` debería devolver `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` debería devolver `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` debería devolver `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` debería devolver `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
