---
id: 68ffb91507a5b645769328c3
title: "Desafío 100: 100 Caracteres"
challengeType: 28
dashedName: challenge-100
---

# --description--

¡Bienvenido al desafío diario de codificación número 100!

Dada una cadena, repite sus caracteres hasta que el resultado tenga exactamente 100 caracteres. Si tus repeticiones superan los 100 caracteres, recorta el exceso para que tenga exactamente 100.

# --hints--

`oneHundred("One hundred ")` debería devolver `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` debería devolver `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` debería devolver `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` debería devolver `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```
