---
id: 68ffb91507a5b645769328c5
title: "Desafío 102: Palabra más larga"
challengeType: 28
dashedName: challenge-102
---

# --description--

Dada una cadena de oración, devuelve la palabra más larga en la oración.

- Las palabras están separadas por un solo espacio.
- Solo las letras (`a-z`, sin distinguir mayúsculas o minúsculas) cuentan para la longitud de la palabra.
- Si hay varias palabras con la misma longitud, devuelve la primera que aparece.
- Devuelve la palabra tal como aparece en la cadena dada, sin puntuación.

# --hints--

`longestWord("The quick red fox")` debería devolver `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` debería devolver `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` debería devolver `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` debería devolver `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` debería devolver `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` debería devolver `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
