---
id: 68ffb91507a5b645769328c9
title: "Desafío 106: Validador de Mensajes"
challengeType: 28
dashedName: challenge-106
---

# --description--

Dada una cadena de mensaje y una cadena de validación, determina si el mensaje es válido.

- Un mensaje es válido si cada palabra en el mensaje comienza con la letra correspondiente en la cadena de validación, en orden.
- Las letras no distinguen entre mayúsculas y minúsculas.
- Las palabras en el mensaje están separadas por espacios simples.

# --hints--

`isValidMessage("hello world", "hw")` debería devolver `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` debería devolver `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` debería devolver `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` debería devolver `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` debería devolver `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
