---
id: 68ffb91507a5b645769328cc
title: "Desafío 109: ¿Cuál es mi edad otra vez?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Dada la fecha de cumpleaños de alguien en el formato `YYYY-MM-DD`, devuelve la edad de la persona al 27 de noviembre de 2025.

- Asume que todos los cumpleaños son fechas válidas antes del 27 de noviembre de 2025.
- Devuelve la edad como un entero.
- Asegúrate de considerar si la persona ya ha cumplido años en 2025.

# --hints--

`calculateAge("2000-11-20")` debería devolver `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` debería devolver `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` debería devolver `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` debería devolver `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` debería devolver `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
