---
id: 68ffb91507a5b645769328c6
title: "Desafío 103: MCM"
challengeType: 28
dashedName: challenge-103
---

# --description--

Dado dos enteros, devuelve el mínimo común múltiplo (LCM) de los dos números.

El MCM de dos números es el entero positivo más pequeño que es múltiplo de ambos números. Por ejemplo, dado `4` y `6`, devuelve `12` porque:

- Los múltiplos de `4` son `4`, `8`, `12` y así sucesivamente.
- Múltiplos de `6` son `6`, `12`, `18` y así sucesivamente.
- `12` es el número más pequeño que es múltiplo de ambos.

# --hints--

`lcm(4, 6)` debería devolver `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` debería devolver `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` debería devolver `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` debería devolver `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` debería devolver `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```
