---
id: 69162d64f96574d9bb629eff
title: "Desafío 114: Camel a Snake"
challengeType: 28
dashedName: challenge-114
---

# --description--

Dada una cadena en camel case, devuelve la versión en snake case de la cadena usando las siguientes reglas:

- La cadena de entrada solo contendrá letras (`A-Z` y `a-z`) y siempre comenzará con una letra minúscula.
- Cada letra mayúscula en la cadena camel case inicia una nueva palabra.
- Convierte todas las letras a minúsculas.
- Separa las palabras con un guion bajo (`_`).

# --hints--

`toSnake("helloWorld")` debería devolver `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` debería devolver `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` debería devolver `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
