---
id: 69162d64f96574d9bb629f00
title: "Desafío 115: Convertidor de Ítems de Lista Ordenada en Markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

Dada una cadena que representa un elemento de lista ordenada en Markdown, devuelve la cadena HTML equivalente.

Un elemento válido de lista ordenada en Markdown debe:

- Comienza con cero o más espacios, seguido de
- Un número (1 o mayor) y un punto (`.`), seguido de
- Al menos un espacio, y luego
- El texto del elemento de la lista.

Si la cadena no tiene el formato exacto anterior, devuelve `"Invalid format"`. De lo contrario, envuelve el texto del elemento de la lista en etiquetas `li` y devuelve la cadena.

Por ejemplo, dado `"1. My item"`, devuelve `"<li>My item</li>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`convertListItem("1. My item")` debería devolver `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` debería devolver `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` debería devolver `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` debería devolver `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` debería devolver `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` debería devolver `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
