---
id: 69162d64f96574d9bb629efe
title: "Desafío 113: Millas a Kilómetros"
challengeType: 28
dashedName: challenge-113
---

# --description--

Dada una distancia en millas como un número, devuelve la distancia equivalente en kilómetros.

- La entrada siempre será un número no negativo.
- 1 milla equivale a 1.60934 kilómetros.
- Redondea el resultado a dos decimales.
- Elimina los ceros finales innecesarios del resultado redondeado.

# --hints--

`convertToKm(1)` debería devolver `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` debería devolver `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` debería devolver `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` debería devolver `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` debería devolver `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
