---
id: 69162d64f96574d9bb629f01
title: "Desafío 116: Conteo de Permutaciones"
challengeType: 28
dashedName: challenge-116
---

# --description--

Dada una cadena, devuelve el número de permutaciones distintas que se pueden formar con sus caracteres.

- Una permutación es cualquier reordenamiento de los caracteres en la cadena.
- No cuentes permutaciones duplicadas.
- Si la cadena contiene caracteres repetidos, los arreglos repetidos solo deben contarse una vez.
- La cadena contendrá solo letras (`A-Z`, `a-z`).

Por ejemplo, dado `"abb"`, devuelve `3` porque hay tres formas únicas de estructurar las letras: `"abb"`, `"bab"` y `"bba"`.

# --hints--

`countPermutations("abb")` debería devolver `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` debería devolver `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` debería devolver `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` debería devolver `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
