---
id: 69162d64f96574d9bb629efd
title: "Desafío 112: Detector de IA"
challengeType: 28
dashedName: challenge-112
---

# --description--

El desafío de hoy está inspirado en el lanzamiento de ChatGPT el 30 de noviembre de 2022.

Dada una cadena de una o más oraciones, determina si probablemente fue generada por IA usando las siguientes reglas:

- Contiene dos o más guiones (`-`).
- Contiene dos o más conjuntos de paréntesis (`()`). El texto puede estar dentro de los paréntesis.
- Contiene tres o más palabras con 7 o más letras.

- Las palabras están separadas por un solo espacio y solo consisten en letras (`A-Z`). No incluyas puntuación ni otros caracteres que no sean letras como parte de una palabra.

Si la oración dada cumple con alguna de las reglas anteriores, devuelve `"AI"`, de lo contrario, devuelve `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` debería devolver `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` debería devolver `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` debería devolver `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` debería devolver `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` debería devolver `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
