---
id: 691b559495c5cb5a37b9b485
title: "Desafío 125: Game of Life"
challengeType: 28
dashedName: challenge-125
---

# --description--

Dada una matriz (arreglo de arreglos) que representa el estado actual en el Juego de la Vida de Conway, devuelve el siguiente estado de la matriz usando estas reglas:

- Cada celda es `1` (viva) o `0` (muerta).
- Los vecinos de una celda son hasta ocho celdas circundantes (vertical, horizontal y diagonalmente).
- Las celdas en los bordes tienen menos de ocho vecinos.

Reglas para actualizar cada celda:

- Cualquier célula viva con menos de dos vecinos vivos muere (subpoblación).
- Cualquier célula viva con dos o tres vecinos vivos continúa viva.
- Cualquier célula viva con más de tres vecinos vivos muere (sobrepoblación).
- Cualquier célula muerta con exactamente tres vecinos vivos se vuelve viva (reproducción).

Por ejemplo, dado:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

retornar:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Cada celda se actualiza según el número de vecinos vivos. Por ejemplo, `[0][0]` permanece muerta (2 vecinos vivos), `[0][1]` permanece viva (2 vecinos vivos), `[0][2]` muere (3 vecinos vivos), y así sucesivamente.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` debería devolver `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` debería devolver `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` debería devolver `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` debería devolver `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
