---
id: 691b559495c5cb5a37b9b487
title: "Desafío 127: Verificación de Velocidad"
challengeType: 28
dashedName: challenge-127
---

# --description--

Dada la velocidad a la que viajas en millas por hora (MPH), y un límite de velocidad en kilómetros por hora (KPH), determina si estás excediendo el límite y si recibirás una advertencia o una multa.

- 1 milla equivale a 1.60934 kilómetros.
- Si viajas a una velocidad menor o igual al límite, devuelve `"Not Speeding"`.
- Si viajas a 5 KPH o menos sobre el límite de velocidad, devuelve `"Warning"`.
- Si viajas a más de 5 KPH sobre el límite de velocidad, devuelve `"Ticket"`.

# --hints--

`speedCheck(30, 70)` debería devolver `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` debería devolver `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` debería devolver `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` debería devolver `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` debería devolver `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` debería devolver `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
