---
id: 691b5597f0f3e85a588a5c94
title: "Desafío 130: Checkerboard"
challengeType: 28
dashedName: challenge-130
---

# --description--

Dado un arreglo con dos números, el primero siendo el número de filas y el segundo el número de columnas, devuelve una matriz (un arreglo de arreglos) llena con caracteres `"X"` y `"O"` del tamaño dado.

- Los caracteres deben alternar como un tablero de ajedrez.
- La celda superior izquierda debe ser siempre `"X"`.

Por ejemplo, dado `[3, 3]`, devuelve:

```json
[
  ["X", "O", "X"],
  ["O", "X", "O"],
  ["X", "O", "X"]
]
```

# --hints--

`createBoard([3, 3])` debería devolver `[["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([3, 3]), [["X", "O", "X"], ["O", "X", "O"], ["X", "O", "X"]]);
```

`createBoard([6, 1])` debería devolver `[["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]`.

```js
assert.deepEqual(createBoard([6, 1]), [["X"], ["O"], ["X"], ["O"], ["X"], ["O"]]);
```

`createBoard([2, 10])` debería devolver `[["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]`.

```js
assert.deepEqual(createBoard([2, 10]), [["X", "O", "X", "O", "X", "O", "X", "O", "X", "O"], ["O", "X", "O", "X", "O", "X", "O", "X", "O", "X"]]);
```

`createBoard([5, 4])` debería devolver `[["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]`.

```js
assert.deepEqual(createBoard([5, 4]), [["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"], ["O", "X", "O", "X"], ["X", "O", "X", "O"]]);
```

# --seed--

## --seed-contents--

```js
function createBoard(dimensions) {
  return dimensions;
}
```

# --solutions--

```js
function createBoard(dimensions) {
  const [rows, cols] = dimensions;
  const board = [];

  for (let r = 0; r < rows; r++) {
    const row = [];
    for (let c = 0; c < cols; c++) {
      row.push((r + c) % 2 === 0 ? "X" : "O");
    }
    board.push(row);
  }

  return board;
}
```
