---
id: 691b559495c5cb5a37b9b483
title: "Desafío 123: Constructor de Números Romanos"
challengeType: 28
dashedName: challenge-123
---

# --description--

Dado un entero, devuelve su valor equivalente en números romanos.

Los números romanos usan estos símbolos:

| Símbolo | Valor |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

Los números romanos se escriben de mayor a menor, de izquierda a derecha, usando las siguientes reglas:

- La adición se usa cuando un símbolo es seguido por otro de valor igual o menor. Por ejemplo, `18` se escribe como `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- La resta se usa cuando un símbolo más pequeño aparece antes de uno más grande, para representar 4 o 9 en cualquier valor posicional. Por ejemplo, 19 se escribe como `XIX` (10 + (10 - 1)).
- Ningún símbolo puede repetirse más de tres veces seguidas. La sustracción se usa cuando de otro modo tendrías que escribir un símbolo más de tres veces seguidas. Así que el número más grande que puedes escribir es 3999.

Aquí tienes un ejemplo más: dado `1464`, devuelve `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` debería devolver `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` debería devolver `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` debería devolver `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` debería devolver `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` debería devolver `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
