---
id: 691b559495c5cb5a37b9b480
title: "Desafío 120: Libras a Kilogramos"
challengeType: 28
dashedName: challenge-120
---

# --description--

Dado un peso en libras como un número, devuelve la cadena `"(lbs) pounds equals (kgs) kilograms."`.

- Reemplaza `"(lbs)"` con el número de entrada.
- Reemplaza `"(kgs)"` con la entrada convertida a kilogramos, redondeada a dos decimales y siempre incluye dos decimales en el valor.
- 1 libra equivale a 0.453592 kilogramos.
- Si la entrada es `1`, usa `"pound"` en lugar de `"pounds"`.
- Si el valor convertido es `1`, usa `"kilogram"` en lugar de `"kilograms"`.

# --hints--

`convertToKgs(1)` debería devolver `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` debería devolver `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` debería devolver `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` debería devolver `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` debería devolver `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
