---
id: 691b559495c5cb5a37b9b488
title: "Desafío 128: Conteo de Consonantes"
challengeType: 28
dashedName: challenge-128
---

# --description--

Dada una cadena y un número objetivo, determina si la cadena contiene exactamente el número objetivo de consonantes.

- Las consonantes son todos los caracteres alfabéticos excepto `"a"`, `"e"`, `"i"`, `"o"` y `"u"` en cualquier caso.
- Ignora dígitos, puntuación, espacios y otros caracteres que no sean letras al contar.

# --hints--

`hasConsonantCount("helloworld", 7)` debería devolver `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` debería devolver `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` debería devolver `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` debería devolver `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` debería devolver `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
