---
id: 691b559495c5cb5a37b9b486
title: "Desafío 126: Pon en mayúscula la primera letra"
challengeType: 28
dashedName: challenge-126
---

# --description--

Dada una cadena title, devuelve una nueva cadena formateada en mayúsculas y minúsculas siguiendo las siguientes reglas:

- Pon en mayúscula la primera letra de cada palabra.
- Haz que todas las demás letras de cada palabra estén en minúsculas.
- Las palabras siempre están separadas por un solo espacio.

# --hints--

`titleCase("hello world")` debería devolver `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` debería devolver `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` debería devolver `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` debería devolver `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
