---
id: 691b559495c5cb5a37b9b482
title: "Desafío 122: Analizador de Negritas en Markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

Dada una cadena que puede incluir texto en **negrita** en Markdown, devuelve la cadena HTML equivalente.

- El texto en negrita en Markdown es cualquier texto que comienza y termina con dos asteriscos (`**`) o dos guiones bajos (`__`).
- No puede haber espacios entre el texto y los asteriscos o guiones bajos, pero sí puede haber espacios dentro del texto mismo.
- Convierte todas las ocurrencias en **negrita** a etiquetas HTML `b` y devuelve la cadena.

Por ejemplo, dado `"**This is bold**"`, devuelve `"<b>This is bold</b>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parseBold("**This is bold**")` debería devolver `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` debería devolver `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` debería devolver `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` debería devolver `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` debería devolver `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
