---
id: 6925e2068081f40f549ced1a
title: "Desafío 136: Analizador de Imágenes en Markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Dada una cadena de una imagen en Markdown, devuelve la cadena HTML equivalente.

Una imagen en Markdown tiene el siguiente formato: `"![alt text](image_url)"`. Donde:

- `alt text` es la descripción de la imagen (el valor del atributo `alt`).
- `image_url` es la URL de origen de la imagen (el valor del atributo `src`).

Devuelve una cadena del tag HTML `img` con el `src` establecido en la URL de la imagen y el `alt` establecido en el texto alternativo.

Por ejemplo, dado `"![Cute cat](cat.png)"` devuelve `'<img src="cat.png" alt="Cute cat">'`;

- Asegúrate de que la etiqueta, el orden de los atributos, el espaciado y el uso de comillas sean iguales a los anteriores.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes — revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parseImage("![Cute cat](cat.png)")` debería devolver `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` debería devolver `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` debería devolver `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
