---
id: 6925e2068081f40f549ced1c
title: "Desafío 138: Suma de Divisores"
challengeType: 28
dashedName: challenge-138
---

# --description--

Dado un entero positivo, devuelve la suma de todos sus divisores.

- Un divisor es cualquier entero que divide el número de manera exacta (el residuo es `0`).
- Cuenta cada divisor solo una vez.

Por ejemplo, dado `6`, devuelve `12` porque los divisores de `6` son `1`, `2`, `3` y `6`, y la suma de esos es `12`.

# --hints--

`sumDivisors(6)` debería devolver `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` debería devolver `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` debería devolver `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` debería devolver `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` debería devolver `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` debería devolver `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
