---
id: 6925e2068081f40f549ced1b
title: "Desafío 137: Generador de Snowflake"
challengeType: 28
dashedName: challenge-137
---

# --description--

Dada una cadena de texto multilínea que usa caracteres de nueva línea (`\n`) para representar un salto de línea, devuelve una nueva cadena donde cada línea está reflejada horizontalmente y se adjunta al final de la línea original.

- Refleja una línea invirtiendo todos sus caracteres, incluidos los espacios.

Por ejemplo, dado `"* \n *\n* "`, que se registra en la consola como:

```sh
* 
 *
* 
```

Devuelve `"*  *\n ** \n*  *"`, que se muestra en la consola como:

```sh
*  *
 ** 
*  *
```

Toma nota cuidadosa de los espacios en blanco en las cadenas dadas y devueltas. Asegúrate de no eliminar ninguno de ellos.

# --hints--

`generateSnowflake("* \n *\n* ")` debería devolver `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` debería devolver `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` debería devolver `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` debería devolver `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` debería devolver `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
