---
id: 6925e2068081f40f549ced1d
title: "Desafío 139: Piedra, Papel o Tijeras"
challengeType: 28
dashedName: challenge-139
---

# --description--

Dadas dos cadenas, la primera representando a Player 1 y la segunda representando a Player 2, determina el ganador de una partida de Rock, Paper, Scissors.

- Las cadenas de entrada siempre serán `"Rock"`, `"Paper"` o `"Scissors"`.
- "Rock" vence a "Scissors".
- `"Paper"` vence a `"Rock"`.
- `"Scissors"` vence a `"Paper"`.

Devuelve:

- "Jugador 1 gana" si Player 1 gana.
- `Player 2 wins` si gana Player 2.
- "Empate" si ambos jugadores eligen la misma opción.

# --hints--

`rockPaperScissors("Rock", "Rock")` debería devolver `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` debería devolver `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` debería devolver `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` debería devolver `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` debería devolver `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` debería devolver `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
