---
id: 69272dcf1c24b44fd79137c6
title: "Desafío 143: Analizador de Cursiva en Markdown"
challengeType: 28
dashedName: challenge-143
---

# --description--

Dada una cadena que puede incluir texto en cursiva en Markdown, devuelve la cadena HTML equivalente.

- El texto en cursiva en Markdown es cualquier texto que comienza y termina con un asterisco simple (`*`) o un guion bajo simple (`_`).
- No puede haber espacios entre el texto y el asterisco o guion bajo, pero puede haber espacios dentro del texto mismo.
- Convierte todas las ocurrencias en cursiva a etiquetas HTML `i` y devuelve la cadena.

Por ejemplo, dado `"*This is italic*"`, devuelve `"<i>This is italic</i>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parseItalic("*This is italic*")` debería devolver `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalic("_This is also italic_")` debería devolver `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalic("*This is not italic *")` debería devolver `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalic("_ This is also not italic_")` debería devolver `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalic("The *quick* brown fox _jumps_ over the *lazy* dog.")` debería devolver `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
