---
id: 69272dcf1c24b44fd79137c3
title: "Desafío 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Dada una cadena que representa un nombre de variable, devuelve el nombre de variable convertido a SCREAMING_SNAKE_CASE.

Los nombres de variables dados se escribirán en uno de los siguientes formatos:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

En los formatos anteriores, las palabras están separadas por un guion bajo (`_`), un guion (`-`), o una nueva palabra comienza con una letra mayúscula.

Para convertir a SCREAMING_SNAKE_CASE:

- Haz que todas las letras estén en mayúsculas
- Separa las palabras con un guion bajo (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` debería devolver `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` debería devolver `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` debería devolver `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` debería devolver `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` debería devolver `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();;
}
```
