---
id: 69272dcf1c24b44fd79137c5
title: "Desafío 142: Suma la Cadena"
challengeType: 28
dashedName: challenge-142
---

# --description--

Dada una cadena que contiene dígitos y otros caracteres, devuelve la suma de todos los números en la cadena.

- Trata los dígitos consecutivos como un solo número. Por ejemplo, `"13"` cuenta como 13, no como 1 + 3.
- Ignora cualquier carácter que no sea un dígito.

# --hints--

`stringSum("3apples2bananas")` debería devolver `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` debería devolver `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` debería devolver `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` debería devolver `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` debería devolver `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
