---
id: 69272dcf1c24b44fd79137c4
title: "Desafío 141: Combustible para el despegue"
challengeType: 28
dashedName: challenge-141
---

# --description--

Dado el número de galones de combustible que tienes actualmente en tu avión, y el número requerido de litros de combustible para llegar a tu destino, determina cuántos galones adicionales de combustible debes agregar.

- 1 galón equivale a 3.78541 litros.
- Si el avión ya tiene suficiente combustible, devuelve `0`.
- Solo puedes agregar galones completos.
- No incluyas decimales en el número de retorno.

# --hints--

`fuelToAdd(0, 1)` debería devolver `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` debería devolver `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` debería devolver `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` debería devolver `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` debería devolver `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
