---
id: 69306364df283fcaff2e1ad8
title: "Desafío 147: Calculadora de Año Bisiesto"
challengeType: 28
dashedName: challenge-147
---

# --description--

Dado un año entero, determina si es un año bisiesto.

Un año es bisiesto si cumple las siguientes reglas:

- El año es divisible uniformemente por 4, y
- El año no es divisible de manera uniforme por 100, a menos que
- El año es divisible exactamente por 400.

# --hints--

`isLeapYear(2024)` debería devolver `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` debería devolver `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` debería devolver `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` debería devolver `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` debería devolver `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` debería devolver `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` debería devolver `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
