---
id: 69306364df283fcaff2e1ada
title: "Desafío 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Dada una cadena, devuelve una nueva cadena donde todas las vocales se convierten en mayúsculas y todos los demás caracteres alfabéticos se convierten en minúsculas.

- Las vocales son `"a"`, `"e"`, `"i"`, `"o"` y `"u"` en cualquier caso.
- Los caracteres no alfabéticos deben permanecer sin cambios.

# --hints--

`vowelCase("vowelcase")` debería devolver `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` debería devolver `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` debería devolver `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` debería devolver `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` debería devolver `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
