---
id: 69306364df283fcaff2e1ad6
title: "Desafío 145: Número Fibonacci N-ésimo"
challengeType: 28
dashedName: challenge-145
---

# --description--

Dado un entero `n`, devuelve el `n`-ésimo número en la secuencia de fibonacci.

La secuencia de Fibonacci es una serie de números donde cada número es la suma de los dos anteriores. Los primeros 10 números en la secuencia son `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` debería devolver `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` debería devolver `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` debería devolver `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` debería devolver `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` debería devolver `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
