---
id: 69306364df283fcaff2e1ad9
title: "Desafío 148: Presión de Neumáticos"
challengeType: 28
dashedName: challenge-148
---

# --description--

Dado un arreglo con cuatro números que representan las presiones de los neumáticos en psi de los cuatro neumáticos de tu vehículo, y otro arreglo de dos números que representan la presión mínima y máxima para tus neumáticos en bar, devuelve un arreglo de cuatro cadenas que describen el estado de cada neumático.

- 1 bar equivale a 14.5038 psi.

Devuelve un arreglo con los siguientes valores para cada neumático:

- "Low" si la presión del neumático está por debajo del mínimo permitido.
- "Bueno" si está entre el mínimo y el máximo permitido.
- "Alto" si está por encima del máximo permitido.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` debería devolver `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` debería devolver `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` debería devolver `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` debería devolver `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` debería devolver `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
