---
id: 69306364df283fcaff2e1ad7
title: "Desafío 146: Asiento para zurdos en la mesa"
challengeType: 28
dashedName: challenge-146
---

# --description--

Dada una matriz 4x2 (arreglo de arreglos) que representa la disposición de los asientos para una comida, determina en cuántos asientos puede sentarse una persona zurda.

- Una persona zurda no puede sentarse donde una persona diestra estaría en el asiento inmediatamente a su izquierda.

En la matriz dada:

- Un `"R"` es un asiento ocupado por una persona diestra.
- Una `"L"` es un asiento ocupado por una persona zurda.
- Una `"U"` es un asiento desocupado.
- Solo los asientos desocupados están disponibles para sentarse.
- Los asientos en la fila superior están orientados hacia "abajo", y los asientos en la fila inferior están orientados hacia "arriba" (como una mesa), por lo que izquierda y derecha son relativos a la orientación del asiento.
- Los asientos de esquina solo tienen un asiento junto a ellos.

Por ejemplo, en la matriz dada:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

El asiento superior izquierdo no se puede ocupar porque hay una persona diestra a la izquierda. Los otros dos asientos libres se pueden ocupar porque no hay una persona diestra a la izquierda.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` debería devolver `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` debería devolver `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` debería devolver `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` debería devolver `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` debería devolver `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
