---
id: 69373793f5a867f769cde135
title: "Desafío 150: Analizador de listas desordenadas en Markdown"
challengeType: 28
dashedName: challenge-150
---

# --description--

Dada la cadena de una lista desordenada válida en Markdown, devuelve la cadena HTML equivalente.

Una lista desordenada consiste en uno o más elementos de lista. Un elemento de lista válido aparece en su propia línea y:

- -
- Al menos un espacio, y luego
- El texto del elemento de la lista.

La lista se proporciona como una sola cadena con líneas nuevas separadas por el carácter de nueva línea (`"\n"`). No incluyas los caracteres de nueva línea en el texto del ítem.

Envuelve cada elemento de la lista en etiquetas HTML `li`, y toda la lista de elementos en etiquetas `ul`.

Por ejemplo, dado `"- Item A\n- Item B"`, devuelve `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Nota: La consola puede no mostrar etiquetas HTML en cadenas al registrar mensajes. Revisa la consola del navegador para ver los registros con las etiquetas incluidas.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` debería devolver `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` debería devolver `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` debería devolver `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` debería devolver `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
