---
id: 69373793f5a867f769cde13a
title: "Desafío 155: Planta la Cosecha"
challengeType: 28
dashedName: challenge-155
---

# --description--

Dado un entero que representa el tamaño de tu campo de cultivo, y `"acres"` o `"hectares"` que representan la unidad para el tamaño de tu campo de cultivo, y un tipo de cultivo, determina cuántas plantas de ese tipo puedes colocar en tu campo.

- 1 acre equivale a 4046.86 metros cuadrados.
- 1 hectárea equivale a 10,000 metros cuadrados.

Aquí hay una lista de cultivos que se darán como entrada y cuánto espacio ocupa una sola planta:

|Cultivo|Espacio por planta|
|-|-|
|`"corn"`|1 metro cuadrado|
|`"wheat"`|0,1 metros cuadrados|
|`"soybeans"`|0,5 metros cuadrados|
|`"tomatoes"`|0,25 metros cuadrados|
|`"lettuce"`|0,2 metros cuadrados|

Devuelve el número de plantas que caben en el campo, redondeado hacia abajo al número entero más cercano.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` debería devolver `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` debería devolver `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` debería devolver `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` debería devolver `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` debería devolver `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
