---
id: 69373793f5a867f769cde139
title: "Desafío 154: Par para el hoyo"
challengeType: 28
dashedName: challenge-154
---

# --description--

Dado dos enteros, el par de un hoyo de golf y el número de golpes que un golfista hizo en ese hoyo, devuelve la puntuación del golfista usando términos de golf.

Devuelve:

- "¡Hoyo en uno!" si tomó un solo golpe.
- "Eagle" si tomó dos golpes menos que el par.
- "Birdie" si tomó un golpe menos que el par.
- "Par" si tomó el mismo número de golpes que par.
- "Bogey" si tomó un golpe más que el par.
- "Double bogey" si tomaste dos golpes más que el par.

# --hints--

`golfScore(3, 3)` debería devolver `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` debería devolver `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` debería devolver `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` debería devolver `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` debería devolver `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` debería devolver `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
