---
id: 69373793f5a867f769cde138
title: "Desafío 153: Tic-Tac-Toe"
challengeType: 28
dashedName: challenge-153
---

# --description--

Dada una matriz 3×3 (un arreglo de arreglos) que representa un juego de Tic-Tac-Toe completado, determina el ganador.

- Cada elemento en la matriz dada es un `"X"` o un `"O"`.

Un jugador gana si tiene tres de sus personajes en fila - horizontalmente, verticalmente o diagonalmente.

Devuelve:

- "X gana" si el jugador X tiene tres en línea.
- "O gana" si el jugador O tiene tres en línea.
- "Empate" si ningún jugador tiene tres en línea.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` debería devolver `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]])` debería devolver `"O wins"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["X", "O", "X"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` debería devolver `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]])` debería devolver `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "X"], ["O", "X", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` debería devolver `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` debería devolver `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```
