---
id: 69373793f5a867f769cde137
title: "Desafío 152: Primo Circular"
challengeType: 28
dashedName: challenge-152
---

# --description--

Dado un entero, determina si es un primo circular.

Un primo circular es un entero donde todas las rotaciones de sus dígitos son primos.

Por ejemplo, `197` es un primo circular porque todas las rotaciones de sus dígitos: `197`, `971` y `719`, son números primos.

# --hints--

`isCircularPrime(197)` debería devolver `true`.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` debería devolver `false`.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` debería devolver `true`.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` debería devolver `false`.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` debería devolver `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
