---
id: 6939b873185d8e00d453563d
title: "Desafío 158: Intercambio de Array"
challengeType: 28
dashedName: challenge-158
---

# --description--

Dado un arreglo con dos valores, devuelve un arreglo con los valores intercambiados.

Por ejemplo, dado `["A", "B"]` devuelve `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` debería devolver `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` debería devolver `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` debería devolver `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` debería devolver `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
