---
id: 6939b873185d8e00d453563e
title: "Desafío 159: Hipotenusa Entera"
challengeType: 28
dashedName: challenge-159
---

# --description--

Dado dos enteros positivos que representan las longitudes de las dos piernas (los dos lados cortos) de un triángulo rectángulo, determina si la hipotenusa es un entero.

La longitud de la hipotenusa se calcula sumando los cuadrados de las longitudes de los dos catetos y luego tomando la raíz cuadrada de ese total (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` debería devolver `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` debería devolver `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` debería devolver `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` debería devolver `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` debería devolver `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` debería devolver `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
