---
id: 6939b873185d8e00d453563b
title: "Desafío 156: ¿Impar o Par?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Dado un entero positivo, devuelve `"Odd"` si es un número impar, y `"Even"` si es par.

# --hints--

`oddOrEven(1)` debería devolver `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` debería devolver `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` debería devolver `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` debería devolver `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` debería devolver `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
