---
id: 694596b0585c11170ac7c7fb
title: "Desafío 163: Caso de Consonantes"
challengeType: 28
dashedName: challenge-163
---

# --description--

Dada una cadena que representa un nombre de variable, conviértela a consonant case usando las siguientes reglas:

- Todas las consonantes deben convertirse a mayúsculas.
- Todas las vocales (`a`, `e`, `i`, `o`, `u` en cualquier caso) deben convertirse a minúsculas.
- Todos los guiones (`-`) deben convertirse en guiones bajos (`_`).

# --hints--

`toConsonantCase("helloworld")` debería devolver `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` debería devolver `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` debería devolver `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` debería devolver `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
