---
id: 694596b0585c11170ac7c7fa
title: "Desafío 162: Consumo de Energía"
challengeType: 28
dashedName: challenge-162
---

# --description--

Dado el número de Calorías quemadas durante un entrenamiento y el número de vatios-hora usados por tus dispositivos electrónicos durante ese entrenamiento, determina cuál usó más energía.

Para compararlos, convierte ambos valores a julios usando las siguientes conversiones:

- 1 Caloría equivale a 4184 julios.
- 1 vatio-hora equivale a 3600 julios.

Devuelve:

- "Workout" si el workout usó más energía.
- "Dispositivos" si el dispositivo usó más energía.
- "Igual" si ambos usaron la misma cantidad de energía.

# --hints--

`compareEnergy(250, 50)` debería devolver `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` debería devolver `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` debería devolver `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` debería devolver `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` debería devolver `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` debería devolver `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
