---
id: 696655d24b614176d4c9b789
title: "Desafío 166: Validador Hexadecimal"
challengeType: 28
dashedName: challenge-166
---

# --description--

Dada una cadena, determina si es un color hexadecimal CSS válido. Un color hexadecimal CSS válido debe:

- Comienza con un `#`, y
- deben ser seguidos por 3 o 6 caracteres hexadecimales.

Los caracteres hexadecimales son números `0` a `9` y letras `a` a `f` (sin distinguir mayúsculas de minúsculas).

# --hints--

`isValidHex("#123")` debería devolver `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` debería devolver `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` debería devolver `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` debería devolver `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` debería devolver `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` debería devolver `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` debería devolver `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` debería devolver `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
