---
id: 696655d24b614176d4c9b78a
title: "Desafío 167: ¡Bingo! Letra"
challengeType: 28
dashedName: challenge-167
---

# --description--

Dado un número, devuelve la letra de bingo asociada con él (en mayúsculas). Los números de bingo se agrupan de la siguiente manera:

| Letter | Number Range |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

# --hints--

`getBingoLetter(75)` debería devolver `"O"`.

```js
assert.equal(getBingoLetter(75), "O");
```

`getBingoLetter(54)` debería devolver `"G"`.

```js
assert.equal(getBingoLetter(54), "G");
```

`getBingoLetter(25)` debería devolver `"I"`.

```js
assert.equal(getBingoLetter(25), "I");
```

`getBingoLetter(38)` debería devolver `"N"`.

```js
assert.equal(getBingoLetter(38), "N");
```

`getBingoLetter(11)` debería devolver `"11"`.

```js
assert.equal(getBingoLetter(11), "B");
```

# --seed--

## --seed-contents--

```js
function getBingoLetter(n) {

  return n;
}
```

# --solutions--

```js
function getBingoLetter(n) {
  if (n >= 1 && n <= 15) return "B";
  if (n >= 16 && n <= 30) return "I";
  if (n >= 31 && n <= 45) return "N";
  if (n >= 46 && n <= 60) return "G";
  if (n >= 61 && n <= 75) return "O";
  return null;
}
```
