---
id: 696655d24b614176d4c9b78c
title: "Desafío 169: FizzBuzz Mini"
challengeType: 28
dashedName: challenge-169
---

# --description--

Dado un entero, devuelve una cadena basada en las siguientes reglas:

- Si el número es divisible por 3, devuelve `"Fizz"`.
- Si el número es divisible por 5, devuelve `"Buzz"`.
- Si el número es divisible tanto por 3 como por 5, devuelve `"FizzBuzz"`.
- De lo contrario, devuelve el número dado como una cadena.

# --hints--

`fizzBuzzMini(3)` debería devolver `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` debería devolver `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` debería devolver `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` debería devolver `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` debería devolver `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
