---
id: 696655d24b614176d4c9b78d
title: "Desafío 170: Día impar o par"
challengeType: 28
dashedName: challenge-170
---

# --description--

Dado un timestamp (número de milisegundos desde la época Unix), devuelve:

- "odd" si el día del mes para ese timestamp es impar.
- "even" si el día del mes para ese timestamp es par.

Por ejemplo, dado `1769472000000`, una marca de tiempo para el 27 de enero de 2026, devuelve `"odd"` porque el número del día (`27`) es un número impar.

Nota: la marca de tiempo está en milisegundos y debes usar la fecha en la zona horaria UTC, no en tu hora local.

# --hints--

`oddOrEvenDay(1769472000000)` debería devolver `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` debería devolver `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` debería devolver `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` debería devolver `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` debería devolver `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
