---
id: 696655d24b614176d4c9b78b
title: "Desafío 168: Imagen escalada"
challengeType: 28
dashedName: challenge-168
---

# --description--

Dada una cadena que representa el ancho y alto de una imagen, y un número para escalar la imagen, devuelve el ancho y alto escalados.

- La cadena de entrada está en el formato `"WxH"`. Por ejemplo, `"800x600"`.
- La escala es un número para multiplicar el ancho y el alto.

Devuelve las dimensiones escaladas en el mismo formato `"WxH"`.

# --hints--

`scaleImage("800x600", 2)` debería devolver `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` debería devolver `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` debería devolver `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` debería devolver `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
