---
id: 69738771fb5a7b8b24cca29f
title: "Desafío 173: Movimientos válidos de peón"
challengeType: 28
dashedName: challenge-173
---

# --description--

Dada la posición de uno de tus peones en un tablero de ajedrez, devuelve un arreglo con todas las casillas válidas a las que puede moverse en orden ascendente.

Un tablero de ajedrez estándar es de 8x8, con columnas etiquetadas de `A` a `H` (de izquierda a derecha) y filas etiquetadas del `1` al `8` (de abajo hacia arriba). Se ve así:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Para este desafío:

- Eres el jugador en la parte inferior del tablero.
- Los peones solo pueden moverse un cuadro "hacia arriba".
- A menos que el peón esté en la fila inicial (fila 2), entonces puede moverse una o dos casillas hacia arriba.

Por ejemplo, dado `"D4"`, devuelve `["D5"]`, la única casilla a la que tu peón puede moverse.  
Dado `"B2"`, devuelve `["B3", "B4"]`, porque está en la fila inicial y debe estar en orden ascendente.

# --hints--

`findPawnMoves("D4")` debería devolver `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` debería devolver `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` debería devolver `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` debería devolver `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` debería devolver `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
