---
id: 69738771fb5a7b8b24cca2a5
title: "Desafío 179: Cambio de bolsillo"
challengeType: 28
dashedName: challenge-179
---

# --description--

Dado un arreglo de enteros que representan las monedas en tu bolsillo, donde cada entero es el valor de una moneda en centavos, devuelve el monto total en el formato `"$D.CC"`.

- 100 centavos equivalen a 1 dólar.
- En el valor de retorno, incluye un cero inicial para cantidades menores a un dólar y siempre exactamente dos dígitos para los centavos.

# --hints--

`countChange([25, 10, 5, 1])` debería devolver `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` debería devolver `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` debería devolver `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` debería devolver `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` debería devolver `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` debería devolver `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```
