---
id: 69738771fb5a7b8b24cca2a3
title: "Desafío 177: Espejo de Cadenas"
challengeType: 28
dashedName: challenge-177
---

# --description--

Dada una cadena, devuelve una nueva cadena que consiste en la cadena dada con una copia invertida de sí misma añadida al final.

# --hints--

`mirror("freeCodeCamp")` debería devolver `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` debería devolver `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` debería devolver `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` debería devolver `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
