---
id: 69738771fb5a7b8b24cca2a4
title: "Desafío 178: Trunca el Texto"
challengeType: 28
dashedName: challenge-178
---

# --description--

Dada una cadena, devuélvela tal cual si tiene 20 caracteres o menos. Si es más larga que 20 caracteres, trúncala a los primeros 17 caracteres y añade `"..."` al final (para que tenga 20 caracteres en total) y devuelve el resultado.

# --hints--

`truncateText("Hello, world!")` debería devolver `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` debería devolver `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` debería devolver `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` debería devolver `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
