---
id: 69738771fb5a7b8b24cca2a2
title: "Desafío 176: Día de la Marmota"
challengeType: 28
dashedName: challenge-176
---

# --description--

Hoy es el Día de la Marmota, en el que una marmota predice el clima según si ve o no su sombra.

Dado un valor que representa la apariencia de la marmota, devuelve la predicción correcta:

- Si el valor dado es el booleano `true` (la marmota vio su sombra), devuelve `"Parece que tendremos seis semanas más de invierno."`.
- Si el valor es el booleano `false` (la marmota no vio su sombra), devuelve `"It's going to be an early spring."`.
- Si el valor es cualquier otro (la marmota no apareció), devuelve `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` debería devolver `"Parece que tendremos seis semanas más de invierno."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` debería devolver `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` debería devolver `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` debería devolver `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` debería devolver `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
