---
id: 697a49e9860d24853adef67c
title: "Desafío 191: Juegos de Invierno 2026 Día 12: Bobsleigh"
challengeType: 28
dashedName: challenge-191
---

# --description--

Dado un arreglo que representa los pesos de los atletas en un equipo de bobsled y un número que representa el peso del bobsled, determina si el equipo es elegible para competir.

- La longitud del array determina el tamaño del equipo: equipos de 1, 2 o 4 personas.
- Todos los valores de peso dados están en kilogramos (kg).

El trineo (trineo por sí mismo) debe tener un peso mínimo de:

- 162 kg para un equipo de 1 persona
- 170 kg para un equipo de 2 personas
- 210 kg para un equipo de 4 personas

El peso total del trineo (atletas más trineo) no debe exceder:

- 247 kg para un equipo de 1 persona
- 390 kg para un equipo de 2 personas
- 630 kg para un equipo de 4 personas

Devuelve `"Eligible"` si el equipo cumple con todos los requisitos, o `"Not Eligible"` si el equipo no cumple con uno o más de los requisitos.

# --hints--

`checkEligibility([78], 165)` debería devolver `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` debería devolver `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` debería devolver `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` debería devolver `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
