---
id: 697a49e6ff50d756c9b6935f
title: "Desafío 182: Juegos de Invierno 2026 Día 3: Biatlón"
challengeType: 28
dashedName: challenge-182
---

# --description--

Dado un arreglo de enteros, donde cada valor representa el número de objetivos alcanzados en una sola ronda de un biatlón, devuelve la distancia total de penalización que el atleta debe esquiar.

- Cada ronda consiste en 5 objetivos.
- Cada objetivo fallado resulta en una penalización de 150 metros en el circuito.

# --hints--

`calculatePenaltyDistance([4, 4])` debería devolver `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` debería devolver `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` debería devolver `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` debería devolver `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` debería devolver `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
