---
id: 697a49e6ff50d756c9b69365
title: "Desafío 188: Juegos de Invierno 2026 Día 9: Skeleton"
challengeType: 28
dashedName: challenge-188
---

# --description--

Dada una cadena que representa las curvas en una pista esquelética, determina la dificultad de la pista.

- La cadena dada solo consistirá en las letras:
  - "L" para un giro a la izquierda
  - `"R"` para un giro a la derecha
  - "S" para un segmento recto

- Cada cambio de dirección añade 15 puntos (una `"L"` seguida de una `"R"` o viceversa).
- Todas las demás curvas suman 5 puntos cada una (todos los demás caracteres `"L"` o `"R"`).
- Los segmentos rectos añaden 0 puntos.

La dificultad de la pista se basa en la puntuación total. Devuelve:

- "Fácil" si el total es 0 - 100
- "Medio" si el total es 101-200
- "Difícil" si el total es mayor a 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` debería devolver `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` debería devolver `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` debería devolver `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` debería devolver `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` debería devolver `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` debería devolver `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
