---
id: 697a49e9860d24853adef67b
title: "Desafío 190: Juegos de Invierno 2026 Día 11: Hockey sobre Hielo"
challengeType: 28
dashedName: challenge-190
---

# --description--

Dado un arreglo de 6 equipos de hockey sobre hielo y sus registros después de los juegos de todos contra todos, determina los enfrentamientos para la ronda semifinal.

- Each array item will have a team and their record in the format `"TEAM: W-OTW-OTL-L"`. Where:
  - "W" es el número de victorias en tiempo reglamentario, que vale 3 puntos cada una
  - "OTW" es el número de victorias en tiempo extra, que valen 2 puntos cada una
  - "OTL" es el número de derrotas en tiempo extra, que valen 1 punto cada una
  - "L" es el número de derrotas, que vale 0 puntos cada una
  
Por ejemplo, `"FIN: 2-2-1-0"` tendría 11 puntos después de sumar su récord.
  
Encuentra el número total de puntos para cada equipo y devuelve `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. Por ejemplo, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` debería devolver `"Los partidos de semifinal serán FIN contra SWE y CAN contra USA."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` debería devolver `"Los partidos de semifinal serán USA contra CZE y CAN contra FIN."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` debería devolver `"Los partidos de semifinal serán CAN contra ITA y USA contra CZE."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` debería devolver `"Los partidos de semifinal serán JPN contra ITA y AUT contra KOR."`.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
