---
id: 697a49e6ff50d756c9b69366
title: "Desafío 189: Juegos de Invierno 2026 Día 10: Esquí Alpino"
challengeType: 28
dashedName: challenge-189
---

# --description--

Dado el desnivel vertical, la distancia horizontal y el tipo de una pista de esquí, determina la calificación de dificultad de la pista.

Para determinar la calificación:

- Calcula la inclinación de la colina tomando la caída dividida por la distancia.
- Luego, calcula la inclinación ajustada según el tipo de colina:
  - "Downhill" multiplica la pendiente por 1.2
  - "Slalom": multiplica la pendiente por 0.9
  - "Giant Slalom": multiplica la pendiente por 1.0

Devuelve:

- "Green" si la pendiente ajustada es menor o igual a 0.1
- "Blue" si la pendiente ajustada es mayor que 0.1 y menor o igual a 0.25
- "Negro" si la pendiente ajustada es mayor que 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` debería devolver `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` debería devolver `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` debería devolver `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` debería devolver `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` debería devolver `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` debería devolver `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
