---
id: 697a49e6ff50d756c9b69360
title: "Desafío 183: Juegos de Invierno 2026 Día 4: Salto de Esquí"
challengeType: 28
dashedName: challenge-183
---

# --description--

Dado puntos de distancia, puntos de estilo, un valor de compensación por viento y un valor de bonificación del punto K, calcula tu puntuación para el salto de esquí y determina si ganaste una medalla o no.

- Tu puntuación se calcula sumando los cuatro valores anteriores.

Las puntuaciones totales actuales de los otros saltadores son:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Si tu puntuación es la mejor, devuelve `"Gold"`
- Si es el segundo mejor, devuelve `"Silver"`
- Si es el tercero mejor, devuelve `"Bronze"`
- De lo contrario, devuelve `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` debería devolver `"Gold"`.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` debería devolver `"Bronce"`.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` debería devolver `"Silver"`.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` debería devolver `"No Medal"`.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` debería devolver `"Gold"`.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` debería devolver `"No Medal"`.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```
