---
id: 697a49e9860d24853adef67f
title: "Desafío 194: Juegos de Invierno 2026 Día 15: Esquí Acrobático"
challengeType: 28
dashedName: challenge-194
---

# --description--

Dado un nombre de truco que consiste en dos palabras, determina si es un nombre válido de truco de esquí estilo libre.

Un truco es válido si la primera palabra está en la lista de primeras palabras válidas, y la segunda palabra está en la lista de segundas palabras válidas.

- Las dos palabras estarán separadas por un solo espacio.

Primeras palabras válidas:

|`"Misty"`|
|-|
|`"Ghost"`|
|`"Thunder"`|
|`"Solar"`|
|`"Sky"`|
|`"Phantom"`|
|`"Frozen"`|
|`"Polar"`|

Palabras segundas válidas:

|`"Twister"`|
|-|
|`"Icequake"`|
|`"Avalanche"`|
|`"Vortex"`|
|`"Snowstorm"`|
|`"Frostbite"`|
|`"Blizzard"`|
|`"Shadow"`|

# --hints--

`isValidTrick("Polar Vortex")` debería devolver `true`.

```js
assert.isTrue(isValidTrick("Polar Vortex"));
```

`isValidTrick("Solar Icequake")` debería devolver `true`.

```js
assert.isTrue(isValidTrick("Solar Icequake"));
```

`isValidTrick("Thunder Blizzard")` debería devolver `true`.

```js
assert.isTrue(isValidTrick("Thunder Blizzard"));
```

`isValidTrick("Phantom Frostbite")` debería devolver `true`.

```js
assert.isTrue(isValidTrick("Phantom Frostbite"));
```

`isValidTrick("Ghost Avalanche")` debería devolver `true`.

```js
assert.isTrue(isValidTrick("Ghost Avalanche"));
```

`isValidTrick("Snowstorm Shadow")` debería devolver `false`.

```js
assert.isFalse(isValidTrick("Snowstorm Shadow"));
```

`isValidTrick("Solar Sky")` debería devolver `false`.

```js
assert.isFalse(isValidTrick("Solar Sky"));
```

# --seed--

## --seed-contents--

```js
function isValidTrick(trickName) {

  return trickName;
}
```

# --solutions--

```js
function isValidTrick(trickName) {
  const validFirst = ["Misty", "Ghost", "Thunder", "Solar", "Sky", "Phantom", "Frozen", "Polar"];
  const validSecond = ["Twister", "Icequake", "Avalanche", "Vortex", "Snowstorm", "Frostbite", "Blizzard", "Shadow"];

  const words = trickName.split(" ");
  const [first, second] = words;
  return validFirst.includes(first) && validSecond.includes(second);
}
```
