---
id: 697a49e6ff50d756c9b6935e
title: "Desafío 181: Juegos de Invierno 2026 Día 2: Snowboarding"
challengeType: 28
dashedName: challenge-181
---

# --description--

Dada la postura inicial de un snowboarder y una rotación en grados, determina su postura de aterrizaje.

- La postura de un snowboarder es `"Regular"` o `"Goofy"`.
- Las rotaciones de trucos son múltiplos de 90 grados. Un valor positivo indica rotación en el sentido de las agujas del reloj, y un valor negativo indica rotación en sentido contrario.
- La postura de aterrizaje gira cada 180 grados de rotación.

Por ejemplo, dado `"Regular"` y `90`, devuelve `"Regular"`. Dado `"Regular"` y `180` grados, devuelve `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` debería devolver `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` debería devolver `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` debería devolver `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` debería devolver `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` debería devolver `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` debería devolver `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
