---
id: 698a1a73ade5ac0e19180fa1
title: "Desafío 198: Conteo de Días Hábiles"
challengeType: 28
dashedName: challenge-198
---

# --description--

Dada una fecha de inicio y una fecha de fin, devuelve el número de días hábiles entre ambas.

- Las fechas dadas están en el formato `"YYYY-MM-DD"`.
- Los días laborables son días hábiles (de lunes a viernes).
- Los fines de semana no son días hábiles (sábado y domingo).
- Incluye tanto la fecha de inicio como la de finalización al contar.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` debe devolver `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` debe devolver `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` debe devolver `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` debería devolver `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` debería devolver `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
