---
id: 698a1a73ade5ac0e19180fa4
title: "Desafío 201: Desplazamiento de Matriz"
challengeType: 28
dashedName: challenge-201
---

# --description--

Dada una matriz (arreglo de arreglos) de números y un entero, desplaza todos los valores en la matriz por la cantidad dada.

- Un desplazamiento positivo mueve los valores hacia la derecha.
- Un desplazamiento negativo mueve los valores hacia la izquierda.

Los valores deben ajustarse:

- Trata la matriz como una secuencia continua de valores.
- Cuando un valor pasa del final de una fila, continúa al principio de la siguiente fila.
- Cuando un valor pasa de la última posición en la matriz, se envuelve a la primera posición.
- Lo mismo aplica a la inversa al desplazar hacia la izquierda.

Por ejemplo, dado:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

con un desplazamiento de `1`, mueve todos los números una posición a la derecha:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)` debe devolver `[[6, 1, 2], [3, 4, 5]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)` debe devolver `[[2, 3, 4], [5, 6, 1]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` debería devolver `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` debe devolver `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` debe devolver `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` debe devolver `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
