---
id: 698a1a73ade5ac0e19180fa7
title: "Desafío 204: Suma las Letras"
challengeType: 28
dashedName: challenge-204
---

# --description--

Dada una cadena, devuelve la suma de sus letras.

- Las letras son A-Z en mayúsculas o minúsculas
- Los valores de las letras son: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26
- Las letras mayúsculas y minúsculas tienen el mismo valor.
- Ignora todos los caracteres que no sean letras.

# --hints--

`sumLetters("Hello")` debería devolver `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` debería devolver `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` debe devolver `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` debe devolver `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` debería devolver `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```
