---
id: 698a1a73ade5ac0e19180fa9
title: "Desafío 206: Valores de las Cartas de Juego"
challengeType: 28
dashedName: challenge-206
---

# --description--

Dado un arreglo de cartas, devuelve un nuevo arreglo con el valor numérico de cada carta.

Valores de la carta:

- Un As (`"A"`) tiene un valor de 1.
- Las cartas numeradas (`"2"` - `"10"`) tienen su valor facial: 2 - 10, respectivamente.
- Las cartas con figuras: Jack (`"J"`), Queen (`"Q"`), y King (`"K"`) valen 10 cada una.

Trajes:

- Cada carta tiene un palo: Espadas (`"S"`), Tréboles (`"C"`), Diamantes (`"D"`), o Corazones (`"H"`).

Formato de tarjeta:

- Cada carta se representa como una cadena: `"valueSuit"`. Por ejemplo: `"AS"` es el As de Espadas, `"10H"` es el Diez de Corazones, y `"QC"` es la Reina de Tréboles.

# --hints--

`cardValues(["3H", "4D", "5S"])` debería devolver `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` debería devolver `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` debe devolver `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` debería devolver `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` debería devolver `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
