---
id: 698a1a73ade5ac0e19180fa6
title: "Desafío 203: Aplanado"
challengeType: 28
dashedName: challenge-203
---

# --description--

Dado un arreglo, determina si es plano.

- Un arreglo es plano si ninguno de sus elementos es un arreglo.

# --hints--

`isFlat([1, 2, 3, 4])` debería devolver `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` debería devolver `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` debe devolver `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` debería devolver `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` debería devolver `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```
