---
id: 698a1a863194f1f4e63f645e
title: "Desafío 207: Brecha más pequeña"
challengeType: 28
dashedName: challenge-207
---

# --description--

Dada una cadena, devuelve la subcadena entre los dos caracteres idénticos que tienen el menor número de caracteres entre ellos (brecha más pequeña).

- Siempre habrá al menos un par de caracteres coincidentes.
- La subcadena devuelta debe excluir los caracteres coincidentes.
- Si dos o más espacios tienen la misma longitud, devuelve los caracteres del primero.

Por ejemplo, dado `"ABCDAC"`, devuelve `"DA"` porque:

- Solo `"A"` y `"C"` se repiten en la cadena.
- El número de caracteres entre los dos caracteres `"A"` es 3, y entre los caracteres `"C"` es 2.
- Entonces devuelve la cadena entre los dos caracteres `"C"`.

# --hints--

`smallestGap("ABCDAC")` debe devolver `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` debe devolver `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` debe devolver `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` debería devolver `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` debe devolver `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
