---
id: 698a1a73ade5ac0e19180fa3
title: "Desafío 200: Conteo de Letras y Números"
challengeType: 28
dashedName: challenge-200
---

# --description--

Dada una cadena, devuelve un mensaje con el conteo de cuántas letras y números contiene.

- Las letras son `A-Z` y `a-z`.
- Los números son `0-9`.
- Ignora todos los demás caracteres.

Devuelve `"The string has X letters and Y numbers."`, donde `"X"` es el conteo de letras y `"Y"`` es el conteo de números. Si alguno de los conteos es 1, usa la forma singular para ese elemento. Por ejemplo: `"1 letter"` en lugar de `"1 letters"` y `"1 number"` en lugar de `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` debería devolver `"The string has 10 letters and 3 numbers."`.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` debería devolver `"The string has 5 letters and 2 numbers."`.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` debe devolver `"The string has 1 letter and 1 number."`.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` debe devolver `"The string has 0 letters and 5 numbers."`.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` debería devolver `"The string has 8 letters and 0 numbers."`.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
