---
id: 698a1a73ade5ac0e19180fa8
title: "Desafío 205: Conteo de Cubos Perfectos"
challengeType: 28
dashedName: challenge-205
---

# --description--

Dado dos enteros, determina cuántos cubos perfectos existen en el rango entre y que incluyen a los dos números.

- El número menor no está garantizado como el primer argumento.
- Un número es un cubo perfecto si existe un entero (`n`) donde `n * n * n = number`. Por ejemplo, 27 es un cubo perfecto porque `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` debería devolver `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` debe devolver `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` debería devolver `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` debería devolver `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` debe devolver `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
