---
id: 698a1a73ade5ac0e19180fa0
title: "Desafío 197: Compatibilidad de Tipo de Sangre"
challengeType: 28
dashedName: challenge-197
---

# --description--

Dado un tipo de sangre de donante y un tipo de sangre de receptor, determina si el donante puede donar sangre al receptor.

Cada tipo de sangre consiste en:

- Una letra: `"A"`, `"B"`, `"AB"`, o `"O"`
- Y un factor Rh: `"+"` o `"-"`

Los tipos de sangre serán una de las letras válidas seguida de un factor Rh. Por ejemplo, `"AB+"` y `"O-"` son tipos de sangre válidos.

Reglas de letras:

- `"O"` puede donar a otro tipo de letra.
- `"A"` puede donar a `"A"` y `"AB"`.
- `"B"` puede donar a `"B"` y `"AB"`.
- `"AB"` solo puede donar a `"AB"`.

Reglas Rh:

- Negativo (`"-"`) puede donar tanto a `"-"` como a `"+"`.
- Positivo (`"+"`) solo puede donar a `"+"`.

Tanto la regla de letra como la regla Rh deben cumplirse para que un donante pueda donar al receptor.

# --hints--

`canDonate("B+", "B+")` debería devolver `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` debería devolver `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` debería devolver `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` debería devolver `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` debería devolver `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` debería devolver `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` debería devolver `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` debe devolver `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` debería devolver `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` debería devolver `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
