---
id: 6994cff2290543b3aec9f510
title: "Desafío 211: Suma de Array"
challengeType: 28
dashedName: challenge-211
---

# --description--

Dado un arreglo de números, devuelve la suma de todos los números.

# --hints--

`sumArray([1, 2, 3, 4, 5])` debe devolver `15`.

```js
assert.equal(sumArray([1, 2, 3, 4, 5]), 15);
```

`sumArray([42])` debería devolver `42`.

```js
assert.equal(sumArray([42]), 42);
```

`sumArray([5, -2, 7, -3])` debe devolver `7`.

```js
assert.equal(sumArray([5, -2, 7, -3]), 7);
```

`sumArray([203, 145, -129, 6293, 523, -919, 845, 2434])` debe devolver `9395`.

```js
assert.equal(sumArray([203, 145, -129, 6293, 523, -919, 845, 2434]), 9395);
```

`sumArray([0, 0])` debería devolver `0`.

```js
assert.equal(sumArray([0, 0]), 0);
```

# --seed--

## --seed-contents--

```js
function sumArray(numbers) {

  return numbers;
}
```

# --solutions--

```js
function sumArray(numbers) {

  return numbers.reduce((a, c) => c + a, 0);
}
```
