---
id: 6994cff2290543b3aec9f50f
title: "Desafío 210: Validador HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Dada una cadena, determina si es un valor válido de color CSS `hsl()`.

- Un valor HSL válido debe estar en el formato `"hsl(h, s%, l%)"`, donde:
  - `h` (hue) debe ser un número entre 0 y 360 (inclusive).
  - `s` (saturation) debe ser un porcentaje entre 0% y 100%.
  - `l` (lightness) debe ser un porcentaje entre 0% y 100%.

- Los espacios solo están permitidos:
  - Después del paréntesis de apertura
  - Antes y/o después de las comas
  - Antes y/o después del paréntesis de cierre

- Opcionalmente, el valor puede terminar con un punto y coma (`";"`).

Por ejemplo, `"hsl(240, 50%, 50%)"` es un valor HSL válido.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` debería devolver `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` debe devolver `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` debería devolver `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` debería devolver `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` debe devolver `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` debería devolver `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` debe devolver `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` debería devolver `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` debería devolver `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
