---
id: 699c8e045ee7cb94ed2322d6
title: "Desafío 215: Calculadora de tarifa de estacionamiento"
challengeType: 28
dashedName: challenge-215
---

# --description--

Dadas dos cadenas que representan la hora en que estacionaste tu auto y la hora en que lo recogiste, calcula la tarifa de estacionamiento.

- Las cadenas dadas estarán en el formato `"HH:MM"` usando un reloj de 24 horas. Por ejemplo, `"14:00"` es 2pm.
- La primera cadena será la hora en que estacionaste tu coche, y la segunda será la hora en que lo recogiste.
- Si la hora de recogida es anterior a la hora de entrada, significa que el estacionamiento se extendió más allá de la medianoche hasta el día siguiente.

Reglas de tarifas:

- Cada hora estacionada cuesta $3.
- Las horas parciales se redondean hacia arriba a la siguiente hora completa.
- Si el estacionamiento abarca toda la noche (más allá de la medianoche), se aplica una tarifa adicional de $10 por la noche.
- Hay una tarifa mínima de $5 (solo se usa si el total sería menor a $5).

Devuelve el costo total en el formato `"$cost"`, por ejemplo, `"$5"`.

# --hints--

`calculateParkingFee("09:00", "11:00")` debe devolver `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` debería devolver `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` debería devolver `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` debe devolver `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` debería devolver `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` debe devolver `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
