---
id: 699c8e045ee7cb94ed2322dd
title: "Desafío 222: Sombras del Equinoccio"
challengeType: 28
dashedName: challenge-222
---

# --description--

Hoy es el equinoccio, cuando el sol está directamente sobre el ecuador y perfectamente en lo alto al mediodía. Dado un tiempo, determina la sombra proyectada por un poste vertical de 4 pies.

- La hora será una cadena en formato de 24 horas `"HH:MM"` (por ejemplo, `"15:00"` es 3pm).
- Solo se te dará una hora en incrementos de 30 minutos.

Reglas:

- El sol sale a las 6am directamente `"east"`, y se pone a las 6pm directamente `"west"`.
- Una sombra siempre apunta en dirección opuesta al sol.
- La longitud de la sombra (en pies) es el número de horas desde el mediodía, elevado al cubo.
- No hay sombra antes del amanecer (antes de las 6am), después del atardecer (6pm o más tarde), o al mediodía.

Devuelve:

- Si existe una sombra, devuelve `"(length)ft (direction)"`. Por ejemplo, `"8ft west"`.
- De lo contrario, devuelve `"No shadow"`.

Por ejemplo, dado `"10:00"`, devuelve `"8ft west"` porque las 10am están a 2 horas del mediodía, así que 2<sup>3</sup> = 8 pies, y la sombra apunta al oeste porque el sol está en el este a las 10am.

# --hints--

`getShadow("10:00")` debería devolver `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` debería devolver `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` debe devolver `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` debería devolver `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` debería devolver `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` debería devolver `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` debe devolver `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` debería devolver `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` debe devolver `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
