---
id: 699c8e045ee7cb94ed2322db
title: "Desafío 220: Número más grande"
challengeType: 28
dashedName: challenge-220
---

# --description--

Dada una cadena de números separados por varios signos de puntuación, devuelve el número más grande.

- La cadena dada solo contendrá números y separadores.
- Los separadores pueden ser comas (`","`), signos de exclamación (`"!"`), signos de interrogación (`"?"`), dos puntos (`":"`) o puntos y coma (`";"`).

# --hints--

`largestNumber("1,2")` debería devolver `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` debe devolver `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` debería devolver `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` debe devolver `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
