---
id: 699c8e045ee7cb94ed2322d5
title: "Desafío 214: Validador de Cadena de Dominó"
challengeType: 28
dashedName: challenge-214
---

# --description--

Dado un arreglo 2D que representa una secuencia de dominós, determina si forma una cadena válida.

- Cada elemento en el arreglo representa un dominó y será un arreglo de dos números del 1 al 6, (inclusive).
- Para que la cadena sea válida, el segundo número de cada dominó debe coincidir con el primer número del siguiente dominó.
- El primer número del primer dominó y el último número del último dominó no necesitan coincidir con nada.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` debería devolver `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` debe devolver `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` debe devolver `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` debe devolver `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` debe devolver `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
