---
id: 699c8e045ee7cb94ed2322da
title: "Desafío 219: Hitos de Aniversario"
challengeType: 28
dashedName: challenge-219
---

# --description--

Dado un entero que representa el número de años que una pareja ha estado casada, devuelve su hito de aniversario más reciente según esta tabla:

| Years Married | Milestone    |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Si no han alcanzado el primer hito, devuelve `"Newlyweds"`.

# --hints--

`getMilestone(0)` debe devolver `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` debe devolver `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` debería devolver `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` debería devolver `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` debería devolver `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` debería devolver `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` debería devolver `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` debe devolver `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` debería devolver `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
