---
id: 69a890af247de743333bd4cd
title: "Desafío 224: Detector de Tostado de Café"
challengeType: 28
dashedName: challenge-224
---

# --description--

Dada una cadena que representa los granos usados para hacer una taza de café, determina el tueste de la taza.

- La cadena dada contendrá los siguientes caracteres, cada uno representando un tipo de frijol:

  - Un apóstrofe (`'`) es un grano de tueste ligero que vale 1 punto cada uno.
  - Un guion (`-`) es un grano de tueste medio que vale 2 puntos cada uno.
  - Un punto (`.`) es un grano de tueste oscuro que vale 3 puntos cada uno.

- El nivel de tueste se determina por el promedio de todos los granos.

Devuelve:

- `"Light"` si el promedio es menor que 1.75.
- `"Medium"` si el promedio está entre 1.75 y 2.5.
- `"Dark"` si el promedio es mayor que 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` debe devolver `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` debe devolver `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` debería devolver `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` debería devolver `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` debería devolver `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` debe devolver `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` debe devolver `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
