---
id: 69a890af247de743333bd4ce
title: "Desafío 225: Sin Repeticiones Consecutivas"
challengeType: 28
dashedName: challenge-225
---

# --description--

Dada una cadena, determina si no tiene caracteres repetidos.

- Una cadena no tiene repeticiones si no tiene el mismo carácter dos o más veces seguidas.

# --hints--

`hasNoRepeats("hi world")` debería devolver `true`.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` debería devolver `false`.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` debería devolver `true`.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` debería devolver `false`.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` debe devolver `true`.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` debe devolver `false`.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```
