---
id: 69a890af247de743333bd4d0
title: "Desafío 227: Tiempo de enfriamiento"
challengeType: 28
dashedName: challenge-227
---

# --description--

Dadas dos marcas de tiempo, la primera que representa cuándo un usuario terminó un examen, y la segunda que representa el tiempo actual, determina si el usuario puede volver a tomar un examen.

- Ambas marcas de tiempo tendrán el formato: `"YYYY-MM-DDTHH:MM:SS"`, por ejemplo `"2026-03-25T14:00:00"`. Ten en cuenta que la hora está en formato de 24 horas.
- Un usuario debe esperar al menos 48 horas antes de volver a presentar un examen.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` debe devolver `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` debe devolver `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` debe devolver `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` debe devolver `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
