---
id: 69a890af247de743333bd4cc
title: "Desafío 223: Decodificador QR"
challengeType: 28
dashedName: challenge-223
---

# --description--

Dada una matriz 6x6 (arreglo de arreglos), que representa un código QR, devuelve la cadena de datos binarios en el código.

- El código QR puede estar en cualquier rotación en incrementos de 90 grados.
- Un código correctamente orientado tiene un grupo 2x2 de `1`'s (marcadores de orientación) en las esquinas inferior izquierda, superior izquierda y superior derecha.
- Los tres marcadores de orientación 2x2 no forman parte de los datos binarios.
- Los datos binarios se leen de izquierda a derecha, de arriba hacia abajo (como un libro) cuando el código QR está correctamente orientado.
- Un código siempre tendrá exactamente una orientación válida.

Por ejemplo, dado:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

o dado el mismo código con una orientación diferente:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Devuelve `"000000000000000000000001"`, todos los datos binarios excluyendo los tres marcadores de orientación 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` debe devolver `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` debe devolver `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` debe devolver `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` debe devolver `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` debe devolver `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
