---
id: 69a890af247de743333bd4d1
title: "Desafío 228: Noche de Película"
challengeType: 28
dashedName: challenge-228
---

# --description--

Dado un string para el día de la semana, otro string para un horario, y un número entero de boletos, devuelve el costo total de las entradas de cine para esa función.

El día dado será uno de:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

La hora se mostrará en el formato `"H:MMam"` o `"H:MMpm"`. Por ejemplo, `"10:00am"` o `"10:00pm"`.

Devuelve el costo total en el formato `"$D.CC"` usando estas reglas:

- Fin de semana (viernes - domingo): $12.00 por boleto.
- Día laborable (lunes - jueves): $10.00 por boleto.
- Matinée (antes de las 5:00pm): resta $2.00 por boleto (excepto los martes).
- Los martes: todos los boletos cuestan $5.00 cada uno.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` debería devolver `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` debería devolver `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` debería devolver `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` debería devolver `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` debe devolver `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` debe devolver `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` debería devolver `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
