---
id: 69b1028d6e265413d0198a2c
title: "Desafío 233: alarma para despertar"
challengeType: 28
dashedName: challenge-233
---

# --description--

Dado un string que representa la hora a la que configuraste tu alarma y un string que representa la hora a la que realmente despertaste, determina si te despertaste temprano, a tiempo o tarde.

- Ambas horas se darán en formato `"HH:MM"` de 24 horas.

Devuelve:

- `"early"` si te despertaste antes de la hora de la alarma.
- `"on time"` si te despertaste a la hora de la alarma o dentro de los 10 minutos posteriores a la alarma (ventana de repetición).
- `"late"` si te despertaste más de 10 minutos después de la hora de la alarma.

Ambas horas corresponden al mismo día.

# --hints--

`alarmCheck("07:00", "06:45")` debería devolver `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` debería devolver `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` debería devolver `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` debería devolver `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` debería devolver `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` debería devolver `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` debería devolver `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
