---
id: 69b1028d6e265413d0198a2a
title: "Desafío 231: validador de ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Dado un string, determina si es un ISBN-10 válido.

Un ISBN-10 consiste en guiones (`"-"`) y 10 caracteres más. Después de eliminar los guiones (`"-"`):

- Los primeros 9 caracteres deben ser dígitos, y
- El último carácter puede ser un dígito o la letra `"X"`, que representa el número 10.

Para validarlo:

- Multiplica cada dígito (o valor) por su posición (multiplica el primer dígito por 1, el segundo por 2, y así sucesivamente).
- Suma todos los resultados.
- Si el total es divisible por 11, es válido.

# --hints--

`isValidIsbn10("0-306-40615-2")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` debería devolver `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` debería devolver `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
