---
id: 69b1028d6e265413d0198a2f
title: "Desafío 236: historial del navegador"
challengeType: 28
dashedName: challenge-236
---

# --description--

Dado un arreglo de comandos del navegador, devuelve un arreglo con dos valores: el historial como un arreglo de URLs y el índice de la página actual.

Los comandos válidos son:

- `"URL"` - donde URL es una dirección web (por ejemplo, `"freecodecamp.org"`). Navega a la URL dada, la añade al historial en la siguiente posición y descarta cualquier historial hacia adelante.
- `"Back"` - se mueve a la página anterior en el historial, o permanece en la página actual si no hay una anterior.
- `"Forward"` - se mueve a la página siguiente en el historial, o permanece en la página actual si no hay una siguiente.

Por ejemplo, dado `["freecodecamp.org", "freecodecamp.org/learn", "Back"]`, devuelve `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` debería devolver `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` debería devolver `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` debería devolver `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` debería devolver `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` debería devolver `[["example.com", "example.com/about"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` debería devolver `[["example.com", "example.com/about"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
