---
id: 69b559d2903b9e4afe9075f8
title: "Desafío 238: escape por rotación de dígitos"
challengeType: 28
dashedName: challenge-238
---

# --description--

Dado un número entero positivo, determina si él, o alguna de sus rotaciones, es divisible de manera exacta por la cantidad de sus dígitos.

Una rotación significa mover el primer dígito al final. Por ejemplo, después de 1 rotación, 123 se convierte en 231.

- Revisa primero la rotación `0` (el número dado).
- Los números dados no contendrán ceros.
- Devuelve el primer número rotado si se encuentra uno, o `"none"` si no.

# --hints--

`getRotation(123)` debería devolver `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` debería devolver `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` debería devolver `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` debería devolver `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
