---
id: 69b559d2903b9e4afe9075f9
title: "Desafío 239: ¿Qué día es?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Dado un timestamp Unix en milisegundos, devuelve el día de la semana.

Los días válidos para devolver son:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Asegúrate de ignorar las zonas horarias.

# --hints--

`getDayOfWeek(1775492249000)` debería devolver `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` debería devolver `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` debería devolver `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` debería devolver `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` debería devolver `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
