---
id: 69b58ce40693f140c84c8559
title: "Desafío 240: caracteres de palíndromo"
challengeType: 28
dashedName: challenge-240
---

# --description--

Dado un string, determina si es un palíndromo y devuelve el carácter del medio (si tiene longitud impar) o los dos caracteres del medio (si tiene longitud par).

- Un palíndromo es un string que es igual de adelante hacia atrás y de atrás hacia adelante.
- Si no es un palíndromo, devuelve `"none"`.

# --hints--

`palindromeLocator("racecar")` debería devolver `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` debería devolver `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` debería devolver `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` debería devolver `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` debería devolver `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
