---
id: 69b58ce40693f140c84c855a
title: "Desafío 241: validador de FizzBuzz"
challengeType: 28
dashedName: challenge-241
---

# --description--

Dado un arreglo de enteros secuenciales, con múltiplos de 3 y 5 reemplazados, determina si es una secuencia válida de FizzBuzz.

En una secuencia válida de FizzBuzz:

- Los múltiplos de 3 se reemplazan con "Fizz".
- Los múltiplos de 5 se reemplazan con "Buzz".
- Los múltiplos de 3 y 5 se reemplazan con "FizzBuzz".
- Todos los demás números permanecen como enteros.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz"])` debe devolver `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz"]));
```

`isFizzBuzz([13, 14, "FizzBuzz", 16, 17])` debe devolver `true`.

```js
assert.isTrue(isFizzBuzz([13, 14, "FizzBuzz", 16, 17]));
```

`isFizzBuzz([1, 2, "Fizz", 4, 5])` debe devolver `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, 5]));
```

`isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"])` debe devolver `true`.

```js
assert.isTrue(isFizzBuzz(["FizzBuzz", 16, 17, "Fizz", 19, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", "Buzz", 5])` debe devolver `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", "Buzz", 5]));
```

`isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103])` debe devolver `false`.

```js
assert.isFalse(isFizzBuzz([97, 98, "Buzz", "Fizz", 101, "Fizz", 103]));
```

`isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"])` debe devolver `true`.

```js
assert.isTrue(isFizzBuzz(["Fizz", "Buzz", 101, "Fizz", 103, 104, "FizzBuzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(arr) {

  return arr;
}
```

# --solutions--

```js
function isFizzBuzz(arr) {
  let start = null;
  for (let i = 0; i < arr.length; i++) {
    if (typeof arr[i] === "number") {
      start = arr[i] - i;
      break;
    }
  }

  if (start === null) return false;

  for (let i = 0; i < arr.length; i++) {
    const n = start + i;
    const expected = n % 15 === 0 ? "FizzBuzz" : n % 3 === 0 ? "Fizz" : n % 5 === 0 ? "Buzz" : n;
    if (arr[i] !== expected) return false;
  }

  return true;
}
```
