---
id: 69b5b2be76ec8135a7fbe973
title: "Desafío 244: ataque de torre y alfil"
challengeType: 28
dashedName: challenge-244
---

# --description--

Dadas dos cadenas que indican la ubicación de una torre y un alfil en un tablero de ajedrez, determina si una pieza puede atacar a la otra.

Un tablero de ajedrez estándar es de 8x8, con columnas etiquetadas `A` hasta `H` (de izquierda a derecha) y filas etiquetadas `1` hasta `8` (de abajo hacia arriba). Se ve así:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

- Las torres pueden moverse tantas casillas como quieran en dirección horizontal o vertical.
- Los alfiles pueden moverse tantas casillas como quieran en cualquier dirección diagonal.
- Una pieza puede atacar a otra si puede moverse a la ubicación de esa pieza.

Devuelve:

- `"rook"` si la torre puede atacar al alfil.
- `"bishop"` si el alfil puede atacar a la torre.
- `"neither"` si ninguna de las dos piezas puede atacarse mutuamente.

# --hints--

`rookBishopAttack("A1", "A5")` debería devolver `"rook"`.

```js
assert.equal(rookBishopAttack("A1", "A5"), "rook");
```

`rookBishopAttack("C3", "F6")` debería devolver `"bishop"`.

```js
assert.equal(rookBishopAttack("C3", "F6"), "bishop");
```

`rookBishopAttack("D4", "D7")` debería devolver `"rook"`.

```js
assert.equal(rookBishopAttack("D4", "D7"), "rook");
```

`rookBishopAttack("B7", "H1")` debería devolver `"bishop"`.

```js
assert.equal(rookBishopAttack("B7", "H1"), "bishop");
```

`rookBishopAttack("B3", "C5")` debería devolver `"neither"`.

```js
assert.equal(rookBishopAttack("B3", "C5"), "neither");
```

`rookBishopAttack("G3", "E8")` debería devolver `"neither"`.

```js
assert.equal(rookBishopAttack("G3", "E8"), "neither");
```

# --seed--

## --seed-contents--

```js
function rookBishopAttack(rook, bishop) {

  return rook;
}
```

# --solutions--

```js
function rookBishopAttack(rook, bishop) {
  const colDiff = Math.abs(rook.charCodeAt(0) - bishop.charCodeAt(0));
  const rowDiff = Math.abs(rook[1] - bishop[1]);

  if (rook[0] === bishop[0] || rook[1] === bishop[1]) return "rook";
  if (colDiff === rowDiff) return "bishop";
  return "neither";
}
```
