---
id: 69b559d2903b9e4afe9075f7
title: "Desafío 237: validación de ecuaciones"
challengeType: 28
dashedName: challenge-237
---

# --description--

Dada una cadena que representa una ecuación matemática, determina si es correcta.

- El lado izquierdo puede contener hasta tres enteros positivos y los operadores `+`, `-`, `*` y `/`.
- La ecuación se dará en el formato: `"number operator number = number"` (con dos o tres números a la izquierda). Por ejemplo: `"2 + 2 = 4"` o `"2 + 3 - 1 = 4"`.
- El lado derecho siempre será un solo entero.

Sigue el orden estándar de operaciones: la multiplicación y división se evalúan antes que la suma y resta, de izquierda a derecha.

# --hints--

`isValidEquation("2 + 2 = 4")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` debería devolver `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` debería devolver `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` debería devolver `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` debería devolver `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
