---
id: 69b58ce40693f140c84c855c
title: "Desafío 243: ataque de torre"
challengeType: 28
dashedName: challenge-243
---

# --description--

Dadas dos cadenas que indican la ubicación de dos torres en un tablero de ajedrez, determina si pueden atacarse entre sí.

Un tablero de ajedrez estándar es de 8x8, con columnas etiquetadas `A` hasta `H` (de izquierda a derecha) y filas etiquetadas `1` hasta `8` (de abajo hacia arriba). Se ve así:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Las torres pueden moverse tantas casillas como quieran en dirección horizontal o vertical. Por eso, si están en la misma fila o columna, pueden atacarse entre sí.

# --hints--

`rookAttack("A1", "A8")` debería devolver `true`.

```js
assert.isTrue(rookAttack("A1", "A8"));
```

`rookAttack("B4", "F4")` debería devolver `true`.

```js
assert.isTrue(rookAttack("B4", "F4"));
```

`rookAttack("E3", "D4")` debería devolver `false`.

```js
assert.isFalse(rookAttack("E3", "D4"));
```

`rookAttack("H7", "F6")` debería devolver `false`.

```js
assert.isFalse(rookAttack("H7", "F6"));
```

# --seed--

## --seed-contents--

```js
function rookAttack(rook1, rook2) {

  return rook1;
}
```

# --solutions--

```js
function rookAttack(rook1, rook2) {
  return rook1[0] === rook2[0] || rook1[1] === rook2[1];
}
```
