---
id: 69bc6cb30c1d112a2e110a06
title: "Desafío 249: matemáticas con cadenas"
challengeType: 28
dashedName: challenge-249
---

# --description--

Dada una cadena con números y otros caracteres, realiza operaciones matemáticas con los números según la cantidad de caracteres no numéricos que hay entre ellos.

- Si la cantidad de caracteres que separan dos números es par, usa suma.
- Si es impar, usa resta.
- Los dígitos consecutivos forman un solo número.
- Las operaciones se aplican de izquierda a derecha.
- Ignora los caracteres iniciales y finales que no sean dígitos.

Por ejemplo, dado `"3ab10c8"`, devuelve `5`. Suma 3 y 10 para obtener 13 porque hay un número par de caracteres entre ellos. Luego resta 8 de 13 porque hay un número impar de caracteres entre el resultado y 8.

# --hints--

`doMath("3ab10c8")` debería devolver `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` debería devolver `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` debería devolver `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` debería devolver `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` debería devolver `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
