---
id: 69bc6cb30c1d112a2e110a04
title: "Desafío 247: última letra"
challengeType: 28
dashedName: challenge-247
---

# --description--

Dada una cadena, devuelve la letra de la cadena que aparece última en el alfabeto.

- Si dos o más letras empatan por ser las últimas en el alfabeto, devuelve la primera de ellas.
- Ignora todos los caracteres que no sean letras.

# --hints--

`getLastLetter("world")` debería devolver `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` debería devolver `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` debería devolver `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` debería devolver `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` debería devolver `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
