---
id: 69bc6cb30c1d112a2e110a05
title: "Desafío 248: intercambio en arreglo ordenado"
challengeType: 28
dashedName: challenge-248
---

# --description--

Dado un arreglo de enteros, devuelve un nuevo arreglo usando las siguientes reglas:

1. Ordena los enteros en orden ascendente
2. Luego intercambia todos los valores cuyo índice sea múltiplo de 3 con el valor anterior a él.

# --hints--

`sortAndSwap([3, 1, 2, 4, 6, 5])` debería devolver `[1, 2, 4, 3, 5, 6]`.

```js
assert.deepEqual(sortAndSwap([3, 1, 2, 4, 6, 5]), [1, 2, 4, 3, 5, 6]);
```

`sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8])` debería devolver `[1, 2, 4, 3, 5, 7, 6, 8, 9]`.

```js
assert.deepEqual(sortAndSwap([9, 7, 5, 3, 1, 2, 4, 6, 8]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9])` debería devolver `[1, 2, 4, 3, 5, 7, 6, 8, 9]`.

```js
assert.deepEqual(sortAndSwap([1, 2, 3, 4, 5, 6, 7, 8, 9]), [1, 2, 4, 3, 5, 7, 6, 8, 9]);
```

`sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11])` debería devolver `[1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]`.

```js
assert.deepEqual(sortAndSwap([12, 5, 8, 1, 3, 10, 2, 7, 6, 4, 9, 11]), [1, 2, 4, 3, 5, 7, 6, 8, 10, 9, 11, 12]);
```

`sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25])` debería devolver `[-75, -50, 0, -25, 25, 75, 50, 100]`.

```js
assert.deepEqual(sortAndSwap([100, -50, 0, 75, -25, 50, -75, 25]), [-75, -50, 0, -25, 25, 75, 50, 100]);
```

`sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2])` debería devolver `[-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]`.

```js
assert.deepEqual(sortAndSwap([5, 9, 13, 77, 88, 313, -10, -65, 0, 8, 99, 101, -4, 2]), [-65, -10, 0, -4, 2, 8, 5, 9, 77, 13, 88, 101, 99, 313]);
```

# --seed--

## --seed-contents--

```js
function sortAndSwap(arr) {

  return arr;
}
```

# --solutions--

```js
function sortAndSwap(arr) {
  const result = [...arr].sort((a, b) => a - b);

  for (let i = 3; i < result.length; i += 3) {
    [result[i], result[i - 1]] = [result[i - 1], result[i]];
  }

  return result;
}
```
