---
id: 69bc6cb30c1d112a2e110a03
title: "Desafío 246: iniciales del nombre"
challengeType: 28
dashedName: challenge-246
---

# --description--

Dado un nombre completo como cadena, devuelve sus iniciales.

- Los nombres a inicializar están separados por un espacio.
- Las iniciales deben estar en mayúsculas.
- Las iniciales deben estar separadas por puntos.

Por ejemplo, `"Tommy Millwood"` devuelve `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` debería devolver `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` debería devolver `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` debería devolver `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` debería devolver `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` debería devolver `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
