---
id: 69bc6cb30c1d112a2e110a08
title: "Desafío 251: buscador de suma en un array"
challengeType: 28
dashedName: challenge-251
---

# --description--

Dado un array de números y un número objetivo, devuelve el primer subconjunto de dos o más números que sumen el objetivo.

- El "primer" subconjunto es aquel cuyos elementos tienen los índices más bajos posibles, priorizando primero el índice más temprano.
- Cada número en el array solo puede usarse una vez.
- Si no existe un subconjunto válido, devuelve `"Sum not found"`.

Devuelve los números que coincidan como un array en el orden en que aparecen en el array original.

# --hints--

`findSum([1, 3, 5, 7], 6)` debería devolver `[1, 5]`.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` debería devolver `[1, 4]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` debería devolver `[1, 2, 3]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` debería devolver `[-1, -2, 4]`.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` debería devolver `[3, 1, 4, 2]`.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` debería devolver `[1, 2, 3, 5, 9]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` debería devolver `"Sum not found"`.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```
