---
id: 69c5f3d787b1725d5f00c8ba
title: "Desafío 256: dirección de tiempo más cercana"
challengeType: 28
dashedName: challenge-256
---

# --description--

Dado dos horas, determina si puedes llegar de la primera a la segunda más rápido moviéndote hacia adelante o hacia atrás.

- Las horas se dan en formato de 24 horas (`"HH:MM"`)
- El reloj da la vuelta (las 23:59 pasan a 00:00 al avanzar, y las 00:00 pasan a 23:59 al retroceder)

Devuelve:

- `"forward"` si avanzar es más corto
- `"backward"` si retroceder es más corto
- `"equal"` si ambas direcciones toman el mismo tiempo

# --hints--

`getDirection("10:00", "12:00")` debe devolver `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` debe devolver `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` debe devolver `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` debe devolver `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` debe devolver `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` debe devolver `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
