---
id: 69c5f3d787b1725d5f00c8bc
title: "Descompresor de palabras desafío 258"
challengeType: 28
dashedName: challenge-258
---

# --description--

Dada una cadena comprimida, devuelve la versión descomprimida usando las siguientes reglas:

- La cadena dada está compuesta por palabras y números separados por espacios.
- Deja las palabras sin cambios.
- Reemplaza los números por la palabra en esa posición, donde la primera palabra está en la posición 1.

Por ejemplo, dado `"practice makes perfect and 3 1 2 3"`, devuelve `"practice makes perfect and perfect practice makes perfect"`.

# --hints--

`decompress("practice makes perfect and 3 1 2 3")` debería devolver `"practice makes perfect and perfect practice makes perfect"`.

```js
assert.equal(decompress("practice makes perfect and 3 1 2 3"), "practice makes perfect and perfect practice makes perfect");
```

`decompress("hello 1 1")` debería devolver `"hello hello hello"`.

```js
assert.equal(decompress("hello 1 1"), "hello hello hello");
```

`decompress("the cat sat on 1 mat 4 which 1 2 3")` debería devolver `"the cat sat on the mat on which the cat sat"`.

```js
assert.equal(decompress("the cat sat on 1 mat 4 which 1 2 3"), "the cat sat on the mat on which the cat sat");
```

`decompress("the more you know 1 2 3 realize 3 don't 4")` debería devolver `"the more you know the more you realize you don't know"`.

```js
assert.equal(decompress("the more you know 1 2 3 realize 3 don't 4"), "the more you know the more you realize you don't know");
```

`decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10")` debería devolver `"lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero"`.

```js
assert.equal(decompress("lorem ipsum dolor sit per elit donec 4 nostra libero 5 7 ligula 4 gravida at 6 vitae a 6 sodales 7 en 7 16 3 nam 13 dignissim risus 16 13 5 27 2 2 15 23 6 5 2 13 23 15 5 21 4 16 27 1 4 5 10 23 2 6 4 21 4 30 6 30 2 6 16 15 18 23 29 27 4 18 sollicitudin 5 9 5 4 10"), "lorem ipsum dolor sit per elit donec sit nostra libero per donec ligula sit gravida at elit vitae a elit sodales donec en donec at dolor nam ligula dignissim risus at ligula per nam ipsum ipsum gravida en elit per ipsum ligula en gravida per sodales sit at nam lorem sit per libero en ipsum elit sit sodales sit risus elit risus ipsum elit at gravida vitae en dignissim nam sit vitae sollicitudin per nostra per sit libero");
```

# --seed--

## --seed-contents--

```js
function decompress(str) {

  return str;
}
```

# --solutions--

```js
function decompress(str) {
  const tokens = str.split(" ");
  const words = [];

  for (const token of tokens) {
    if (isNaN(parseInt(token))) {
      words.push(token)
    } else {
      words.push(tokens[token - 1])
    }
  }

  return words.join(" ")
}
```
