---
id: 69c5f3d787b1725d5f00c8bd
title: "Desafío 259: explosión fizzbuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--

Dado un número entero, devuelve la cantidad de pasos que se necesitan para convertir la palabra `"fizzbuzz"` en una cadena con al menos el número dado de `"z"`'s usando las siguientes reglas:

- Comienza con la cadena `"fizzbuzz"`.
- En cada paso, aplica las reglas estándar de FizzBuzz usando la posición de la letra en la cadena (la primera `"f"` está en la posición 1).
  - Si la posición de la letra es divisible por 3, reemplaza la letra con `"fizz"`
  - Si es divisible por 5, reemplaza la letra con `"buzz"`
  - Si es divisible por 3 y 5, reemplaza la letra con `"fizzbuzz"`

Así que después de 1 paso, `"fizzbuzz"` se convierte en `"fifizzzbuzzfizzzz"`, que tiene 9 `"z"`'s.

# --hints--

`explodeFizzbuzz(9)` debería devolver `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` debería devolver `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` debería devolver `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` debería devolver `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` debería devolver `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` debería devolver `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` debería devolver `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` debería devolver `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
