---
id: 69c5f3d787b1725d5f00c8b7
title: "Desafío 253: buscador de acrónimos"
challengeType: 28
dashedName: challenge-253
---

# --description--

Dado un string que representa un acrónimo, devuelve el nombre completo de la organización a la que pertenece según la lista siguiente:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

Cada letra del acrónimo dado debe coincidir con la primera letra de cada palabra en la organización a la que pertenece, en el mismo orden.

# --hints--

`findOrg("NASA")` debería devolver `"National Avocado Storage Authority"`.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` debería devolver `"Cats Infiltration Agency"`.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` debería devolver `"Fluffy Beanbag Inspectors"`.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` debería devolver `"Department Of Jelly"`.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` debería devolver `"Wild Honey Organization"`.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` debería devolver `"Eating Pancakes Administration"`.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
