---
id: 69c5f3d787b1725d5f00c8b9
title: "Desafío 255: equipo de limpieza del día de la tierra"
challengeType: 28
dashedName: challenge-255
---

# --description--

Hoy es el día de la Tierra. Dado un arreglo de objetos que limpiaste, devuelve tu puntaje total de limpieza según las reglas siguientes.

Los objetos dados serán uno de:

| Objeto | Valor base |
| - | - |
| `"bottle"` | 10 |
| `"can"` | 6 |
| `"bag"` | 8 |
| `"tire"` | 35 |
| `"straw"` | 4 |
| `"cardboard"` | 3 |
| `"newspaper"` | 3 |
| `"shoe"` | 12 |
| `"electronics"` | 25 |
| `"battery"` | 18 |
| `"mattress"` | 38 |

Un objeto raro está representado como `["rare", value]`. Por ejemplo, `["rare", 80]`. Los objetos raros no reciben bonificación por racha.

- Bonificación por racha: si el mismo objeto aparece consecutivamente, recibe puntos de bonificación crecientes.
  - Primera ocurrencia consecutiva: valor base
  - Segunda: valor base + 1
  - Tercera: valor base + 2
  - etc.
- Multiplicador en el quinto objeto: cada quinto objeto recogido recibe un multiplicador.
  - Quinto objeto: *2
  - Décimo objeto: *3
  - etc.

- Aplica el multiplicador después de calcular cualquier bonificación.

# --hints--

`getCleanupScore(["bottle", "straw", "shoe", "battery"])` debería devolver `44`.

```js
assert.equal(getCleanupScore(["bottle", "straw", "shoe", "battery"]), 44);
```

`getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"])` debería devolver `58`.

```js
assert.equal(getCleanupScore(["electronics", "straw", "newspaper", "bottle", "bag"]), 58);
```

`getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"])` debería devolver `79`.

```js
assert.equal(getCleanupScore(["shoe", "can", "can", "can", "bottle", "bottle", "straw", "straw", "straw"]), 79);
```

`getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]])` debería devolver `358`.

```js
assert.equal(getCleanupScore(["mattress", ["rare", 80], "tire", "tire", "tire", ["rare", 95]]), 358);
```

`getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"])` debería devolver `383`.

```js
assert.equal(getCleanupScore(["bottle", "can", "can", "shoe", "shoe", ["rare", 56], "bottle", "bottle", "can", "can", "electronics", "bottle", ["rare", 48], "bottle", "can", "can", "can", "can", "can", "can", "can"]), 383);
```

# --seed--

## --seed-contents--

```js
function getCleanupScore(items) {

  return items;
}
```

# --solutions--

```js
function getCleanupScore(items) {
  const values = {
    bottle: 10,
    can: 6,
    bag: 8,
    tire: 35,
    straw: 4,
    cardboard: 3,
    newspaper: 3,
    shoe: 12,
    electronics: 25,
    battery: 18,
    mattress: 38
  };

  let total = 0;
  let prevItem = null;
  let streak = 0;

  for (let i = 0; i < items.length; i++) {
    let item = items[i];
    let base;
    let isRare = Array.isArray(item) && item[0] === "rare";

    if (isRare) {
      base = item[1];
      prevItem = null;
      streak = 0;
    } else {
      base = values[item];
      streak = item === prevItem ? streak + 1 : 0;
      prevItem = item;
    }

    let points = base;
    if (!isRare) points += streak;

    if ((i + 1) % 5 === 0) {
      const multiplier = (i + 1) / 5 + 1;
      points *= multiplier;
    }

    total += points;
  }

  return total;
}
```
