---
id: 69c5f3d787b1725d5f00c8b8
title: "Desafío 254: palabras con número impar de letras"
challengeType: 28
dashedName: challenge-254
---

# --description--

Dada una cadena de palabras, devuelve solo las palabras que tengan un número impar de letras.

- Las palabras en la cadena dada estarán separadas por un solo espacio.
- Devuelve las palabras separadas por un solo espacio.

# --hints--

`getOddWords("This is a super good test")` debería devolver `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` debería devolver `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` debería devolver `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` debería devolver `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
