---
id: 69c6ff713a52713463aa7929
title: "Desafío 263: crucigrama binario"
challengeType: 28
dashedName: challenge-263
---

# --description--

Dado un carácter, determina si su representación binaria de 8 bits se puede encontrar en la siguiente cuadrícula, horizontal o verticalmente en cualquiera de las dos direcciones:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Por ejemplo, `"A"` tiene la representación binaria `01000001`, que aparece en la primera fila de izquierda a derecha.

# --hints--

`isInCrossword("I")` debería devolver `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` debería devolver `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` debería devolver `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` debería devolver `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` debería devolver `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` debería devolver `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` debería devolver `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` debería devolver `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
