---
id: 69c6ff713a52713463aa7926
title: "Desafío 260: puntuación de palabras"
challengeType: 28
dashedName: challenge-260
---

# --description--

Dada una palabra, devuelve su puntuación usando una tabla estándar de valores por letra:

| Letra | Valor |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Las letras mayúsculas y minúsculas tienen el mismo valor.

# --hints--

`getWordScore("hi")` debería devolver `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` debería devolver `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` debería devolver `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` debería devolver `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
