---
id: 69c6ff713a52713463aa7927
title: "Desafío 261: palabras de números"
challengeType: 28
dashedName: challenge-261
---

# --description--

Dado un número entero del 0 al 99, devuelve su representación en palabras en inglés.

- `0` devuelve `"zero"`.
- Los números del 1 al 19 tienen nombres únicos (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Los múltiplos de 10 del 20 al 90 tienen sus propios nombres (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Los números del 21 al 99 que no son múltiplos de 10 se escriben como dos palabras unidas por un guion. Por ejemplo `"forty-two"` y `"fifty-three"`.

# --hints--

`getNumberWords(0)` debería devolver `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` debería devolver `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` debería devolver `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` debería devolver `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` debería devolver `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` debería devolver `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` debería devolver `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` debería devolver `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` debería devolver `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` debería devolver `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
