---
id: 69c6ff713a52713463aa7928
title: "Desafío 262: analizador de consultas URL"
challengeType: 28
dashedName: challenge-262
---

# --description--

Dada una URL que contiene una cadena de consulta, analiza la cadena de consulta en un objeto (o diccionario) de pares clave-valor.

- La cadena de consulta comienza después de `"?"`,
- cada parámetro está separado por `"&"`,
- cada par clave/valor está separado por `"="`

Por ejemplo, dada `"https://example.com/search?name=Alice&age=30"`, devuelve:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Todos los valores deben devolverse como cadenas. 

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` debería devolver `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` debería devolver `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` debería devolver `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` debería devolver `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
