---
id: 69cfca90e8a0a6d4d6871c51
title: "Desafío 267: conversor de parsec"
challengeType: 28
dashedName: challenge-267
---

# --description--

En una galaxia lejana, los parsecs se usan para medir tanto el tiempo como la distancia. Dado un número entero de parsecs, devuelve su equivalente en tiempo o distancia.

- Si el entero dado es impar, representa tiempo. Si es par, representa distancia.

Usa estas tasas de conversión:

| Parsecs | Tiempo/Distancia |
| - | - |
| 1 | 2 horas |
| 2 | 6 años luz |

Devuelve el valor convertido como un entero.

# --hints--

`convertParsecs(1)` debería devolver `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` debería devolver `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` debería devolver `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` debería devolver `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` debería devolver `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` debería devolver `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
