---
id: 69cfca90e8a0a6d4d6871c52
title: "Desafío 268: número narcisista"
challengeType: 28
dashedName: challenge-268
---

# --description--

Dado un número entero positivo, determina si es un número narcisista.

- Un número es narcisista si la suma de cada uno de sus dígitos elevados a la potencia del total de dígitos es igual al mismo número.

Por ejemplo, 153 tiene 3 dígitos, y 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, por lo que es narcisista.

# --hints--

`isNarcissistic(153)` debería devolver `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` debería devolver `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` debería devolver `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` debería devolver `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` debería devolver `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` debería devolver `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` debería devolver `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` debería devolver `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
