---
id: 69cfca90e8a0a6d4d6871c50
title: "Desafío 266: buen día"
challengeType: 28
dashedName: challenge-266
---

# --description--

Dada una cadena de tiempo en formato `"HH:MM"` (reloj de 24 horas), devuelve:

- `"Good morning"` para horas `05:00` a `11:59`
- `"Good afternoon"` para horas `12:00` a `17:59`
- `"Good evening"` para horas `18:00` a `21:59`
- `"Good night"` para horas `22:00` a `04:59`

# --hints--

`getGreeting("06:30")` debería devolver `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` debería devolver `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` debería devolver `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` debería devolver `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` debería devolver `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
