---
id: 69e2383af7832c8032603b94
title: "Desafío 278: analizador de pedidos de café"
challengeType: 28
dashedName: challenge-278
---

# --description--

Dada una cadena con un pedido de café, identifica los artículos del menú y devuelve un pedido formateado.

Usa los siguientes artículos del menú y precios:

| Artículo | Precio | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

Devuelve una cadena con los artículos coincidentes unidos por `" + "`, seguidos de dos puntos y un espacio (`": "`), y el precio total.

Por ejemplo, dado `"I'd like an oat latte with vanilla syrup and an extra shot please."`, devuelve `"oat latte + vanilla syrup + extra shot: $6.25"`

Los artículos deben aparecer en el orden en que aparecen en el menú y el precio total siempre debe tener dos decimales.

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` debería devolver `"oat latte + vanilla syrup + extra shot: $6.25"`.

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` debería devolver `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`.

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` debería devolver `"cold brew + extra shot + cream: $5.75"`.

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` debería devolver `"espresso: $3.00"`.

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` debería devolver `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`.

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
