---
id: 69e2383af7832c8032603b90
title: "Desafío 274: persona más vieja"
challengeType: 28
dashedName: challenge-274
---

# --description--

Dado un arreglo de objetos, cada uno con una propiedad `"name"` y `"age"`, devuelve un arreglo que contenga el nombre de la persona más vieja.

Si varias personas comparten la edad más avanzada, devuelve todos sus nombres en el orden en que aparecen en la entrada.

# --hints--

`getOldest([{ name: "Brenda", age: 40 }])` debería devolver `["Brenda"]`.

```js
assert.deepEqual(getOldest([{ name: "Brenda", age: 40 }]), ["Brenda"]);
```

`getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }])` debería devolver `["Alice"]`.

```js
assert.deepEqual(getOldest([{ name: "Alice", age: 30 }, { name: "Bob", age: 25 }]), ["Alice"]);
```

`getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }])` debería devolver `["Bill", "Carol"]`.

```js
assert.deepEqual(getOldest([{ name: "Allison", age: 25 }, { name: "Bill", age: 30 }, { name: "Carol", age: 30 }]), ["Bill", "Carol"]);
```

`getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }])` debería devolver `["George", "Holly", "Zach"]`.

```js
assert.deepEqual(getOldest([{ name: "George", age: 50 }, { name: "Shirley", age: 42 }, { name: "Beth", age: 48 }, { name: "Holly", age: 50 }, { name: "Kevin", age: 44 }, { name: "Frank", age: 47 }, { name: "Zach", age: 50 }, { name: "Jennifer", age: 43 }]), ["George", "Holly", "Zach"]);
```

# --seed--

## --seed-contents--

```js
function getOldest(people) {

  return people;
}
```

# --solutions--

```js
function getOldest(people) {
  const maxAge = Math.max(...people.map(p => p.age));
  return people.filter(p => p.age === maxAge).map(p => p.name);
}
```
