---
id: 69e2383af7832c8032603b8f
title: "Desafío 273: validador de ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

Dado un string, determina si es un número ISBN-13 válido.

Un ISBN-13 válido:

- Contiene solo dígitos y guiones
- Tiene exactamente 13 dígitos después de eliminar los guiones
- Pasa la siguiente verificación:
  1. Multiplica cada dígito por 1 o 3, alternando (multiplica el primer dígito por 1, el segundo por 3, el tercero por 1, y así sucesivamente).
  2. La suma de los resultados debe ser divisible por 10.

# --hints--

`isValidIsbn13("9780306406157")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` debería devolver `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` debería devolver `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` debería devolver `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
