---
id: 69f35a5bb823ed620fcb7cbb
title: "Desafío 282: deuda de sueño"
challengeType: 28
dashedName: challenge-282
---

# --description--

Dado un arreglo con las horas dormidas cada noche hasta hoy, y un número objetivo de horas por noche, devuelve cuántas horas necesitas dormir esta noche para eliminar tu deuda de sueño.

- Incluye las horas de esta noche en el total necesario para ponerte al día.
- Si ya has dormido lo suficiente para cubrir o superar el objetivo de esta noche, devuelve `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` debería devolver `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` debería devolver `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` debería devolver `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` debería devolver `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` debería devolver `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
