---
id: 69f35a5bb823ed620fcb7cbc
title: "Desafío 283: entrelazar cadenas"
challengeType: 28
dashedName: challenge-283
---

# --description--

Dadas dos cadenas, devuelve una nueva cadena que entrelace sus caracteres uno a uno. Si una cadena es más larga, añade los caracteres restantes al final.

Comienza con el primer carácter de la primera cadena.

# --hints--

`zipStrings("abc", "123")` debería devolver `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` debería devolver `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` debería devolver `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` debería devolver `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` debería devolver `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
