---
id: 69f35a5bb823ed620fcb7cbd
title: "Desafío 284: i antes de e"
challengeType: 28
dashedName: challenge-284
---

# --description--

Dada una palabra o frase, devuelve una versión corregida donde cada palabra siga la regla "i antes de e excepto después de c".

- Si una palabra contiene `"ei"` que no está precedida por `"c"`, reemplázala por `"ie"`.
- Si una palabra contiene `"ie"` precedida por `"c"`, reemplázala por `"ei"`.
- Todas las demás palabras se dejan sin cambios.

# --hints--

`iBeforeE("beleive")` debería devolver `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` debería devolver `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` debería devolver `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` debería devolver `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` debería devolver `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
