---
id: 69f8c998d78ad3171a0713bf
title: "Desafío 292: relación de aspecto más ancha"
challengeType: 28
dashedName: challenge-292
---

# --description--

Dadas dos cadenas que representan diferentes dimensiones de imagen, devuelve la relación de aspecto de la imagen con una mayor proporción ancho-alto.

- Las cadenas dadas estarán en el formato `"WxH"`, por ejemplo, `"1920x1080"`.
- La relación de aspecto es la proporción entre el ancho y el alto, reducida a los números enteros más bajos. Por ejemplo, `"1920x1080"` se reduce a `"16:9"`.
- Devuelve una cadena en formato `"W:H"`, por ejemplo, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` debería devolver `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` debería devolver `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` debería devolver `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` debería devolver `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` debería devolver `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` debería devolver `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
