---
id: 69f8c998d78ad3171a0713bb
title: "Desafío 288: número secreto"
challengeType: 28
dashedName: challenge-288
---

# --description--

Dado un número secreto y una suposición, determina si la suposición es correcta.

Devuelve:

- `"higher"` si el número secreto es mayor que la suposición.
- `"lower"` si el número secreto es menor que la suposición.
- `"you got it!"` si la suposición es correcta.

# --hints--

`guessNumber(50, 30)` debe devolver `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` debe devolver `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` debe devolver `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` debe devolver `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` debe devolver `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` debe devolver `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
