---
id: 69f8c998d78ad3171a0713bd
title: "Desafío 290: fiesta de pizza"
challengeType: 28
dashedName: challenge-290
---

# --description--

Dado un arreglo con las horas trabajadas hoy por persona, devuelve la cantidad de pizzas que se deben pedir para una fiesta de pizza.

- Divide las horas trabajadas de cada persona entre 3 para obtener la cantidad de porciones.
- No puedes comer una porción parcial, así que redondea hacia arriba la cantidad de porciones de cada persona al número entero más cercano.
- Cada persona recibe un mínimo de dos porciones.
- Cada pizza tiene 8 porciones. Redondea hacia arriba el total de pizzas al número entero más cercano.

# --hints--

`getPizzasToOrder([8, 8, 8])` debería devolver `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` debería devolver `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` debería devolver `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` debería devolver `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` debería devolver `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` debería devolver `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
