---
id: 69f8c998d78ad3171a0713b9
title: "Desafío 286: issues abiertas"
challengeType: 28
dashedName: challenge-286
---

# --description--

Dado un arreglo de números de issues y otro arreglo de números de pull requests (PR), devuelve un arreglo con los issues que permanecen abiertos después de que todos los PR se hayan fusionado.

- Un PR cierra un issue si sus dígitos son una rotación uno del otro. Por ejemplo, el issue 123 sería cerrado por el PR 231 o 312.
- Un PR no cierra un issue con el mismo número exacto. Por ejemplo, el PR 123 no cierra el issue 123. Así que un issue con todos los mismos números no puede cerrarse.
- Cualquiera de los números puede tener ceros a la izquierda eliminados. Por ejemplo, el PR 201 cerraría el issue 12 (012, una rotación de 201). De manera similar, el issue 201 sería cerrado por el PR 12.

Devuelve los issues abiertos restantes en el orden en que fueron dados.

# --hints--

`getOpenIssues([123, 234], [231])` debería devolver `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` debería devolver `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` debería devolver `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` debería devolver `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` debería devolver `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
