---
id: 69f8c998d78ad3171a0713be
title: "Desafío 291: conteo de fizzbuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--

Dado un número inicial y uno final, cuenta la cantidad de apariciones de fizz y buzz en el rango (inclusive).

- Los números divisibles por 3 cuentan como fizz.
- Los números divisibles por 5 cuentan como buzz.
- Los números divisibles tanto por 3 como por 5 cuentan como fizz y buzz.

Devuelve un objeto o diccionario con los conteos en el formato: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` debería devolver `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` debería devolver `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` debería devolver `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` debería devolver `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` debería devolver `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
