---
id: 69f8c998d78ad3171a0713ba
title: "Desafío 287: corregir números romanos"
challengeType: 28
dashedName: challenge-287
---

# --description--

Dado un string de números romanos mal formados, devuelve el valor en notación romana estándar.

La entrada solo usará notación aditiva, por lo que cada símbolo suma su valor al total. Como recordatorio, aquí están los símbolos y sus valores:

| Símbolo | Valor |
| ------ | ----- |
| `"I"`  | 1     |
| `"V"`  | 5     |
| `"X"`  | 10    |
| `"L"`  | 50    |
| `"C"`  | 100   |
| `"D"`  | 500   |
| `"M"`  | 1000  |

Al volver a codificar, usa el símbolo más grande posible en cada paso, usando pares sustractivos (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) cuando sea necesario.

# --hints--

`fixNumerals("XIIIII")` debería devolver `"XV"`.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` debería devolver `"LXIV"`.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` debería devolver `"XL"`.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` debería devolver `"MDCCXCIX"`.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` debería devolver `"MCDLXIV"`.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` debería devolver `"MMCMLXXXIV"`.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
