---
id: 69f90c1329a94b37e2a2086d
title: "Desafío 294: combinaciones de paréntesis"
challengeType: 28
dashedName: challenge-294
---

# --description--

Dado un entero, `n`, devuelve el número de combinaciones válidas de `n` pares de paréntesis.

- Una combinación válida es una cadena donde cada paréntesis de apertura tiene un paréntesis de cierre correspondiente, y ningún paréntesis de cierre aparece antes que su paréntesis de apertura correspondiente.

Por ejemplo, dado `2`, hay 2 combinaciones válidas:

```json
(())
()()
```

# --hints--

`getCombinations(2)` debería devolver `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` debería devolver `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` debería devolver `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` debería devolver `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` debería devolver `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
