---
id: 69f90c1329a94b37e2a2086c
title: "Desafío 293: mejor mano"
challengeType: 28
dashedName: challenge-293
---

# --description--

Dado un arreglo de cinco cadenas que representan cartas de juego, devuelve el nombre de la mejor mano.

- Cada carta está representada como una cadena de dos caracteres: el rango seguido del palo, `"2h"` por ejemplo.
  - Los rangos, de menor a mayor, son: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"` y `"A"`.
  - Los palos son: `"h"`, `"d"`, `"c"` y `"s"`.
- Los ases (`"A"`) pueden usarse como altos o bajos en una escalera.

Las manos, en orden de peor a mejor, son:

| Nombre | Descripción |
| - | - |
| `"High Card"` | Sin pareja ni mejor combinación |
| `"Pair"` | Dos cartas del mismo rango |
| `"Two Pair"` | Dos cartas de un rango y dos de otro |
| `"Three of a Kind"` | Tres cartas del mismo rango |
| `"Straight"` | Cinco rangos consecutivos |
| `"Flush"` | Cinco cartas del mismo palo |
| `"Full House"` | Tres cartas de un rango y dos de otro |
| `"Four of a Kind"` | Cuatro cartas del mismo rango |
| `"Straight Flush"` | Cinco rangos consecutivos del mismo palo |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` del mismo palo |

Devuelve el nombre de la mejor mano.

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` debería devolver `"Three of a Kind"`.

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` debería devolver `"Full House"`.

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` debería devolver `"Flush"`.

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` debería devolver `"Four of a Kind"`.

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` debería devolver `"Two Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` debería devolver `"Straight Flush"`.

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` debería devolver `"High Card"`.

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` debería devolver `"Straight"`.

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` debería devolver `"Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` debería devolver `"Royal Flush"`.

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```
