---
id: 6a0dcd03ee4e68698080ef6d
title: "Desafío 309: ordenar números"
challengeType: 28
dashedName: challenge-309
---

# --description--

Dada una cadena de números separados por comas, devuelve un arreglo con los números ordenados de menor a mayor.

# --hints--

`sortNumbers("3,1,2")` debería devolver `[1, 2, 3]`.

```js
assert.deepEqual(sortNumbers("3,1,2"), [1, 2, 3]);
```

`sortNumbers("5,3,8,1,9,2")` debería devolver `[1, 2, 3, 5, 8, 9]`.

```js
assert.deepEqual(sortNumbers("5,3,8,1,9,2"), [1, 2, 3, 5, 8, 9]);
```

`sortNumbers("12,61,49,80,19,50,77,38")` debería devolver `[12, 19, 38, 49, 50, 61, 77, 80]`.

```js
assert.deepEqual(sortNumbers("12,61,49,80,19,50,77,38"), [12, 19, 38, 49, 50, 61, 77, 80]);
```

`sortNumbers("0,6,-19,44,-2,7,0")` debería devolver `[-19, -2, 0, 0, 6, 7, 44]`.

```js
assert.deepEqual(sortNumbers("0,6,-19,44,-2,7,0"), [-19, -2, 0, 0, 6, 7, 44]);
```

# --seed--

## --seed-contents--

```js
function sortNumbers(str) {

  return str;
}
```

# --solutions--

```js
function sortNumbers(str) {
  return str.split(",").map(Number).sort((a, b) => a - b);
}
```
