---
id: 6a0dcd03ee4e68698080ef6c
title: "Desafío 308: validador de tarjeta de crédito"
challengeType: 28
dashedName: challenge-308
---

# --description--

Dada una cadena de dígitos para un número de tarjeta de crédito, determina si es un número válido usando el siguiente método:

- Comenzando desde el penúltimo dígito, duplica cada otro dígito moviéndote hacia la izquierda.
- Si al duplicar un dígito el resultado es un número mayor que 9, réstale 9.
- Suma todos los dígitos (duplicados y no duplicados).
- Si el total es divisible por 10, el número es válido.

# --hints--

`isValidCard("4532015112830366")` debería devolver `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` debería devolver `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` debería devolver `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` debería devolver `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` debería devolver `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` debería devolver `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` debería devolver `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
