---
id: 6a0dcd03ee4e68698080ef66
title: "Desafío 302: Jet lag"
challengeType: 28
dashedName: challenge-302
---

# --description--

Dadas una ciudad de salida, una ciudad de llegada, una duración de vuelo en horas y una dirección de viaje, devuelve el número de horas de jet lag que experimenta el viajero.

Las ciudades dadas serán de la siguiente lista que incluye su diferencia horaria UTC:

| Ciudad | Diferencia |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Para calcular las horas de jet lag:

1. Encuentra la diferencia de zona horaria en horas entre las dos ciudades.
2. Determina el multiplicador de dirección. Si viajas `"east"`, es 1.5; de lo contrario, es 1.0.
3. Obtén las horas de jet lag con la fórmula: diferencia de zona horaria + (duración del vuelo * 0.1) * multiplicador de dirección

Devuelve las horas de jet lag redondeadas a un decimal.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` debería devolver `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` debería devolver `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` debería devolver `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` debería devolver `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` debería devolver `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` debería devolver `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` debería devolver `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
