---
id: 6a0dcc730cb92a616f86f0bf
title: "Desafío 295: validador de esquema parte 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Dado un objeto (JavaScript) o diccionario (Python), determina si coincide con el siguiente esquema:

```json
{
  username: string
}
```

- Se permiten claves adicionales

# --hints--

`isValidSchema({ username: "bob" })` debería devolver `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` debería devolver `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` debería devolver `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` debería devolver `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` debería devolver `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
