---
id: 6a0dcd03ee4e68698080ef6b
title: "Desafío 307: regulaciones de zonificación"
challengeType: 28
dashedName: challenge-307
---

# --description--

Dada una rejilla 2D (arreglo de arreglos) que representa el diseño de edificios de una ciudad, devuelve las coordenadas de todos los edificios que están violando las reglas de zonificación.

Cada celda en la rejilla contiene una de las etiquetas de la tabla a continuación. Un edificio está en violación si cualquiera de sus (hasta) 4 vecinos, horizontales o verticales, es de un tipo al que no puede estar adyacente.

| Etiqueta | Tipo | No puede estar adyacente a |
| - | - | - |
| `"i"` | industrial | `"R"`, `"I"` |
| `"A"` | agrícola | `"C"` |
| `"R"` | residencial | `"i"`, `"C"` |
| `"I"` | institucional | `"i"` |
| `"C"` | comercial | `"R"`, `"A"` |
| `""` (cadena vacía) | sin desarrollar | sin restricciones |

Devuelve las coordenadas de todas las celdas que violan las reglas como un arreglo de pares `[row, col]`, en cualquier orden. Si no existen violaciones, devuelve un arreglo vacío.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` debería devolver `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` debería devolver `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` debería devolver `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` debería devolver `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` debería devolver `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
