---
id: 6a0dcd03ee4e68698080ef68
title: "Desafío 304: arreglos de itinerarios"
challengeType: 28
dashedName: challenge-304
---

# --description--

Dado un arreglo de al menos dos paradas opcionales para una excursión de un día, devuelve el número de arreglos válidos del itinerario.

El itinerario siempre incluye `"breakfast"`, `"lunch"` y `"dinner"`, estos no se pasan como argumentos. Las paradas opcionales pueden colocarse en cualquier parte del itinerario, sujeto a las siguientes reglas:

- `"breakfast"` siempre es la primera, con al menos una parada antes de `"lunch"`.
- `"lunch"` debe aparecer antes de `"dinner"`, con al menos una parada entre ellas.
- Como máximo, una parada opcional puede aparecer después de `"dinner"`.

Devuelve el número de arreglos válidos.

# --hints--

`getItineraryCount(["library", "park"])` debería devolver `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` debería devolver `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` debería devolver `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` debería devolver `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` debería devolver `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
