---
id: 6a0dcc730cb92a616f86f0c5
title: "Desafío 301: última carga"
challengeType: 28
dashedName: challenge-301
---

# --description--

Dado el número de cucharadas de detergente para ropa que te quedan y un arreglo con la cantidad de cucharadas que usaste en cada uno de los días anteriores, devuelve el número de días completos de detergente que te quedan.

Calcula tu uso diario promedio a partir del historial de uso y asume esa cantidad de uso cada día en adelante.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` debería devolver `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` debería devolver `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` debería devolver `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` debería devolver `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` debería devolver `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
