---
id: 6a15cadf5f240d05a264955c
title: "Desafío 318: mutaciones en el ADN"
challengeType: 28
dashedName: challenge-318
---

# --description--

Dadas dos cadenas de ADN de igual longitud, devuelve un arreglo con los índices donde las cadenas difieren (mutaciones).

- Las cadenas de ADN son cadenas formadas por los caracteres `"A"`, `"T"`, `"C"` y `"G"`
- Devuelve los índices en orden ascendente
- Si no hay mutaciones, devuelve un arreglo vacío

# --hints--

`detectMutations("ATCG", "ATGG")` debería devolver `[2]`.

```js
assert.deepEqual(detectMutations("ATCG", "ATGG"), [2]);
```

`detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` debería devolver `[4, 9, 11]`.

```js
assert.deepEqual(detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11]);
```

`detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` debería devolver `[]`.

```js
assert.deepEqual(detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), []);
```

`detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` debería devolver `[15, 16, 17, 18]`.

```js
assert.deepEqual(detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18]);
```

`detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` debería devolver `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`.

```js
assert.deepEqual(detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]);
```

# --seed--

## --seed-contents--

```js
function detectMutations(strand1, strand2) {

  return strand1;
}
```

# --solutions--

```js
function detectMutations(strand1, strand2) {
  return strand1
    .split("")
    .reduce((mutations, base, i) => {
      if (base !== strand2[i]) mutations.push(i);
      return mutations;
    }, []);
}
```
