---
id: 6a151d32d772271dd2448c2d
title: "Desafío 311: lanzador de hechizos"
challengeType: 28
dashedName: challenge-311
---

# --description--

Dada una cadena de códigos de hechizos que estás lanzando, calcula la puntuación total.

Cada carácter en la cadena representa un hechizo:

| Código | Hechizo | Categoría | Puntuación base |
| - | - | - | - |
| `"f"` | Fuego | Destrucción | 3 |
| `"l"` | Rayo | Destrucción | 3 |
| `"i"` | Hielo | Control | 2 |
| `"w"` | Viento | Control | 2 |
| `"h"` | Curar | Restauración | 1 |
| `"s"` | Escudo | Restauración | 1 |

Se aplica un multiplicador de combo según cuántos hechizos consecutivos se hayan lanzado de diferentes categorías:

- El primer hechizo siempre puntúa con su valor base.
- Cada hechizo consecutivo de una categoría diferente a la anterior aumenta el multiplicador en 1.
- Lanzar un hechizo de la misma categoría que el anterior reinicia el multiplicador a 1.
- La puntuación de cada hechizo es su puntuación base multiplicada por el multiplicador actual.

Devuelve la puntuación total de la secuencia de hechizos.

# --hints--

`cast("fihwl")` debería devolver `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` debería devolver `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` debería devolver `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` debería devolver `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` debería devolver `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
