---
id: 6a15cadf5f240d05a264955b
title: "Desafío 317: calculadora de IMC"
challengeType: 28
dashedName: challenge-317
---

# --description--

Dado un peso en libras y una altura en pulgadas, devuelve el IMC (índice de masa corporal) redondeado a un decimal.

Para obtener el IMC: divide el peso por la altura al cuadrado, luego multiplica el resultado por 703.

# --hints--

`calculateBmi(180, 70)` debe devolver `25.8`.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` debe devolver `24.0`.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` debe devolver `19.5`.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` debe devolver `39.1`.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` debe devolver `22.8`.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
