---
id: 6a15cadf5f240d05a2649558
title: "Desafío 314: factorización en factores primos"
challengeType: 28
dashedName: challenge-314
---

# --description--

Dado un entero mayor que 1, devuelve su factorización en factores primos como un arreglo de números en orden ascendente.

Una factorización en factores primos es el conjunto de números primos que se multiplican para producir el entero dado. Cada número tiene exactamente un conjunto. Por ejemplo, la factorización en factores primos de 20 es `[2, 2, 5]` porque 2 * 2 * 5 = 20.

Si el entero dado es primo, devuélvelo en un arreglo de un solo elemento.

# --hints--

`primeFactorization(20)` debería devolver `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` debería devolver `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` debería devolver `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` debería devolver `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` debería devolver `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` debería devolver `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` debería devolver `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
