---
id: 6a15cadf5f240d05a2649557
title: "Desafío 313: costo del alquiler"
challengeType: 28
dashedName: challenge-313
---

# --description--

Dado un sello de tiempo de alquiler, un sello de tiempo de devolución y un nivel de alquiler, devuelve el costo total del alquiler incluyendo cualquier cargo por retraso.

- Los sellos de tiempo dados son cadenas ISO UTC, por ejemplo: `"2026-06-18T18:30:00Z"`.
- El nivel de alquiler es el número de días antes de que el alquiler deba ser devuelto: `1`, `3` o `7`.
- Los alquileres deben devolverse a más tardar a las 12:00 PM UTC del último día del período de alquiler. Por ejemplo, un alquiler de 1 día que se saque en cualquier momento del 15 de marzo debe devolverse antes de las 12:00 PM UTC del 16 de marzo.
- Cada día después de la fecha y hora de devolución genera un cargo por retraso.

Los precios son los siguientes:

| Nivel | Costo base | Cargo por retraso por día |
| - | - | - |
| 1 día | $4.99 | $3.99 |
| 3 días | $3.99 | $2.99 |
| 7 días | $2.99 | $0.99 |

Devuelve el costo total redondeado a dos decimales en el formato `"$D.CC"`.

# --hints--

`getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1)` debe devolver `"$4.99"`.

```js
assert.equal(getRentalCost("2026-06-18T18:30:00Z", "2026-06-19T10:30:00Z", 1), "$4.99");
```

`getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1)` debe devolver `"$12.97"`.

```js
assert.equal(getRentalCost("2026-06-18T14:30:00Z", "2026-06-20T12:30:00Z", 1), "$12.97");
```

`getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3)` debe devolver `"$3.99"`.

```js
assert.equal(getRentalCost("2026-06-18T10:15:00Z", "2026-06-18T19:45:00Z", 3), "$3.99");
```

`getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3)` debe devolver `"$9.97"`.

```js
assert.equal(getRentalCost("2026-06-18T15:20:00Z", "2026-06-23T08:10:00Z", 3), "$9.97");
```

`getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7)` debe devolver `"$2.99"`.

```js
assert.equal(getRentalCost("2026-06-18T12:00:00Z", "2026-06-25T12:00:00Z", 7), "$2.99");
```

`getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7)` debe devolver `"$358.40"`.

```js
assert.equal(getRentalCost("2026-06-18T08:00:00Z", "2027-06-18T14:00:00Z", 7), "$358.40");
```

# --seed--

## --seed-contents--

```js
function getRentalCost(rented, returned, tier) {

  return rented;
}
```

# --solutions--

```js
function getRentalCost(rented, returned, tier) {
  const pricing = {
    1: { base: 4.99, late: 3.99 },
    3: { base: 3.99, late: 2.99 },
    7: { base: 2.99, late: 0.99 },
  };

  const rentedDate = new Date(rented);
  const returnDate = new Date(returned);
  const due = new Date(rentedDate);
  due.setUTCDate(due.getUTCDate() + tier);
  due.setUTCHours(12, 0, 0, 0);
  const { base, late } = pricing[tier];

  if (returnDate <= due) {
    return `$${base.toFixed(2)}`;
  }

  const msPerDay = 1000 * 60 * 60 * 24;
  const daysLate = Math.ceil((returnDate - due) / msPerDay);
  const total = base + daysLate * late;

  return `$${total.toFixed(2)}`;
}
```
