---
id: 6a151d32d772271dd2448c2c
title: "Desafío 310: de inglés británico a americano"
challengeType: 28
dashedName: challenge-310
---

# --description--

Dada una oración, convierte cualquier ortografía del inglés británico a su equivalente en inglés americano usando la siguiente tabla de consulta y devuelve la oración actualizada:

| British | American |
| - | - |
| `"colour"` | `"color"` |
| `"flavour"` | `"flavor"` |
| `"honour"` | `"honor"` |
| `"neighbour"` | `"neighbor"` |
| `"labour"` | `"labor"` |
| `"humour"` | `"humor"` |
| `"centre"` | `"center"` |
| `"fibre"` | `"fiber"` |
| `"defence"` | `"defense"` |
| `"offence"` | `"offense"` |
| `"organise"` | `"organize"` |
| `"recognise"` | `"recognize"` |
| `"analyse"` | `"analyze"` |

- Los reemplazos deben ser insensibles a mayúsculas y minúsculas. Por ejemplo, `"Colour"` debería convertirse en `"Color"`.
- La entrada puede contener palabras que se basan en la ortografía exacta de una raíz en la tabla y que también necesitan ser cambiadas. Por ejemplo, `"colouring"` debería convertirse en `"coloring"`, y `"disorganised"` debería convertirse en `"disorganized"`.

# --hints--

`britishToAmerican("I love the colour blue.")` debería devolver `"I love the color blue."`

```js
assert.equal(britishToAmerican("I love the colour blue."), "I love the color blue.");
```

`britishToAmerican("The fibre optic cable is new.")` debería devolver `"The fiber optic cable is new."`

```js
assert.equal(britishToAmerican("The fibre optic cable is new."), "The fiber optic cable is new.");
```

`britishToAmerican("It's an honour to meet someone with such humour.")` debería devolver `"It's an honor to meet someone with such humor."`

```js
assert.equal(britishToAmerican("It's an honour to meet someone with such humour."), "It's an honor to meet someone with such humor.");
```

`britishToAmerican("The unrecognised artist analysed his colour palette at the centre.")` debería devolver `"The unrecognized artist analyzed his color palette at the center."`

```js
assert.equal(britishToAmerican("The unrecognised artist analysed his colour palette at the centre."), "The unrecognized artist analyzed his color palette at the center.");
```

`britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful.")` debería devolver `The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.`

```js
assert.equal(britishToAmerican("The offence analysed, with organisation, the defence centre and recognised that the neighbouring labouror was humourous, flavourful, and colourful."), "The offense analyzed, with organisation, the defense center and recognized that the neighboring laboror was humorous, flavorful, and colorful.");
```

# --seed--

## --seed-contents--

```js
function britishToAmerican(sentence) {

  return sentence;
}
```

# --solutions--

```js
function britishToAmerican(sentence) {
  const lookup = {
    colour: "color",
    flavour: "flavor",
    honour: "honor",
    neighbour: "neighbor",
    labour: "labor",
    humour: "humor",
    centre: "center",
    fibre: "fiber",
    defence: "defense",
    offence: "offense",
    organise: "organize",
    recognise: "recognize",
    analyse: "analyze"
  };

  let result = sentence;

  for (const [british, american] of Object.entries(lookup)) {
    const regex = new RegExp(british, "gi");
    result = result.replace(regex, match => {
      return match[0] === match[0].toUpperCase()
        ? american[0].toUpperCase() + american.slice(1)
        : american;
    });
  }

  return result;
}
```
