---
id: 6a19b1062d1b153d8ac76d72
title: "Desafío 324: conteo de caracteres duplicados"
challengeType: 28
dashedName: challenge-324
---

# --description--

Dadas dos cadenas, devuelve un conteo de caracteres de la segunda cadena que se pueden encontrar en la primera.

- Los caracteres duplicados en la segunda cadena se cuentan por separado.

# --hints--

`duplicateCharacterCount("aloha", "hei")` debería devolver `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` debería devolver `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` debería devolver `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` debería devolver `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` debería devolver `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` debería devolver `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` debería devolver `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
