---
id: 6a1d9f98e819ed70a0e994dd
title: "Desafío 331: múltiplo más cercano"
challengeType: 28
dashedName: challenge-331
---

# --description--

Dado dos enteros, redondea el primero al múltiplo más cercano del segundo.

# --hints--

`roundToNearestMultiple(5, 3)` debería devolver `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` debería devolver `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` debería devolver `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` debería devolver `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` debería devolver `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
