---
id: 6a1d9f98e819ed70a0e994e0
title: "Desafío 334: Cambio exacto"
challengeType: 28
dashedName: challenge-334
---

# --description--

Dado un monto entero en centavos, devuelve el número de formas distintas de hacer cambio exacto usando pennies (1 centavo), nickels (5 centavos), dimes (10 centavos) y quarters (25 centavos).

# --hints--

`exactChange(3)` debería devolver `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` debería devolver `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` debería devolver `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` debería devolver `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` debería devolver `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` debería devolver `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
