---
id: 6a1d9f98e819ed70a0e994dc
title: "Desafío 330: palabras en minúsculas"
challengeType: 28
dashedName: challenge-330
---

# --description--

Dada una cadena, devuelve solo las palabras que están completamente en minúsculas, en su orden original y con un espacio entre cada palabra.

# --hints--

`getLowercaseWords("hello GOOD world")` debería devolver `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` debería devolver `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` debería devolver `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` debería devolver `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` debería devolver `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
