---
id: 6a1d9f98e819ed70a0e994d9
title: "Desafío 327: migración de base de datos"
challengeType: 28
dashedName: challenge-327
---

# --description--

Dado dos objetos de base de datos, devuelve el segundo objeto con las propiedades faltantes del primero completadas.

- Los campos que ya existen en el registro no deben ser sobrescritos.

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` debería devolver `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 })` debería devolver `{ username: "camper", posts: 5 }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", posts: 5 }), { username: "camper", posts: 5 });
```

`migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" })` debería devolver `{ username: "camper", posts: 0, verified: false }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0, verified: false }, { username: "camper" }), { username: "camper", posts: 0, verified: false });
```

`migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" })` debería devolver `{ username: "camper", role: "admin", posts: 0 }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { username: "camper", role: "admin" }), { username: "camper", role: "admin", posts: 0 });
```

`migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" })` debería devolver `{ username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false }`.

```js
assert.deepEqual(migrateRecord({ username: "", email: "", posts: 0, verified: false, role: "user", banned: false }, { username: "camper", email: "camper@freecodecamp.org", role: "admin" }), { username: "camper", email: "camper@freecodecamp.org", role: "admin", posts: 0, verified: false, banned: false });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
