---
id: 6a1d9f98e819ed70a0e994da
title: "Desafío 328: la rutina de Kaprekar"
challengeType: 28
dashedName: challenge-328
---

# --description--

Dado un número de 4 dígitos, devuelve la cantidad de veces que necesitas aplicar la rutina de Kaprekar hasta llegar a 6174.

La rutina de Kaprekar funciona así:

- Ordena los dígitos en orden descendente para formar el número más grande
- Ordena los dígitos en orden ascendente para formar el número más pequeño (rellena con ceros a la izquierda si es necesario)
- Resta el número más pequeño del más grande
- Repite con el nuevo número

# --hints--

`kaprekar(1234)` debería devolver `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` debería devolver `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` debería devolver `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` debería devolver `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` debería devolver `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
