---
id: 6a1d9f98e819ed70a0e994e2
title: "Desafío 336: Compatibilidad del horóscopo"
challengeType: 28
dashedName: challenge-336
---

# --description--

Dadas dos cadenas de signos zodiacales, devuelve su porcentaje de compatibilidad.

Los signos están organizados en una rueda de 12 posiciones en este orden: `"Aries"`, `"Taurus"`, `"Gemini"`, `"Cancer"`, `"Leo"`, `"Virgo"`, `"Libra"`, `"Scorpio"`, `"Sagittarius"`, `"Capricorn"`, `"Aquarius"`, `"Pisces"`, volviendo a `"Aries"` después de `"Pisces"`. Encuentra la distancia más corta entre los dos signos y devuelve la compatibilidad:

| Distancia | Compatibilidad |
|----------|---------------|
| 0 | `"100%"` |
| 1 | `"40%"` |
| 2 | `"80%"` |
| 3 | `"30%"` |
| 4 | `"90%"` |
| 5 | `"20%"` |
| 6 | `"50%"` |

# --hints--

`horoscopeMatch("Libra", "Sagittarius")` debería devolver `"80%"`.

```js
assert.equal(horoscopeMatch("Libra", "Sagittarius"), "80%");
```

`horoscopeMatch("Gemini", "Scorpio")` debería devolver `"20%"`.

```js
assert.equal(horoscopeMatch("Gemini", "Scorpio"), "20%");
```

`horoscopeMatch("Pisces", "Aries")` debería devolver `"40%"`.

```js
assert.equal(horoscopeMatch("Pisces", "Aries"), "40%");
```

`horoscopeMatch("Capricorn", "Cancer")` debería devolver `"50%"`.

```js
assert.equal(horoscopeMatch("Capricorn", "Cancer"), "50%");
```

`horoscopeMatch("Aquarius", "Aquarius")` debería devolver `"100%"`.

```js
assert.equal(horoscopeMatch("Aquarius", "Aquarius"), "100%");
```

`horoscopeMatch("Virgo", "Taurus")` debería devolver `"90%"`.

```js
assert.equal(horoscopeMatch("Virgo", "Taurus"), "90%");
```

`horoscopeMatch("Leo", "Scorpio")` debería devolver `"30%"`.

```js
assert.equal(horoscopeMatch("Leo", "Scorpio"), "30%");
```

# --seed--

## --seed-contents--

```js
function horoscopeMatch(sign1, sign2) {

  return sign1;
}
```

# --solutions--

```js
function horoscopeMatch(sign1, sign2) {
  const signs = ["Aries", "Taurus", "Gemini", "Cancer", "Leo", "Virgo", "Libra", "Scorpio", "Sagittarius", "Capricorn", "Aquarius", "Pisces"];
  
  const a = signs.indexOf(sign1) + 1;
  const b = signs.indexOf(sign2) + 1;
  const diff = Math.abs(a - b);
  const distance = Math.min(diff, 12 - diff);

  const compatibility = { 0: "100%", 1: "40%", 2: "80%", 3: "30%", 4: "90%", 5: "20%", 6: "50%" };
  return compatibility[distance];
}
```
