---
id: 6a1d9f98e819ed70a0e994e1
title: "Desafío 335: cinco dados"
challengeType: 28
dashedName: challenge-335
---

# --description--

Dado un arreglo de cinco dados con valores del 1 al 6, devuelve la mejor mano posible.

Aquí están las manos ordenadas de menor a mayor:

| Mano | Descripción |
|------|-------------|
| `"no pair"` | Sin pares ni mejor |
| `"pair"` | Dos dados con el mismo valor |
| `"two pair"` | Dos pares diferentes |
| `"three of a kind"` | Tres dados con el mismo valor |
| `"small straight"` | Cuatro valores consecutivos |
| `"large straight"` | Cinco valores consecutivos |
| `"full house"` | Tres iguales y un par |
| `"four of a kind"` | Cuatro dados con el mismo valor |
| `"five of a kind"` | Los cinco dados con el mismo valor |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` debería devolver `"five of a kind"`.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` debería devolver `"four of a kind"`.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` debería devolver `"large straight"`.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` debería devolver `"three of a kind"`.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` debería devolver `"pair"`.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` debería devolver `"no pair"`.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` debería devolver `"small straight"`.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` debería devolver `"full house"`.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` debería devolver `"two pair"`.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
