---
id: 6a1d9f98e819ed70a0e994df
title: "Desafío 333: Clasificación de problemas 2"
challengeType: 28
dashedName: challenge-333
---

# --description--

Dado un título de problema y un arreglo de etiquetas actuales, devuelve un arreglo actualizado de etiquetas según las siguientes reglas:

Si el problema no tiene ninguna etiqueta, agrega:

- `"bug"` y `"needs triage"` si el título contiene `"error"` o `"bug"`
- `"enhancement"` y `"discussing"` si el título contiene `"feature"` o `"add"`

De lo contrario, si las etiquetas dadas contienen:

- `"needs triage"` y el título contiene `"simple"` o `"easy"`, elimina `"needs triage"` y agrega `"good first issue"`
- `"discussing"` y el título contiene `"planned"` o `"next"`, elimina `"discussing"` y agrega `"on the roadmap"`
- De lo contrario, si `"needs triage"` o `"discussing"` está presente, elimínalo y agrega `"help wanted"`

Si el título contiene:

- `"security"`, agrega una etiqueta `"critical"`

# --hints--

`triageIssue("app crashes with error", [])` debería devolver `["bug", "needs triage"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", []).sort(), ["bug", "needs triage"].sort());
```

`triageIssue("app crashes with error", ["bug", "needs triage"])` debería devolver `["bug", "help wanted"]`.

```js
assert.deepEqual(triageIssue("app crashes with error", ["bug", "needs triage"]).sort(), ["bug", "help wanted"].sort());
```

`triageIssue("add dark mode", [])` debería devolver `["enhancement", "discussing"]`.

```js
assert.deepEqual(triageIssue("add dark mode", []).sort(), ["enhancement", "discussing"].sort());
```

`triageIssue("add dark mode", ["enhancement", "discussing"])` debería devolver `["enhancement", "help wanted"]`.

```js
assert.deepEqual(triageIssue("add dark mode", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted"].sort());
```

`triageIssue("xss security bug", [])` debería devolver `["bug", "needs triage", "critical"]`.

```js
assert.deepEqual(triageIssue("xss security bug", []).sort(), ["bug", "needs triage", "critical"].sort());
```

`triageIssue("security vulnerability in auth", [])` debería devolver `["critical"]`.

```js
assert.deepEqual(triageIssue("security vulnerability in auth", []).sort(), ["critical"].sort());
```

`triageIssue("easy a11y fix", ["bug", "needs triage"])` debería devolver `["bug", "good first issue"]`.

```js
assert.deepEqual(triageIssue("easy a11y fix", ["bug", "needs triage"]).sort(), ["bug", "good first issue"].sort());
```

`triageIssue("planned api migration", ["enhancement", "discussing"])` debería devolver `["enhancement", "on the roadmap"]`.

```js
assert.deepEqual(triageIssue("planned api migration", ["enhancement", "discussing"]).sort(), ["enhancement", "on the roadmap"].sort());
```

`triageIssue("improve security", ["enhancement", "discussing"])` debería devolver `["enhancement", "help wanted", "critical"]`.

```js
assert.deepEqual(triageIssue("improve security", ["enhancement", "discussing"]).sort(), ["enhancement", "help wanted", "critical"].sort());
```

# --seed--

## --seed-contents--

```js
function triageIssue(title, labels) {

  return title;
}
```

# --solutions--

```js
function triageIssue(title, labels) {
  const result = [...labels];
  const t = title.toLowerCase();

  const has = label => result.includes(label);
  const add = label => { if (!has(label)) result.push(label); };
  const remove = label => { const i = result.indexOf(label); if (i > -1) result.splice(i, 1); };

  if (result.length === 0) {
    if (t.includes("error") || t.includes("bug")) {
      add("bug"); add("needs triage");
    } else if (t.includes("feature") || t.includes("add")) {
      add("enhancement"); add("discussing");
    }
  } else {
    if (has("needs triage") && (t.includes("simple") || t.includes("easy"))) {
      remove("needs triage"); add("good first issue");
    } else if (has("discussing") && (t.includes("planned") || t.includes("next"))) {
      remove("discussing"); add("on the roadmap");
    } else if (has("needs triage") || has("discussing")) {
      remove("needs triage"); remove("discussing"); add("help wanted");
    }
  }

  if (t.includes("security")) {
    add("critical");
  }

  return result;
}
```
