---
id: 6a1d9f98e819ed70a0e994de
title: "Desafío 332: Clasificación de problemas"
challengeType: 28
dashedName: challenge-332
---

# --description--

Dado un número de milisegundos desde la última publicación en un problema y el último mensaje publicado en el problema, determina qué deberías hacer con el problema según estas reglas:

- Si el último mensaje fue hace menos de 7 días, devuelve `"leave it"`
- Si el último mensaje fue hace 7 o más días y su contenido contiene `"bump"` (sin distinguir mayúsculas o minúsculas), devuelve `"close it"`
- De lo contrario, devuelve `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` debería devolver `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` debería devolver `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` debería devolver `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` debería devolver `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` debería devolver `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` debería devolver `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
