---
id: 6a2037a68a0bc2aef0006003
title: "Desafío 341: cuenta regresiva para el cumpleaños"
challengeType: 28
dashedName: challenge-341
---

# --description--

Dada la fecha de hoy y un cumpleaños, devuelve el número de días hasta el próximo cumpleaños de la persona.

- La fecha de hoy se da como una cadena en formato `"YYYY-MM-DD"`, con ceros a la izquierda, por ejemplo: `"2026-07-16"`.
- El cumpleaños se da como una cadena en formato `"M/D"`, sin ceros a la izquierda, por ejemplo: `"9/7"`.
- Si hoy es su cumpleaños, devuelve el número de días hasta su próximo cumpleaños (no `0`).
- Se deben tener en cuenta los años bisiestos.

# --hints--

`daysUntilBirthday("2026-07-16", "9/7")` debería devolver `53`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "9/7"), 53);
```

`daysUntilBirthday("2026-07-16", "3/22")` debería devolver `249`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "3/22"), 249);
```

`daysUntilBirthday("2026-07-16", "7/16")` debería devolver `365`.

```js
assert.equal(daysUntilBirthday("2026-07-16", "7/16"), 365);
```

`daysUntilBirthday("2024-02-28", "3/1")` debería devolver `2`.

```js
assert.equal(daysUntilBirthday("2024-02-28", "3/1"), 2);
```

`daysUntilBirthday("2023-04-24", "12/30")` debería devolver `250`.

```js
assert.equal(daysUntilBirthday("2023-04-24", "12/30"), 250);
```

`daysUntilBirthday("2024-03-01", "2/29")` debería devolver `1460`.

```js
assert.equal(daysUntilBirthday("2024-03-01", "2/29"), 1460);
```

`daysUntilBirthday("2096-03-01", "2/29")` debería devolver `2920`.

```js
assert.equal(daysUntilBirthday("2096-03-01", "2/29"), 2920);
```

# --seed--

## --seed-contents--

```js
function daysUntilBirthday(today, birthday) {

  return today;
}
```

# --solutions--

```js
function daysUntilBirthday(today, birthday) {
  const [year, month, day] = today.split("-").map(Number);
  const [bMonth, bDay] = birthday.split("/").map(Number);

  const todayDate = new Date(year, month - 1, day);

  function isLeapYear(y) {
    return (y % 4 === 0 && y % 100 !== 0) || y % 400 === 0;
  }

  function nextValidBirthday(fromYear) {
    for (let y = fromYear; y <= fromYear + 8; y++) {
      if (bMonth === 2 && bDay === 29 && !isLeapYear(y)) continue;
      const date = new Date(y, bMonth - 1, bDay);
      if (date > todayDate) return date;
    }
  }

  const nextBirthday = nextValidBirthday(year);
  const diff = nextBirthday - todayDate;
  return Math.round(diff / (1000 * 60 * 60 * 24));
}
```
