---
id: 6a2037a68a0bc2aef0005fff
title: "Desafío 337: contador de marcas"
challengeType: 28
dashedName: challenge-337
---

# --description--

Dada una cadena de marcas de conteo, devuelve el total representado.

- Cada barra `"|"` representa un conteo.
- Cada quinta marca se representa con una barra diagonal `"/"`, completando un grupo de cinco (`"||||/"`).
- Los grupos están separados por un espacio.

# --hints--

`getTallyCount("||||")` debería devolver `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` debería devolver `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` debería devolver `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` debería devolver `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` debería devolver `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
