---
id: 6a2037a68a0bc2aef0006000
title: "Desafío 338: calculadora de edad de mascotas"
challengeType: 28
dashedName: challenge-338
---

# --description--

Dado un tipo de mascota y su edad en años humanos, devuelve la edad equivalente en años de mascota usando la siguiente tabla de conversión:

| Mascota | Multiplicador |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` debería devolver `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` debería devolver `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` debería devolver `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` debería devolver `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` debería devolver `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` debería devolver `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` debería devolver `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
