---
id: 6a2037a68a0bc2aef0006007
title: "Desafío 345: Mezclador de palabras"
challengeType: 28
dashedName: challenge-345
---

# --description--

Dadas dos palabras, devuelve una nueva palabra combinando la primera mitad de la primera palabra con la segunda mitad de la segunda palabra.

- Para palabras de longitud impar, la primera mitad es la mitad más corta.

# --hints--

`blendWords("turtle", "toucan")` debería devolver `"turcan"`.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` debería devolver `"chipingo"`.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` debería devolver `"falican"`.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` debería devolver `"hyana"`.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` debería devolver `"scorilla"`.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` debería devolver `"platerine"`.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
