---
id: 6a2037a68a0bc2aef0006002
title: "Desafío 340: convertidor a Pig Latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

Dada una cadena, conviértela a Pig Latin usando las siguientes reglas:

- Si una palabra comienza con una vocal (`"a"`, `"e"`, `"i"`, `"o"` o `"u"`), agrega `"way"` al final. Por ejemplo, `"universe"` se convierte en `"universeway"`.
- Si una palabra comienza con una o más consonantes, muévelas al final y agrega `"ay"`. Por ejemplo, `"hello"` se convierte en `"ellohay"`.
- Conserva la mayúscula de la primera letra. Por ejemplo, `"Hello"` se convierte en `"Ellohay"`.

# --hints--

`pigLatin("universe")` debería devolver `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` debería devolver `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` debería devolver `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` debería devolver `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` debería devolver `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` debería devolver `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
