---
id: 6a2037a68a0bc2aef0006001
title: "Desafío 339: Dividir arreglos en partes"
challengeType: 28
dashedName: challenge-339
---

# --description--

Dado un arreglo y un tamaño de parte, devuelve el arreglo dividido en subarreglos de ese tamaño.

- La última parte puede ser más pequeña si el arreglo no se divide de manera uniforme.

# --hints--

`chunkArray([1, 2, 3, 4, 5, 6], 3)` debería devolver `[[1, 2, 3], [4, 5, 6]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5, 6], 3), [[1, 2, 3], [4, 5, 6]]);
```

`chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2)` debería devolver `[[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]`.

```js
assert.deepEqual(chunkArray([1, "two", 3, "four", 5, "six", 7, "eight"], 2), [[1, "two"], [3, "four"], [5, "six"], [7, "eight"]]);
```

`chunkArray([1, 2, 3, 4, 5], 3)` debería devolver `[[1, 2, 3], [4, 5]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3, 4, 5], 3), [[1, 2, 3], [4, 5]]);
```

`chunkArray(["a", "b", "c", "d", "e"], 1)` debería devolver `[["a"], ["b"], ["c"], ["d"], ["e"]]`.

```js
assert.deepEqual(chunkArray(["a", "b", "c", "d", "e"], 1), [["a"], ["b"], ["c"], ["d"], ["e"]]);
```

`chunkArray([1, 2, 3], 5)` debería devolver `[[1, 2, 3]]`.

```js
assert.deepEqual(chunkArray([1, 2, 3], 5), [[1, 2, 3]]);
```

# --seed--

## --seed-contents--

```js
function chunkArray(arr, size) {

  return arr;
}
```

# --solutions--

```js
function chunkArray(arr, size) {
  const result = [];
  for (let i = 0; i < arr.length; i += size) {
    result.push(arr.slice(i, i + size));
  }
  return result;
}
```
