---
id: 6a22d77ddf034bc4e35b1d58
title: "Desafío 351: número prónico"
challengeType: 28
dashedName: challenge-351
---

# --description--

Dado un número, determina si es un número prónico.

Un número prónico es el producto de dos enteros consecutivos. Por ejemplo, 6 es prónico porque 2 * 3 = 6.

# --hints--

`isPronic(6)` debería devolver `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` debería devolver `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` debería devolver `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` debería devolver `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` debería devolver `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` debería devolver `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
