---
id: 6a22d77ddf034bc4e35b1d5b
title: "Desafío 354: Calificación de contraste 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

Dado dos arreglos que representan valores RGB y un booleano que indica si el texto es grande, devuelve la calificación de contraste WCAG usando el siguiente método:

Primero, convierte cada valor RGB a luminancia relativa:

- Divide cada canal `[R, G, B]` entre 255 para obtener un valor entre `0` y `1`
- Aplica la fórmula de corrección gamma a cada canal:
  - Si el valor del canal es menor o igual a `0.04045`: `channel / 12.92`
  - De lo contrario: `((channel + 0.055) / 1.055) ^ 2.4`
- Calcula la luminancia: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Luego, calcula la relación de contraste sumando 0.05 a cada valor de luminancia, y luego dividiendo el más claro entre el más oscuro. El más claro siempre será el primer argumento.

Devuelve la calificación basada en la relación de contraste usando la siguiente tabla:

| Calificación | Texto normal | Texto grande |
|-------------|--------------|--------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | por debajo de 4.5 | por debajo de 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` debería devolver `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` debería devolver `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` debería devolver `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` debería devolver `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` debería devolver `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` debería devolver `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
