---
id: 6a22d77ddf034bc4e35b1d5a
title: "Desafío 353: clasificación de contraste 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

Dadas dos valores de luminancia relativa y un booleano que indica si el texto es grande, devuelve la clasificación de contraste WCAG usando el siguiente método:

Calcula la relación de contraste agregando 0.05 a cada valor de luminancia, luego dividiendo el más claro entre el más oscuro. El más claro siempre será el primer argumento.

Devuelve la clasificación basada en la relación de contraste usando la siguiente tabla:

| Clasificación | Texto normal | Texto grande |
|--------|-------------|------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | por debajo de 4.5 | por debajo de 3.0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` debería devolver `"AAA"`.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` debería devolver `"AA"`.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` debería devolver `"Fail"`.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` debería devolver `"AAA"`.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` debería devolver `"AA"`.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` debería devolver `"Fail"`.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
