---
id: 6a22d77ddf034bc4e35b1d54
title: "Desafío 347: teoría de juegos"
challengeType: 28
dashedName: challenge-347
---

# --description--

Dadas dos cadenas de igual longitud que representan las estrategias de dos jugadores para un juego, devuelve las puntuaciones como un arreglo `[player1, player2]`.

- Las cadenas dadas solo contendrán una de dos letras: `"C"` (cooperar) o `"D"` (traicionar).
- Cada carácter representa una ronda, puntuándose de la siguiente manera:
  - Si ambos jugadores cooperan, cada uno obtiene 3 puntos.
  - Si ambos jugadores traicionan, cada uno obtiene 1 punto.
  - Si un jugador traiciona y el otro coopera, el traidor obtiene 5 puntos y el cooperador 0.

# --hints--

`playGame("CCCC", "CCCC")` debería devolver `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` debería devolver `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` debería devolver `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` debería devolver `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` debería devolver `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
