---
id: 6a22d77ddf034bc4e35b1d57
title: "Desafío 350: distancia entre letras"
challengeType: 28
dashedName: challenge-350
---

# --description--

Dadas dos cadenas de igual longitud, devuelve la suma de las distancias más cortas entre cada par de caracteres.

- La entrada solo contendrá letras minúsculas
- El alfabeto se trata como un círculo, por lo que la distancia entre `a` y `z` es 1.

# --hints--

`letterDistance("abc", "bcd")` debería devolver `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` debería devolver `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` debería devolver `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` debería devolver `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` debería devolver `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` debería devolver `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
