---
id: 6a26df3e2988bcdded204897
title: "Desafío 361: hora hablada"
challengeType: 28
dashedName: challenge-361
---

# --description--

Dadas las posiciones en grados de las manecillas de la hora y los minutos de un reloj analógico (en sentido horario desde el 12), devuelve la hora en inglés hablado.

Convierte el ángulo de la manecilla de los minutos a minutos (360° = 60 minutos), luego usa las siguientes reglas:

| Minutos | Hablado |
|---------|---------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (excluyendo 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (excluyendo 45) | `"X minutes to Z"` (donde `X` es 60 - minutos) |

Donde `Y` es la hora actual y `Z` es la siguiente hora, ambas derivadas del ángulo de la manecilla de la hora (360° = 12 horas).

Nota: Los ángulos de las manecillas pueden no coincidir exactamente con un número, considera redondearlos de alguna manera.

# --hints--

`getSpokenTime(90, 0)` debería devolver `"3 o'clock"`.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` debería devolver `"20 minutes past 5"`.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` debería devolver `"half past 8"`.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` debería devolver `"quarter past 2"`.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` debería devolver `"20 minutes to 7"`.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` debería devolver `"quarter to 11"`.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` debería devolver `"5 minutes to 4"`.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
