---
id: 6a26df3e2988bcdded204896
title: "Desafío 360: duración hablada"
challengeType: 28
dashedName: challenge-360
---

# --description--

Dado un número de segundos, devuelve la duración en inglés hablado.

- Divide la duración en horas, minutos y segundos.
- Omite cualquier valor cero.
- Usa singular o plural según corresponda (`"1 hour"`, `"2 hours"`).
- Si están presentes, une las dos últimas unidades con `"and"`, y la segunda y tercera desde el final con una coma (`"1 hour, 2 minutes and 3 seconds"`).

# --hints--

`getSpokenDuration(3723)` debería devolver `"1 hour, 2 minutes and 3 seconds"`.

```js
assert.equal(getSpokenDuration(3723), "1 hour, 2 minutes and 3 seconds");
```

`getSpokenDuration(7295)` debería devolver `"2 hours, 1 minute and 35 seconds"`.

```js
assert.equal(getSpokenDuration(7295), "2 hours, 1 minute and 35 seconds");
```

`getSpokenDuration(8521)` debería devolver `"2 hours, 22 minutes and 1 second"`.

```js
assert.equal(getSpokenDuration(8521), "2 hours, 22 minutes and 1 second");
```

`getSpokenDuration(435)` debería devolver `"7 minutes and 15 seconds"`.

```js
assert.equal(getSpokenDuration(435), "7 minutes and 15 seconds");
```

`getSpokenDuration(14455)` debería devolver `"4 hours and 55 seconds"`.

```js
assert.equal(getSpokenDuration(14455), "4 hours and 55 seconds");
```

`getSpokenDuration(72000)` debería devolver `"20 hours"`.

```js
assert.equal(getSpokenDuration(72000), "20 hours");
```

`getSpokenDuration(1)` debería devolver `"1 second"`.

```js
assert.equal(getSpokenDuration(1), "1 second");
```

# --seed--

## --seed-contents--

```js
function getSpokenDuration(seconds) {

  return seconds;
}
```

# --solutions--

```js
function getSpokenDuration(seconds) {
  const hours = Math.floor(seconds / 3600);
  const minutes = Math.floor((seconds % 3600) / 60);
  const secs = seconds % 60;

  const parts = [];
  if (hours) parts.push(`${hours} ${hours === 1 ? "hour" : "hours"}`);
  if (minutes) parts.push(`${minutes} ${minutes === 1 ? "minute" : "minutes"}`);
  if (secs) parts.push(`${secs} ${secs === 1 ? "second" : "seconds"}`);

  if (parts.length === 1) return parts[0];
  if (parts.length === 2) return `${parts[0]} and ${parts[1]}`;
  return `${parts[0]}, ${parts[1]} and ${parts[2]}`;
}
```
