---
id: 6a26df3e2988bcdded204895
title: "Desafío 359: Calculadora de hándicap de golf"
challengeType: 28
dashedName: challenge-359
---

# --description--

Dado un arreglo de puntuaciones de golf y un arreglo correspondiente de valores par del campo, devuelve el índice de hándicap del golfista usando el siguiente método:

- Calcula la diferencia para cada ronda restando el par a la puntuación, luego devuelve el promedio de todas las diferencias redondeado a un decimal.

# --hints--

`calculateHandicap([72, 72, 72], [72, 72, 72])` debería devolver `0`.

```js
assert.equal(calculateHandicap([72, 72, 72], [72, 72, 72]), 0);
```

`calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72])` debería devolver `6`.

```js
assert.equal(calculateHandicap([80, 76, 78, 78], [72, 72, 72, 72]), 6);
```

`calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36])` debería devolver `8.3`.

```js
assert.equal(calculateHandicap([42, 45, 46, 44], [36, 36, 36, 36]), 8.3);
```

`calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71])` debería devolver `8.8`.

```js
assert.equal(calculateHandicap([85, 80, 76, 79, 82], [72, 72, 72, 71, 71]), 8.8);
```

`calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37])` debería devolver `11.7`.

```js
assert.equal(calculateHandicap([41, 50, 48, 52, 46, 49], [35, 37, 35, 37, 35, 37]), 11.7);
```

# --seed--

## --seed-contents--

```js
function calculateHandicap(scores, pars) {

  return scores;
}
```

# --solutions--

```js
function calculateHandicap(scores, pars) {
  const differentials = scores.map((score, i) => score - pars[i]);
  const average = differentials.reduce((sum, d) => sum + d, 0) / differentials.length;
  return Math.round(average * 10) / 10;
}
```
