---
id: 6a26df95efa55a2524399745
title: "Desafío 364: entre dos cubetas"
challengeType: 28
dashedName: challenge-364
---

# --description--

Dadas dos cubetas de pintura, cada una con un color RGB y un nivel de llenado, devuelve el color RGB mezclado como un arreglo de tres enteros.

- Cada cubeta es un objeto (JavaScript) o diccionario (Python) con una propiedad `color` (un arreglo de tres enteros `[r, g, b]`) y una propiedad `fullness` (0–100).
- El color mezclado es un promedio ponderado de cada canal en los dos colores basado en el nivel de llenado, con cada canal redondeado al entero más cercano.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` debería devolver `[125, 125, 125]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` debería devolver `[200, 200, 200]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` debería devolver `[100, 150, 200]`.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` debería devolver `[114, 184, 170]`.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` debería devolver `[68, 163, 123]`.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
