---
id: 686386e76197458b7dfd52eb
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Crea un elemento `div` debajo del `h1`. Dentro de ese `div`, crea otro elemento `div` y anida otro elemento `div` dentro de este.

Y dentro del `div` más interno, crea un elemento `h2` con el texto `Listener`. Esto representará el primer plan de precios.

# --hints--

Debes crear un elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Debes crear otro elemento `div` dentro del `div` exterior.

```js
assert.exists(document.querySelector("div > div"));
```

Debes crear otro elemento `div` dentro del `div` interno.

```js
assert.exists(document.querySelector("div > div > div"));
```

Debes crear un `h2` con el texto `Listener` dentro del elemento `div` más interno.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
