---
id: 686386e76197458b7dfd52ef
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Ahora crearás un botón CTA. Un Call to Action (CTA) es un botón o enlace diseñado para incitar a un usuario a realizar una acción específica, en este caso, comenzar a escuchar.

Fuera del `div` más interno, crea un elemento anchor con un `href` de `#` y el texto `Start listening`.

Eso concluye el HTML para la primera tarjeta de precios.

# --hints--

Debes crear un elemento `anchor`.

```js
assert.exists(document.querySelector("a"))
```

Tu elemento anchor debe tener un `href` establecido en `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Tu elemento anchor debe tener el texto `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
