---
id: 686386e76197458b7dfd52ed
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Debajo del elemento `p` que acabas de agregar, crea otro elemento `p` con el texto `Start exploring millions of songs with basic features and ads.`. Debajo de ese, crea un elemento `ul` vacío.

# --hints--

Debes crear otro elemento `p` con el texto `Start exploring millions of songs with basic features and ads.`

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)
assert.exists(targetP)
```

Debes crear un elemento `ul` vacío debajo del elemento `p`.

```js
const paragraphs = document.querySelectorAll("p")

const targetP = Array.from(paragraphs).find(p =>
  p.textContent === "Start exploring millions of songs with basic features and ads."
)

const ul = targetP.nextElementSibling
assert.equal(ul?.tagName, "UL")
assert.isEmpty(ul?.innerHTML.trim())
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
