---
id: 686386e76197458b7dfd52ee
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro del `ul`, crea dos elementos de lista. Dentro del primer elemento de lista, crea un elemento `span` y agrega la entidad HTML `&#10003;` como su texto. Justo después de ese `span` pero dentro del elemento de lista, agrega el texto `Ad-supported streaming`.

Dentro del segundo elemento de la lista, crea un elemento `span` y agrega la entidad `&#10003;` como su texto. Justo después de ese `span`, pero dentro del elemento de la lista, agrega el texto `Curated playlists`.

# --hints--

Debes tener dos elementos de lista dentro de tu elemento `ul`.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Tu primer elemento `li` debe tener un elemento `span` con la entidad de carácter HTML `&#10003;` dentro.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Tu primer elemento `li` debe tener el texto `Ad-supported streaming` justo después del elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Tu segundo elemento `li` debe tener un elemento `span` con la entidad de carácter HTML `&#10003;` dentro.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

Tu segundo elemento `li` debe tener el texto `Curated playlists` justo después del elemento `span`.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
