---
id: 686383abac6c625a84458ea9
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, aprenderás más sobre Tailwind CSS usando sus clases de utilidad de grid para construir un componente de precios.

Se ha proporcionado la estructura básica de HTML para ti. Comienza creando un elemento `h1` con el texto `Choose your listening plan`.

# --hints--

Debes crear un elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Tu elemento `h1` debe tener el texto `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
