---
id: 6864df7ec0fc508ac3828ab5
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

En CSS regular, aprendiste que asignar un display de `grid` a un elemento no hace nada hasta que especificas las filas y columnas de la grid. Así es como funciona también en Tailwind.

En Tailwind, para asignar una fila de la cuadrícula y una columna de la cuadrícula, usarías sus respectivas clases de utilidad así:

```html
<div class="grid grid-cols-[num] grid-rows-[num]">
  <div>Grid Item</div>
  <div>Grid Item</div>
  <div>Grid Item</div>
</div>
```

Asigna al `div` contenedor las clases de utilidad `grid-cols-1` y `md:grid-cols-3`.

La clase `grid-cols-1` organiza los tres elementos en una sola columna, apilándolos verticalmente. En pantallas medianas y más grandes, `md:grid-cols-3` cambia la estructura a tres columnas, mostrando los elementos uno al lado del otro.

# --hints--

Tu elemento `div` debe tener la clase `grid-cols-1`.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("grid-cols-1"))
```

Tu elemento `div` debe tener la clase `md:grid-cols-3`.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("md:grid-cols-3"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
--fcc-editable-region--
      <div class="grid ">
--fcc-editable-region--
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
