---
id: 615f2d4150fe0d4cbd0f2628
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أسفل عنصر `h1`، أضف عنصر `p` يحتوي على النص `8 servings per container`.

# --hints--

يجب إضافة عنصر `p` جديد.

```js
assert.exists(document.querySelector('p'));
```

يجب أن يكون عنصر `p` داخل عنصر `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

يجب أن يأتي عنصر `p` بعد عنصر `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

يجب أن يحتوي عنصر `p` على النص `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
    
</body>
</html>
--fcc-editable-region--
```

```css

```
