---
id: 615f357957e370510f21ea16
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أنشئ محدِّد `body` وأعطه خاصية `font-family` مضبوطة على `Open Sans` مع نظام بديل `sans-serif`.

تذكّر أن الخطوط التي تحتوي أسماؤها على فراغات يجب أن تُحاط بعلامات اقتباس في CSS.

# --hints--

يجب أن يكون لديك محدِّد `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب أن يحتوي محدِّد `body` على خاصية `font-family` مضبوطة على `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
