---
id: 615f2abbe7d18d49a1e0e1c8
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

لقد وفرنا لك قالب جاهز أساسي لـ HTML.

أنشئ عنصر `h1` داخل عنصر `body` وأعطه النص `Nutrition Facts`.

# --hints--

يجب أن تضيف عنصر `h1` جديد.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يكون عنصر `h1` داخل عنصر `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

يجب أن يحتوي عنصر `h1` على النص `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
