---
id: 615f34ecc1091b4fd5a8a484
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل عنصر `head`، أضف عنصر `link` مع خاصية `rel` مضبوطة على `stylesheet` وخاصية `href` مضبوطة على `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

هذا سيستورد عائلة الخطوط `Open Sans`، مع قيم وزن الخط `400` و`700` و`800`.

أضف أيضًا عنصر `link` لربط ملف `styles.css` الخاص بك.

# --hints--

يجب أن يحتوي شفرتك على عنصرين `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

عناصر `link` يجب أن تكون عناصر فارغة، لا يجب أن تحتوي على وسم إغلاق `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

يجب أن يكون العنصران `link` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

يجب أن تحتوي كلا عنصري `link` على خاصية `rel` مضبوطة على `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

يجب أن يحتوي أحد عنصري `link` على خاصية `href` مضبوطة على `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

يجب أن يحتوي أحد عنصري `link` على خاصية `href` مضبوطة على `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
