---
id: 699b5b7cb68b6d35afca61b7
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Хоча можна продовжувати додавати властивості таким чином до об’єкта `profile`, було б добре обмежити типи властивостей, які мають бути дозволені. Краще заздалегідь визначити форму об’єкта `profile`.

У попередніх уроках ви навчились працювати з типами об’єктів так:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

У цьому прикладі `composer` — це об’єкт із рядковим типом об’єкта. Цей об’єкт має три обов’язкові властивості `name`, `instrument` і `isActive`. Перші дві властивості мають бути типу `string`, а `isActive` має бути типу `boolean`.

Оновіть свій існуючий об’єкт `profile`, щоб включити рядковий тип об’єкта. `username` має бути типізований як `string`, `age` — як `number`, а `isLoggedIn` — як `boolean`.

# --hints--

Ваш об’єкт `profile` має мати рядковий тип об’єкта з властивістю `username`, встановленою в `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ваш об’єкт `profile` має мати рядковий тип об’єкта з властивістю `age`, встановленою в `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Ваш об’єкт `profile` має мати рядковий тип об’єкта з властивістю `isLoggedIn`, встановленою в `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
