---
id: 699b61ca8c9cb9bc1968364e
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Якщо ви відкриєте консоль, то побачите таке повідомлення про помилку:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

У чистому JavaScript ви могли додавати будь-які властивості та значення до об’єкта `profile`. Але в TypeScript компілятор перевіряє, що об’єкт має лише ті властивості, які ви визначили (у цьому випадку `username`, `age` і `isLoggedIn`). Оскільки `mood` не було включено до типу, TypeScript вважає це недійсною властивістю і видає помилку.

Видаліть `mood: null` зі свого коду.

# --hints--

У вашому коді більше не повинно бути `mood: null`.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
