---
id: 699b68f94a0d154c2a7a5ce0
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Тепер, коли об’єкт `profile` завершено, настав час створити новий об’єкт, який представлятиме ролі користувачів.

Створіть нову змінну з назвою `userRoles` і призначте їй об’єкт. Ваш об’єкт має містити такі властивості та значення:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

У вас має бути змінна з назвою `userRoles`.

```js
assert.isDefined(userRoles);
```

Ваша змінна `userRoles` має бути об’єктом.

```js
assert.isObject(userRoles);
```

Ваш об’єкт `userRoles` має містити властивість `admin`.

```js
assert.property(userRoles, "admin");
```

Властивість `admin` має мати значення `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Ваш об’єкт `userRoles` має містити властивість `editor`.

```js
assert.property(userRoles, "editor");
```

Властивість `editor` має мати значення `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Ваш об’єкт `userRoles` має містити властивість `viewer`.

```js
assert.property(userRoles, "viewer");
```

Властивість `viewer` має мати значення `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
