---
id: 699b6dd1a2961bd38491ef2e
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Об’єкт `profile` мав суворі ключі та типи. Якщо ви намагалися додати властивість, яка не була визначена, або призначити несумісний тип, TypeScript видавав помилку.

Було б добре, якби об’єкт `userRoles` не був таким суворим і мав більш гнучкі ключі.

Як ви пам’ятаєте з попередніх уроків, цього можна досягти, використовуючи тип `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

У цьому прикладі `studentGrades` TypeScript не потрібно знати всі предмети наперед. Це дозволяє додавати більше предметів, як `Art` або `Music`, пізніше. TypeScript піклується лише про те, щоб значення для кожного ключа були числом.

Оновіть свій об’єкт `userRoles`, додавши рядковий тип `Record`. Призначте його `<string, string>`, щоб кожен ключ міг бути будь-яким рядком, а кожне значення мало бути рядком.

# --hints--

Ваш об’єкт `userRoles` має мати рядковий тип `Record`, призначений `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
