---
id: 699b62c99ec8c63470313365
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер настав час додати ще одну властивість до вашого об’єкта `profile`.

Почніть з оновлення рядкового типу об’єкта, щоб включити властивість `bio` з типом, встановленим у `string`.

# --hints--

Ви повинні оновити рядковий тип об’єкта, щоб включити властивість `bio`, встановлену в `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
