---
id: 699b5370a43e299b4673bcaa
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Було б корисно бачити властивості об’єкта `profile` у консолі під час його створення.

Додайте `console.log`, який виводить `profile` у консоль.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Ви повинні вивести `profile` у консоль.

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
