---
id: 699b64096301fd75e3efeb1b
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Якщо ви знову відкриєте консоль, то побачите інше повідомлення про помилку:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Як ви пам’ятаєте з попередніх уроків, ви можете позначати властивості як необов’язкові, використовуючи символ `?` поруч із назвою властивості ось так:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

У цьому прикладі `numberOfAlbums` позначено як необов’язкову, тому TypeScript не видасть помилку в цьому випадку.

Оновіть свою властивість `bio`, щоб вона була позначена як необов’язкова.

# --hints--

Ваша властивість `bio` має бути позначена як необов’язкова.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
