---
id: 699b4e284291adcfcb90df47
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви попрактикуєтеся у роботі з примітивними типами, а також з об’єктами та масивами, створюючи типобезпечний профіль користувача.

Почніть зі створення змінної з назвою `profile` і призначте їй об’єкт. Ваш об’єкт має містити ці три властивості та значення:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

У вас має бути змінна з назвою `profile`.

```js
assert.exists(profile);
```

Ваша змінна `profile` має бути об’єктом.

```js
assert.isObject(profile);
```

У вашому об’єкті `username` має бути властивість `profile`.

```js
assert.property(profile, "username");
```

Властивість `username` має мати значення `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Ваш об’єкт `profile` має містити властивість `age`.

```js
assert.property(profile, "age");
```

Властивість `age` має мати значення `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Ваш об’єкт `profile` має містити властивість `isLoggedIn`.

```js
assert.property(profile, "isLoggedIn");
```

Властивість `isLoggedIn` має мати значення `false`.

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
