---
id: 699b71d98c820ac7ab9eea2a
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Якщо ви знову відкриєте консоль, то побачите нове повідомлення про помилку:

```md
Type 'number' is not assignable to type 'string'.
```

Це відбувається тому, що ви намагаєтеся додати властивість `guest` зі значенням типу число.

Щоб усунути помилку TypeScript, змініть значення властивості `guest` на `"read-only"`.

Після цього останнього кроку ваше практичне заняття з профілю користувача буде завершене!

# --hints--

Ваша властивість `guest` для об’єкта `userRoles` має мати значення `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
