---
id: 699b65caf17a8e8d340af27d
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Тепер час додати останню властивість до об’єкта `profile`.

У попередніх уроках ви навчились працювати з типами масивів. Ви дізналися про різні типи масивів, такі як `string[]` або `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

У цьому прикладі масив `scores` приймає лише числа. Якщо ви спробуєте додати інший тип, наприклад `string` або `null`, TypeScript видасть помилку.

Оновіть свій об’єкт `profile`, щоб додати властивість `programmingLanguages`, типізовану як `string[]`. Усередині літерала об’єкта встановіть цю властивість у `["JavaScript", "Python", "C++"]`.

# --hints--

У вашому об’єкті `programmingLanguages` має бути властивість `profile`.

```js
assert.property(profile, "programmingLanguages");
```

Властивість `programmingLanguages` має мати значення `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Властивість `programmingLanguages` має бути типізована як `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
