---
id: 5cd9a70215d3c4e65518328f
title: Побудувати зворотний відлік
challengeType: 26
dashedName: build-a-countdown
---

# --description--

У цій лабораторній роботі ви створите функцію `countdown`, яка прийматиме початкове число і генеруватиме список зворотного відліку.

Коли ви викликаєте `countdown(n)`, вона має повертати масив, який починається з `n` і рахує вниз з кроком `1` до `1`.

Наприклад:

- `countdown(5)` має повертати `[5, 4, 3, 2, 1]`
- `countdown(1)` має повертати `[1]`

Якщо початкове число менше за `1`, відліку робити нема чого, тому поверніть порожній масив:

- `countdown(0)` має повертати `[]`
- `countdown(-3)` має повертати `[]`

Щоб завершити лабораторну роботу, потрібно побудувати результат за допомогою рекурсії і не використовувати цикли (`for`, `while`), `Array.from()` або методи ітерації масивів (`forEach`, `map`, `filter`, `reduce`). Кожен виклик має створювати власний масив результату (не використовуйте глобальні змінні для збереження відліку).

**Мета:** Виконайте наведені нижче історії користувача і пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви маєте створити функцію з іменем `countdown`.
2. Функція `countdown` має приймати один аргумент `n` (число).
3. Якщо `n` менше за `1`, `countdown` має повертати порожній масив.
4. Інакше `countdown` має повертати масив, що містить цілі числа від `n` до `1` у спадному порядку.
5. Ви маєте використовувати рекурсію (функція має викликати сама себе) і не використовувати жодних циклів (`for`, `while`, `Array.from()`) або методів вищого порядку, таких як `forEach`, `map`, `filter`, `reduce`.
6. Кожен рекурсивний виклик має використовувати менше значення, ніж попередній виклик (наприклад, `n - 1`), щоб функція досягла базового випадку.
7. Рішення не повинно використовувати глобальні змінні для збереження або кешування масиву результату.
8. Виклик `countdown` кілька разів з різними вхідними даними має завжди повертати правильний, незалежний результат.

# --hints--

Ви маєте створити функцію з іменем `countdown`.

```js
assert.isFunction(countdown);
```

`countdown(-1)` має повертати порожній масив.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)` має повертати `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)` має повертати `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

Ваш код не повинен покладатися на будь-які цикли (`for`, `while`, `Array.from()`) або функції вищого порядку, такі як `forEach`, `map`, `filter` і `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

Ви маєте використовувати рекурсію для розв’язання цієї задачі.

```js
assert.isAtLeast(
  __helpers.removeJSComments(countdown.toString()).match(/countdown\s*\(.+\)/g)?.length,
  2
);
```

Глобальні змінні не повинні використовуватися для кешування масиву.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
