---
id: 66c057041df6394ca796bf33
title: Створіть генератор речень
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

У цій лабораторній роботі ви створите дві різні історії, використовуючи шаблон речення. Ви будете використовувати змінні для збереження різних частин історії, а потім виводити історії в консоль.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні оголосити такі змінні, використовуючи `let`:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. Ви повинні призначити цим змінним деякі рядкові значення на свій вибір.

1. Ви повинні оголосити змінну `firstStory`.
   
1. Ви повинні використати наведений шаблон історії для створення першої історії та призначити її змінній `firstStory`: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. Ви повинні вивести вашу першу історію в консоль за допомогою повідомлення `"First story: [firstStory]"`.

1. Ви повинні призначити нові значення змінним `adjective`, `noun`, `verb`, `place`, `adjective2` та `noun2`.

1. Ви повинні оголосити змінну `secondStory`.

1. Створіть іншу історію, використовуючи той самий шаблон, і призначте її змінній `secondStory`.

1. Ви повинні вивести вашу другу історію в консоль за допомогою повідомлення `"Second story: [secondStory]"`.


# --hints--

Ви повинні оголосити змінну `adjective`.

```js
assert.isNotNull(adjective);
```

Ви повинні оголосити змінну `noun`.

```js
assert.isNotNull(noun);
```

Ви повинні оголосити змінну `verb`.

```js
assert.isNotNull(verb);
```

Ви повинні оголосити змінну `place`.

```js
assert.isNotNull(place);
```

Ви повинні оголосити змінну `adjective2`.

```js
assert.isNotNull(adjective2);
```

Ви повинні оголосити змінну `noun2`.

```js
assert.isNotNull(noun2);
```

Ви повинні призначити рядкове значення змінній `adjective`.

```js
assert.isString(adjective);
```

Ви повинні призначити рядкове значення змінній `noun`.

```js
assert.isString(noun);
```

Ви повинні призначити рядкове значення змінній `verb`.

```js
assert.isString(verb);
```

Ви повинні призначити рядкове значення змінній `place`.

```js
assert.isString(place);
```

Ви повинні призначити рядкове значення змінній `adjective2`.

```js
assert.isString(adjective2);
```

Ви повинні призначити рядкове значення змінній `noun2`.

```js
assert.isString(noun2);
```

Ви повинні оголосити змінну `firstStory`.

```js
assert.isNotNull(firstStory);
```

Ви повинні використати правильний формат історії для першої історії: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Зверніть увагу на пробіли.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

Ви повинні зібрати вашу першу історію, використовуючи оголошені змінні у правильному порядку.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Ви повинні вивести вашу першу історію, використовуючи повідомлення `"First story: [firstStory]"`.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

Ви повинні оголосити змінну `secondStory`.

```js
assert.isNotNull(secondStory);
```

Ви повинні переназначити змінну `adjective` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

Ви повинні переназначити змінну `noun` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

Ви повинні переназначити змінну `verb` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

Ви повинні переназначити змінну `place` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

Ви повинні переназначити змінну `adjective2` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

Ви повинні переназначити змінну `noun2` для другої історії.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

Ви повинні використати правильний формат історії для другої історії: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. Зверніть увагу на пробіли.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

Ви повинні зібрати вашу другу історію, використовуючи оголошені змінні у правильному порядку.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

Ви повинні вивести вашу другу історію, використовуючи формат `"Second story: [secondStory]"`.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory` не повинен бути таким самим, як `secondStory`.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
