---
id: 65ef1cda150a59c3b8306944
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Si la diferencia no está dentro de la `tolerance` especificada, crea una estructura `elif` para verificar si `square_mid` es menor que `square_target`.

Asigna el valor de `mid` a `low` ya que la raíz cuadrada ahora estaría entre `low` y `mid`.

# --hints--

Debes crear una estructura `elif` que verifique si `square_mid < square_target`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[1].is_equivalent("square_mid < square_target")`))     
                  
    }

})
```

Debes asignar el valor de `mid` a `low`.

```js
({
    test: () => 
    {
          assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid")`))
    }

})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

--fcc-editable-region--
            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

--fcc-editable-region--
```
