---
id: 65ef1d5e3d2927c5e0f4997b
title: Paso 17
challengeType: 20
dashedName: step-17
---

# --description--

En Python, la palabra clave `is` verifica la identidad del objeto. Se usa para determinar si dos variables apuntan al mismo objeto en la memoria. En contraste con `is`, el operador de igualdad (`==`) determina si los valores de dos objetos son iguales, sin importar si son el mismo objeto en la memoria.

Fuera del bucle for, crea una estructura `if` para verificar si `root is None`. Si es así, imprime el mensaje `'Failed to converge within {max_iterations} iterations.'`. Recuerda formatear el mensaje usando un f-string.

# --hints--

Debes crear una estructura `if` que verifique si `root is None`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].find_conditions()[0].is_equivalent("root is None")`));

    }
})
```

Debes pasar el argumento `f'Failed to converge within {max_iterations} iterations.'` a la función `print`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_ifs()[0].is_equivalent("if root is None:\\n    print(f'Failed to converge within {max_iterations} iterations.')")`));
       
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if abs(square_mid - square_target) < tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

--fcc-editable-region--

--fcc-editable-region--

```
