---
id: 65ef1be2b958cabf712296cd
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora reducirás repetidamente el intervalo encontrando el punto medio del intervalo actual y comparando el cuadrado del punto medio con el valor objetivo.

Para eso, dentro del bloque `else`, crea un ciclo `for` que se ejecute hasta `max_iterations` veces.

Para tu ciclo, usa la función `range`, que genera una secuencia de números sobre la que puedes iterar. La sintaxis es `range(start, stop, step)`, donde `start` es el entero inicial (inclusive), `stop` es el último entero (no inclusive), y `step` es la diferencia entre un número y el anterior en la secuencia.

Además, usa `_` como variable de bucle. El `_` actúa como un marcador de posición y es útil cuando necesitas usar una variable pero en realidad no necesitas su valor.

# --hints--

Debes crear un ciclo `for` para iterar sobre `range(max_iterations)`. Usa `_` como la variable del ciclo.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--

--fcc-editable-region--
```
