---
id: 65ef190c6b51e9b5a5f7ed29
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Da a la función `square_root_bisection` los siguientes parámetros:

- `square_target`: El número para el cual quieres encontrar la raíz cuadrada.
- `tolerance` (opcional): La diferencia aceptable entre el cuadrado del valor aproximado de la raíz y el valor objetivo real (el valor predeterminado es `1e-7`). La tolerancia `1e-7` implica que la solución será precisa dentro de `0.0000001` del valor verdadero y es una buena opción predeterminada que equilibra precisión y rendimiento.
- `max_iterations` (opcional): El número máximo de iteraciones a realizar (el valor predeterminado es `100`). Si el método no converge dentro de este límite, asumirás que la solución no se encontró.
  
# --hints--

Tu función debe tener estos parámetros: `square_target`, `tolerance = 1e-7` y `max_iterations = 100`. El orden importa.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

