---
id: 65ef198fde24dfb7ff675b42
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

La sentencia `raise` te permite forzar una excepción específica para ocurrir. Consiste en la palabra clave `raise` seguida del tipo de excepción, y le permite proporcionar un mensaje de error personalizado:

```py
raise ValueError("Invalid value")
```

Cuando el código anterior se ejecuta, se eleva un `ValueError` y se muestra el mensaje `"Invalid value"` al usuario.

Si el `square_target` es menor que `0`, no se puede calcular ninguna raíz cuadrada con valor real. Por lo tanto, genere un `ValueError` con el mensaje `'Square root of negative number is not defined in real numbers'`. No olvides eliminar la palabra clave `pass`.

# --hints--

Debes reemplazar la palabra clave `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

Debes generar un `ValueError` con el mensaje `'Square root of negative number is not defined in real numbers'` dentro del cuerpo `if`.

```js
({
    test: () => {
        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
