---
id: 65ef181dad3d96b14a21a03e
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este proyecto, encontrarás la raíz cuadrada aproximada de un número dado usando el método de bisección.

El método de bisección es una técnica para encontrar las raíces de una función de valor real. Funciona reduciendo un intervalo donde se encuentra la raíz cuadrada hasta que converge a un valor dentro de una tolerancia especificada.

Comienza creando una función llamada `square_root_bisection`. Deja los parámetros vacíos por ahora.

# --hints--

Debes definir la función `square_root_bisection` sin parámetros.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
