---
id: 65ef1a50049cf9bada13266f
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Crea una estructura `elif` para verificar si `square_target` es igual a `0`. Si es así, asigna el valor `0` a la variable `root`. Además, imprime el mensaje `'The square root of {square_target} is 0'`. Recuerda formatear el mensaje usando una f-string.

# --hints--

Debes tener una estructura `elif` para verificar la condición `square_target == 0`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_conditions()[1].is_equivalent("square_target == 0")`));
    }
})

```

Debes asignar el valor `0` a la variable `root` y pasar el argumento `f'The square root of {square_target} is 0'` a una llamada `print`.


```js

({
    test: () => 
    {        
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[1].is_equivalent("root = 0\\nprint(f'The square root of {square_target} is 0')")`));
    }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    
--fcc-editable-region--
```
