---
id: 65ef19425d1b27b6c930bae6
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Si el número para el que quieres encontrar la raíz cuadrada es negativo, el código debe generar un error ya que la raíz cuadrada de un número negativo no está definida en los números reales.

Elimina la sentencia `pass` y crea una sentencia `if` para verificar si `square_target` es menor que `0`.

# --hints--

Debes reemplazar la palabra clave `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").has_pass()`))
    }
})
```

Debes crear una estructura `if` para verificar si `square_target < 0`.

```js
({
    test: () => {
       assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_if("square_target < 0").is_empty()`))

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance = 1e-7, max_iterations = 100):
    pass
--fcc-editable-region--
```
