---
id: 65ef1c34a541afc12dbb849d
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

La función `abs()` devuelve el valor absoluto de un número, que siempre es positivo, sin importar el signo del número. La usarás para verificar que la raíz cuadrada estimada esté lo suficientemente cerca del valor real.

Ahora, crea una estructura `if` para verificar si el valor absoluto de la diferencia entre `square_mid` y `square_target` está dentro de la `tolerance` especificada.

# --hints--

Debes tener una estructura `if` con la condición `abs(square_mid - square_target) < tolerance:` dentro del cuerpo del ciclo `for`.

```js
({
    test: () =>
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("abs(square_mid - square_target) < tolerance")`))
    } 
})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None
        
--fcc-editable-region--
        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2
            
--fcc-editable-region--
```
