---
id: 683ec95f2f8781355556ff84
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`elif`의 조건이 참이면 `pass` 키워드를 제거하고 `low`을 `1` 변수에 더해 `mid` 변수의 값을 업데이트하세요.

이렇게 하면 리스트에서 현재 탐색 영역의 오른쪽 절반으로 탐색 범위가 확장됩니다. `value`가 `value_at_middle`보다 크다는 것은 `value`이 현재 탐색 영역의 오른쪽 절반에 있어야 함을 의미하기 때문입니다.

# --hints--

`pass` 키워드를 제거하고 `low` 변수를 `mid + 1`으로 업데이트하세요.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
