---
id: 683ec95f2f8781355556ff80
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

다음으로 할 일은 대상 `value`가 중간에 있는지 확인하는 조건을 만드는 것입니다.

`while` 반복문 안에 대상 `if`가 `value`와 같은지 확인하는 `value_at_middle` 문을 만드세요. 같으면 `path_to_target` 변수를 반환하세요.

# --hints--

대상 `if`가 `value`와 같은지 확인하는 `value_at_middle` 문을 만드세요.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

`path_to_target` 문 안에서 `if` 변수를 반환해야 합니다.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
