---
id: 683ec95f2f8781355556ff82
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

함수는 `search_list`라는 리스트와 찾고자 하는 `value`을 받는다는 점을 기억하세요.

이제 `binary_search([1, 2, 3, 4, 5], 3)`와 `binary_search([1, 2, 3, 4, 5, 9], 4)`으로 함수를 호출하고 바로 호출 결과를 출력하세요.

# --hints--

함수를 `binary_search([1, 2, 3, 4, 5], 3)`로 호출하고 바로 출력해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5], 3))")
    `))
})
```

함수를 `binary_search([1, 2, 3, 4, 5, 9], 4)`로 호출하고 바로 출력해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_call("print(binary_search([1, 2, 3, 4, 5, 9], 4))")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        break
    return []
          
--fcc-editable-region--

--fcc-editable-region--
```
