---
id: 683ec95f2f8781355556ff7c
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

검색 범위를 정의했으니, 이제 이진 탐색 알고리즘을 수행할 반복문을 만드세요.

이진 탐색은 검색 공간을 반복해서 좁혀 나가면서 작동합니다. 이 과정은 검사할 유효한 요소 범위가 있을 때 계속됩니다.

`while` 반복문을 사용해 `low` 포인터가 `high` 포인터보다 작거나 같을 때 계속되도록 표현할 수 있습니다.

만약 `low`가 `high`보다 커지면, 검색 공간이 비어 있다는 뜻이고 `value`는 목록에 없습니다.

`binary_search` 함수 안에 `while` 반복문을 만들고, 조건으로 `low`가 high보다 작거나 같은지 확인하세요. 반복문 안에는 일단 `pass` 키워드를 자리 표시자로 추가하세요.

# --hints--

`while`이 `low`보다 작거나 같을 동안 실행되는 `high` 반복문을 만들어야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
--fcc-editable-region--
    
--fcc-editable-region--
```
