---
id: 683ec95f2f8781355556ff7e
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`mid` 인덱스를 계산했으니, 이제 `search_list`에서 해당 인덱스에 있는 값을 실제로 가져와야 합니다. 이 값이 `value` 매개변수, 즉 검색하려는 대상과 비교할 값입니다.

`value_at_middle` 변수를 만들고 `search_list`에서 `mid` 인덱스에 위치한 요소를 할당하세요.

# --hints--

`value_at_middle` 변수를 만들고 `search_list[mid]`로 설정하세요.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2

        
--fcc-editable-region--
```
