---
id: 683ec95f2f8781355556ff7f
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 `value_at_middle`를 `path_to_target` 목록에 포함하여, 대상인지 여부와 상관없이 탐색 중에 거친 단계를 추적하세요.

# --hints--

`append()` 메서드를 사용해 `value_at_middle`를 `path_to_target` 목록에 추가하세요.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid] \\npath_to_target.append(value_at_middle)")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        
--fcc-editable-region--
```
