---
id: 683ec95f2f8781355556ff7a
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

작성하는 이진 탐색 알고리즘은 값이 발견되었는지 여부를 나타내기 위해 `True` 또는 `False`만 반환하지 않습니다. 또한 그 값을 찾기 위해 거치는 경로도 반환합니다. 이를 위해 경로를 추적하는 `path_to_target` 변수가 필요합니다.

`binary_search` 함수 안에서 `pass` 키워드를 `path_to_target` 변수로 바꾸고, 우선 빈 리스트로 초기화하세요.

# --hints--

`path_to_target` 함수 내에 빈 리스트로 설정된 `binary_search` 변수를 만들어야 합니다. `pass` 키워드는 꼭 제거하세요.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
