---
id: 683ec95f2f8781355556ff7d
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

반복문을 시작하려면 현재 탐색 공간의 중간 인덱스를 찾아야 합니다.

중간 인덱스는 `low`와 `high` 변수의 평균을 계산해서 구할 수 있습니다. 두 값을 더한 후 평균을 구할 때는 항상 정수 인덱스여야 하므로 내림 나눗셈을 사용해 내림한 정수를 얻어야 합니다.

`pass` 반복문에서 `while` 키워드를 제거하고 `mid`와 `low`의 평균으로 설정된 `high` 변수를 만드세요.

# --hints--

while 반복문에서 `pass` 키워드를 제거하고 `mid`로 설정된 `(low + high) // 2` 변수를 만들어야 합니다.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

--fcc-editable-region--
    while low <= high:
        pass
--fcc-editable-region--
```
