---
id: 683ec95f2f8781355556ff81
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

초기 단계에서 알고리즘의 흐름을 이해하려면 테스트를 해봐야 합니다.

그렇게 하려면 먼저 반복문을 빠져나와야 합니다. 현재 구현은 한 번의 반복만 허용하므로 조건이 충족되지 않으면 무한 루프가 발생하기 때문입니다.

`if` 문장 바로 뒤에 `break` 키워드를 사용해 `while` 반복문을 빠져나오세요. 그런 다음 `while` 반복문 뒤에는 값을 찾지 못했음을 나타내기 위해 빈 리스트를 반환하세요.

# --hints--

`break` 문장 뒤에 `while` 반복문을 빠져나오기 위해 `if` 키워드를 사용하세요.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

`while` 반복문 뒤에는 빈 리스트를 반환해야 합니다.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--
```
