---
id: 683ec95f2f8781355556ff7b
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이진 탐색에서는 가장 낮은 인덱스와 가장 높은 인덱스가 있습니다.

가장 낮은 인덱스는 현재 탐색 중인 값들의 구간에서 가장 왼쪽 경계를 나타내며, 이 경우 `search_list`입니다. 반면 가장 높은 인덱스는 탐색 중인 값들에서 가장 오른쪽 경계를 나타냅니다.

이 두 가지를 고려하려면 탐색 중인 전체 리스트를 생각해야 합니다. 그래서 `binary_search` 함수 안에 `low` 변수를 `0` 값으로 정의하고, 탐색 중인 리스트의 마지막 인덱스를 나타내는 값을 가진 `high` 변수를 정의하세요.

# --hints--

`low` 함수 안에 `binary_search`라는 변수가 있어야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

`low` 변수는 `0`로 설정되어야 합니다.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

`high` 함수 안에 `binary_search`라는 변수가 있어야 합니다.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

`high` 변수는 `len(search_list) - 1`로 설정되어야 합니다.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []
    
--fcc-editable-region--
```
