---
id: 683ec95f2f8781355556ff83
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

첫 번째 호출이 `[3]`를 반환하는 것을 볼 수 있습니다. 이는 `3`이 `[1, 2, 3, 4, 5]`에서 중간값이기 때문입니다.

반면 두 번째 호출은 빈 리스트를 반환합니다. 이는 중간값을 확인한 후, 현재 루프가 `low` 또는 `high` 변수를 업데이트하여 검색 범위를 좁히는 로직이 없기 때문입니다. 단순히 중간값만 확인하고 추가 검색 없이 종료됩니다.

이진 검색이 계속해서 검색 범위를 좁힐 수 있도록, `elif`이 `value`보다 큰지 확인하는 `value_at_middle` 블록을 추가하세요. 우선 `pass` 안에 `elif` 키워드를 추가하세요.

# --hints--

`elif`이 `value`보다 클 때를 위한 `value_at_middle` 블록을 추가해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

--fcc-editable-region--
        if value == value_at_middle:
            return path_to_target

--fcc-editable-region--          
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
