---
id: 683ec39565dbc7e3375263b1
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이전 수업에서 이진 탐색과 선형 탐색의 차이점을 배웠습니다.

이번 워크숍에서는 이진 탐색 알고리즘을 직접 만들어 보면서 배운 내용을 실습합니다.

먼저, `binary_search` 함수를 `search_list`과 `value` 매개변수로 만드세요.

`search_list`은 함수가 탐색할 요소들의 리스트이고, `value`는 함수가 `search_list` 내에서 찾으려는 대상입니다.

테스트가 통과하도록 함수 안에 `pass` 키워드를 일단 추가하세요.

# --hints--

`binary_search` 함수를 만들어야 합니다.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

`binary_search` 함수는 `search_list`과 `value` 매개변수를 가져야 합니다.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
