---
id: 684169b00a073c9b9f0a703c
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

두 번째 함수 호출에 `3, 5, 4` 값이 목록에 있는 것을 보셨나요? 이것은 이진 탐색이 의도한 대로 작동하고 있음을 나타냅니다.

이렇게 된 과정은 다음과 같습니다: 알고리즘이 먼저 `3`을 초기 목록의 중간으로 확인했습니다. `4`가 `3`보다 크기 때문에 탐색이 오른쪽 절반으로 이동했습니다. 그다음 새로운 중간인 `5`를 조사했습니다. `4`가 `5`보다 작아서 탐색이 왼쪽으로 이동했고, 결국 `4`을 최종 범위의 중간으로 확인했습니다.

함수를 다시 테스트하려면 `[1, 3, 5, 9, 14, 22], 10`로 호출하고 바로 출력을 하세요. 이 경우에는 값이 발견되지 않을 것입니다.

# --hints--

함수를 `binary_search([1, 3, 5, 9, 14, 22], 10)`로 호출하고 바로 출력해야 합니다.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

    return []
    
print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
