---
id: 68416aab85dd00a7852cdeb6
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

마지막 호출이 빈 리스트를 출력하는 것을 볼 수 있습니다. 검색한 리스트에 값이 없다는 것을 더 명확히 지정하려면 `return` 문장을 업데이트하여 빈 리스트와 메시지 `Value not found`를 반환하세요.

# --hints--

`return` 문장을 업데이트하여 `[], 'Value not found'`를 반환해야 합니다.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").has_return("[], 'Value not found'")`
            )
        )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1

--fcc-editable-region--
    return []
--fcc-editable-region--

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
print(binary_search([1, 3, 5, 9, 14, 22], 10))
```
