---
id: 690ddb5e6ee94eb73ed172b0
title: 偵錯捐款表單
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

一個區域的慈善機構建立了一個捐款表單網站，但有幾個問題需要修正。該表單無法存取，且有一些 HTML 語法錯誤。

你的工作是修正所有錯誤，使表單能正確運作並讓所有使用者都能存取。完成以下使用者故事中的項目，然後點擊「Run the Tests」以檢查你是否修正了所有錯誤。

**使用者故事：**

1. `input` 元素是空元素，不應該有結束標籤。請從表單中移除所有 `</input>` 結束標籤。
2. 為你的每個表單輸入欄添加 `label` 元素，讓使用者知道每個欄位的用途。標籤文字應該與目前每個輸入欄旁邊的文字相符。
3. `Email Address:` 輸入型別應該是 `email`，而非 `text`。
4. 你應該使用 `label` 上的 `for` 屬性和 `input` 上相符的 `id` 屬性，將每個 `label` 元素與其對應的 `input` 元素關聯起來。
5. 為你的文字、電子郵件和數字輸入欄添加 `required` 屬性（但不包含核取方塊或送出按鈕），以確保使用者填寫必填資訊。

# --hints--

你的程式碼中不應該有任何 `</input>` 結束標籤。

```js
assert.notMatch(code, /<\/input>/);
```

你的表單中應該有正好五個 `input` 元素。

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

你的表單中應該有正好四個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

你的第一個 `label` 應該有文字 `Full Name:`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

你的第一個 `input` 應該有一個 `required` 屬性。

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

你的第二個 `label` 應該有 `Email Address:` 的文字。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

你的第二個 `input` 應該有一個 `required` 屬性。

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

你的第三個 `label` 應該有文字 `Donation Amount ($):`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

你的第三個 `input` 應該有一個 `required` 屬性。

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

你的第四個 `label` 應該有文字 `Subscribe`。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

你應該使用 `for` 和 `id` 屬性將每個 `input` 元素與 `label` 元素關聯起來。

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

你的 `input` 電子郵件位址應該是 `email` 型別。

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

你的核取方塊 `input` 不應該有 `required` 屬性。

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

你的提交 `button` 不應該有 `required` 屬性。

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
