---
id: 61329d52e5010e08d9b9d66b
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Outro elemento `meta` importante para a acessibilidade e para a SEO Ă© a `description`. O valor do atributo `content` Ă© usado pelos mecanismos de busca para fornecer uma descriĂ§ĂŁo de sua pĂĄgina.

Adicione um elemento `meta` com o atributo `name` definido como `description`, e dĂȘ a ele um atributo `content`.

# --hints--

VocĂȘ deve adicionar um novo elemento `meta` Ă  `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

VocĂȘ deve dar ao elemento `meta` um atributo `name` de `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

VocĂȘ deve dar ao elemento `meta` um atributo `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

O valor do atributo `content` nĂŁo deve ter mais de 165 caracteres. *Este Ă© o comprimento mĂĄximo da descriĂ§ĂŁo do Google.*

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
    
    --fcc-editable-region--
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
