---
id: 614202874ca576084fca625f
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

A funĂ§ĂŁo de cada `region` requer um rĂłtulo, o que ajuda os usuĂĄrios de leitores de tela a entenderem o propĂłsito da regiĂŁo. Um mĂ©todo para adicionar um rĂłtulo Ă© adicionar um elemento de tĂ­tulo dentro da regiĂŁo e, em seguida, referenciĂĄ-lo com o atributo `aria-labelledby`.

Adicione os seguintes atributos `aria-labelledby` aos elementos `section`:

- `student-info`
- `html-questions`
- `css-questions`

EntĂŁo, dentro de cada elemento `section`, aninhe um elemento `h2` com um `id` que corresponda ao atributo `aria-labelledby` correspondente. DĂȘ a cada `h2` conteĂșdo de texto adequado.

# --hints--

VocĂȘ deve dar ao primeiro elemento `section` um atributo `aria-labelledby` de `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

VocĂȘ deve dar ao segundo elemento `section` um atributo `aria-labelledby` de `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

VocĂȘ deve dar ao terceiro elemento `section` um atributo `aria-labelledby` de `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

VocĂȘ deve aninhar um elemento `h2` dentro do primeiro elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

VocĂȘ deve aninhar um elemento `h2` dentro do segundo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

VocĂȘ deve aninhar um elemento `h2` dentro do terceiro elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

VocĂȘ deve dar ao primeiro elemento `h2` um atributo `id` de `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

VocĂȘ deve dar ao segundo elemento `h2` um atributo `id` de `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

VocĂȘ deve dar ao terceiro elemento `h2` um atributo `id` de `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

VocĂȘ deve fornecer ao primeiro elemento `h2` um conteĂșdo de texto adequado. *Dica: Eu teria escolhido `Student Info`*

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

VocĂȘ deve fornecer ao segundo elemento `h2` um conteĂșdo de texto adequado. *Dica: Eu teria escolhido `HTML`*

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

VocĂȘ deve fornecer ao terceiro elemento `h2` um conteĂșdo de texto adequado. *Dica: Eu teria escolhido `CSS`*

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        --fcc-editable-region--
        <section role="region">
          
        </section>
        <section role="region">
          
        </section>
        <section role="region">
          
        </section>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
