---
id: 614ccc21ea91ef1736b9b578
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Para este workshop, você recebeu um boilerplate básico.

Inicie esta jornada de acessibilidade fornecendo um atributo `lang` ao seu elemento `html`. Isto ajudará os leitores de tela a identificar o idioma da página.

Além disso, lembre-se de vincular sua folha de estilos à página.

# --hints--

Você deve dar ao elemento `html` um atributo `lang`. *Dica: Você pode usar o valor `en` para inglês.*

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Você deve ter um elemento `link` no seu elemento `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Seu elemento `link` deve ter `rel="stylesheet"` e `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
