---
id: 699f4b903bf5427a599f791d
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

In den letzten Schritten werden Sie eine weitere Funktion erstellen, um weiterhin das Annotieren von Typen bei Funktionen zu üben.

Beginnen Sie damit, eine Funktion namens `raiseTo` zu erstellen, die zwei Parameter namens `base` und `exponent` entgegennimmt. Beide Parameter sollten vom Typ `number` sein. Die Funktion soll das Ergebnis zurückgeben, wenn `base` zur Potenz von `exponent` erhoben wird. Der Rückgabetyp der Funktion sollte ebenfalls als `number` annotiert sein.

# --hints--

Sie sollten eine `raiseTo`-Funktion haben.

```js
assert.isFunction(raiseTo);
```

Ihre `raiseTo`-Funktion sollte zwei Parameter haben, `base` und `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*(?::\s*\w+\s*)?\s*=>/
);
```

Ihre Parameter `base` und `exponent` sollten beide vom Typ `number` sein.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Ihre `raiseTo`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(raiseTo(2, 3));
```

Ihre `raiseTo`-Funktion sollte das Ergebnis zurückgeben, wenn `base` zur Potenz von `exponent` erhoben wird.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Ihre `raiseTo`-Funktion sollte eine Rückgabetyp-Annotation von `number` haben.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
