---
id: 699f5063d2393c674b1f45dc
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Wie Sie sich aus einem vorherigen Schritt erinnern, wirft TypeScript einen Fehler, wenn Sie eine Funktion mit der falschen Anzahl von Argumenten aufrufen.

Es gibt eine Möglichkeit, Parameter als optional zu kennzeichnen, indem Sie nach dem Parameternamen ein `?` hinzufügen. Dadurch können Sie die Funktion mit weniger Argumenten als Parametern aufrufen.

Hier ist ein Beispiel:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

In diesem Beispiel, wenn der `greeting`-Parameter beim Aufrufen der `greet`-Funktion nicht angegeben wird, ist er `undefined`, und die Funktion gibt eine Standardbegrüßung zurück.

Aktualisieren Sie Ihre `raiseTo`-Funktion, um den `exponent`-Parameter optional zu machen.

# --hints--

Ihre `raiseTo`-Funktion sollte den `exponent`-Parameter als optional gekennzeichnet haben.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
