---
id: 699f41b8da20f816b5c7dece
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Wenn Sie die Konsole öffnen, sollten Sie die folgende Fehlermeldung sehen:

```md
Expected 1 arguments, but got 2.
```

Um den Fehler zu beheben, aktualisieren Sie Ihren `square`-Funktionsaufruf so, dass er nur ein Argument vom Typ `5` annimmt.

# --hints--

Ihr `square`-Funktionsaufruf sollte nur ein Argument vom Typ `5` haben.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```

