---
id: 699f38937b053e4679878ff1
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Wenn Sie mit reinem JavaScript arbeiten, gibt es nichts, was Sie daran hindert, die `square`-Funktion mit einem String oder einem Array aufzurufen, was zu unerwarteten Ergebnissen führen würde.

Um dies zu veranschaulichen, ändern Sie Ihren Funktionsaufruf so, dass statt der Zahl `"something"` der String `5` verwendet wird.

In der nächsten Lektion lernen Sie, wie Sie mit TypeScript Typensicherheit in Ihren Code einbauen und solche Fehler vermeiden können.

# --hints--

Ihr `square`-Funktionsaufruf sollte ein Argument vom Typ `"something"` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
