---
id: 699f4ec939f10da62fb56c3e
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Jetzt ist es Zeit, Ihre Funktion zu testen.

Erstellen Sie eine Variable namens `raisedResult` und weisen Sie ihr das Ergebnis der Ausführung von `raiseTo` mit den Argumenten `2` und `3` zu.

Geben Sie anschließend `raisedResult` in der Konsole aus.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Variable namens `raisedResult` haben.

```js
assert.isDefined(raisedResult);
```

Ihre `raisedResult`-Variable sollte eine Zahl sein.

```js
assert.isNumber(raisedResult);
```

Sie sollten `raisedResult` in der Konsole ausgeben.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
