---
id: 699f1d72e906e4616155290d
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie weiterhin, wie Typannotationen funktionieren und wie sie mit Funktionen zusammenwirken, indem Sie ein typsicheres Mathe-Toolkit erstellen.

Beginnen Sie damit, eine Funktion namens `square` zu erstellen, die einen Parameter akzeptiert und das Quadrat einer Zahl zurückgibt.

# --hints--

Sie sollten eine Funktion namens `square` haben.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Ihre `square`-Funktion sollte einen Parameter haben.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Ihre `square`-Funktion sollte das Quadrat einer Zahl zurückgeben.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
