---
id: 699f4861dd5b5f164e1bda4a
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Ein weiterer Vorteil der Verwendung von TypeScript ist, dass sichergestellt wird, dass die Rückgabetypen eingehalten werden und Ihren Erwartungen entsprechen. Manchmal kann TypeScript den Rückgabetyp einer Funktion ableiten, aber Sie können ihn auch explizit angeben.

Um dies zu veranschaulichen, ändern Sie den Rückgabetyp der `getAverage`-Funktion in den String `"something"`.

# --hints--

Ihre `getAverage`-Funktion sollte einen String zurückgeben.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Ihre `getAverage`-Funktion sollte den String `"something"` zurückgeben.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
