---
id: 699f3cb8d057f29954205d7d
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Wie Sie sich aus früheren Lektionen erinnern, können Sie Typannotationen zu Funktionsparametern wie folgt hinzufügen:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

In diesem Beispiel zeigt TypeScript einen Fehler an, wenn Sie versuchen, die `add`-Funktion mit Argumenten aufzurufen, die keine Zahlen sind.

Aktualisieren Sie Ihre `square`-Funktion, um eine Typannotation für den `num`-Parameter hinzuzufügen. Der Typ sollte `number` sein.

# --hints--

Sie sollten eine Typannotation für den `num`-Parameter in der `square`-Funktion hinzufügen. Der Typ sollte ein `number` sein.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
--fcc-editable-region--
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
