---
id: 699f4639ae0016caaf81b7c6
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Jetzt ist es Zeit, Ihre Funktion zu testen.

Beginnen Sie damit, eine Variable namens `avgResult` zu erstellen und ihr den Funktionsaufruf von `getAverage` zuzuweisen. Das Argument für diesen Funktionsaufruf sollte `[2, 14, 26, 8]` sein.

Geben Sie dann den Wert von `avgResult` in der Konsole aus.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Variable namens `avgResult` haben.

```js
assert.isDefined(avgResult);
```

Sie sollten `getAverage([2, 14, 26, 8])` Ihrer Variable `avgResult` zuweisen.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Ihre Variable `avgResult` sollte eine Zahl sein.

```js
assert.isNumber(avgResult);
```

Sie sollten `avgResult` in der Konsole ausgeben.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
