---
id: 699f52220de88c8091032076
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

Wenn Sie die Konsole öffnen, ist das Ergebnis `NaN`, weil das zweite Argument für die `exponent` nie angegeben wurde. Für diese `raiseTo`-Funktion sind ein paar Änderungen nötig, die das beheben.

Die erste Änderung besteht darin, `base ** 2` zurückzugeben.

Im nächsten Schritt sehen Sie, wie Sie mit dem optionalen `exponent`-Parameter umgehen.

# --hints--

Ihre `raiseTo`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(raiseTo(4));
```

Ihre `raiseTo`-Funktion sollte den Wert von `base` potenziert mit `2` zurückgeben.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
--fcc-editable-region--
  return base ** exponent;
--fcc-editable-region--
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
