---
id: 699f412114c87221f16a99e3
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Ein weiterer Vorteil von TypeScript ist, dass Fehler erkannt werden, wenn Sie zu viele oder zu wenige Argumente an einen Funktionsaufruf übergeben.

Um dies zu veranschaulichen, aktualisieren Sie Ihren `square`-Funktionsaufruf so, dass er zwei Argumente `5` und `10` erhält.

# --hints--

Ihr `square`-Funktionsaufruf sollte zwei Argumente `5` und `10` haben.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
