---
id: 699f52fac083214b6bf51de6
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

In diesem letzten Schritt müssen Sie eine bedingte Prüfung für den `exponent`-Parameter in der `raiseTo`-Funktion hinzufügen.

Wenn das `exponent`-Argument angegeben ist, soll die Funktion `base` potenziert mit `exponent` zurückgeben. Wenn das `exponent`-Argument nicht angegeben ist, soll sie `base ** 2` zurückgeben.

Mit diesem letzten Schritt ist Ihr typsicheres Mathe-Werkzeug komplett!

# --hints--

Ihre `raiseTo`-Funktion sollte `base ** exponent` zurückgeben, wenn das `exponent`-Argument angegeben ist.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Ihre `raiseTo`-Funktion sollte `base ** 2` zurückgeben, wenn das `exponent`-Argument nicht angegeben ist.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
