---
id: 699f4a3e7b4c6bfa538cc4f7
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Wenn Sie die Konsole öffnen, sehen Sie die folgende Fehlermeldung:

```md
Type 'string' is not assignable to type 'number'.
```

Um diesen Fehler zu beheben, müssen Sie die Rückgabe wieder auf `sum / numbers.length` ändern.

# --hints--

Ihre `getAverage`-Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Ihre `getAverage`-Funktion sollte den Durchschnitt der Zahlen im Array zurückgeben.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

