---
id: 699f3738921fb79a651ded56
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Jetzt ist es Zeit, Ihre Funktion aufzurufen.

Erstellen Sie eine Variable namens `result` und weisen Sie ihr den Funktionsaufruf von `square` mit dem Argument `5` zu.

Geben Sie darunter den Wert von `result` in der Konsole aus.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Variable namens `result` haben.

```js
assert.isDefined(result);
```

Ihre `result`-Variable sollte eine Zahl sein.

```js
assert.isNumber(result);
```

Ihre `result`-Variable sollte das Quadrat von `5` sein.

```js
assert.strictEqual(result, 25);
```

Sie sollten die `result`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
