---
id: 699f49669ea51a31d2ba89d8
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Wie Sie sich aus vorherigen Lektionen erinnern, können Sie den Rückgabetyp einer Funktion explizit so angeben:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

In diesem Beispiel ist der Rückgabetyp der `circleArea`-Funktion `number`. Wenn Sie versuchen, einen anderen Typ zurückzugeben, zeigt TypeScript einen Fehler an.

Aktualisieren Sie Ihre `getAverage`-Funktion, sodass der Rückgabetyp explizit als `number` angegeben ist.

# --hints--

Ihre `getAverage`-Funktion sollte einen Rückgabetyp von `number` haben.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
--fcc-editable-region--
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
