---
id: 699f3f636aa840b7201aa37c
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Wenn Sie die Konsole öffnen, sollten Sie die folgende TypeScript-Fehlermeldung sehen:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Da der Parameter jetzt explizit als Zahl typisiert ist, kann TypeScript den Fehler erkennen, dass versucht wird, ein String-Argument an die `square`-Funktion zu übergeben.

Ändern Sie das an die `square`-Funktion übergebene Argument in die Zahl `5`, um den Fehler zu beheben und die korrekte Ausgabe in der Konsole zu sehen.

# --hints--

Ihr `square`-Funktionsaufruf sollte ein Argument von `5` anstelle von `"something"` haben.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```

