---
id: 699f434e294a9e11cf4f65f0
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Jetzt ist es Zeit, Ihrem Toolkit-Projekt einen weiteren mathematischen Helfer hinzuzufügen.

Beginnen Sie damit, eine Funktion namens `getAverage` zu erstellen, die einen Parameter entgegennimmt. Ihre Funktion soll den Durchschnitt für ein Array von Zahlen zurückgeben.

In den nächsten Schritten fügen Sie die Typannotationen für diese Funktion hinzu.

# --hints--

Sie sollten eine Funktion namens `getAverage` haben.

```js
assert.isFunction(getAverage);
```

Ihre Funktion sollte einen Parameter entgegennehmen.

```js
assert.strictEqual(getAverage.length, 1);
```

Ihre Funktion sollte eine Zahl zurückgeben.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Ihre Funktion sollte für jedes Array von Zahlen das korrekte Ergebnis zurückgeben.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
