---
id: 699f458e47fb88c756fdba81
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Im Moment hat der `numbers`-Parameter einen impliziten `any`-Typ. Um das zu beheben, können Sie den Parameter explizit als Array von Zahlen typisieren.

Erinnern Sie sich aus vorherigen Lektionen, dass Sie ein Array von Zahlen so typisieren können:

```ts
const numbers: number[] = [1, 2, 3];
```

Aktualisieren Sie Ihre `getAverage`-Funktion, um den `numbers`-Parameter explizit als Array von Zahlen zu typisieren.

# --hints--

Ihre `getAverage`-Funktion sollte einen Parameter namens `numbers` haben, der explizit als Array von Zahlen typisiert ist.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
--fcc-editable-region--
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
