---
id: 66ad8d150264db3926eccfeb
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

In the previous lectures, you learned about the `let` keyword and how to declare variables in JavaScript.

Aquí tienes un recordatorio de cómo declarar una variable usando la palabra clave `let`:

```js
let greeting;
```

Usa la palabra clave `let` para declarar una variable llamada `bot`.

**NOTA**: Estás usando `let` aquí porque más adelante en el taller, estarás cambiando el valor de la variable `bot`.

# --hints--

Debes tener la palabra clave `let` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Debes declarar una variable llamada `bot`.

```js
assert.isNotNull(bot);
```

Debes usar la palabra clave `let` para declarar la variable.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
