---
id: 66ad998285df023b23dacdd3
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Ahora, es el momento de asignar algunos valores a sus variables `bot` y `botLocation`.

In the previous lectures, you learned how to assign values to variables like this:

```js
variableName = "Here is the value";
```

Recuerda que lo que está en el lado derecho del operador de asignación `=` es el valor que estás asignando a la variable en el lado izquierdo.

Asigna la cadena `"teacherBot"` a la variable `bot` y la cadena `"the universe"` a la variable `botLocation`.

# --hints--

Debes asignar una cadena a tu variable `bot`.

```js
assert.isString(bot);
```

Tu variable `bot` debe tener asignada la cadena `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

Debes asignar una cadena a tu variable `botLocation`.

```js
assert.isString(botLocation);
```

Tu variable `botLocation` debe tener asignada la cadena `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

--fcc-editable-region--

--fcc-editable-region--
```
