---
id: 66ad8294a0ad902f1b31b612
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, aprenderás a trabajar con los fundamentos de JavaScript construyendo un bot de saludo.

En este primer paso, querrás mostrar un mensaje en la consola desde el bot de saludo.

Remember that you learned about `console.log()` and strings in the previous lectures.

Aquí tienes un recordatorio de cómo usar `console.log()` con cadenas de texto:

```js
console.log("Hello, World!");
```

Agrega una declaración `console.log()` que muestre la cadena `"Hi there!"` en la consola. ¡No olvides tus comillas alrededor del mensaje!


# --hints--

Debes tener una declaración `console.log()` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

Tu declaración `console` debe tener el mensaje `"Hi there!"`. No olvides las comillas.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
