---
id: 66adc6046acc18453decf577
title: Paso 12
challengeType: 1
dashedName: step-12
---

# --description--

Para ver el nuevo apodo del bot, necesitarás registrar un nuevo mensaje en la consola.

Crea una nueva variable llamada `newNicknameGreeting`.

Luego usa concatenación de cadenas para unir la cadena `"I love my nickname but I wish people would call me "` con la variable `bot` seguida de un punto.

Asignar el resultado a la variable `newNicknameGreeting`.

Luego, registra el valor de `newNicknameGreeting` en la consola.

# --hints--

Debes tener una variable llamada `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

Tu variable `newNicknameGreeting` debe ser una cadena.

```js
assert.isString(newNicknameGreeting);
```

Debes usar concatenación de cadenas para unir la cadena `"I love my nickname but I wish people would call me "` con la variable `bot` seguida de un punto.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Debes tener una declaración `console` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Debes registrar el valor de `newNicknameGreeting` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
