---
id: 66adcf383276814776aba3ca
title: Paso 14
challengeType: 1
dashedName: step-14
---

# --description--

A continuación, crea una variable llamada `favoriteSubjectSentence`.

Usa la concatenación de cadenas para unir la cadena `"Mi asignatura favorita es "` con la variable `favoriteSubject` seguida de un punto.

Asigna el resultado a la variable `favoriteSubjectSentence`.

Luego, registra el valor de `favoriteSubjectSentence` en la `consola`.

# --hints--

Debes tener una variable llamada `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Tu variable `favoriteSubjectSentence` debe ser una cadena.

```js
assert.isString(favoriteSubjectSentence);
```

Debes usar la concatenación de cadenas para unir la cadena `"Mi asignatura favorita es "` con la variable `favoriteSubject` seguida de un punto. Ten en cuenta los espacios.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Debes tener una instrucción `console` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Debes registrar el valor de `favoriteSubjectSentence` en la `consola`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
