---
id: 66ad9c633fa26d3c1475eae3
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Ahora, es momento de añadir otro mensaje del bot.

Añade otra declaración de `console` al código que registre el mensaje `"Allow me to introduce myself."`.

# --hints--

Debe haber una tercera declaración de `console.log()` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Tu declaración de `console` debería tener el mensaje `"Allow me to introduce myself."`. No olvides las comillas.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
