---
id: 66adbc8a9793d64250f7e609
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

En los videos anteriores de la clase, aprendiste cómo reasignar valores a las variables como esto:

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

Usando la reasignación, asigna la cadena `"professorBot"` a la variable `bot`.

# --hints--

Debes asignar la cadena `"professorBot"` a la variable `bot`.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
