---
id: 68ed661e9f463c155bb677a1
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este taller practicarás trabajar con CSS flexbox diseñando un conjunto de cajas de colores.

Para comenzar, dentro del elemento `body`, agrega un elemento `header`. Anida un elemento `h1` dentro del elemento `header` con el texto `Colored Boxes Layout`.

# --hints--

Debes tener un elemento `header` dentro del elemento `body`.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Tu elemento `header` debe tener un elemento `h1` dentro de él.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Tu elemento `h1` debe tener el texto `Colored Boxes Layout` dentro de él.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
