---
id: 68ef0b3237fef991c92d9c19
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Ahora agrega un elemento `div` con la clase `flex-container` debajo de tu `header`. Luego crea seis elementos `div` con la clase `box` dentro del `div` que acabas de crear.

Recuerda que un atributo class se usa a menudo para apuntar a un nombre de clase en una hoja de estilos. En este caso, todos los elementos `div` se estilizarán por igual según la definición de estilo `.flex-container` y `.box`.

# --hints--

Debes tener un elemento `div` debajo del elemento `header`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.exists(divELBelowHeader);
```

Tu elemento `div` debajo del elemento `header` debe tener la clase `flex-container`.

```js
const divELBelowHeader = document.querySelector('header + div');
assert.equal(divELBelowHeader.className, 'flex-container');
```

Debes tener seis elementos `div` anidados dentro del elemento `.flex-container`.

```js
const divElsInsideContainer = document.querySelectorAll('div.flex-container > div');
assert.lengthOf(divElsInsideContainer, 6);
```

Los seis elementos `div` dentro del elemento `.flex-container` deben tener la clase `box`.

```js
const divElsInsideContainerWithBox = document.querySelectorAll('div.flex-container > div.box');
assert.lengthOf(divElsInsideContainerWithBox, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```

```css
h1 {
	text-align: center;
	margin-bottom: 10px;
}
```
