---
id: 68f91df5776e693049fb8942
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Ahora es momento de agregar algunas reglas CSS al archivo `styles.css`. Comienza creando un selector para el elemento `h1`. 

Centra el contenido del elemento `h1` configurando su propiedad `text-align` a `center`. También crea una propiedad `margin-bottom` con el valor `10px` para añadir margen entre la parte inferior del elemento `h1` y cualquier elemento HTML que esté debajo.

# --hints--

Deberías tener un selector `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

El selector `h1` debe darle a la propiedad `text-align` el valor `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Tu selector `h1` debe establecer la propiedad `margin-bottom` en `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

