---
id: 68f932c803f4d3d9f511b436
title: Paso 20
challengeType: 0
dashedName: step-20
---

# --description--

La propiedad `flex-wrap` define si los elementos flex se obligan a estar en una sola línea o pueden ajustarse en varias líneas.

Por defecto, los elementos flex tienen un valor de `nowrap`, lo que significa que permanecen en una sola línea incluso si desbordan el contenedor. El valor `wrap` en cambio hace que los elementos flex se ajusten a nuevas líneas cuando no caben en el ancho del contenedor, creando un diseño más flexible y responsivo.

Para continuar, agrega la propiedad `flex-wrap` con el valor `wrap` al selector de clase `.flex-container`.

# --hints--

Tu selector `.flex-container` debe tener la propiedad `flex-wrap` con el valor `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('flex-wrap'), 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
