---
id: 692f2c8ea9aa5a593758ccff
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

La propiedad `display` controla cómo se renderiza un elemento en la página. Cuando la defines con el valor `flex`, esto habilita el modo de diseño flexbox, que te da control para organizar y alinear los elementos dentro del contenedor.

Ahora agrega la propiedad `display` con el valor `flex` al selector `.box`.

# --hints--

Tu selector `.box` debe tener la propiedad `display` con el valor `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('display'), 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
