---
id: 692f22933ba1cb86cd75a70c
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Ahora es momento de crear el nombre del color para las cajas. Para hacerlo, agrega un elemento `p` debajo de cada elemento `h2` con estos nombres de color en orden: `Red`, `Orange`, `Yellow`, `Green`, `Blue` y `Indigo`. Un color para cada caja.

# --hints--

Debes agregar un elemento `p` con el texto `Red` debajo del elemento `h2` dentro del primer elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

Debes agregar un elemento `p` con el texto `Orange` debajo del elemento `h2` dentro del segundo elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

Debes agregar un elemento `p` con el texto `Yellow` debajo del elemento `h2` dentro del tercer elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

Debes agregar un elemento `p` con el texto `Green` debajo del elemento `h2` dentro del cuarto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

Debes agregar un elemento `p` con el texto `Blue` debajo del elemento `h2` dentro del quinto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

Debes agregar un elemento `p` con el texto `Indigo` debajo del elemento `h2` dentro del sexto elemento `.box`.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
    <div class="box">
      <h2>Box 1</h2>
			
    </div>
    <div class="box">
      <h2>Box 2</h2>
			
    </div>
    <div class="box">
      <h2>Box 3</h2>
			
    </div>
    <div class="box">
      <h2>Box 4</h2>
			
    </div>
    <div class="box">
      <h2>Box 5</h2>
			
    </div>
    <div class="box">
      <h2>Box 6</h2>
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
