---
id: 692f22a2ad44ae887a24c8f9
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro de cada elemento `.box`, agrega un elemento `h2` con el texto `Box` seguido de un espacio y un número secuencial que comienza desde `1`.

# --hints--

Debes tener un elemento `h2` con el texto `Box 1` dentro del primer `div` con la clase `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

Debes tener un elemento `h2` con el texto `Box 2` dentro del segundo `div` con la clase `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

Debes tener un elemento `h2` con el texto `Box 3` dentro del tercer `div` con la clase `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

Debes tener un elemento `h2` con el texto `Box 4` dentro del cuarto `div` con la clase `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

Debes tener un elemento `h2` con el texto `Box 5` dentro del quinto `div` con la clase `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

Debes tener un elemento `h2` con el texto `Box 6` dentro del último `div` con la clase `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
--fcc-editable-region--
		<div class="box">
			
		</div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
    <div class="box">
			
    </div>
--fcc-editable-region--
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
