---
id: 692f2d37f34ed866ed42fcd6
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

La propiedad `flex-direction` controla cómo se organizan los elementos flexibles dentro de su contenedor. El valor predeterminado es `row`, que organiza los elementos horizontalmente de izquierda a derecha. Esta propiedad solo funciona en elementos con `display: flex` o `display: inline-flex`.

Con el valor `column`, los elementos se apilan verticalmente de arriba hacia abajo. Otros valores posibles son `row`, `row-reverse` y `column-reverse`.

Ahora agrega la propiedad `flex-direction` con el valor `column` al selector `.box`.

# --hints--

Tu selector `.box` debe tener la propiedad `flex-direction` con el valor `column`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex-direction'), 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
--fcc-editable-region--
  
--fcc-editable-region--  
}
```
