---
id: 587d7fab367417b2b2512bd9
title: Ongeza lebo kwa miduara ya chati ya kutawanyika
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

Unaweza kuongeza maandishi kuunda lebo kwa pointi katika chati ya kutawanyika.

Lengo ni kuonyesha thamani zilizoandikwa kwa koma kwa sehemu ya kwanza (`x`) na ya pili (`y`) ya kila kipengele katika `dataset`.

Nodi za `text` zinahitaji sifa za `x` na `y` ili kuziweka kwenye SVG. Katika changamoto hii, thamani ya `y` (inayoamua urefu) inaweza kutumia thamani ile ile ambayo `circle` hutumia kwa sifa yake ya `cy`. Thamani ya `x` inaweza kuwa kidogo kubwa kuliko thamani ya `cx` ya `circle`, ili lebo ionekane. Hii itasukuma lebo upande wa kulia wa pointi iliyopangwa.

# --instructions--

Leba kila pointi kwenye chati ya kutawanyika kwa kutumia vipengele vya `text`. Maandishi ya lebo yanapaswa kuwa thamani mbili zilizoachwa kwa koma na nafasi. Kwa mfano, lebo ya pointi ya kwanza ni `34, 78`. Weka sifa ya `x` ili iwe na thamani ya vitengo `5` zaidi kuliko thamani uliyotumia kwa sifa ya `cx` kwenye `circle`. Weka sifa ya `y` kwa njia ile ile inayotumika kwa thamani ya `cy` kwenye `circle`.

# --hints--

Msimbo wako unapaswa kuwa na vipengele 10 vya `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Lebo ya kwanza inapaswa kuwa na maandishi ya `34, 78`, thamani ya `x` ya `39`, na thamani ya `y` ya `422`.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

Lebo ya pili inapaswa kuwa na maandishi ya `109, 280`, thamani ya `x` ya `114`, na thamani ya `y` ya `220`.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

Lebo ya tatu inapaswa kuwa na maandishi ya `310, 120`, thamani ya `x` ya `315`, na thamani ya `y` ya `380`.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

Lebo ya nne inapaswa kuwa na maandishi ya `79, 411`, thamani ya `x` ya `84`, na thamani ya `y` ya `89`.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

Lebo ya tano inapaswa kuwa na maandishi ya `420, 220`, thamani ya `x` ya `425`, na thamani ya `y` ya `280`.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

Lebo ya sita inapaswa kuwa na maandishi ya `233, 145`, thamani ya `x` ya `238`, na thamani ya `y` ya `355`.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

Lebo ya saba inapaswa kuwa na maandishi ya `333, 96`, thamani ya `x` ya `338`, na thamani ya `y` ya `404`.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

Lebo ya nane inapaswa kuwa na maandishi ya `222, 333`, thamani ya `x` ya `227`, na thamani ya `y` ya `167`.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

Lebo ya tisa inapaswa kuwa na maandishi ya `78, 320`, thamani ya `x` ya `83`, na thamani ya `y` ya `180`.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

Lebo ya kumi inapaswa kuwa na maandishi ya `21, 123`, thamani ya `x` ya `26`, na thamani ya `y` ya `377`.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
