---
id: 587d7faa367417b2b2512bd4
title: Ongeza athari ya kupaa kwa kipengele cha D3
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

Inawezekana kuongeza athari zinazobainisha baa wakati mtumizi anapopaa kipanya juu yake. Hadi sasa, mtindo wa mstatili umewekwa kwa kutumia njia zilizojengwa ndani za D3 na SVG, lakini unaweza pia kutumia CSS.

Unaweka darasa la vitu la CSS kwenye vipengele vya SVG kwa njia ya `attr()`. Kisha darasa la bandia la `:hover` kwa darasa lako jipya lina sheria za mtindo kwa athari yoyote ya kupaa.

# --instructions--

Tumia njia ya `attr()` kuongeza darasa la `bar` kwa vipengele vyote vya `rect`. Hii hubadilisha rangi ya `fill` ya baa kuwa ya kahawia wakati unapopaa kipanya juu yake.

# --hints--

Vipengele vyako vya `rect` vinapaswa kuwa na darasa la `bar`.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```
