---
id: 587d7fa8367417b2b2512bc9
title: Sahihisha urefu wa kipengele kwa njia ya mabadiliko
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

Changamoto zilizopita zilijumuisha jinsi ya kuonyesha data kutoka kwenye safu ya data na jinsi ya kuongeza madarasa ya CSS. Unaweza kuunganisha mafundisho haya kuunda chati rahisi ya mstari. Kuna hatua mbili za kufanya hivi:

1. Unda `div` kwa kila kipengele cha data katika safu ya data

2. Toa urefu wa `div` kwa njia ya mabadiliko, ukitumia kitendakazi cha mwitiko katika njia ya `style()` ambayo inaweka urefu sawa na thamani ya data

Kumbuka muundo wa kuweka mtindo ukitumia kitendakazi cha mwitiko:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

Ongeza njia ya `style()` kwenye msimbo katika mhariri ili kuweka sifa ya `height` kwa kila kipengele. Tumia kitendakazi cha mwitiko kurudisha thamani ya kipengele cha data pamoja na mfuatano wa herufi `px` ulioungwa mkono.

# --hints--

`div` ya kwanza inapaswa kuwa na `height` ya pikseli `12`.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

`div` ya pili inapaswa kuwa na `height` ya pikseli `31`.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

`div` ya tatu inapaswa kuwa na `height` ya pikseli `22`.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

`div` ya nne inapaswa kuwa na `height` ya pikseli `17`.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

`div` ya tano inapaswa kuwa na `height` ya pikseli `25`.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

`div` ya sita inapaswa kuwa na `height` ya pikseli `18`.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

`div` ya saba inapaswa kuwa na `height` ya pikseli `29`.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

`div` ya nane inapaswa kuwa na `height` ya pikseli `14`.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

`div` ya tisa inapaswa kuwa na `height` ya pikseli `9`.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
