---
id: 587d7fa7367417b2b2512bc8
title: Ongeza madarasa na D3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

Kutumia mitindo mingi ndani ya mstari kwenye vipengele vya HTML huwa vigumu kusimamia, hata kwa programu ndogo ndogo. Ni rahisi zaidi kuongeza darasa kwa vipengele na kuunda mtindo wa darasa hilo mara moja kwa kutumia sheria za CSS. D3 ina njia ya `attr()` kuongeza sifa yoyote ya HTML kwa kipengele, ikiwa ni pamoja na jina la darasa.

Njia ya `attr()` hufanya kazi sawa na `style()`. Inachukua thamani zilizoandikwa kwa koma, na inaweza kutumia kitendakazi cha mwitiko. Hapa kuna mfano wa kuongeza darasa la `container` kwa uteuzi:

```js
selection.attr('class', 'container');
```

Kumbuka kuwa kigezo cha `class` kitaendelea kuwa sawa kila unapotaka kuongeza darasa na kigezo cha `container` ndicho pekee kitakachobadilika.

# --instructions--

Ongeza njia ya `attr()` kwenye msimbo ulioko kwenye mhariri na weka darasa la `bar` kwenye vipengele vya `div`.

# --hints--

Vipengele vyako vya `div` vinapaswa kuwa na darasa la `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Msimbo wako unapaswa kutumia njia ya `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
