---
id: 587d7fa7367417b2b2512bc5
title: Fanya kazi na data inayobadilika katika D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

Changamoto mbili za mwisho zinahusu misingi ya kuonyesha data inayobadilika kwa kutumia D3 kwa njia za `data()` na `enter()`. Njia hizi zinachukua seti ya data na, pamoja na njia ya `append()`, huunda kipengele kipya cha DOM kwa kila kipengee katika seti ya data.

Katika changamoto iliyopita, uliunda kipengele kipya cha `h2` kwa kila kipengee katika safu ya data ya `dataset`, lakini vyote vilikuwa na maandishi sawa, `New Title`. Hii ni kwa sababu hujatumia data iliyounganishwa na kila kipengele cha `h2`.

Njia ya D3 ya `text()` inaweza kuchukua mfuatano wa herufi au kitendakazi cha mwitiko kama hoja:

```js
selection.text(d => d);
```

Katika mfano hapo juu, kigezo cha `d` kinahusu kipengee kimoja katika seti ya data ambayo uteuzi umeunganishwa nayo.

Kutumia mfano huu kama muktadha, kipengele cha kwanza cha `h2` kimeunganishwa na 12, kipengele cha pili cha `h2` kimeunganishwa na 31, kipengele cha tatu cha `h2` kimeunganishwa na 22, na kadhalika.

# --instructions--

Badilisha njia ya `text()` ili kila kipengele cha `h2` kiwe na thamani inayolingana kutoka katika safu ya data ya `dataset` pamoja na nafasi moja na mfuatano wa herufi `USD`. Kwa mfano, kichwa cha kwanza kinapaswa kuwa `12 USD`.

# --hints--

Kichwa cha kwanza cha `h2` kinapaswa kuwa na maandishi `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

Kichwa cha pili cha `h2` kinapaswa kuwa na maandishi `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

Kichwa cha tatu cha `h2` kinapaswa kuwa na maandishi `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

Kichwa cha nne cha `h2` kinapaswa kuwa na maandishi `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

Kichwa cha tano cha `h2` kinapaswa kuwa na maandishi `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

Kichwa cha sita cha `h2` kinapaswa kuwa na maandishi `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

Kichwa cha saba cha `h2` kinapaswa kuwa na maandishi `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

Kichwa cha nane cha `h2` kinapaswa kuwa na maandishi `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

Kichwa cha tisa cha `h2` kinapaswa kuwa na maandishi `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
