---
id: 587d7fa8367417b2b2512bca
title: Badilisha Muonekano wa Chati ya mstari
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

Changamoto ya mwisho iliunda chati ya mstari, lakini kuna mabadiliko machache ya muundo yanayoweza kuiboresha:

1. Ongeza nafasi kati ya kila mstari ili kuwatenganisha kwa macho, ambayo hufanyika kwa kuongeza ukingo kwenye CSS ya darasa la `bar`

2. Ongeza urefu wa mistari ili kuonyesha tofauti ya thamani vizuri zaidi, ambayo hufanyika kwa kuzidisha thamani kwa nambari fulani ili kupanua urefu

# --instructions--

Kwanza, ongeza `margin` ya `2px` kwenye darasa la `bar` katika lebo ya `style`. Kisha, badilisha kitendakazi cha mwitiko katika njia ya `style()` ili irudishe thamani `10` mara thamani halisi ya data (pamoja na `px`).

**Kumbuka:** Kuzidisha kila nukta ya data kwa _kiasi sawa_ hubadilisha tu kiwango. Ni kama kupanua picha, na hakubadilishi maana ya data iliyopo chini.

# --hints--

`div` ya kwanza inapaswa kuwa na ukingo wa `height` wa pikseli `120` na ukingo wa `margin` wa pikseli `2`.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

`div` ya pili inapaswa kuwa na ukingo wa `height` wa pikseli `310` na ukingo wa `margin` wa pikseli `2`.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

`div` ya tatu inapaswa kuwa na ukingo wa `height` wa pikseli `220` na ukingo wa `margin` wa pikseli `2`.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

`div` ya nne inapaswa kuwa na ukingo wa `height` wa pikseli `170` na ukingo wa `margin` wa pikseli `2`.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

`div` ya tano inapaswa kuwa na ukingo wa `height` wa pikseli `250` na ukingo wa `margin` wa pikseli `2`.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

`div` ya sita inapaswa kuwa na ukingo wa `height` wa pikseli `180` na ukingo wa `margin` wa pikseli `2`.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

`div` ya saba inapaswa kuwa na ukingo wa `height` wa pikseli `290` na ukingo wa `margin` wa pikseli `2`.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

`div` ya nane inapaswa kuwa na ukingo wa `height` wa pikseli `140` na ukingo wa `margin` wa pikseli `2`.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

`div` ya tisa inapaswa kuwa na ukingo wa `height` wa pikseli `90` na ukingo wa `margin` wa pikseli `2`.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
