---
id: 587d7fab367417b2b2512bd7
title: Tengeneza chati ya kutawanyika kwa mizunguko ya SVG
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

Chati ya kutawanyika ni aina nyingine ya uwasilishaji wa data. Kawaida hutumia mizunguko kuonyesha pointi za data, ambazo zina thamani mbili kila moja. Thamani hizi zinahusiana na mhimili wa `x` na `y`, na hutumika kuweka mzunguko katika uwasilishaji.

SVG ina lebo ya `circle` kutengeneza umbo la mzunguko. Inafanya kazi kwa njia kama ile ya vipengele vya `rect` ulivyotumia kwa chati ya mstari.

# --instructions--

Tumia njia za `data()`, `enter()`, na `append()` kufunga `dataset` kwa vipengele vipya vya `circle` vinavyoongezwa kwenye SVG.

**Kumbuka:** Mizunguko haitakuwa inayoonekana kwa sababu bado hatujaweka sifa zao za kipengele. Tutafanya hivyo katika changamoto inayofuata.

# --hints--

Msimbo wako unapaswa kuwa na vipengele 10 vya `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
