---
id: 587d7fac367417b2b2512bdc
title: >-
  Tumia kitendakazi cha d3.max na d3.min kupata viwango vya kiwango cha chini na cha juu katika seti ya data
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

Njia za D3 `domain()` na `range()` huweka taarifa hiyo kwa kiwango chako kulingana na data. Kuna njia chache za kufanya hivyo kwa urahisi zaidi.

Mara nyingi unapoweka domain, utataka kutumia viwango vya kiwango cha chini na cha juu ndani ya seti ya data. Kujaribu kupata viwango hivi kwa mkono, hasa katika seti kubwa ya data, kunaweza kusababisha makosa.

D3 ina njia mbili - `min()` na `max()` kurudisha taarifa hii. Hapa kuna mfano:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

Seti ya data inaweza kuwa na safu za data zilizopangwa ndani, kama jozi za kuratibu za `[x, y]` zilizokuwa katika mfano wa mchoro wa alama. Katika hali hiyo, unahitaji kusema kwa D3 jinsi ya kuhesabu kiwango cha juu na cha chini. Kwa bahati nzuri, njia za `min()` na `max()` zote zinachukua kitendakazi cha mwitiko. Katika mfano huu, hoja ya kitendakazi cha mwitiko `d` ni kwa ajili ya safu ya ndani ya sasa. Kitendakazi cha mwitiko kinahitaji kurudisha kipengele kutoka kwenye safu ya ndani (thamani ya `x` au `y`) ambacho unataka kuhesabu kiwango cha juu au cha chini. Hapa kuna mfano wa jinsi ya kupata viwango vya kiwango cha chini na cha juu kwa safu ya safu za data:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` ingekuwa na thamani ya `1`.

# --instructions--

Safu ya `positionData` ina safu ndogo za kuratibu za x, y, na z. Tumia njia ya D3 kupata thamani ya juu zaidi ya kuratibu ya z (thamani ya tatu) kutoka kwa safu hizo na uiweke katika kigezo cha `output`.

# --hints--

Maandishi katika `h2` yanapaswa kuwa `8`.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

Msimbo wako unapaswa kutumia njia ya `max()`.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
