---
id: 587d7fa6367417b2b2512bc3
title: Chagua kundi la vipengele kwa kutumia D3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3 pia ina njia ya `selectAll()` kuchagua kundi la vipengele. Inarudisha safu ya nodi za HTML kwa vitu vyote katika hati vinavyolingana na mfuatano wa herufi wa ingizo. Hapa kuna mfano wa kuchagua lebo zote za tagi ya nanga katika hati:

```js
const anchors = d3.selectAll('a');
```

Kama njia ya `select()`, `selectAll()` inaunga mkono uunganishaji wa mfululizo wa njia, na unaweza kuitumia pamoja na njia nyingine.

# --instructions--

Chagua lebo zote za `li` katika hati, na badilisha maandishi yao kuwa mfuatano wa herufi `list item` kwa kuunganisha njia ya `.text()`.

# --hints--

Kuna vipengele 3 vya `li` kwenye ukurasa, na maandishi katika kila kimo lazima yasiseme `list item`. Uandishi wa herufi kubwa na nafasi lazima ulingane kabisa.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Msimbo wako unapaswa kupata upatikanaji wa kitu cha `d3`.

```js
assert.match(code, /d3/g);
```

Msimbo wako unapaswa kutumia njia ya `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
