---
id: 587d7fa7367417b2b2512bc6
title: Ongeza mtindo ndani ya mstari kwa vipengele
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3 inakuwezesha kuongeza mitindo ya CSS ndani ya mstari kwa vipengele vinavyobadilika kwa kutumia njia ya `style()`.

Njia ya `style()` inachukua jozi ya funguo na thamani zilizoandikwa kwa koma kama hoja. Hapa kuna mfano wa kuweka rangi ya maandishi ya uteuzi kuwa buluu:

```js
selection.style('color', 'blue');
```

# --instructions--

Ongeza njia ya `style()` kwenye msimbo ulioko kwenye mhariri ili kufanya maandishi yote yanayoonyeshwa kuwa na `font-family` ya `verdana`.

# --hints--

Vipengele vyako vya `h2` vinapaswa kuwa na `font-family` ya `verdana`.

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

Msimbo wako unapaswa kutumia njia ya `style()`.

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
