---
id: 587d7fa8367417b2b2512bcb
title: Jifunze kuhusu SVG katika D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> inamaanisha <dfn>Scalable Vector Graphics</dfn>.

Hapa "scalable" inamaanisha kwamba, ukiongeza au kupunguza ukubwa wa kitu, hakitaonekana kama picha yenye alama za pikseli. Inapima kulingana na mfumo wa kuonyesha, iwe ni kwenye skrini ndogo ya simu au kwenye skrini kubwa ya televisheni.

SVG hutumika kutengeneza maumbo ya kijiometri ya kawaida. Kwa kuwa D3 huwakilisha data kwa picha, hutumia SVG kutengeneza maumbo ya uwakilishi huo. Maumbo ya SVG kwa ukurasa wa mtandao lazima yawe ndani ya lebo ya HTML `svg`.

CSS inaweza kuwa scalable wakati mitindo inatumia vipimo vinavyohusiana (kama `vh`, `vw`, au asilimia), lakini kutumia SVG ni rahisi zaidi kujenga uwakilishi wa data kwa picha.

# --instructions--

Ongeza nodi ya `svg` kwenye `body` kwa kutumia `append()`. Ipe sifa ya `width` iliyowekwa kwa thabiti ya `w` na sifa ya `height` iliyowekwa kwa thabiti ya `h` kwa kutumia njia za `attr()` au `style()` kwa kila moja. Utaiona kwenye matokeo kwa sababu kuna `background-color` ya rangi ya waridi iliyotumika kwenye lebo ya `style`.

**Kumbuka:** Unapotumia sifa za upana na urefu za `attr()` hazina vipimo. Hii ni msingi wa kupima ukubwa - kipengele kitakuwa na uwiano wa upana kwa urefu wa 5:1, bila kujali kiwango cha kuongezwa au kupunguzwa ukubwa.

# --hints--

Nyaraka yako inapaswa kuwa na kipengele 1 cha `svg`.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

Kipengele cha `svg` kinapaswa kuwa na sifa ya `width` iliyowekwa kwa `500` au kupambwa kuwa na upana wa `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

Kipengele cha `svg` kinapaswa kuwa na sifa ya `height` iliyowekwa kwa `100` au kupambwa kuwa na urefu wa `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
