---
id: 587d7fa9367417b2b2512bd0
title: Geuza vipengele vya SVG
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

Huenda umeona chati ya mstari inaonekana kama imegeuzwa kichwa chini, au imegeuzwa. Hii ni kwa sababu ya jinsi SVG inavyotumia kuratibu za (x, y).

Katika SVG, sehemu ya mwanzo ya kuratibu iko kona ya juu-kushoto. Kuratibu ya `x` yenye thamani 0 huweka umbo kwenye kingo ya kushoto ya eneo la SVG. Kuratibu ya `y` yenye thamani 0 huweka umbo kwenye kingo ya juu ya eneo la SVG. Thamani kubwa za `x` husukuma mstatili kwenda kulia. Thamani kubwa za `y` husukuma mstatili chini.

Ili kufanya mistari iwe sawa na upande wa juu, unahitaji kubadilisha jinsi kuratibu ya `y` inavyopimwa. Inahitaji kuzingatia urefu wa mstari na urefu wa jumla wa eneo la SVG.

Urefu wa eneo la SVG ni 100. Ikiwa una kipimo cha data cha 0 katika seti, ungependa mstari uanze chini ya eneo la SVG (si juu). Kufanya hivyo, kuratibu ya `y` inahitaji thamani ya 100. Ikiwa thamani ya kipimo cha data ni 1, ungeanza na kuratibu ya `y` yenye thamani 100 kuweka mstari chini. Kisha unahitaji kuzingatia urefu wa mstari wa 1, hivyo kuratibu ya mwisho ya `y` itakuwa 99.

Kuratibu ya `y` ambayo ni `y = heightOfSVG - heightOfBar` itafanya mistari iwe sawa na upande wa juu.

# --instructions--

Badilisha kitendakazi cha mwitiko cha sifa ya `y` ili kuweka mistari sawa na upande wa juu. Kumbuka kwamba `height` ya mstari ni mara 3 ya thamani ya data `d`.

**Kumbuka:** Kwa ujumla, uhusiano ni `y = h - m * d`, ambapo `m` ni thabiti inayopima vipimo vya data.

# --hints--

`rect` ya kwanza inapaswa kuwa na thamani ya `y` ya `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

`rect` ya pili inapaswa kuwa na thamani ya `y` ya `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

`rect` ya tatu inapaswa kuwa na thamani ya `y` ya `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

`rect` ya nne inapaswa kuwa na thamani ya `y` ya `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

`rect` ya tano inapaswa kuwa na thamani ya `y` ya `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

`rect` ya sita inapaswa kuwa na thamani ya `y` ya `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

`rect` ya saba inapaswa kuwa na thamani ya `y` ya `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

`rect` ya nane inapaswa kuwa na thamani ya `y` ya `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

`rect` ya tisa inapaswa kuwa na thamani ya `y` ya `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
