---
id: 587d7fa9367417b2b2512bce
title: Weka kwa nguvu mwelekeo wa kila mstari
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

Changamoto ya mwisho ilitengeneza na kuongeza mstatili kwenye kipengele cha `svg` kwa kila nukta katika `dataset` kuwakilisha mstari. Kwa bahati mbaya, yote yalikuwa yamewekwa juu ya kila mmoja.

Mahali pa kuweka mstatili huhusishwa na sifa za `x` na `y`. Zinamuambia D3 wapi kuanza kuchora umbo katika eneo la `svg`. Changamoto ya mwisho ilizipanga zote kuwa 0, hivyo kila mstari uliwekwa kona ya juu-kushoto.

Kwa chati ya mstari, mistari yote inapaswa kuwa katika kiwango kimoja wima, ambayo inamaanisha thamani ya `y` inabaki sawa (kwa 0) kwa mistari yote. Hata hivyo, thamani ya `x` inahitaji kubadilika unapoongeza mistari mipya. Kumbuka kuwa thamani kubwa za `x` husukuma vitu zaidi upande wa kulia. Unapopita juu ya vipengele vya safu ya data katika `dataset`, thamani ya `x` inapaswa kuongezeka.

Njia ya `attr()` katika D3 inakubali kitendakazi cha mwitiko kuweka sifa hiyo kwa nguvu. Kitendakazi cha mwitiko kinapokea hoja mbili, moja kwa nukta ya data yenyewe (kawaida `d`) na moja kwa kielezo cha nukta ya data katika safu ya data. Hoja ya pili ya kielezo ni hiari. Huu ni muundo:

```js
selection.attr('property', (d, i) => {});
```

Ni muhimu kuelewa kuwa HAUHITAJI kuandika mzunguko wa `for` wala kutumia `forEach()` kurudia juu ya vitu katika seti ya data. Kumbuka kuwa njia ya `data()` huchanganuzi seti ya data, na njia yoyote inayofuata `data()` hutekelezwa mara moja kwa kila kipengee katika seti ya data.

# --instructions--

Badilisha kitendakazi cha mwitiko cha sifa ya `x` ili kirudishe kielezo mara 30.

**Kumbuka:** Kila mstari una upana wa 25, hivyo kuongeza thamani ya `x` kwa 30 huongeza nafasi kati ya mistari. Thamani yoyote kubwa kuliko 25 itafanya kazi katika mfano huu.

# --hints--

`rect` ya kwanza inapaswa kuwa na thamani ya `x` ya `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

`rect` ya pili inapaswa kuwa na thamani ya `x` ya `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

`rect` ya tatu inapaswa kuwa na thamani ya `x` ya `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

`rect` ya nne inapaswa kuwa na thamani ya `x` ya `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

`rect` ya tano inapaswa kuwa na thamani ya `x` ya `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

`rect` ya sita inapaswa kuwa na thamani ya `x` ya `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

`rect` ya saba inapaswa kuwa na thamani ya `x` ya `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

`rect` ya nane inapaswa kuwa na thamani ya `x` ya `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

`rect` ya tisa inapaswa kuwa na thamani ya `x` ya `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
