---
id: 587d7fa8367417b2b2512bcc
title: Onyesha Maumbo kwa kutumia SVG
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

Changamoto ya mwisho iliunda kipengele cha `svg` chenye upana na urefu uliotolewa, ambacho kilionekana kwa sababu kilikuwa na `background-color` kilichotumika ndani ya lebo ya `style`. Msimbo ulifanya nafasi kwa upana na urefu uliotolewa.

Hatua inayofuata ni kuunda umbo la kuweka katika eneo la `svg`. Kuna maumbo kadhaa yanayoungwa mkono katika SVG, kama vile mstatili na mduara. Yanitumika kuonyesha data. Kwa mfano, umbo la mstatili (`<rect>`) la SVG linaweza kuunda mstari katika chati ya mstari.

Unapoweka umbo katika eneo la `svg`, unaweza kubainisha mahali linapowekwa kwa kutumia kuratibu za `x` na `y`. Sehemu ya mwanzo ya (0, 0) iko kona ya juu-kushoto. Thamani chanya za `x` husukuma umbo kwenda kulia, na thamani chanya za `y` husukuma umbo chini kutoka sehemu ya mwanzo.

Ili kuweka umbo katikati ya `svg` ya upana 500 na urefu 100 kutoka changamoto ya mwisho, kuratibu ya `x` itakuwa 250 na kuratibu ya `y` itakuwa 50.

Lebo ya SVG ya `rect` ina sifa nne. Kuna kuratibu za `x` na `y` za mahali inapo wekwa katika eneo la `svg`. Pia ina `height` na `width` kubainisha ukubwa.

# --instructions--

Ongeza umbo la `rect` katika `svg` kwa kutumia `append()`, na umpe sifa ya `width` yenye thamani ya `25` na sifa ya `height` yenye thamani ya `100`. Pia, mpe `rect` sifa za `x` na `y` kila moja ikiwa na thamani ya `0`.

# --hints--

Hati yako inapaswa kuwa na kipengele kimoja cha `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

Kipengele cha `rect` kinapaswa kuwa na sifa ya `width` yenye thamani ya `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

Kipengele cha `rect` kinapaswa kuwa na sifa ya `height` yenye thamani ya `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

Kipengele cha `rect` kinapaswa kuwa na sifa ya `x` yenye thamani ya `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

Kipengele cha `rect` kinapaswa kuwa na sifa ya `y` yenye thamani ya `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
