---
id: 587d7fac367417b2b2512bde
title: Tumia kipimo kilichowekwa awali kuweka vipengele
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

Baada ya kuweka vipimo, ni wakati wa kuchora tena mchoro wa alama. Vipimo ni kama njia za usindikaji zinazobadilisha data ghafi ya `x` na `y` kuwa thamani zinazofaa na kuonyeshwa kwa usahihi kwenye SVG. Zinahakikisha data iko ndani ya eneo la mchoro kwenye skrini.

Unaweka thamani za sifa za kuratibu kwa umbo la SVG kwa kutumia kitendakazi cha kupima. Hii ni pamoja na sifa za `x` na `y` kwa vipengele vya `rect` au `text`, au `cx` na `cy` kwa `circles`. Hapa kuna mfano:

```js
shape.attr('x', d => xScale(d[0]));
```

Vipimo huweka sifa za kuratibu za umbo ili kuweka alama za data kwenye SVG. Huhitaji kutumia vipimo unapoonyesha thamani halisi ya data, kwa mfano, katika njia ya `text()` kwa kidokezo cha muda au lebo.

# --instructions--

Tumia `xScale` na `yScale` kuweka umbo la `circle` na `text` kwenye SVG. Kwa `circles`, tumia vipimo kuweka sifa za `cx` na `cy`. Pia waweke kipenyo cha `5`.

Kwa vipengele vya `text`, tumia vipimo kuweka sifa za `x` na `y`. Lebo ziwe zimepishana kidogo upande wa kulia wa alama. Ili kufanya hivyo, ongeza `10` kwa thamani ya data ya `x` kabla ya kuipatia `xScale`.

# --hints--

Msimbo wako unapaswa kuwa na vipengele 10 vya `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Kipengele cha kwanza cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `91` na thamani ya `cy` takriban `368` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

Kipengele cha pili cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `159` na thamani ya `cy` takriban `181` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

Kipengele cha tatu cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `340` na thamani ya `cy` takriban `329` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

Kipengele cha nne cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `131` na thamani ya `cy` takriban `60` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

Kipengele cha tano cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `440` na thamani ya `cy` takriban `237` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

Kipengele cha sita cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `271` na thamani ya `cy` takriban `306` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

Kipengele cha saba cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `361` na thamani ya `cy` takriban `351` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

Kipengele cha nane cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `261` na thamani ya `cy` takriban `132` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

Kipengele cha tisa cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `131` na thamani ya `cy` takriban `144` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

Kipengele cha kumi cha `circle` kinapaswa kuwa na thamani ya `cx` takriban `79` na thamani ya `cy` takriban `326` baada ya kutumia vipimo. Pia kinapaswa kuwa na thamani ya `r` ya `5`.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

Msimbo wako unapaswa kuwa na vipengele 10 vya `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Lebo ya kwanza inapaswa kuwa na thamani ya `x` takriban `100` na thamani ya `y` takriban `368` baada ya kutumia vipimo.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

Lebo ya pili inapaswa kuwa na thamani ya `x` takriban `168` na thamani ya `y` takriban `181` baada ya kutumia vipimo.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

Lebo ya tatu inapaswa kuwa na thamani ya `x` takriban `350` na thamani ya `y` takriban `329` baada ya kutumia vipimo.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

Lebo ya nne inapaswa kuwa na thamani ya `x` takriban `141` na thamani ya `y` takriban `60` baada ya kutumia vipimo.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

Lebo ya tano inapaswa kuwa na thamani ya `x` takriban `449` na thamani ya `y` takriban `237` baada ya kutumia vipimo.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

Lebo ya sita inapaswa kuwa na thamani ya `x` takriban `280` na thamani ya `y` takriban `306` baada ya kutumia vipimo.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

Lebo ya saba inapaswa kuwa na thamani ya `x` takriban `370` na thamani ya `y` takriban `351` baada ya kutumia vipimo.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

Lebo ya nane inapaswa kuwa na thamani ya `x` takriban `270` na thamani ya `y` takriban `132` baada ya kutumia vipimo.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

Lebo ya tisa inapaswa kuwa na thamani ya `x` takriban `140` na thamani ya `y` takriban `144` baada ya kutumia vipimo.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

Lebo ya kumi inapaswa kuwa na thamani ya `x` takriban `88` na thamani ya `y` takriban `326` baada ya kutumia vipimo.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
