---
id: 587d7fab367417b2b2512bd8
title: Ongeza sifa za kipengele kwa vipengele vya mduara
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

Changamoto ya mwisho iliunda vipengele vya `circle` kwa kila nukta katika `dataset`, na kuviambatanisha na SVG. Lakini D3 inahitaji taarifa zaidi kuhusu nafasi na ukubwa wa kila `circle` ili kuvionyesha kwa usahihi.

`circle` katika SVG ina sifa kuu tatu. Sifa za `cx` na `cy` ni kuratibu. Zinawaambia D3 wapi kuweka _kituo_ cha umbo kwenye SVG. Radius (sifa ya `r`) hutoa ukubwa wa `circle`.

Kama vile kuratibu za `rect` `y`, sifa ya `cy` kwa `circle` hupimwa kutoka juu ya SVG, si kutoka chini.

Sifa zote tatu zinaweza kutumia kitendakazi cha mwitiko kuweka thamani zao kwa njia ya mabadiliko. Kumbuka kuwa njia zote zilizofuatana baada ya `data(dataset)` zinaendeshwa mara moja kwa kila kipengee katika `dataset`. Kigezo cha `d` katika kitendakazi cha mwitiko kinahusu kipengee cha sasa katika `dataset`, ambayo ni safu ya data kwa kila nukta. Unatumia uandishi wa mabano, kama `d[0]`, kupata thamani katika safu hiyo.

# --instructions--

Ongeza sifa za `cx`, `cy`, na `r` kwa vipengele vya `circle`. Thamani ya `cx` inapaswa kuwa nambari ya kwanza katika safu ya kila kipengee katika `dataset`. Thamani ya `cy` inapaswa kutegemea nambari ya pili katika safu, lakini hakikisha kuonyesha chati kwa mwelekeo wa kawaida na si kugeuzwa. Thamani ya `r` inapaswa kuwa `5` kwa mduara yote.

# --hints--

Msimbo wako unapaswa kuwa na vipengele 10 vya `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Kipengele cha kwanza cha `circle` kinapaswa kuwa na thamani ya `cx` ya `34`, thamani ya `cy` ya `422`, na thamani ya `r` ya `5`.

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

Kipengele cha pili cha `circle` kinapaswa kuwa na thamani ya `cx` ya `109`, thamani ya `cy` ya `220`, na thamani ya `r` ya `5`.

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

Kipengele cha tatu cha `circle` kinapaswa kuwa na thamani ya `cx` ya `310`, thamani ya `cy` ya `380`, na thamani ya `r` ya `5`.

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

Kipengele cha nne cha `circle` kinapaswa kuwa na thamani ya `cx` ya `79`, thamani ya `cy` ya `89`, na thamani ya `r` ya `5`.

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

Kipengele cha tano cha `circle` kinapaswa kuwa na thamani ya `cx` ya `420`, thamani ya `cy` ya `280`, na thamani ya `r` ya `5`.

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

Kipengele cha sita cha `circle` kinapaswa kuwa na thamani ya `cx` ya `233`, thamani ya `cy` ya `355`, na thamani ya `r` ya `5`.

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

Kipengele cha saba cha `circle` kinapaswa kuwa na thamani ya `cx` ya `333`, thamani ya `cy` ya `404`, na thamani ya `r` ya `5`.

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

Kipengele cha nane cha `circle` kinapaswa kuwa na thamani ya `cx` ya `222`, thamani ya `cy` ya `167`, na thamani ya `r` ya `5`.

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

Kipengele cha tisa cha `circle` kinapaswa kuwa na thamani ya `cx` ya `78`, thamani ya `cy` ya `180`, na thamani ya `r` ya `5`.

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

Kipengele cha kumi cha `circle` kinapaswa kuwa na thamani ya `cx` ya `21`, thamani ya `cy` ya `377`, na thamani ya `r` ya `5`.

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
